Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.integer
% (31333)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 779
% Passive clauses: 3554
% Generated clauses: 15525
% Final active clauses: 708
% Final passive clauses: 2563
% Input formulas: 261
% Initial clauses: 559
% 
% Duplicate literals: 161
% Trivial inequalities: 42
% Fw subsumption resolutions: 1101
% Fw demodulations: 1630
% Bw demodulations: 69
% 
% Simple tautologies: 287
% Equational tautologies: 14
% Forward subsumptions: 8610
% 
% Binary resolution: 3811
% Factoring: 445
% Forward superposition: 4452
% Backward superposition: 3253
% Self superposition: 24
% Equality factoring: 2
% Equality resolution: 8
% 
% Split clauses: 220
% Split components: 443
% SAT solver clauses: 339
% SAT solver unit clauses: 13
% SAT solver binary clauses: 78
% 
% Sat splits: 220
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 339
% TWLsolver calls for satisfiability: 261
% 
% Memory used [KB]: 7419
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.integer
% (31367)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 530
% Passive clauses: 2695
% Generated clauses: 15849
% Final active clauses: 496
% Final passive clauses: 1993
% Input formulas: 261
% Initial clauses: 559
% 
% Duplicate literals: 92
% Trivial inequalities: 35
% Fw subsumption resolutions: 708
% Fw demodulations: 2571
% Bw demodulations: 206
% Condensations: 55
% Global subsumptions: 25
% 
% Simple tautologies: 332
% Equational tautologies: 9
% Forward subsumptions: 6273
% 
% Binary resolution: 2256
% Factoring: 64
% Forward superposition: 3952
% Backward superposition: 2542
% Self superposition: 37
% Equality factoring: 2
% Equality resolution: 12
% 
% Split clauses: 19
% Split components: 39
% SAT solver clauses: 3623
% SAT solver unit clauses: 653
% SAT solver binary clauses: 175
% SAT solver learnt clauses: 50
% SAT solver learnt literals: 14
% 
% Sat splits: 19
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 3523
% TWLsolver calls for satisfiability: 3978
% 
% Memory used [KB]: 15095
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.integer
% (31389)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1070
% Passive clauses: 7669
% Generated clauses: 64118
% Final active clauses: 831
% Final passive clauses: 3834
% Input formulas: 261
% Initial clauses: 559
% Discarded non-redundant clauses: 3850
% 
% Duplicate literals: 18
% Trivial inequalities: 42
% Fw subsumption resolutions: 1685
% Fw demodulations: 21451
% Bw demodulations: 597
% Condensations: 42
% Global subsumptions: 81
% 
% Simple tautologies: 287
% Equational tautologies: 52
% Forward subsumptions: 28518
% Fw demodulations to eq. taut.: 160
% Bw demodulations to eq. taut.: 2
% 
% Unit resulting resolution: 20271
% Factoring: 5
% Forward superposition: 10862
% Backward superposition: 7952
% Self superposition: 83
% Equality factoring: 2
% Equality resolution: 13
% 
% Split clauses: 40
% Split components: 81
% SAT solver clauses: 9434
% SAT solver unit clauses: 2889
% SAT solver binary clauses: 4661
% 
% Sat splits: 40
% Sat splitting refutations: 1602
% 
% Memory used [KB]: 121533
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 7 next slice time: 7
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.integer
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.integer
% SZS output start Proof for thm.integer
fof(f40773,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f642,f641,f640,f988,f984,f994,f996,f997,f998,f999,f1000,f1001,f1002,f1004,f1005,f1096,f1148,f1149,f1150,f1151,f1152,f1153,f1155,f1250,f1258,f1460,f1431,f1415,f1562,f1563,f1564,f1565,f1566,f1567,f1568,f1569,f1570,f1571,f1581,f1587,f1588,f1589,f1590,f1591,f1592,f1594,f1689,f1869,f1870,f1882,f1884,f1886,f1887,f1888,f1889,f1890,f1891,f1893,f1988,f2038,f2040,f2041,f2042,f2043,f2044,f2045,f2046,f2048,f2049,f2140,f2207,f2208,f2209,f2210,f2211,f2212,f2214,f2309,f2317,f2531,f2378,f2379,f2380,f2381,f2382,f2383,f2385,f2480,f2545,f2493,f2552,f2557,f2563,f2565,f2568,f2573,f2576,f2579,f2584,f2589,f2345,f2596,f2603,f2609,f2610,f2611,f2612,f2613,f2614,f2616,f2711,f2752,f2754,f2756,f2757,f2758,f2759,f2760,f2761,f2763,f2858,f3043,f3044,f3045,f2897,f1125,f3016,f2921,f2919,f2918,f2917,f2916,f2915,f2914,f2908,f3065,f3066,f3067,f3068,f3069,f3070,f3071,f3073,f3074,f3165,f3385,f3232,f3233,f3234,f3235,f3236,f3237,f3239,f3334,f3393,f3347,f3394,f3397,f3401,f3403,f3404,f3405,f3406,f3407,f3412,f3415,f3449,f3451,f3453,f3454,f3455,f3456,f3457,f3458,f3460,f3555,f3612,f3618,f3619,f3620,f3621,f3622,f3623,f3625,f3720,f3761,f3763,f3765,f3766,f3767,f3768,f3769,f3770,f3772,f3867,f3941,f3945,f3946,f3947,f3948,f3949,f3951,f3952,f4043,f4086,f4088,f4089,f4090,f4091,f4092,f4093,f4094,f4096,f4097,f4188,f4255,f4256,f4257,f4258,f4259,f4260,f4262,f4357,f4435,f4437,f4439,f4440,f4441,f4442,f4443,f4444,f4446,f4541,f4716,f4718,f4721,f4722,f4723,f4724,f4725,f4727,f4728,f4819,f4434,f4433,f5125,f5131,f5137,f5141,f5253,f5256,f5259,f5261,f5265,f5269,f5271,f5273,f5362,f5371,f5415,f5421,f5422,f5423,f5424,f5425,f5426,f5428,f5523,f5576,f5578,f5580,f5581,f5582,f5583,f5584,f5585,f5587,f5682,f5894,f5759,f5760,f5761,f5762,f5763,f5764,f5765,f5767,f5768,f5859,f5925,f5929,f5930,f5931,f5932,f5933,f5935,f5936,f6027,f6066,f6068,f6071,f6072,f6073,f6074,f6075,f6077,f6078,f6169,f6237,f6241,f6242,f6243,f6244,f6245,f6247,f6248,f6339,f6400,f6402,f6404,f6405,f6406,f6407,f6408,f6409,f6411,f6506,f6580,f6581,f6582,f6583,f6584,f6585,f6587,f6682,f6964,f6811,f6812,f6813,f6814,f6815,f6816,f6818,f6913,f6972,f6926,f6973,f6976,f6980,f6982,f6983,f6984,f6985,f6986,f6991,f6994,f7036,f7038,f7041,f7042,f7043,f7044,f7045,f7047,f7048,f7139,f7170,f7176,f7177,f7178,f7179,f7180,f7181,f7183,f7278,f7382,f7381,f7380,f7385,f7658,f7633,f7631,f7629,f7612,f7775,f7776,f7777,f7778,f7779,f7780,f7781,f7782,f7891,f7956,f7962,f7963,f7964,f7965,f7966,f7967,f7969,f8064,f8155,f8158,f8419,f8421,f8425,f8429,f8541,f8543,f8545,f8547,f8551,f8553,f8555,f8557,f8646,f8649,f8702,f8703,f8704,f8705,f8706,f8707,f8708,f8710,f8711,f8802,f8858,f8862,f8863,f8864,f8865,f8866,f8868,f8869,f8960,f9159,f9023,f9024,f9025,f9026,f9027,f9028,f9029,f9031,f9032,f9123,f9130,f9197,f9199,f9201,f9202,f9203,f9204,f9205,f9206,f9208,f9303,f9516,f9363,f9364,f9365,f9366,f9367,f9368,f9370,f9465,f9524,f9478,f9525,f9528,f9532,f9534,f9535,f9536,f9537,f9538,f9543,f9546,f9674,f9676,f9679,f9680,f9681,f9682,f9683,f9685,f9686,f9777,f9843,f9849,f9850,f9851,f9852,f9853,f9854,f9856,f9951,f10454,f10455,f10456,f10457,f10458,f10459,f10460,f10462,f10463,f10554,f10620,f10621,f10622,f10623,f10624,f10625,f10627,f10722,f10773,f10779,f10780,f10781,f10782,f10783,f10784,f10786,f10881,f10939,f10940,f10941,f10942,f10943,f10945,f10946,f11037,f11094,f10770,f11129,f11130,f11141,f11142,f11145,f11151,f11152,f11153,f11154,f11155,f11156,f11158,f11253,f11294,f11296,f11298,f11299,f11300,f11301,f11302,f11303,f11305,f11400,f11439,f11443,f11444,f11445,f11446,f11447,f11449,f11450,f11541,f11580,f11582,f11585,f11586,f11587,f11588,f11589,f11591,f11592,f11683,f11719,f11720,f11721,f11722,f11723,f11724,f11725,f11727,f11728,f11819,f12203,f12204,f12205,f12206,f12207,f12208,f12210,f12305,f11712,f11293,f11292,f11442,f11441,f1126,f12624,f5129,f12626,f5133,f12630,f12634,f5126,f12746,f12748,f12750,f5254,f12752,f5257,f12756,f12758,f2548,f5267,f12760,f12762,f5266,f12851,f12854,f5364,f12903,f12904,f12905,f12906,f12907,f12908,f12909,f12911,f12912,f13003,f13210,f13223,f11826,f11854,f13270,f13274,f13275,f13276,f13277,f13278,f13280,f13281,f13372,f13571,f13435,f13436,f13437,f13438,f13439,f13440,f13441,f13443,f13444,f13535,f13542,f1296,f13765,f2535,f2533,f13612,f13613,f13614,f13615,f13616,f13617,f13619,f13714,f13773,f13727,f1121,f13774,f13777,f13781,f13783,f13784,f13785,f13786,f13787,f13792,f13795,f13805,f13807,f13809,f13810,f13811,f13812,f13813,f13814,f13816,f13911,f14086,f14088,f14091,f14092,f14093,f14094,f14095,f14097,f14098,f14189,f14255,f14261,f14262,f14263,f14264,f14265,f14266,f14268,f14363,f14767,f14768,f14769,f14770,f14771,f14772,f14773,f14775,f14776,f14867,f14980,f14982,f14983,f14984,f14985,f14986,f14987,f14988,f14990,f14991,f15082,f15124,f15135,f986,f15170,f15171,f15182,f15183,f15186,f15321,f15294,f15199,f15197,f15196,f15195,f15194,f15193,f15192,f15337,f15339,f15341,f15342,f15343,f15344,f15345,f15346,f15348,f15443,f15482,f15486,f15487,f15488,f15489,f15490,f15492,f15493,f15584,f15623,f15625,f15628,f15629,f15630,f15631,f15632,f15634,f15635,f15726,f15900,f15906,f15907,f15908,f15909,f15910,f15911,f15913,f16008,f16130,f16131,f16132,f16133,f16134,f16135,f16137,f16232,f16372,f16353,f10585,f2551,f16519,f16538,f16539,f16540,f16541,f16542,f16543,f16544,f16546,f16547,f16638,f16692,f15897,f16749,f16750,f16751,f16752,f16753,f16755,f16756,f16847,f16888,f16889,f16890,f16891,f16892,f16893,f16895,f16990,f17044,f17046,f17048,f17049,f17050,f17051,f17052,f17053,f17055,f17150,f17325,f17327,f17330,f17331,f17332,f17333,f17334,f17336,f17337,f17428,f17472,f2534,f17478,f17479,f17480,f17481,f17482,f17483,f17485,f17580,f17630,f17631,f17632,f17633,f17634,f17636,f17637,f17728,f17917,f17918,f17919,f17920,f17921,f17922,f17924,f18019,f16745,f18080,f18081,f18082,f18083,f18084,f18085,f18086,f18088,f18089,f18180,f18562,f18563,f18564,f18565,f18566,f18567,f18568,f18570,f18571,f18662,f10587,f18703,f18705,f18706,f18707,f18708,f18709,f18710,f18711,f18713,f18714,f18805,f18862,f18863,f18864,f18865,f18866,f18867,f18869,f18964,f18997,f19013,f19014,f3588,f10586,f19015,f19016,f19017,f19018,f19019,f19021,f16670,f19022,f19113,f19152,f18073,f1122,f19169,f19006,f3058,f18216,f19204,f19205,f19216,f19217,f19220,f19226,f19227,f19228,f19229,f19230,f19231,f19233,f19328,f19369,f19371,f2020,f19373,f19374,f19375,f19376,f19377,f19378,f19380,f19475,f19514,f2547,f19518,f19519,f19520,f19521,f19522,f19524,f19525,f19616,f19655,f19657,f6197,f3586,f19660,f19661,f19662,f19663,f19664,f19666,f19667,f19758,f19932,f19938,f19939,f19940,f19941,f19942,f8102,f19943,f19945,f20040,f20134,f20135,f20136,f20137,f20138,f20139,f20141,f20236,f20385,f20386,f20387,f20388,f20389,f20390,f20392,f20487,f20544,f20545,f20546,f20547,f20548,f20549,f20550,f20552,f20553,f20644,f20682,f20715,f20716,f20717,f20718,f20719,f20526,f20720,f20722,f20817,f20945,f20946,f20947,f20948,f20949,f20950,f20952,f21047,f21145,f21146,f21147,f21148,f21149,f21150,f21151,f21153,f21154,f21245,f21329,f21330,f21331,f21332,f21333,f21334,f21336,f21431,f19011,f21499,f21500,f21501,f21502,f21503,f21504,f21505,f21507,f21508,f21599,f21652,f21653,f21654,f21655,f21656,f2539,f21657,f21659,f21754,f2538,f21815,f21821,f21822,f21823,f21824,f21825,f21830,f2554,f22059,f22060,f22193,f22195,f22061,f22062,f22063,f22064,f22065,f22067,f22068,f22159,f22196,f22214,f22215,f22216,f22217,f22218,f21466,f22219,f22221,f22316,f21465,f22366,f22367,f3057,f22372,f22373,f22374,f22375,f22376,f22377,f22379,f22474,f22515,f20825,f22546,f19929,f989,f987,f20853,f1297,f22548,f22050,f16374,f22605,f22606,f22607,f22608,f22609,f22611,f1129,f22612,f22703,f1127,f22744,f22745,f22746,f22747,f22748,f1289,f22749,f22751,f22846,f22900,f3587,f22902,f3589,f22904,f22905,f22906,f22907,f22908,f3593,f22909,f22911,f23006,f3592,f23181,f2021,f23183,f4847,f23186,f23187,f23188,f23189,f23190,f23192,f2024,f23193,f23284,f4848,f23334,f23335,f23336,f23337,f23338,f3056,f23339,f23341,f23436,f3055,f23487,f23488,f23489,f23490,f23491,f23493,f11078,f23494,f23585,f23686,f23687,f23688,f23689,f23690,f18062,f23691,f23693,f23788,f24502,f24508,f24509,f24510,f24511,f24512,f24646,f24513,f24515,f24610,f24660,f24662,f24663,f2019,f15116,f24664,f24665,f24666,f24667,f24668,f24670,f12339,f24671,f24762,f24807,f24808,f24809,f24810,f24811,f8840,f24812,f24814,f24909,f21081,f25003,f25120,f25266,f25267,f25268,f25282,f25430,f25432,f25263,f25434,f25293,f25294,f25295,f25296,f25297,f25437,f25298,f25300,f25395,f25436,f25534,f25677,f25546,f25549,f25550,f25551,f25552,f25553,f25555,f25684,f25556,f25647,f25755,f25756,f25757,f25758,f25759,f25761,f25895,f25762,f25853,f25894,f25975,f25976,f25977,f25978,f25979,f25981,f26110,f25982,f26073,f26193,f26194,f26195,f26196,f26197,f26198,f26200,f26295,f26512,f26513,f26520,f26521,f26542,f26543,f26544,f26545,f26546,f26338,f26547,f26549,f26644,f26337,f26838,f26707,f26710,f26711,f26712,f26713,f26714,f26716,f26717,f26808,f26853,f27049,f26840,f26918,f26843,f26921,f26922,f26923,f26924,f26925,f26927,f26845,f26928,f27019,f26844,f27057,f27201,f27202,f27203,f27204,f27205,f27206,f27209,f27210,f27361,f27363,f27225,f25681,f27227,f27228,f27229,f27230,f27231,f27366,f27232,f27234,f27329,f27365,f27386,f27387,f27388,f27389,f27390,f27392,f27520,f27393,f27484,f27519,f27596,f27597,f27598,f27599,f27600,f27602,f27761,f27603,f27694,f27760,f28559,f28423,f28425,f28426,f28427,f28428,f28429,f1130,f28564,f28430,f28432,f28527,f29370,f29372,f29373,f29375,f29377,f29472,f29511,f29736,f26683,f28560,f26531,f27216,f2017,f26907,f3591,f29781,f1124,f25892,f25678,f25741,f1123,f29785,f29786,f29787,f29788,f29789,f29790,f29791,f29589,f29591,f29592,f29593,f29596,f29687,f29792,f29746,f29750,f2567,f29716,f29718,f29754,f29794,f29829,f29830,f29841,f29842,f8157,f29845,f29851,f29853,f29854,f29856,f29858,f29953,f1720,f30130,f1298,f4069,f29995,f29996,f5896,f29997,f29999,f30000,f30001,f30004,f30095,f5897,f30144,f2018,f2016,f30148,f30150,f30151,f30153,f30155,f30250,f2023,f30293,f30295,f30296,f30297,f30300,f30391,f4072,f30435,f30437,f30438,f30439,f30442,f30533,f6198,f31032,f31034,f31035,f31037,f31039,f31134,f12338,f31325,f30430,f6195,f3590,f31375,f31377,f31378,f31380,f31382,f31477,f1288,f31604,f31605,f2015,f31759,f31606,f31608,f31609,f31610,f31613,f31704,f31760,f31778,f2572,f31780,f2570,f31866,f31868,f31869,f31870,f31873,f31964,f11077,f32024,f32026,f32027,f32029,f32031,f32126,f18061,f32442,f24643,f21485,f32448,f32450,f32451,f32453,f32455,f32550,f24645,f33615,f33621,f33633,f33639,f33647,f33770,f33772,f33785,f33787,f33789,f33786,f33910,f33912,f33921,f34126,f34127,f34128,f34129,f32983,f1719,f1717,f34175,f8096,f15114,f34179,f34181,f34182,f34183,f34186,f34277,f15117,f34385,f34387,f34388,f34389,f34392,f34483,f34720,f25679,f34588,f25682,f34591,f34593,f34594,f34595,f34598,f34689,f25683,f34901,f34902,f34907,f34908,f34909,f34910,f34978,f34980,f34981,f34982,f34985,f35076,f26109,f35772,f3585,f28561,f35636,f4845,f26842,f35638,f35640,f35641,f35643,f35645,f35740,f28563,f36749,f33979,f36764,f27385,f27384,f25705,f36768,f25288,f25287,f8838,f2549,f37111,f37271,f37276,f37390,f37395,f37091,f5366,f27553,f37428,f27534,f37429,f25399,f25264,f5367,f37433,f37768,f5127,f5368,f37771,f38033,f8099,f8097,f38039,f38041,f38042,f38044,f38046,f38141,f8101,f38246,f8836,f25428,f38248,f36766,f38376,f38250,f38252,f38253,f38255,f38257,f38352,f38377,f38395,f38397,f38398,f38399,f38402,f38493,f34521,f38565,f21486,f985,f38569,f38571,f38572,f38573,f38576,f38667,f16669,f38766,f38768,f38769,f38771,f38773,f38868,f20525,f38965,f38967,f38968,f38969,f38972,f39063,f8839,f39281,f39283,f39284,f1128,f39285,f1132,f39288,f1131,f39379,f1133,f39414,f39809,f38394,f38393,f16373,f2022,f40128,f40119,f40118,f40116,f40114,f39979,f33773,f39978,f2559,f33771,f39855,f2558,f33640,f39851,f33643,f33641,f39847,f33635,f2586,f39835,f39834,f2578,f39832,f2575,f39830,f33619,f33617,f40475,f40495])).
fof(f40495,plain,(
  $false | ($spl5 | $spl178)),
  inference(subsumption_resolution,[],[f40226,f1146])).
fof(f1146,plain,(
  s('type.integer.int',sK14) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))) | $spl5),
  inference(resolution,[],[f989,f906])).
fof(f906,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) != s('type.integer.int',X1)) )),
  inference(cnf_transformation,[],[f617])).
fof(f617,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1)) & ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & s('type.integer.int',X0) != s('type.integer.int',X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(flattening,[],[f616])).
fof(f616,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1))) & ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & s('type.integer.int',X0) != s('type.integer.int',X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f361])).
fof(f361,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1)))),
  inference(rectify,[],[f148])).
fof(f148,axiom,(
  ! [X9,X11] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X9),s('type.integer.int',X11)))) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X9),s('type.integer.int',X11)))) | s('type.integer.int',X9) = s('type.integer.int',X11)))),
  file('thm.integer.INT_ABS_LE|split|2','thm.integer.INT_LE_LT')).
fof(f40226,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))) | $spl178),
  inference(superposition,[],[f962,f16373])).
fof(f962,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))) )),
  inference(cnf_transformation,[],[f400])).
fof(f400,plain,(
  ! [X0] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X9] : s('type.integer.int',X9) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9)))))),
  file('thm.integer.INT_ABS_LE|split|2','thm.integer.INT_NEGNEG')).
fof(f40475,plain,(
  $false | ($spl61 | $spl178)),
  inference(subsumption_resolution,[],[f40165,f888])).
fof(f888,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f349])).
fof(f349,plain,(
  ! [X0] : p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X0))))),
  inference(rectify,[],[f210])).
fof(f210,axiom,(
  ! [X9] : p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X9),s('type.integer.int',X9))))),
  file('thm.integer.INT_ABS_LE|split|2','thm.integer.INT_LE_REFL')).
fof(f40165,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK15)))) | ($spl61 | $spl178)),
  inference(backward_demodulation,[],[f16373,f2547])).
fof(f33617,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | $spl290),
  inference(cnf_transformation,[],[f33617_D])).
fof(f33617_D,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f33619,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl292),
  inference(cnf_transformation,[],[f33619_D])).
fof(f33619_D,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f39830,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | ($spl31 | $spl286)),
  inference(forward_demodulation,[],[f39829,f31759])).
fof(f39829,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | ($spl31 | $spl286)),
  inference(forward_demodulation,[],[f39828,f31759])).
fof(f39828,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | ($spl31 | $spl286)),
  inference(subsumption_resolution,[],[f39491,f1296])).
fof(f39491,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl286),
  inference(superposition,[],[f31759,f952])).
fof(f952,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46(X0)))))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45(X0))))) )),
  inference(cnf_transformation,[],[f635])).
fof(f635,plain,(
  ! [X0] : ((s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45(X0)))) & s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45(X0))) | (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46(X0)))))) & s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46(X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK45,sK46])],[f393])).
fof(f393,plain,(
  ! [X0] : (? [X2] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X2))) & s('type.num.num','const.num.|48|_0') != s('type.num.num',X2)) | ? [X1] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))) & s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(flattening,[],[f392])).
fof(f392,plain,(
  ! [X0] : (? [X2] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X2))) & ~s('type.num.num','const.num.|48|_0') = s('type.num.num',X2)) | ? [X1] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))) & ~s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : (? [X10] : (s('type.integer.int',X6) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10))) & ~s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) | ? [X10] : (s('type.integer.int',X6) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10))))) & ~s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) | s('type.integer.int',X6) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.integer.INT_ABS_LE|split|2','thm.integer.INT_NUM_CASES')).
fof(f2575,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl77),
  inference(cnf_transformation,[],[f2575_D])).
fof(f2575_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) <=> ~$spl77),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl77])])).
fof(f39832,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ($spl31 | $spl286)),
  inference(forward_demodulation,[],[f39831,f31759])).
fof(f39831,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ($spl31 | $spl286)),
  inference(subsumption_resolution,[],[f39492,f1296])).
fof(f39492,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl286),
  inference(superposition,[],[f31759,f953])).
fof(f953,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46(X0)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45(X0))) )),
  inference(cnf_transformation,[],[f635])).
fof(f2578,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl79),
  inference(cnf_transformation,[],[f2578_D])).
fof(f2578_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) <=> ~$spl79),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl79])])).
fof(f39834,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ($spl31 | $spl286)),
  inference(forward_demodulation,[],[f39833,f31759])).
fof(f39833,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | ($spl31 | $spl286)),
  inference(subsumption_resolution,[],[f39493,f1296])).
fof(f39493,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl286),
  inference(superposition,[],[f31759,f954])).
fof(f954,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46(X0)) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45(X0))))) )),
  inference(cnf_transformation,[],[f635])).
fof(f39835,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ($spl31 | $spl286)),
  inference(subsumption_resolution,[],[f39494,f1296])).
fof(f39494,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl286),
  inference(superposition,[],[f31759,f955])).
fof(f955,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46(X0)) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45(X0))) )),
  inference(cnf_transformation,[],[f635])).
fof(f2586,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl80),
  inference(cnf_transformation,[],[f2586_D])).
fof(f2586_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f33635,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | $spl294),
  inference(cnf_transformation,[],[f33635_D])).
fof(f33635_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f39847,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ($spl65 | $spl286)),
  inference(subsumption_resolution,[],[f39539,f2551])).
fof(f39539,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl286),
  inference(superposition,[],[f956,f31759])).
fof(f956,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f394])).
fof(f394,plain,(
  ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))))),
  inference(rectify,[],[f157])).
fof(f157,axiom,(
  ! [X9] : (s('type.integer.int',X9) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X9)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9)))))))),
  file('thm.integer.INT_ABS_LE|split|2','thm.integer.INT_LT_NEGTOTAL')).
fof(f33641,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl297),
  inference(cnf_transformation,[],[f33641_D])).
fof(f33641_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) <=> ~$spl297),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl297])])).
fof(f33643,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK43('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | $spl298),
  inference(cnf_transformation,[],[f33643_D])).
fof(f33643_D,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK43('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f39851,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK43('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl286),
  inference(forward_demodulation,[],[f39850,f31759])).
fof(f39850,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK43('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | $spl286),
  inference(forward_demodulation,[],[f39546,f921])).
fof(f921,plain,(
  ( ! [X0] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))) )),
  inference(cnf_transformation,[],[f373])).
fof(f373,plain,(
  ! [X0] : s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))),
  inference(rectify,[],[f140])).
fof(f140,axiom,(
  ! [X9] : s(bool,'const.integer.int_le_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_le_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9)))))),
  file('thm.integer.INT_ABS_LE|split|2','thm.integer.INT_LE_NEGR')).
fof(f39546,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK43('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | $spl286),
  inference(superposition,[],[f881,f31759])).
fof(f881,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK43(X0)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(cnf_transformation,[],[f607])).
fof(f607,plain,(
  ! [X0] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK43(X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK43])],[f449])).
fof(f449,plain,(
  ! [X0] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ? [X1] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))))),
  inference(ennf_transformation,[],[f340])).
fof(f340,plain,(
  ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) => ? [X1] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))))),
  inference(rectify,[],[f129])).
fof(f129,axiom,(
  ! [X17] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X17),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) => ? [X10] : s('type.integer.int',X17) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10))))))),
  file('thm.integer.INT_ABS_LE|split|2','thm.integer.NUM_NEGINT_EXISTS')).
fof(f33640,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl296),
  inference(cnf_transformation,[],[f33640_D])).
fof(f33640_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f2558,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl68),
  inference(cnf_transformation,[],[f2558_D])).
fof(f2558_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f39855,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl286),
  inference(forward_demodulation,[],[f39854,f919])).
fof(f919,plain,(
  ( ! [X0] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))) )),
  inference(cnf_transformation,[],[f371])).
fof(f371,plain,(
  ! [X0] : s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X9] : s(bool,'const.integer.int_le_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9)))))),
  file('thm.integer.INT_ABS_LE|split|2','thm.integer.INT_NEG_GE|48|')).
fof(f39854,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl286),
  inference(forward_demodulation,[],[f39853,f921])).
fof(f39853,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl286),
  inference(forward_demodulation,[],[f39548,f919])).
fof(f39548,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl286),
  inference(superposition,[],[f885,f31759])).
fof(f885,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f345])).
fof(f345,plain,(
  ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))))),
  inference(rectify,[],[f130])).
fof(f130,axiom,(
  ! [X9] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X9)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9)))))))),
  file('thm.integer.INT_ABS_LE|split|2','thm.integer.INT_LE_NEGTOTAL')).
fof(f33771,plain,(
  ( ! [X131] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK15),s('type.integer.int',X131)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X131))))) ) | $spl300),
  inference(cnf_transformation,[],[f33771_D])).
fof(f33771_D,plain,(
  ( ! [X131] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK15),s('type.integer.int',X131)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X131))))) ) <=> ~$spl300),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl300])])).
fof(f2559,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl69),
  inference(cnf_transformation,[],[f2559_D])).
fof(f2559_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl69),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl69])])).
fof(f39978,plain,(
  ( ! [X131] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK15),s('type.integer.int',X131)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X131))))) ) | $spl286),
  inference(forward_demodulation,[],[f39640,f919])).
fof(f39640,plain,(
  ( ! [X131] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK15),s('type.integer.int',X131)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X131)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))))) ) | $spl286),
  inference(superposition,[],[f910,f31759])).
fof(f910,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f457])).
fof(f457,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))))),
  inference(flattening,[],[f456])).
fof(f456,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))))),
  inference(ennf_transformation,[],[f364])).
fof(f364,plain,(
  ! [X0,X1] : ((p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1))))) => p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))))),
  inference(rectify,[],[f151])).
fof(f151,axiom,(
  ! [X9,X11] : ((p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X9)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X11))))) => p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X9),s('type.integer.int',X11)))))))),
  file('thm.integer.INT_ABS_LE|split|2','thm.integer.INT_LE_MUL')).
fof(f33773,plain,(
  ( ! [X132] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X132),s('type.integer.int',sK15)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X132))))) ) | $spl302),
  inference(cnf_transformation,[],[f33773_D])).
fof(f33773_D,plain,(
  ( ! [X132] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X132),s('type.integer.int',sK15)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X132))))) ) <=> ~$spl302),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl302])])).
fof(f39979,plain,(
  ( ! [X132] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X132),s('type.integer.int',sK15)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X132))))) ) | $spl286),
  inference(forward_demodulation,[],[f39641,f919])).
fof(f39641,plain,(
  ( ! [X132] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X132),s('type.integer.int',sK15)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X132))))) ) | $spl286),
  inference(superposition,[],[f910,f31759])).
fof(f40114,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | ($spl31 | $spl286)),
  inference(forward_demodulation,[],[f40113,f31759])).
fof(f40113,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | ($spl31 | $spl286)),
  inference(forward_demodulation,[],[f40112,f31759])).
fof(f40112,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | ($spl31 | $spl286)),
  inference(subsumption_resolution,[],[f39744,f1296])).
fof(f39744,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl286),
  inference(superposition,[],[f952,f31759])).
fof(f40116,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ($spl31 | $spl286)),
  inference(forward_demodulation,[],[f40115,f31759])).
fof(f40115,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ($spl31 | $spl286)),
  inference(subsumption_resolution,[],[f39745,f1296])).
fof(f39745,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl286),
  inference(superposition,[],[f953,f31759])).
fof(f40118,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ($spl31 | $spl286)),
  inference(forward_demodulation,[],[f40117,f31759])).
fof(f40117,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | ($spl31 | $spl286)),
  inference(subsumption_resolution,[],[f39746,f1296])).
fof(f39746,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl286),
  inference(superposition,[],[f954,f31759])).
fof(f40119,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ($spl31 | $spl286)),
  inference(subsumption_resolution,[],[f39747,f1296])).
fof(f39747,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl286),
  inference(superposition,[],[f955,f31759])).
fof(f40128,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | ($spl31 | $spl286)),
  inference(subsumption_resolution,[],[f40127,f1296])).
fof(f40127,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | $spl286),
  inference(forward_demodulation,[],[f40126,f31759])).
fof(f40126,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) | $spl286),
  inference(forward_demodulation,[],[f40125,f958])).
fof(f958,plain,(
  ( ! [X0] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(cnf_transformation,[],[f396])).
fof(f396,plain,(
  ! [X0] : s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X9] : s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X9))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.integer.INT_ABS_LE|split|2','thm.integer.INT_NEG_LT|48|')).
fof(f40125,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) | $spl286),
  inference(forward_demodulation,[],[f39749,f959])).
fof(f959,plain,(
  ( ! [X0] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))) )),
  inference(cnf_transformation,[],[f397])).
fof(f397,plain,(
  ! [X0] : s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X9] : s(bool,'const.integer.int_lt_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9)))))),
  file('thm.integer.INT_ABS_LE|split|2','thm.integer.INT_NEG_GT|48|')).
fof(f39749,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) | $spl286),
  inference(superposition,[],[f956,f31759])).
fof(f2022,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl46),
  inference(cnf_transformation,[],[f2022_D])).
fof(f2022_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f16373,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) | $spl178),
  inference(cnf_transformation,[],[f16373_D])).
fof(f16373_D,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f38393,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl36),
  inference(resolution,[],[f1719,f904])).
fof(f904,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) = s('type.integer.int',X1) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f617])).
fof(f38394,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) | $spl36),
  inference(resolution,[],[f1719,f909])).
fof(f909,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f619])).
fof(f619,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & s('type.integer.int',X0) != s('type.integer.int',X1))) & (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(flattening,[],[f618])).
fof(f618,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & s('type.integer.int',X0) != s('type.integer.int',X1))) & ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f363])).
fof(f363,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & s('type.integer.int',X0) != s('type.integer.int',X1)))),
  inference(flattening,[],[f362])).
fof(f362,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & ~s('type.integer.int',X0) = s('type.integer.int',X1)))),
  inference(rectify,[],[f150])).
fof(f150,axiom,(
  ! [X9,X11] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X9),s('type.integer.int',X11)))) <=> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X9),s('type.integer.int',X11)))) & ~s('type.integer.int',X9) = s('type.integer.int',X11)))),
  file('thm.integer.INT_ABS_LE|split|2','thm.integer.INT_LT_LE')).
fof(f39809,plain,(
  $false | ($spl62 | $spl179 | $spl286)),
  inference(subsumption_resolution,[],[f39808,f2548])).
fof(f39808,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | ($spl179 | $spl286)),
  inference(forward_demodulation,[],[f39807,f922])).
fof(f922,plain,(
  ( ! [X0] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f374])).
fof(f374,plain,(
  ! [X0] : s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X0)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X9] : s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X9))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))),s('type.integer.int',X9)))),
  file('thm.integer.INT_ABS_LE|split|2','thm.integer.INT_LE_NEGL')).
fof(f39807,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | ($spl179 | $spl286)),
  inference(subsumption_resolution,[],[f39473,f16374])).
fof(f39473,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | $spl286),
  inference(superposition,[],[f31759,f917])).
fof(f917,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f621])).
fof(f621,plain,(
  ! [X0] : ((s('type.integer.int',X0) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))) & (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))),
  inference(nnf_transformation,[],[f369])).
fof(f369,plain,(
  ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) <=> p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6] : (s('type.integer.int',X6) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))) <=> p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X6)))))),
  file('thm.integer.INT_ABS_LE|split|2','thm.integer.INT_ABS_EQ_ID')).
fof(f39414,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | $spl322),
  inference(cnf_transformation,[],[f39414_D])).
fof(f39414_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) <=> ~$spl322),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl322])])).
fof(f1133,plain,(
  ( ! [X13] : (p(s(bool,X13))) ) | $spl24),
  inference(cnf_transformation,[],[f1133_D])).
fof(f1133_D,plain,(
  ( ! [X13] : (p(s(bool,X13))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f39379,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | $spl152),
  inference(resolution,[],[f8836,f976])).
fof(f976,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f647])).
fof(f647,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f485])).
fof(f485,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f484])).
fof(f484,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f416])).
fof(f416,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f415])).
fof(f415,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f264])).
fof(f264,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f222])).
fof(f222,axiom,(
  ! [X9,X26,X11,X27] : ((s(bool,X9) = s(bool,X26) & (p(s(bool,X26)) => s(bool,X11) = s(bool,X27))) => ((p(s(bool,X9)) => p(s(bool,X11))) <=> (p(s(bool,X26)) => p(s(bool,X27)))))),
  file('thm.integer.INT_ABS_LE|split|2','thm.bool.IMP_CONG')).
fof(f1131,plain,(
  ( ! [X8,X9] : (s(bool,X8) != s(bool,X9) | ~p(s(bool,X8))) ) | $spl20),
  inference(cnf_transformation,[],[f1131_D])).
fof(f1131_D,plain,(
  ( ! [X8,X9] : (s(bool,X8) != s(bool,X9) | ~p(s(bool,X8))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f39288,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14))) | s(bool,X23) != s(bool,X24)) ) | $spl152),
  inference(resolution,[],[f8836,f650])).
fof(f650,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f485])).
fof(f1132,plain,(
  ( ! [X14] : (~p(s(bool,X14))) ) | $spl22),
  inference(cnf_transformation,[],[f1132_D])).
fof(f1132_D,plain,(
  ( ! [X14] : (~p(s(bool,X14))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f39285,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | $spl152),
  inference(resolution,[],[f8836,f649])).
fof(f649,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f485])).
fof(f1128,plain,(
  ( ! [X3,X1] : (s(bool,X1) != s(bool,X3) | p(s(bool,X1))) ) | $spl14),
  inference(cnf_transformation,[],[f1128_D])).
fof(f1128_D,plain,(
  ( ! [X3,X1] : (s(bool,X1) != s(bool,X3) | p(s(bool,X1))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f39284,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | $spl152),
  inference(resolution,[],[f8836,f648])).
fof(f648,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f485])).
fof(f39283,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14))) | s(bool,X8) != s(bool,X9)) ) | $spl152),
  inference(resolution,[],[f8836,f646])).
fof(f646,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f485])).
fof(f39281,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | $spl152),
  inference(resolution,[],[f8836,f644])).
fof(f644,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f485])).
fof(f8839,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | $spl156),
  inference(cnf_transformation,[],[f8839_D])).
fof(f8839_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f39063,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | $spl62),
  inference(resolution,[],[f2548,f976])).
fof(f38972,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14))) | s(bool,X23) != s(bool,X24)) ) | $spl62),
  inference(resolution,[],[f2548,f650])).
fof(f38969,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | $spl62),
  inference(resolution,[],[f2548,f649])).
fof(f38968,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | $spl62),
  inference(resolution,[],[f2548,f648])).
fof(f38967,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14))) | s(bool,X8) != s(bool,X9)) ) | $spl62),
  inference(resolution,[],[f2548,f646])).
fof(f38965,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | $spl62),
  inference(resolution,[],[f2548,f644])).
fof(f20525,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl190),
  inference(cnf_transformation,[],[f20525_D])).
fof(f20525_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f38868,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl61),
  inference(resolution,[],[f2547,f976])).
fof(f38773,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X21) != s(bool,X23)) ) | $spl61),
  inference(resolution,[],[f2547,f650])).
fof(f38771,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl61),
  inference(resolution,[],[f2547,f649])).
fof(f38769,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl61),
  inference(resolution,[],[f2547,f648])).
fof(f38768,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X6) != s(bool,X8)) ) | $spl61),
  inference(resolution,[],[f2547,f646])).
fof(f38766,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl61),
  inference(resolution,[],[f2547,f644])).
fof(f16669,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl180),
  inference(cnf_transformation,[],[f16669_D])).
fof(f16669_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f38667,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl54),
  inference(resolution,[],[f2535,f976])).
fof(f38576,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) | s(bool,X23) != s(bool,X24)) ) | $spl54),
  inference(resolution,[],[f2535,f650])).
fof(f38573,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl54),
  inference(resolution,[],[f2535,f649])).
fof(f38572,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl54),
  inference(resolution,[],[f2535,f648])).
fof(f38571,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) | s(bool,X8) != s(bool,X9)) ) | $spl54),
  inference(resolution,[],[f2535,f646])).
fof(f38569,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl54),
  inference(resolution,[],[f2535,f644])).
fof(f985,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) | $spl1),
  inference(cnf_transformation,[],[f985_D])).
fof(f985_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f21486,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl202),
  inference(cnf_transformation,[],[f21486_D])).
fof(f21486_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f38565,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) | $spl54),
  inference(resolution,[],[f2535,f901])).
fof(f901,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f614])).
fof(f614,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) | (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))))) & (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1))))))),
  inference(flattening,[],[f613])).
fof(f613,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) | (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))))) & ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f358])).
fof(f358,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) <=> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0))))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6,X7] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))),s('type.integer.int',X7)))) <=> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X7))),s('type.integer.int',X6))))))),
  file('thm.integer.INT_ABS_LE|split|2','thm.integer.INT_ABS_LE|split|0')).
fof(f34521,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) ) | $spl306),
  inference(cnf_transformation,[],[f34521_D])).
fof(f34521_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) ) <=> ~$spl306),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl306])])).
fof(f38493,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) ) | $spl36),
  inference(resolution,[],[f1719,f976])).
fof(f38402,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))))) | s(bool,X23) != s(bool,X24)) ) | $spl36),
  inference(resolution,[],[f1719,f650])).
fof(f38399,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) ) | $spl36),
  inference(resolution,[],[f1719,f649])).
fof(f38398,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) ) | $spl36),
  inference(resolution,[],[f1719,f648])).
fof(f38397,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))))) | s(bool,X8) != s(bool,X9)) ) | $spl36),
  inference(resolution,[],[f1719,f646])).
fof(f38395,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) ) | $spl36),
  inference(resolution,[],[f1719,f644])).
fof(f38377,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) ) | $spl320),
  inference(cnf_transformation,[],[f38377_D])).
fof(f38377_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) ) <=> ~$spl320),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl320])])).
fof(f38352,plain,(
  ( ! [X187,X188] : (p(s(bool,X187)) | ~p(s(bool,X188)) | s(bool,X188) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) ) | $spl221),
  inference(resolution,[],[f25264,f976])).
fof(f38257,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14))) | s(bool,X22) != s(bool,X24)) ) | $spl221),
  inference(resolution,[],[f25264,f650])).
fof(f38255,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) ) | $spl221),
  inference(resolution,[],[f25264,f649])).
fof(f38253,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) ) | $spl221),
  inference(resolution,[],[f25264,f648])).
fof(f38252,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14))) | s(bool,X7) != s(bool,X9)) ) | $spl221),
  inference(resolution,[],[f25264,f646])).
fof(f38250,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) ) | $spl221),
  inference(resolution,[],[f25264,f644])).
fof(f38376,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl319),
  inference(cnf_transformation,[],[f38376_D])).
fof(f38376_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) <=> ~$spl319),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl319])])).
fof(f36766,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl310),
  inference(cnf_transformation,[],[f36766_D])).
fof(f36766_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) <=> ~$spl310),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl310])])).
fof(f38248,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl221),
  inference(resolution,[],[f25264,f933])).
fof(f933,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))))) )),
  inference(cnf_transformation,[],[f629])).
fof(f629,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))))) & ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) & ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))))))),
  inference(flattening,[],[f628])).
fof(f628,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))) | (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))))) & ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) & ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))))))),
  inference(nnf_transformation,[],[f380])).
fof(f380,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X6,X7] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6)))))) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int',X6)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X7))))))))),
  file('thm.integer.INT_ABS_LE|split|2','thm.integer.INT_ABS_LT|split|1')).
fof(f25428,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK14)))) | $spl225),
  inference(cnf_transformation,[],[f25428_D])).
fof(f25428_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK14)))) <=> ~$spl225),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl225])])).
fof(f8836,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | $spl152),
  inference(cnf_transformation,[],[f8836_D])).
fof(f8836_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f38246,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK14)))) | $spl221),
  inference(resolution,[],[f25264,f927])).
fof(f927,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f625])).
fof(f625,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) & ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) & ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))),s('type.integer.int',X1))))))),
  inference(flattening,[],[f624])).
fof(f624,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))),s('type.integer.int',X1)))) | (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))) & ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) & ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f378])).
fof(f378,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))),s('type.integer.int',X1)))) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6,X7] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))))),s('type.integer.int',X7)))) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X7))),s('type.integer.int',X6)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int',X7))))))),
  file('thm.integer.INT_ABS_LE|split|2','thm.integer.INT_ABS_LT|split|2')).
fof(f8101,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl146),
  inference(cnf_transformation,[],[f8101_D])).
fof(f8101_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f38141,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl7),
  inference(resolution,[],[f1121,f976])).
fof(f38046,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X21) != s(bool,X23)) ) | $spl7),
  inference(resolution,[],[f1121,f650])).
fof(f38044,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl7),
  inference(resolution,[],[f1121,f649])).
fof(f38042,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl7),
  inference(resolution,[],[f1121,f648])).
fof(f38041,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X6) != s(bool,X8)) ) | $spl7),
  inference(resolution,[],[f1121,f646])).
fof(f38039,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl7),
  inference(resolution,[],[f1121,f644])).
fof(f8097,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl143),
  inference(cnf_transformation,[],[f8097_D])).
fof(f8097_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) <=> ~$spl143),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl143])])).
fof(f8099,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl144),
  inference(cnf_transformation,[],[f8099_D])).
fof(f8099_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f38033,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl7),
  inference(resolution,[],[f1121,f896])).
fof(f896,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))))) )),
  inference(cnf_transformation,[],[f612])).
fof(f612,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))))) & ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) & ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))))))),
  inference(flattening,[],[f611])).
fof(f611,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))) | (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))))) & ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) & ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))))))),
  inference(nnf_transformation,[],[f357])).
fof(f357,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))) <=> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6,X7] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6)))))) <=> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X7),s('type.integer.int',X6)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X6),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X7))))))))),
  file('thm.integer.INT_ABS_LE|split|2','thm.integer.INT_ABS_LE|split|1')).
fof(f37771,plain,(
  $false | ($spl122 | $spl124)),
  inference(subsumption_resolution,[],[f37512,f5364])).
fof(f37512,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl122),
  inference(superposition,[],[f965,f5266])).
fof(f965,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(cnf_transformation,[],[f403])).
fof(f403,plain,(
  ! [X0] : ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(flattening,[],[f402])).
fof(f402,plain,(
  ! [X0] : ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X6] : ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  file('thm.integer.INT_ABS_LE|split|2','thm.integer.INT_ABS_LT|48|')).
fof(f5368,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl128),
  inference(cnf_transformation,[],[f5368_D])).
fof(f5368_D,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f5127,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl113),
  inference(cnf_transformation,[],[f5127_D])).
fof(f5127_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) <=> ~$spl113),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl113])])).
fof(f37768,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl122),
  inference(forward_demodulation,[],[f37494,f921])).
fof(f37494,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl122),
  inference(superposition,[],[f914,f5266])).
fof(f914,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(cnf_transformation,[],[f620])).
fof(f620,plain,(
  ! [X0] : ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))),
  inference(nnf_transformation,[],[f368])).
fof(f368,plain,(
  ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> s('type.integer.int',X6) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.integer.INT_ABS_LE|split|2','thm.integer.INT_ABS_LE|48|')).
fof(f37433,plain,(
  $false | ($spl113 | $spl154)),
  inference(subsumption_resolution,[],[f37432,f888])).
fof(f37432,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) | ($spl113 | $spl154)),
  inference(forward_demodulation,[],[f5127,f8838])).
fof(f5367,plain,(
  s('type.integer.int',sK14) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl129),
  inference(cnf_transformation,[],[f5367_D])).
fof(f5367_D,plain,(
  s('type.integer.int',sK14) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl129),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl129])])).
fof(f25264,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) | $spl221),
  inference(cnf_transformation,[],[f25264_D])).
fof(f25264_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) <=> ~$spl221),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl221])])).
fof(f25399,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) | s('type.integer.int',sK14) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl153),
  inference(superposition,[],[f8835,f951])).
fof(f951,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f634])).
fof(f634,plain,(
  ! [X0] : ((s('type.integer.int',X0) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))),
  inference(nnf_transformation,[],[f391])).
fof(f391,plain,(
  ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) <=> s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f173])).
fof(f173,axiom,(
  ! [X6] : (s('type.integer.int',X6) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X6))) <=> s('type.integer.int',X6) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.integer.INT_ABS_LE|split|2','thm.integer.INT_NEG_SAME_EQ')).
fof(f8835,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | $spl153),
  inference(cnf_transformation,[],[f8835_D])).
fof(f8835_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) <=> ~$spl153),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl153])])).
fof(f37429,plain,(
  ( ! [X0] : (sP12(X0,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) ) | $spl316),
  inference(cnf_transformation,[],[f37429_D])).
fof(f37429_D,plain,(
  ( ! [X0] : (sP12(X0,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) ) <=> ~$spl316),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl316])])).
fof(f27534,plain,(
  ( ! [X0] : (sP12(X0,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK14))) | s('type.integer.int',sK14) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl112),
  inference(superposition,[],[f27469,f951])).
fof(f27469,plain,(
  ( ! [X167] : (sP12(X167,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) ) | $spl112),
  inference(resolution,[],[f5126,f780])).
fof(f780,plain,(
  ( ! [X0,X1] : (sP12(X0,X1) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f546])).
fof(f546,plain,(
  ! [X0,X1] : ((~sP12(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP12(X0,X1)))),
  inference(flattening,[],[f545])).
fof(f545,plain,(
  ! [X0,X1] : ((~sP12(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP12(X0,X1)))),
  inference(nnf_transformation,[],[f477])).
fof(f477,plain,(
  ! [X0,X1] : (sP12(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP12])])).
fof(f37428,plain,(
  ( ! [X0] : (sP12('const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)),X0)) ) | $spl314),
  inference(cnf_transformation,[],[f37428_D])).
fof(f37428_D,plain,(
  ( ! [X0] : (sP12('const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)),X0)) ) <=> ~$spl314),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl314])])).
fof(f27553,plain,(
  ( ! [X0] : (sP12('const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)),X0) | s('type.integer.int',sK14) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl112),
  inference(superposition,[],[f27470,f951])).
fof(f27470,plain,(
  ( ! [X168] : (sP12('const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))),X168)) ) | $spl112),
  inference(resolution,[],[f5126,f781])).
fof(f781,plain,(
  ( ! [X0,X1] : (sP12(X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f546])).
fof(f5366,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | $spl126),
  inference(cnf_transformation,[],[f5366_D])).
fof(f5366_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f37091,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl154),
  inference(duplicate_literal_removal,[],[f36835])).
fof(f36835,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl154),
  inference(superposition,[],[f956,f8838])).
fof(f37395,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | $spl154),
  inference(forward_demodulation,[],[f37394,f8838])).
fof(f37394,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl154),
  inference(duplicate_literal_removal,[],[f37393])).
fof(f37393,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl154),
  inference(forward_demodulation,[],[f37392,f958])).
fof(f37392,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl154),
  inference(forward_demodulation,[],[f37034,f959])).
fof(f37034,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl154),
  inference(superposition,[],[f956,f8838])).
fof(f37390,plain,(
  $false | ($spl129 | $spl154)),
  inference(subsumption_resolution,[],[f37389,f5367])).
fof(f37389,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl154),
  inference(forward_demodulation,[],[f37388,f8838])).
fof(f37388,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl154),
  inference(subsumption_resolution,[],[f37031,f8838])).
fof(f37031,plain,(
  s('type.integer.int',sK14) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl154),
  inference(superposition,[],[f950,f8838])).
fof(f950,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f634])).
fof(f37276,plain,(
  $false | ($spl123 | $spl154)),
  inference(subsumption_resolution,[],[f37275,f888])).
fof(f37275,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) | ($spl123 | $spl154)),
  inference(forward_demodulation,[],[f37274,f8838])).
fof(f37274,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | ($spl123 | $spl154)),
  inference(forward_demodulation,[],[f37273,f921])).
fof(f37273,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl123 | $spl154)),
  inference(forward_demodulation,[],[f37272,f919])).
fof(f37272,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | ($spl123 | $spl154)),
  inference(subsumption_resolution,[],[f36948,f5267])).
fof(f36948,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl154),
  inference(superposition,[],[f917,f8838])).
fof(f37271,plain,(
  $false | ($spl123 | $spl154)),
  inference(subsumption_resolution,[],[f37270,f5267])).
fof(f37270,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | $spl154),
  inference(forward_demodulation,[],[f37269,f8838])).
fof(f37269,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) | $spl154),
  inference(forward_demodulation,[],[f37268,f961])).
fof(f961,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))) )),
  inference(cnf_transformation,[],[f399])).
fof(f399,plain,(
  ! [X0] : s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X6)))))),
  file('thm.integer.INT_ABS_LE|split|2','thm.integer.INT_ABS_NEG')).
fof(f37268,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))))) | $spl154),
  inference(subsumption_resolution,[],[f37267,f888])).
fof(f37267,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))))) | $spl154),
  inference(forward_demodulation,[],[f37266,f8838])).
fof(f37266,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))))) | $spl154),
  inference(forward_demodulation,[],[f36947,f922])).
fof(f36947,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))))) | $spl154),
  inference(superposition,[],[f917,f8838])).
fof(f37111,plain,(
  $false | ($spl129 | $spl154)),
  inference(subsumption_resolution,[],[f37092,f5367])).
fof(f37092,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl154),
  inference(trivial_inequality_removal,[],[f36834])).
fof(f36834,plain,(
  s('type.integer.int',sK14) != s('type.integer.int',sK14) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl154),
  inference(superposition,[],[f950,f8838])).
fof(f2549,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | $spl63),
  inference(cnf_transformation,[],[f2549_D])).
fof(f2549_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) <=> ~$spl63),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl63])])).
fof(f8838,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) | $spl154),
  inference(cnf_transformation,[],[f8838_D])).
fof(f8838_D,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f25287,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | $spl153),
  inference(resolution,[],[f8835,f904])).
fof(f25288,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | $spl153),
  inference(resolution,[],[f8835,f909])).
fof(f36768,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int',sK14)))) | $spl313),
  inference(cnf_transformation,[],[f36768_D])).
fof(f36768_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int',sK14)))) <=> ~$spl313),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl313])])).
fof(f25705,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int',sK14)))) | $spl77),
  inference(forward_demodulation,[],[f25694,f922])).
fof(f25694,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | $spl77),
  inference(superposition,[],[f2575,f917])).
fof(f27384,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl112),
  inference(resolution,[],[f5126,f904])).
fof(f27385,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) | $spl112),
  inference(resolution,[],[f5126,f909])).
fof(f36764,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int',sK14)))) | $spl309),
  inference(cnf_transformation,[],[f36764_D])).
fof(f36764_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int',sK14)))) <=> ~$spl309),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl309])])).
fof(f33979,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int',sK14)))) | $spl79),
  inference(forward_demodulation,[],[f33968,f922])).
fof(f33968,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | $spl79),
  inference(superposition,[],[f2578,f917])).
fof(f36749,plain,(
  $false | ($spl3 | $spl63 | $spl286)),
  inference(subsumption_resolution,[],[f36748,f889])).
fof(f889,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f350])).
fof(f350,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.integer.INT_ABS_LE|split|2',aHL_FALSITY)).
fof(f36748,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl63 | $spl286)),
  inference(forward_demodulation,[],[f36747,f29945])).
fof(f29945,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | $spl3),
  inference(resolution,[],[f987,f848])).
fof(f848,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f592])).
fof(f592,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X0)) & (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f591])).
fof(f591,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X0)) & (s(bool,'F_0') != s(bool,X0) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f442])).
fof(f442,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X0))),
  inference(ennf_transformation,[],[f321])).
fof(f321,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X0))),
  inference(rectify,[],[f74])).
fof(f74,axiom,(
  ! [X12] : ((p(s(bool,X12)) => p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X12))),
  file('thm.integer.INT_ABS_LE|split|2','thm.bool.IMP_F_EQ_F')).
fof(f36747,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | ($spl63 | $spl286)),
  inference(forward_demodulation,[],[f36746,f31759])).
fof(f36746,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl63),
  inference(subsumption_resolution,[],[f36610,f888])).
fof(f36610,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl63),
  inference(resolution,[],[f20936,f896])).
fof(f20936,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | $spl63),
  inference(resolution,[],[f2549,f900])).
fof(f900,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f614])).
fof(f28563,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl276),
  inference(cnf_transformation,[],[f28563_D])).
fof(f28563_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f35740,plain,(
  ( ! [X187,X188] : (p(s(bool,X187)) | ~p(s(bool,X188)) | s(bool,X188) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl34),
  inference(resolution,[],[f14977,f976])).
fof(f14977,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | $spl34),
  inference(resolution,[],[f1716,f893])).
fof(f893,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f609])).
fof(f609,plain,(
  ! [X0,X1] : ((p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) & (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f354])).
fof(f354,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(flattening,[],[f353])).
fof(f353,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f146])).
fof(f146,axiom,(
  ! [X9,X11] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X9),s('type.integer.int',X11)))) <=> p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X11),s('type.integer.int',X9)))))),
  file('thm.integer.INT_ABS_LE|split|2','thm.integer.INT_NOT_LE')).
fof(f1716,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl34),
  inference(cnf_transformation,[],[f1716_D])).
fof(f1716_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f35645,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X22) != s(bool,X24)) ) | $spl34),
  inference(resolution,[],[f14977,f650])).
fof(f35643,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl34),
  inference(resolution,[],[f14977,f649])).
fof(f35641,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl34),
  inference(resolution,[],[f14977,f648])).
fof(f35640,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X7) != s(bool,X9)) ) | $spl34),
  inference(resolution,[],[f14977,f646])).
fof(f35638,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl34),
  inference(resolution,[],[f14977,f644])).
fof(f26842,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl255),
  inference(cnf_transformation,[],[f26842_D])).
fof(f26842_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) <=> ~$spl255),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl255])])).
fof(f4845,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))))) | $spl106),
  inference(cnf_transformation,[],[f4845_D])).
fof(f4845_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))))) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f35636,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl34),
  inference(resolution,[],[f14977,f933])).
fof(f28561,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl274),
  inference(cnf_transformation,[],[f28561_D])).
fof(f28561_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f3585,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))),s('type.integer.int',sK14)))) | $spl91),
  inference(cnf_transformation,[],[f3585_D])).
fof(f3585_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))),s('type.integer.int',sK14)))) <=> ~$spl91),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl91])])).
fof(f35772,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl34),
  inference(forward_demodulation,[],[f35634,f968])).
fof(f968,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))) )),
  inference(cnf_transformation,[],[f405])).
fof(f405,plain,(
  ! [X0] : s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X6] : s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6)))))),
  file('thm.integer.INT_ABS_LE|split|2','thm.integer.INT_ABS_ABS')).
fof(f35634,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))))),s('type.integer.int',sK14)))) | $spl34),
  inference(resolution,[],[f14977,f927])).
fof(f26109,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) | $spl244),
  inference(cnf_transformation,[],[f26109_D])).
fof(f26109_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f35076,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) | $spl88),
  inference(resolution,[],[f3058,f976])).
fof(f34985,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))))) | s(bool,X23) != s(bool,X24)) ) | $spl88),
  inference(resolution,[],[f3058,f650])).
fof(f34982,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) | $spl88),
  inference(resolution,[],[f3058,f649])).
fof(f34981,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) | $spl88),
  inference(resolution,[],[f3058,f648])).
fof(f34980,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))))) | s(bool,X8) != s(bool,X9)) ) | $spl88),
  inference(resolution,[],[f3058,f646])).
fof(f34978,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) | $spl88),
  inference(resolution,[],[f3058,f644])).
fof(f34910,plain,(
  $false | ($spl5 | $spl89)),
  inference(subsumption_resolution,[],[f34762,f1140])).
fof(f1140,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl5),
  inference(resolution,[],[f989,f892])).
fof(f892,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f609])).
fof(f34762,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl89),
  inference(resolution,[],[f3057,f907])).
fof(f907,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f619])).
fof(f34909,plain,(
  $false | ($spl5 | $spl89)),
  inference(subsumption_resolution,[],[f34760,f1140])).
fof(f34760,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl89),
  inference(resolution,[],[f3057,f905])).
fof(f905,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f617])).
fof(f34908,plain,(
  $false | ($spl5 | $spl89)),
  inference(subsumption_resolution,[],[f34759,f1145])).
fof(f1145,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | $spl5),
  inference(resolution,[],[f989,f905])).
fof(f34759,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | $spl89),
  inference(resolution,[],[f3057,f903])).
fof(f903,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f615])).
fof(f615,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) & (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f360])).
fof(f360,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(flattening,[],[f359])).
fof(f359,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X9,X11] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X9),s('type.integer.int',X11)))) <=> ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X11),s('type.integer.int',X9)))))),
  file('thm.integer.INT_ABS_LE|split|2','thm.integer.int_le')).
fof(f34907,plain,(
  $false | ($spl3 | $spl34 | $spl89)),
  inference(subsumption_resolution,[],[f34906,f1716])).
fof(f34906,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | ($spl3 | $spl89)),
  inference(forward_demodulation,[],[f34905,f961])).
fof(f34905,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))))) | ($spl3 | $spl89)),
  inference(subsumption_resolution,[],[f34904,f880])).
fof(f880,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f448])).
fof(f448,plain,(
  ! [X0] : ((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f339])).
fof(f339,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(flattening,[],[f338])).
fof(f338,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(rectify,[],[f164])).
fof(f164,axiom,(
  ! [X12] : ((p(s(bool,X12)) => p(s(bool,'F_0'))) => ~p(s(bool,X12)))),
  file('thm.integer.INT_ABS_LE|split|2','thm.bool.IMP_F')).
fof(f34904,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))))) | ($spl3 | $spl89)),
  inference(forward_demodulation,[],[f34903,f29945])).
fof(f34903,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))))) | $spl89),
  inference(forward_demodulation,[],[f34757,f913])).
fof(f913,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) )),
  inference(cnf_transformation,[],[f367])).
fof(f367,plain,(
  ! [X0,X1] : s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X9,X11] : s(bool,'const.integer.int_le_2'(s('type.integer.int',X11),s('type.integer.int',X9))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X11)))))),
  file('thm.integer.INT_ABS_LE|split|2','thm.integer.INT_LE_NEG')).
fof(f34757,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))))) | $spl89),
  inference(resolution,[],[f3057,f896])).
fof(f34902,plain,(
  $false | ($spl5 | $spl89)),
  inference(subsumption_resolution,[],[f34756,f1145])).
fof(f34756,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | $spl89),
  inference(resolution,[],[f3057,f894])).
fof(f894,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f610])).
fof(f610,plain,(
  ! [X0,X1] : ((p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) & (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f356])).
fof(f356,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(flattening,[],[f355])).
fof(f355,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X9,X11] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X9),s('type.integer.int',X11)))) <=> p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X11),s('type.integer.int',X9)))))),
  file('thm.integer.INT_ABS_LE|split|2','thm.integer.INT_NOT_LT')).
fof(f34901,plain,(
  $false | ($spl5 | $spl89)),
  inference(subsumption_resolution,[],[f34755,f1145])).
fof(f34755,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | $spl89),
  inference(resolution,[],[f3057,f892])).
fof(f25683,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl236),
  inference(cnf_transformation,[],[f25683_D])).
fof(f25683_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f34689,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl44),
  inference(resolution,[],[f2019,f976])).
fof(f34598,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))) | s(bool,X23) != s(bool,X24)) ) | $spl44),
  inference(resolution,[],[f2019,f650])).
fof(f34595,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl44),
  inference(resolution,[],[f2019,f649])).
fof(f34594,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl44),
  inference(resolution,[],[f2019,f648])).
fof(f34593,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))) | s(bool,X8) != s(bool,X9)) ) | $spl44),
  inference(resolution,[],[f2019,f646])).
fof(f34591,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl44),
  inference(resolution,[],[f2019,f644])).
fof(f25682,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl234),
  inference(cnf_transformation,[],[f25682_D])).
fof(f25682_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f34588,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))),s('type.integer.int',sK14)))) | $spl44),
  inference(resolution,[],[f2019,f938])).
fof(f938,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f631])).
fof(f631,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) | (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))))) & (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1))))))),
  inference(flattening,[],[f630])).
fof(f630,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) | (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))))) & ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f381])).
fof(f381,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1)))) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0))))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X6,X7] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))),s('type.integer.int',X7)))) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X7))),s('type.integer.int',X6))))))),
  file('thm.integer.INT_ABS_LE|split|2','thm.integer.INT_ABS_LT|split|0')).
fof(f25679,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl233),
  inference(cnf_transformation,[],[f25679_D])).
fof(f25679_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) <=> ~$spl233),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl233])])).
fof(f34720,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl44),
  inference(forward_demodulation,[],[f34586,f968])).
fof(f34586,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl44),
  inference(resolution,[],[f2019,f932])).
fof(f932,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))))) )),
  inference(cnf_transformation,[],[f627])).
fof(f627,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))))) | (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))) & (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))))))))),
  inference(flattening,[],[f626])).
fof(f626,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))))) | (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))) & ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))))))))),
  inference(nnf_transformation,[],[f379])).
fof(f379,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))))) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6,X7] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6)))))))) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X7)))))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int',X6))))))),
  file('thm.integer.INT_ABS_LE|split|2','thm.integer.INT_ABS_LT|split|3')).
fof(f34483,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) ) | $spl36),
  inference(resolution,[],[f1719,f976])).
fof(f34392,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))))) | s(bool,X23) != s(bool,X24)) ) | $spl36),
  inference(resolution,[],[f1719,f650])).
fof(f34389,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) ) | $spl36),
  inference(resolution,[],[f1719,f649])).
fof(f34388,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) ) | $spl36),
  inference(resolution,[],[f1719,f648])).
fof(f34387,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))))) | s(bool,X8) != s(bool,X9)) ) | $spl36),
  inference(resolution,[],[f1719,f646])).
fof(f34385,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) ) | $spl36),
  inference(resolution,[],[f1719,f644])).
fof(f15117,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl176),
  inference(cnf_transformation,[],[f15117_D])).
fof(f15117_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f34277,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl34),
  inference(resolution,[],[f1716,f976])).
fof(f34186,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))) | s(bool,X23) != s(bool,X24)) ) | $spl34),
  inference(resolution,[],[f1716,f650])).
fof(f34183,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl34),
  inference(resolution,[],[f1716,f649])).
fof(f34182,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl34),
  inference(resolution,[],[f1716,f648])).
fof(f34181,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))) | s(bool,X8) != s(bool,X9)) ) | $spl34),
  inference(resolution,[],[f1716,f646])).
fof(f34179,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl34),
  inference(resolution,[],[f1716,f644])).
fof(f15114,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))),s('type.integer.int',sK14)))) | $spl173),
  inference(cnf_transformation,[],[f15114_D])).
fof(f15114_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))),s('type.integer.int',sK14)))) <=> ~$spl173),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl173])])).
fof(f8096,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl142),
  inference(cnf_transformation,[],[f8096_D])).
fof(f8096_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f34175,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))),s('type.integer.int',sK14)))) | $spl34),
  inference(resolution,[],[f1716,f901])).
fof(f1717,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl35),
  inference(cnf_transformation,[],[f1717_D])).
fof(f1717_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f1719,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl36),
  inference(cnf_transformation,[],[f1719_D])).
fof(f1719_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f32983,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl3),
  inference(subsumption_resolution,[],[f32675,f851])).
fof(f851,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f594])).
fof(f594,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f593])).
fof(f593,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f322])).
fof(f322,plain,(
  ! [X0] : ((p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X12] : ((p(s(bool,'F_0')) | p(s(bool,X12))) <=> p(s(bool,X12)))),
  file('thm.integer.INT_ABS_LE|split|2','thm.bool.OR_CLAUSES|split|2')).
fof(f32675,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl3),
  inference(superposition,[],[f896,f29945])).
fof(f34129,plain,(
  $false | ($spl5 | $spl45)),
  inference(subsumption_resolution,[],[f33985,f1287])).
fof(f1287,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | $spl5),
  inference(forward_demodulation,[],[f1143,f961])).
fof(f1143,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))))),s('type.integer.int',sK14)))) | $spl5),
  inference(resolution,[],[f989,f899])).
fof(f899,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f614])).
fof(f33985,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | $spl45),
  inference(resolution,[],[f2020,f903])).
fof(f34128,plain,(
  $false | ($spl5 | $spl45)),
  inference(subsumption_resolution,[],[f33984,f1287])).
fof(f33984,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | $spl45),
  inference(resolution,[],[f2020,f894])).
fof(f34127,plain,(
  $false | ($spl5 | $spl45)),
  inference(subsumption_resolution,[],[f33983,f1287])).
fof(f33983,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | $spl45),
  inference(resolution,[],[f2020,f892])).
fof(f34126,plain,(
  $false | ($spl45 | $spl46)),
  inference(subsumption_resolution,[],[f33981,f2022])).
fof(f33981,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl45),
  inference(resolution,[],[f2020,f935])).
fof(f935,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))))) )),
  inference(cnf_transformation,[],[f629])).
fof(f33921,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | ($spl31 | $spl286)),
  inference(subsumption_resolution,[],[f33920,f1296])).
fof(f33920,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | $spl286),
  inference(forward_demodulation,[],[f33919,f31759])).
fof(f33919,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) | $spl286),
  inference(forward_demodulation,[],[f33918,f958])).
fof(f33918,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) | $spl286),
  inference(forward_demodulation,[],[f33535,f959])).
fof(f33535,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) | $spl286),
  inference(superposition,[],[f956,f31759])).
fof(f33912,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ($spl31 | $spl286)),
  inference(forward_demodulation,[],[f33911,f31759])).
fof(f33911,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | ($spl31 | $spl286)),
  inference(subsumption_resolution,[],[f33532,f1296])).
fof(f33532,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl286),
  inference(superposition,[],[f954,f31759])).
fof(f33910,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | ($spl31 | $spl286)),
  inference(forward_demodulation,[],[f33909,f31759])).
fof(f33909,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | ($spl31 | $spl286)),
  inference(forward_demodulation,[],[f33908,f31759])).
fof(f33908,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | ($spl31 | $spl286)),
  inference(subsumption_resolution,[],[f33530,f1296])).
fof(f33530,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl286),
  inference(superposition,[],[f952,f31759])).
fof(f33786,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))) | $spl304),
  inference(cnf_transformation,[],[f33786_D])).
fof(f33786_D,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))) <=> ~$spl304),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl304])])).
fof(f33789,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))) | $spl286),
  inference(forward_demodulation,[],[f33441,f919])).
fof(f33441,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl286),
  inference(superposition,[],[f917,f31759])).
fof(f33787,plain,(
  s('type.integer.int',sK15) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))) | $spl305),
  inference(cnf_transformation,[],[f33787_D])).
fof(f33787_D,plain,(
  s('type.integer.int',sK15) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))) <=> ~$spl305),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl305])])).
fof(f33785,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',sK15) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))) | $spl286),
  inference(forward_demodulation,[],[f33439,f919])).
fof(f33439,plain,(
  s('type.integer.int',sK15) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl286),
  inference(superposition,[],[f916,f31759])).
fof(f916,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | s('type.integer.int',X0) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f621])).
fof(f33772,plain,(
  ( ! [X132] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X132),s('type.integer.int',sK15)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X132))))) ) | $spl286),
  inference(forward_demodulation,[],[f33427,f919])).
fof(f33427,plain,(
  ( ! [X132] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X132),s('type.integer.int',sK15)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X132))))) ) | $spl286),
  inference(superposition,[],[f910,f31759])).
fof(f33770,plain,(
  ( ! [X131] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK15),s('type.integer.int',X131)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X131))))) ) | $spl286),
  inference(forward_demodulation,[],[f33426,f919])).
fof(f33426,plain,(
  ( ! [X131] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK15),s('type.integer.int',X131)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X131)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))))) ) | $spl286),
  inference(superposition,[],[f910,f31759])).
fof(f33647,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl286),
  inference(forward_demodulation,[],[f33646,f919])).
fof(f33646,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl286),
  inference(forward_demodulation,[],[f33645,f921])).
fof(f33645,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl286),
  inference(forward_demodulation,[],[f33334,f919])).
fof(f33334,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl286),
  inference(superposition,[],[f885,f31759])).
fof(f33639,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK43('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl286),
  inference(forward_demodulation,[],[f33638,f31759])).
fof(f33638,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK43('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | $spl286),
  inference(forward_demodulation,[],[f33332,f921])).
fof(f33332,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK43('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | $spl286),
  inference(superposition,[],[f881,f31759])).
fof(f33633,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ($spl65 | $spl286)),
  inference(subsumption_resolution,[],[f33325,f2551])).
fof(f33325,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl286),
  inference(superposition,[],[f956,f31759])).
fof(f33621,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ($spl31 | $spl286)),
  inference(forward_demodulation,[],[f33620,f31759])).
fof(f33620,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | ($spl31 | $spl286)),
  inference(subsumption_resolution,[],[f33279,f1296])).
fof(f33279,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl286),
  inference(superposition,[],[f31759,f954])).
fof(f33615,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | ($spl31 | $spl286)),
  inference(forward_demodulation,[],[f33614,f31759])).
fof(f33614,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | ($spl31 | $spl286)),
  inference(forward_demodulation,[],[f33613,f31759])).
fof(f33613,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | ($spl31 | $spl286)),
  inference(subsumption_resolution,[],[f33277,f1296])).
fof(f33277,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl286),
  inference(superposition,[],[f31759,f952])).
fof(f24645,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl216),
  inference(cnf_transformation,[],[f24645_D])).
fof(f24645_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f32550,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f1287,f976])).
fof(f32455,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X21) != s(bool,X23)) ) | $spl5),
  inference(resolution,[],[f1287,f650])).
fof(f32453,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f1287,f649])).
fof(f32451,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f1287,f648])).
fof(f32450,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X6) != s(bool,X8)) ) | $spl5),
  inference(resolution,[],[f1287,f646])).
fof(f32448,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f1287,f644])).
fof(f21485,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl203),
  inference(cnf_transformation,[],[f21485_D])).
fof(f21485_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) <=> ~$spl203),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl203])])).
fof(f24643,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))))) | $spl214),
  inference(cnf_transformation,[],[f24643_D])).
fof(f24643_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))))) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f32442,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl5),
  inference(resolution,[],[f1287,f896])).
fof(f18061,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl184),
  inference(cnf_transformation,[],[f18061_D])).
fof(f18061_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f32126,plain,(
  ( ! [X187,X188] : (p(s(bool,X187)) | ~p(s(bool,X188)) | s(bool,X188) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f1145,f976])).
fof(f32031,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X22) != s(bool,X24)) ) | $spl5),
  inference(resolution,[],[f1145,f650])).
fof(f32029,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f1145,f649])).
fof(f32027,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f1145,f648])).
fof(f32026,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X7) != s(bool,X9)) ) | $spl5),
  inference(resolution,[],[f1145,f646])).
fof(f32024,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f1145,f644])).
fof(f11077,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) | $spl164),
  inference(cnf_transformation,[],[f11077_D])).
fof(f11077_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f31964,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) | $spl5),
  inference(resolution,[],[f1140,f976])).
fof(f31873,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))))) | s(bool,X23) != s(bool,X24)) ) | $spl5),
  inference(resolution,[],[f1140,f650])).
fof(f31870,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) | $spl5),
  inference(resolution,[],[f1140,f649])).
fof(f31869,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) | $spl5),
  inference(resolution,[],[f1140,f648])).
fof(f31868,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))))) | s(bool,X8) != s(bool,X9)) ) | $spl5),
  inference(resolution,[],[f1140,f646])).
fof(f31866,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) | $spl5),
  inference(resolution,[],[f1140,f644])).
fof(f2570,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl72),
  inference(cnf_transformation,[],[f2570_D])).
fof(f2570_D,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f31780,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | ($spl0 | $spl67)),
  inference(subsumption_resolution,[],[f31779,f2554])).
fof(f31779,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl0),
  inference(forward_demodulation,[],[f31731,f922])).
fof(f31731,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl0),
  inference(superposition,[],[f984,f954])).
fof(f2572,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | $spl74),
  inference(cnf_transformation,[],[f2572_D])).
fof(f2572_D,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f31778,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | ($spl0 | $spl67)),
  inference(subsumption_resolution,[],[f31777,f2554])).
fof(f31777,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl0),
  inference(forward_demodulation,[],[f31729,f922])).
fof(f31729,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl0),
  inference(superposition,[],[f984,f952])).
fof(f31760,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) ) | $spl288),
  inference(cnf_transformation,[],[f31760_D])).
fof(f31760_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) ) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f31704,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) ) | $spl0),
  inference(resolution,[],[f984,f976])).
fof(f31613,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15))) | s(bool,X23) != s(bool,X24)) ) | $spl0),
  inference(resolution,[],[f984,f650])).
fof(f31610,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) ) | $spl0),
  inference(resolution,[],[f984,f649])).
fof(f31609,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) ) | $spl0),
  inference(resolution,[],[f984,f648])).
fof(f31608,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15))) | s(bool,X8) != s(bool,X9)) ) | $spl0),
  inference(resolution,[],[f984,f646])).
fof(f31606,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) ) | $spl0),
  inference(resolution,[],[f984,f644])).
fof(f31759,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) | $spl286),
  inference(cnf_transformation,[],[f31759_D])).
fof(f31759_D,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f2015,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) | $spl40),
  inference(cnf_transformation,[],[f2015_D])).
fof(f2015_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f31605,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) | $spl0),
  inference(resolution,[],[f984,f909])).
fof(f31604,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) | $spl0),
  inference(resolution,[],[f984,f904])).
fof(f1288,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl26),
  inference(cnf_transformation,[],[f1288_D])).
fof(f1288_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f31477,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f989,f976])).
fof(f31382,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X21) != s(bool,X23)) ) | $spl5),
  inference(resolution,[],[f989,f650])).
fof(f31380,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f989,f649])).
fof(f31378,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f989,f648])).
fof(f31377,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X6) != s(bool,X8)) ) | $spl5),
  inference(resolution,[],[f989,f646])).
fof(f31375,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f989,f644])).
fof(f3590,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl97),
  inference(cnf_transformation,[],[f3590_D])).
fof(f3590_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) <=> ~$spl97),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl97])])).
fof(f6195,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl134),
  inference(cnf_transformation,[],[f6195_D])).
fof(f6195_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f30430,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl3),
  inference(resolution,[],[f1579,f932])).
fof(f1579,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | $spl3),
  inference(resolution,[],[f987,f892])).
fof(f31325,plain,(
  $false | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f31176,f2370])).
fof(f2370,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(resolution,[],[f985,f892])).
fof(f31176,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl4),
  inference(resolution,[],[f5754,f930])).
fof(f930,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))))))) )),
  inference(cnf_transformation,[],[f627])).
fof(f5754,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl4),
  inference(resolution,[],[f988,f893])).
fof(f12338,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl168),
  inference(cnf_transformation,[],[f12338_D])).
fof(f12338_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f31134,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl35),
  inference(resolution,[],[f1717,f976])).
fof(f31039,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))) | s(bool,X21) != s(bool,X23)) ) | $spl35),
  inference(resolution,[],[f1717,f650])).
fof(f31037,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl35),
  inference(resolution,[],[f1717,f649])).
fof(f31035,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl35),
  inference(resolution,[],[f1717,f648])).
fof(f31034,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))) | s(bool,X6) != s(bool,X8)) ) | $spl35),
  inference(resolution,[],[f1717,f646])).
fof(f31032,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl35),
  inference(resolution,[],[f1717,f644])).
fof(f6198,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl138),
  inference(cnf_transformation,[],[f6198_D])).
fof(f6198_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f30533,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl3),
  inference(resolution,[],[f1579,f976])).
fof(f30442,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X23) != s(bool,X24)) ) | $spl3),
  inference(resolution,[],[f1579,f650])).
fof(f30439,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl3),
  inference(resolution,[],[f1579,f649])).
fof(f30438,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl3),
  inference(resolution,[],[f1579,f648])).
fof(f30437,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X8) != s(bool,X9)) ) | $spl3),
  inference(resolution,[],[f1579,f646])).
fof(f30435,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl3),
  inference(resolution,[],[f1579,f644])).
fof(f4072,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl104),
  inference(cnf_transformation,[],[f4072_D])).
fof(f4072_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f30391,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl32),
  inference(resolution,[],[f1297,f976])).
fof(f30300,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X23) != s(bool,X24)) ) | $spl32),
  inference(resolution,[],[f1297,f650])).
fof(f30297,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl32),
  inference(resolution,[],[f1297,f649])).
fof(f30296,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl32),
  inference(resolution,[],[f1297,f648])).
fof(f30295,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X8) != s(bool,X9)) ) | $spl32),
  inference(resolution,[],[f1297,f646])).
fof(f30293,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl32),
  inference(resolution,[],[f1297,f644])).
fof(f2023,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl48),
  inference(cnf_transformation,[],[f2023_D])).
fof(f2023_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f30250,plain,(
  ( ! [X187,X188] : (p(s(bool,X187)) | ~p(s(bool,X188)) | s(bool,X188) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl9),
  inference(resolution,[],[f1123,f976])).
fof(f30155,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X22) != s(bool,X24)) ) | $spl9),
  inference(resolution,[],[f1123,f650])).
fof(f30153,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl9),
  inference(resolution,[],[f1123,f649])).
fof(f30151,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl9),
  inference(resolution,[],[f1123,f648])).
fof(f30150,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X7) != s(bool,X9)) ) | $spl9),
  inference(resolution,[],[f1123,f646])).
fof(f30148,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl9),
  inference(resolution,[],[f1123,f644])).
fof(f2016,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) | $spl41),
  inference(cnf_transformation,[],[f2016_D])).
fof(f2016_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f2018,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | $spl42),
  inference(cnf_transformation,[],[f2018_D])).
fof(f2018_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f30144,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) | $spl9),
  inference(resolution,[],[f1123,f927])).
fof(f5897,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl132),
  inference(cnf_transformation,[],[f5897_D])).
fof(f5897_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f30095,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl4),
  inference(resolution,[],[f988,f976])).
fof(f30004,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X23) != s(bool,X24)) ) | $spl4),
  inference(resolution,[],[f988,f650])).
fof(f30001,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl4),
  inference(resolution,[],[f988,f649])).
fof(f30000,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl4),
  inference(resolution,[],[f988,f648])).
fof(f29999,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X8) != s(bool,X9)) ) | $spl4),
  inference(resolution,[],[f988,f646])).
fof(f29997,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl4),
  inference(resolution,[],[f988,f644])).
fof(f5896,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))) | $spl130),
  inference(cnf_transformation,[],[f5896_D])).
fof(f5896_D,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f29996,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))) | $spl4),
  inference(resolution,[],[f988,f909])).
fof(f29995,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | $spl4),
  inference(resolution,[],[f988,f904])).
fof(f4069,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | $spl102),
  inference(cnf_transformation,[],[f4069_D])).
fof(f4069_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f1298,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | $spl33),
  inference(cnf_transformation,[],[f1298_D])).
fof(f1298_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f30130,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | $spl4),
  inference(forward_demodulation,[],[f30129,f913])).
fof(f30129,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl4),
  inference(forward_demodulation,[],[f29993,f961])).
fof(f29993,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl4),
  inference(resolution,[],[f988,f901])).
fof(f1720,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl38),
  inference(cnf_transformation,[],[f1720_D])).
fof(f1720_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f29953,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl3),
  inference(resolution,[],[f987,f976])).
fof(f29858,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X21) != s(bool,X23)) ) | $spl3),
  inference(resolution,[],[f987,f650])).
fof(f29856,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl3),
  inference(resolution,[],[f987,f649])).
fof(f29854,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl3),
  inference(resolution,[],[f987,f648])).
fof(f29853,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X6) != s(bool,X8)) ) | $spl3),
  inference(resolution,[],[f987,f646])).
fof(f29851,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl3),
  inference(resolution,[],[f987,f644])).
fof(f29845,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl3),
  inference(resolution,[],[f987,f896])).
fof(f8157,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | $spl150),
  inference(cnf_transformation,[],[f8157_D])).
fof(f8157_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f29842,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | p(s(bool,X1))) ) | $spl3),
  inference(subsumption_resolution,[],[f29832,f851])).
fof(f29832,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | p(s(bool,X1))) ) | $spl3),
  inference(resolution,[],[f6155,f779])).
fof(f779,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP12(X0,X1)) )),
  inference(cnf_transformation,[],[f546])).
fof(f6155,plain,(
  ( ! [X168] : (sP12('const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)),X168)) ) | $spl3),
  inference(resolution,[],[f1579,f781])).
fof(f29841,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | p(s(bool,X0))) ) | $spl3),
  inference(subsumption_resolution,[],[f29831,f851])).
fof(f29831,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | p(s(bool,X0))) ) | $spl3),
  inference(resolution,[],[f6155,f783])).
fof(f783,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP12(X0,X1)) )),
  inference(cnf_transformation,[],[f548])).
fof(f548,plain,(
  ! [X0,X1] : ((~sP12(X0,X1) | (p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP12(X0,X1)))),
  inference(flattening,[],[f547])).
fof(f547,plain,(
  ! [X0,X1] : ((~sP12(X0,X1) | ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0')))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP12(X0,X1)))),
  inference(nnf_transformation,[],[f478])).
fof(f478,plain,(
  ! [X0,X1] : (sP12(X0,X1) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(definition_folding,[],[f425,f477])).
fof(f425,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f424])).
fof(f424,plain,(
  ! [X0,X1] : (((p(s(bool,X1)) | p(s(bool,X0))) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(ennf_transformation,[],[f288])).
fof(f288,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(flattening,[],[f287])).
fof(f287,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(rectify,[],[f61])).
fof(f61,axiom,(
  ! [X22,X8] : ((~(p(s(bool,X8)) | p(s(bool,X22))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X8)) => p(s(bool,'F_0'))) => (~p(s(bool,X22)) => p(s(bool,'F_0')))))),
  file('thm.integer.INT_ABS_LE|split|2','thm.sat.sat_9_prove_ltb')).
fof(f29830,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))))) ) | $spl3),
  inference(subsumption_resolution,[],[f29820,f851])).
fof(f29820,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))))) ) | $spl3),
  inference(resolution,[],[f6154,f779])).
fof(f6154,plain,(
  ( ! [X167] : (sP12(X167,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl3),
  inference(resolution,[],[f1579,f780])).
fof(f29829,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))))) ) | $spl3),
  inference(subsumption_resolution,[],[f29819,f851])).
fof(f29819,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))))) ) | $spl3),
  inference(resolution,[],[f6154,f783])).
fof(f29794,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl284),
  inference(cnf_transformation,[],[f29794_D])).
fof(f29794_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f29754,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl1),
  inference(subsumption_resolution,[],[f29753,f965])).
fof(f29753,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl1),
  inference(forward_demodulation,[],[f29752,f959])).
fof(f29752,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(forward_demodulation,[],[f29751,f958])).
fof(f29751,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(forward_demodulation,[],[f29720,f959])).
fof(f29720,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(superposition,[],[f2370,f956])).
fof(f29718,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(superposition,[],[f2370,f954])).
fof(f29716,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(superposition,[],[f2370,f952])).
fof(f2567,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) | $spl71),
  inference(cnf_transformation,[],[f2567_D])).
fof(f2567_D,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) <=> ~$spl71),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl71])])).
fof(f29750,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl1),
  inference(forward_demodulation,[],[f29715,f962])).
fof(f29715,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) | $spl1),
  inference(superposition,[],[f2370,f950])).
fof(f29746,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl1),
  inference(forward_demodulation,[],[f29745,f968])).
fof(f29745,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl1),
  inference(forward_demodulation,[],[f29711,f961])).
fof(f29711,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl1),
  inference(superposition,[],[f2370,f914])).
fof(f29792,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) ) | $spl282),
  inference(cnf_transformation,[],[f29792_D])).
fof(f29792_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) ) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f29687,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) ) | $spl1),
  inference(resolution,[],[f2370,f976])).
fof(f29596,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) | s(bool,X23) != s(bool,X24)) ) | $spl1),
  inference(resolution,[],[f2370,f650])).
fof(f29593,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) ) | $spl1),
  inference(resolution,[],[f2370,f649])).
fof(f29592,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) ) | $spl1),
  inference(resolution,[],[f2370,f648])).
fof(f29591,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) | s(bool,X8) != s(bool,X9)) ) | $spl1),
  inference(resolution,[],[f2370,f646])).
fof(f29589,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) ) | $spl1),
  inference(resolution,[],[f2370,f644])).
fof(f29791,plain,(
  $false | ($spl4 | $spl9 | $spl44)),
  inference(subsumption_resolution,[],[f22514,f2019])).
fof(f22514,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | ($spl4 | $spl9)),
  inference(forward_demodulation,[],[f22513,f961])).
fof(f22513,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))))) | ($spl4 | $spl9)),
  inference(subsumption_resolution,[],[f22512,f1123])).
fof(f22512,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))))) | $spl4),
  inference(forward_demodulation,[],[f22370,f943])).
fof(f943,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) )),
  inference(cnf_transformation,[],[f385])).
fof(f385,plain,(
  ! [X0,X1] : s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X9,X11] : s(bool,'const.integer.int_lt_2'(s('type.integer.int',X11),s('type.integer.int',X9))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X11)))))),
  file('thm.integer.INT_ABS_LE|split|2','thm.integer.INT_LT_NEG')).
fof(f22370,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))))) | $spl4),
  inference(resolution,[],[f5754,f933])).
fof(f29790,plain,(
  $false | ($spl35 | $spl44)),
  inference(subsumption_resolution,[],[f12200,f2019])).
fof(f12200,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl35),
  inference(resolution,[],[f1717,f905])).
fof(f29789,plain,(
  $false | ($spl35 | $spl44)),
  inference(subsumption_resolution,[],[f12202,f2019])).
fof(f12202,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl35),
  inference(resolution,[],[f1717,f907])).
fof(f29788,plain,(
  $false | ($spl35 | $spl44)),
  inference(subsumption_resolution,[],[f16127,f2019])).
fof(f16127,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl35),
  inference(resolution,[],[f1717,f905])).
fof(f29787,plain,(
  $false | ($spl35 | $spl44)),
  inference(subsumption_resolution,[],[f16129,f2019])).
fof(f16129,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl35),
  inference(resolution,[],[f1717,f907])).
fof(f29786,plain,(
  $false | ($spl35 | $spl44)),
  inference(subsumption_resolution,[],[f20131,f2019])).
fof(f20131,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl35),
  inference(resolution,[],[f1717,f905])).
fof(f29785,plain,(
  $false | ($spl35 | $spl44)),
  inference(subsumption_resolution,[],[f20133,f2019])).
fof(f20133,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl35),
  inference(resolution,[],[f1717,f907])).
fof(f1123,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | $spl9),
  inference(cnf_transformation,[],[f1123_D])).
fof(f1123_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f25741,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | $spl46),
  inference(resolution,[],[f2022,f932])).
fof(f25678,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl232),
  inference(cnf_transformation,[],[f25678_D])).
fof(f25678_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f25892,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl46),
  inference(forward_demodulation,[],[f25752,f962])).
fof(f25752,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) | $spl46),
  inference(resolution,[],[f2022,f938])).
fof(f1124,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | $spl8),
  inference(cnf_transformation,[],[f1124_D])).
fof(f1124_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f29781,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | $spl93),
  inference(forward_demodulation,[],[f26681,f961])).
fof(f26681,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))))))),s('type.integer.int',sK15)))) | $spl93),
  inference(forward_demodulation,[],[f26538,f943])).
fof(f26538,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))))))),s('type.integer.int',sK15)))) | $spl93),
  inference(resolution,[],[f3586,f927])).
fof(f3591,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl96),
  inference(cnf_transformation,[],[f3591_D])).
fof(f3591_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f26907,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | $spl94),
  inference(resolution,[],[f3588,f933])).
fof(f2017,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | $spl43),
  inference(cnf_transformation,[],[f2017_D])).
fof(f2017_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f27216,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | $spl109),
  inference(resolution,[],[f4846,f938])).
fof(f4846,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | $spl109),
  inference(cnf_transformation,[],[f4846_D])).
fof(f4846_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) <=> ~$spl109),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl109])])).
fof(f26531,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))),s('type.integer.int',sK14)))) | $spl93),
  inference(resolution,[],[f3586,f927])).
fof(f28560,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl275),
  inference(cnf_transformation,[],[f28560_D])).
fof(f28560_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) <=> ~$spl275),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl275])])).
fof(f26683,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl93),
  inference(forward_demodulation,[],[f26540,f962])).
fof(f26540,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl93),
  inference(resolution,[],[f3586,f933])).
fof(f29736,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f29573,f991])).
fof(f991,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | $spl2),
  inference(resolution,[],[f986,f893])).
fof(f29573,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | $spl1),
  inference(resolution,[],[f2370,f931])).
fof(f931,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))))))) )),
  inference(cnf_transformation,[],[f627])).
fof(f29511,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl280),
  inference(cnf_transformation,[],[f29511_D])).
fof(f29511_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f29472,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl1),
  inference(resolution,[],[f2369,f976])).
fof(f2369,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl1),
  inference(resolution,[],[f985,f900])).
fof(f29377,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) | s(bool,X21) != s(bool,X23)) ) | $spl1),
  inference(resolution,[],[f2369,f650])).
fof(f29375,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl1),
  inference(resolution,[],[f2369,f649])).
fof(f29373,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl1),
  inference(resolution,[],[f2369,f648])).
fof(f29372,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) | s(bool,X6) != s(bool,X8)) ) | $spl1),
  inference(resolution,[],[f2369,f646])).
fof(f29370,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl1),
  inference(resolution,[],[f2369,f644])).
fof(f28527,plain,(
  ( ! [X187,X188] : (p(s(bool,X187)) | ~p(s(bool,X188)) | s(bool,X188) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl34),
  inference(resolution,[],[f14977,f976])).
fof(f28432,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X22) != s(bool,X24)) ) | $spl34),
  inference(resolution,[],[f14977,f650])).
fof(f28430,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl34),
  inference(resolution,[],[f14977,f649])).
fof(f28564,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl278),
  inference(cnf_transformation,[],[f28564_D])).
fof(f28564_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f1130,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | p(s(bool,X4)) | ~p(s(bool,X5))) ) | $spl18),
  inference(cnf_transformation,[],[f1130_D])).
fof(f1130_D,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | p(s(bool,X4)) | ~p(s(bool,X5))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f28429,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X13) != s(bool,X14)) ) | $spl34),
  inference(resolution,[],[f14977,f648])).
fof(f28428,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl34),
  inference(resolution,[],[f14977,f648])).
fof(f28427,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X7) != s(bool,X9)) ) | $spl34),
  inference(resolution,[],[f14977,f646])).
fof(f28426,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X4) != s(bool,X5)) ) | $spl34),
  inference(resolution,[],[f14977,f644])).
fof(f28425,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl34),
  inference(resolution,[],[f14977,f644])).
fof(f28423,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl34),
  inference(resolution,[],[f14977,f933])).
fof(f28559,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl34),
  inference(forward_demodulation,[],[f28421,f968])).
fof(f28421,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))))),s('type.integer.int',sK14)))) | $spl34),
  inference(resolution,[],[f14977,f927])).
fof(f27760,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl270),
  inference(cnf_transformation,[],[f27760_D])).
fof(f27760_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f27694,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl124),
  inference(resolution,[],[f5364,f976])).
fof(f27603,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,X23) != s(bool,X24)) ) | $spl124),
  inference(resolution,[],[f5364,f650])).
fof(f27761,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl272),
  inference(cnf_transformation,[],[f27761_D])).
fof(f27761_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f27602,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl124),
  inference(resolution,[],[f5364,f650])).
fof(f27600,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl124),
  inference(resolution,[],[f5364,f649])).
fof(f27599,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl124),
  inference(resolution,[],[f5364,f648])).
fof(f27598,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,X8) != s(bool,X9)) ) | $spl124),
  inference(resolution,[],[f5364,f646])).
fof(f27597,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl124),
  inference(resolution,[],[f5364,f646])).
fof(f27596,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl124),
  inference(resolution,[],[f5364,f644])).
fof(f27519,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) ) | $spl266),
  inference(cnf_transformation,[],[f27519_D])).
fof(f27519_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) ) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f27484,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) ) | $spl112),
  inference(resolution,[],[f5126,f976])).
fof(f27393,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))))) | s(bool,X23) != s(bool,X24)) ) | $spl112),
  inference(resolution,[],[f5126,f650])).
fof(f27520,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) ) | $spl268),
  inference(cnf_transformation,[],[f27520_D])).
fof(f27520_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) ) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f27392,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) ) | $spl112),
  inference(resolution,[],[f5126,f650])).
fof(f27390,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) ) | $spl112),
  inference(resolution,[],[f5126,f649])).
fof(f27389,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) ) | $spl112),
  inference(resolution,[],[f5126,f648])).
fof(f27388,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))))) | s(bool,X8) != s(bool,X9)) ) | $spl112),
  inference(resolution,[],[f5126,f646])).
fof(f27387,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) ) | $spl112),
  inference(resolution,[],[f5126,f646])).
fof(f27386,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) ) | $spl112),
  inference(resolution,[],[f5126,f644])).
fof(f27365,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl262),
  inference(cnf_transformation,[],[f27365_D])).
fof(f27365_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f27329,plain,(
  ( ! [X187,X188] : (p(s(bool,X187)) | ~p(s(bool,X188)) | s(bool,X188) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl109),
  inference(resolution,[],[f4846,f976])).
fof(f27234,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X22) != s(bool,X24)) ) | $spl109),
  inference(resolution,[],[f4846,f650])).
fof(f27232,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl109),
  inference(resolution,[],[f4846,f649])).
fof(f27366,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl264),
  inference(cnf_transformation,[],[f27366_D])).
fof(f27366_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f27231,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X13) != s(bool,X14)) ) | $spl109),
  inference(resolution,[],[f4846,f648])).
fof(f27230,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl109),
  inference(resolution,[],[f4846,f648])).
fof(f27229,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X7) != s(bool,X9)) ) | $spl109),
  inference(resolution,[],[f4846,f646])).
fof(f27228,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X4) != s(bool,X5)) ) | $spl109),
  inference(resolution,[],[f4846,f644])).
fof(f27227,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl109),
  inference(resolution,[],[f4846,f644])).
fof(f25681,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl235),
  inference(cnf_transformation,[],[f25681_D])).
fof(f25681_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) <=> ~$spl235),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl235])])).
fof(f27225,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl109),
  inference(resolution,[],[f4846,f933])).
fof(f27363,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl260),
  inference(cnf_transformation,[],[f27363_D])).
fof(f27363_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f27361,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl109),
  inference(forward_demodulation,[],[f27223,f968])).
fof(f27223,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))),s('type.integer.int',sK15)))) | $spl109),
  inference(resolution,[],[f4846,f927])).
fof(f27210,plain,(
  $false | ($spl94 | $spl108)),
  inference(subsumption_resolution,[],[f27072,f3588])).
fof(f27072,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl108),
  inference(resolution,[],[f4847,f975])).
fof(f975,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f458])).
fof(f458,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(ennf_transformation,[],[f413])).
fof(f413,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) => ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(flattening,[],[f412])).
fof(f412,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) => ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f171])).
fof(f171,axiom,(
  ! [X9,X11] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X9),s('type.integer.int',X11)))) => ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X11),s('type.integer.int',X9)))))),
  file('thm.integer.INT_ABS_LE|split|2','thm.integer.INT_LT_GT')).
fof(f27209,plain,(
  $false | ($spl94 | $spl108)),
  inference(subsumption_resolution,[],[f27071,f3588])).
fof(f27071,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl108),
  inference(resolution,[],[f4847,f975])).
fof(f27206,plain,(
  $false | ($spl7 | $spl108)),
  inference(subsumption_resolution,[],[f27065,f1121])).
fof(f27065,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | $spl108),
  inference(resolution,[],[f4847,f907])).
fof(f27205,plain,(
  $false | ($spl7 | $spl108)),
  inference(subsumption_resolution,[],[f27064,f1121])).
fof(f27064,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | $spl108),
  inference(resolution,[],[f4847,f905])).
fof(f27204,plain,(
  $false | ($spl54 | $spl108)),
  inference(subsumption_resolution,[],[f27063,f2535])).
fof(f27063,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl108),
  inference(resolution,[],[f4847,f902])).
fof(f902,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f615])).
fof(f27203,plain,(
  $false | ($spl54 | $spl108)),
  inference(subsumption_resolution,[],[f27062,f2535])).
fof(f27062,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl108),
  inference(resolution,[],[f4847,f895])).
fof(f895,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f610])).
fof(f27202,plain,(
  $false | ($spl54 | $spl108)),
  inference(subsumption_resolution,[],[f27061,f2535])).
fof(f27061,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl108),
  inference(resolution,[],[f4847,f893])).
fof(f27201,plain,(
  $false | ($spl5 | $spl108)),
  inference(subsumption_resolution,[],[f27059,f1145])).
fof(f27059,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | $spl108),
  inference(resolution,[],[f4847,f937])).
fof(f937,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f631])).
fof(f27057,plain,(
  $false | ($spl5 | $spl102)),
  inference(subsumption_resolution,[],[f4069,f1287])).
fof(f26844,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl256),
  inference(cnf_transformation,[],[f26844_D])).
fof(f26844_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f27019,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl94),
  inference(resolution,[],[f3588,f976])).
fof(f26928,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) | s(bool,X23) != s(bool,X24)) ) | $spl94),
  inference(resolution,[],[f3588,f650])).
fof(f26845,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl258),
  inference(cnf_transformation,[],[f26845_D])).
fof(f26845_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f26927,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl94),
  inference(resolution,[],[f3588,f650])).
fof(f26925,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl94),
  inference(resolution,[],[f3588,f649])).
fof(f26924,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl94),
  inference(resolution,[],[f3588,f648])).
fof(f26923,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) | s(bool,X8) != s(bool,X9)) ) | $spl94),
  inference(resolution,[],[f3588,f646])).
fof(f26922,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl94),
  inference(resolution,[],[f3588,f646])).
fof(f26921,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl94),
  inference(resolution,[],[f3588,f644])).
fof(f26843,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl254),
  inference(cnf_transformation,[],[f26843_D])).
fof(f26843_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f26918,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) | $spl94),
  inference(resolution,[],[f3588,f938])).
fof(f26840,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl253),
  inference(cnf_transformation,[],[f26840_D])).
fof(f26840_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) <=> ~$spl253),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl253])])).
fof(f27049,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl94),
  inference(forward_demodulation,[],[f26916,f968])).
fof(f26916,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl94),
  inference(resolution,[],[f3588,f932])).
fof(f26853,plain,(
  $false | ($spl7 | $spl160)),
  inference(subsumption_resolution,[],[f10586,f10777])).
fof(f10777,plain,(
  s('type.integer.int',sK15) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | $spl7),
  inference(resolution,[],[f1121,f906])).
fof(f26808,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl94),
  inference(resolution,[],[f3588,f976])).
fof(f26717,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) | s(bool,X23) != s(bool,X24)) ) | $spl94),
  inference(resolution,[],[f3588,f650])).
fof(f26716,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl94),
  inference(resolution,[],[f3588,f650])).
fof(f26714,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl94),
  inference(resolution,[],[f3588,f649])).
fof(f26713,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl94),
  inference(resolution,[],[f3588,f648])).
fof(f26712,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) | s(bool,X8) != s(bool,X9)) ) | $spl94),
  inference(resolution,[],[f3588,f646])).
fof(f26711,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl94),
  inference(resolution,[],[f3588,f646])).
fof(f26710,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl94),
  inference(resolution,[],[f3588,f644])).
fof(f26707,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) | $spl94),
  inference(resolution,[],[f3588,f938])).
fof(f26838,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl94),
  inference(forward_demodulation,[],[f26705,f968])).
fof(f26705,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl94),
  inference(resolution,[],[f3588,f932])).
fof(f26337,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl248),
  inference(cnf_transformation,[],[f26337_D])).
fof(f26337_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f26644,plain,(
  ( ! [X187,X188] : (p(s(bool,X187)) | ~p(s(bool,X188)) | s(bool,X188) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl93),
  inference(resolution,[],[f3586,f976])).
fof(f26549,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X22) != s(bool,X24)) ) | $spl93),
  inference(resolution,[],[f3586,f650])).
fof(f26547,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl93),
  inference(resolution,[],[f3586,f649])).
fof(f26338,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl250),
  inference(cnf_transformation,[],[f26338_D])).
fof(f26338_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f26546,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X13) != s(bool,X14)) ) | $spl93),
  inference(resolution,[],[f3586,f648])).
fof(f26545,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl93),
  inference(resolution,[],[f3586,f648])).
fof(f26544,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X7) != s(bool,X9)) ) | $spl93),
  inference(resolution,[],[f3586,f646])).
fof(f26543,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X4) != s(bool,X5)) ) | $spl93),
  inference(resolution,[],[f3586,f644])).
fof(f26542,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl93),
  inference(resolution,[],[f3586,f644])).
fof(f26521,plain,(
  $false | ($spl46 | $spl92)),
  inference(subsumption_resolution,[],[f26379,f2022])).
fof(f26379,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl92),
  inference(resolution,[],[f3587,f975])).
fof(f26520,plain,(
  $false | ($spl46 | $spl92)),
  inference(subsumption_resolution,[],[f26378,f2022])).
fof(f26378,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl92),
  inference(resolution,[],[f3587,f975])).
fof(f26513,plain,(
  $false | ($spl61 | $spl92)),
  inference(subsumption_resolution,[],[f26372,f2547])).
fof(f26372,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | $spl92),
  inference(resolution,[],[f3587,f907])).
fof(f26512,plain,(
  $false | ($spl61 | $spl92)),
  inference(subsumption_resolution,[],[f26371,f2547])).
fof(f26371,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | $spl92),
  inference(resolution,[],[f3587,f905])).
fof(f26295,plain,(
  ( ! [X187,X188] : (p(s(bool,X187)) | ~p(s(bool,X188)) | s(bool,X188) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl93),
  inference(resolution,[],[f3586,f976])).
fof(f26200,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X22) != s(bool,X24)) ) | $spl93),
  inference(resolution,[],[f3586,f650])).
fof(f26198,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl93),
  inference(resolution,[],[f3586,f649])).
fof(f26197,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X13) != s(bool,X14)) ) | $spl93),
  inference(resolution,[],[f3586,f648])).
fof(f26196,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl93),
  inference(resolution,[],[f3586,f648])).
fof(f26195,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X7) != s(bool,X9)) ) | $spl93),
  inference(resolution,[],[f3586,f646])).
fof(f26194,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X4) != s(bool,X5)) ) | $spl93),
  inference(resolution,[],[f3586,f644])).
fof(f26193,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl93),
  inference(resolution,[],[f3586,f644])).
fof(f26073,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) | $spl88),
  inference(resolution,[],[f3058,f976])).
fof(f25982,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))))) | s(bool,X23) != s(bool,X24)) ) | $spl88),
  inference(resolution,[],[f3058,f650])).
fof(f26110,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) | $spl246),
  inference(cnf_transformation,[],[f26110_D])).
fof(f26110_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f25981,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) | $spl88),
  inference(resolution,[],[f3058,f650])).
fof(f25979,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) | $spl88),
  inference(resolution,[],[f3058,f649])).
fof(f25978,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) | $spl88),
  inference(resolution,[],[f3058,f648])).
fof(f25977,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))))) | s(bool,X8) != s(bool,X9)) ) | $spl88),
  inference(resolution,[],[f3058,f646])).
fof(f25976,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) | $spl88),
  inference(resolution,[],[f3058,f646])).
fof(f25975,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) | $spl88),
  inference(resolution,[],[f3058,f644])).
fof(f25894,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) ) | $spl240),
  inference(cnf_transformation,[],[f25894_D])).
fof(f25894_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) ) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f25853,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) ) | $spl46),
  inference(resolution,[],[f2022,f976])).
fof(f25762,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))))) | s(bool,X23) != s(bool,X24)) ) | $spl46),
  inference(resolution,[],[f2022,f650])).
fof(f25895,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) ) | $spl242),
  inference(cnf_transformation,[],[f25895_D])).
fof(f25895_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) ) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f25761,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) ) | $spl46),
  inference(resolution,[],[f2022,f650])).
fof(f25759,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) ) | $spl46),
  inference(resolution,[],[f2022,f649])).
fof(f25758,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) ) | $spl46),
  inference(resolution,[],[f2022,f648])).
fof(f25757,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))))) | s(bool,X8) != s(bool,X9)) ) | $spl46),
  inference(resolution,[],[f2022,f646])).
fof(f25756,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) ) | $spl46),
  inference(resolution,[],[f2022,f646])).
fof(f25755,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) ) | $spl46),
  inference(resolution,[],[f2022,f644])).
fof(f25647,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl44),
  inference(resolution,[],[f2019,f976])).
fof(f25556,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))) | s(bool,X23) != s(bool,X24)) ) | $spl44),
  inference(resolution,[],[f2019,f650])).
fof(f25684,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl238),
  inference(cnf_transformation,[],[f25684_D])).
fof(f25684_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f25555,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl44),
  inference(resolution,[],[f2019,f650])).
fof(f25553,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl44),
  inference(resolution,[],[f2019,f649])).
fof(f25552,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl44),
  inference(resolution,[],[f2019,f648])).
fof(f25551,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))) | s(bool,X8) != s(bool,X9)) ) | $spl44),
  inference(resolution,[],[f2019,f646])).
fof(f25550,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl44),
  inference(resolution,[],[f2019,f646])).
fof(f25549,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl44),
  inference(resolution,[],[f2019,f644])).
fof(f25546,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))),s('type.integer.int',sK14)))) | $spl44),
  inference(resolution,[],[f2019,f938])).
fof(f25677,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl44),
  inference(forward_demodulation,[],[f25544,f968])).
fof(f25544,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl44),
  inference(resolution,[],[f2019,f932])).
fof(f25534,plain,(
  $false | ($spl5 | $spl42)),
  inference(subsumption_resolution,[],[f2018,f1145])).
fof(f25436,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | $spl228),
  inference(cnf_transformation,[],[f25436_D])).
fof(f25436_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f25395,plain,(
  ( ! [X187,X188] : (p(s(bool,X187)) | ~p(s(bool,X188)) | s(bool,X188) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | $spl153),
  inference(resolution,[],[f8835,f976])).
fof(f25300,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14))) | s(bool,X22) != s(bool,X24)) ) | $spl153),
  inference(resolution,[],[f8835,f650])).
fof(f25298,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | $spl153),
  inference(resolution,[],[f8835,f649])).
fof(f25437,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | $spl230),
  inference(cnf_transformation,[],[f25437_D])).
fof(f25437_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f25297,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14))) | s(bool,X13) != s(bool,X14)) ) | $spl153),
  inference(resolution,[],[f8835,f648])).
fof(f25296,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | $spl153),
  inference(resolution,[],[f8835,f648])).
fof(f25295,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14))) | s(bool,X7) != s(bool,X9)) ) | $spl153),
  inference(resolution,[],[f8835,f646])).
fof(f25294,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14))) | s(bool,X4) != s(bool,X5)) ) | $spl153),
  inference(resolution,[],[f8835,f644])).
fof(f25293,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | $spl153),
  inference(resolution,[],[f8835,f644])).
fof(f25434,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl227),
  inference(cnf_transformation,[],[f25434_D])).
fof(f25434_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) <=> ~$spl227),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl227])])).
fof(f25263,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) | $spl220),
  inference(cnf_transformation,[],[f25263_D])).
fof(f25263_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f25432,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl153),
  inference(forward_demodulation,[],[f25291,f962])).
fof(f25291,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl153),
  inference(resolution,[],[f8835,f933])).
fof(f25430,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) | $spl153),
  inference(forward_demodulation,[],[f25429,f961])).
fof(f25429,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))))))),s('type.integer.int',sK14)))) | $spl153),
  inference(forward_demodulation,[],[f25289,f943])).
fof(f25289,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))))))),s('type.integer.int',sK14)))) | $spl153),
  inference(resolution,[],[f8835,f927])).
fof(f25282,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK14)))) | $spl153),
  inference(resolution,[],[f8835,f927])).
fof(f25268,plain,(
  $false | ($spl63 | $spl152)),
  inference(subsumption_resolution,[],[f25127,f2549])).
fof(f25127,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | $spl152),
  inference(resolution,[],[f8836,f907])).
fof(f25267,plain,(
  $false | ($spl63 | $spl152)),
  inference(subsumption_resolution,[],[f25126,f2549])).
fof(f25126,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | $spl152),
  inference(resolution,[],[f8836,f905])).
fof(f25266,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | $spl222),
  inference(cnf_transformation,[],[f25266_D])).
fof(f25266_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f25120,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK14)))) | $spl152),
  inference(resolution,[],[f8836,f938])).
fof(f25003,plain,(
  $false | ($spl2 | $spl150)),
  inference(subsumption_resolution,[],[f8157,f991])).
fof(f21081,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | $spl196),
  inference(cnf_transformation,[],[f21081_D])).
fof(f21081_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f24909,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | $spl63),
  inference(resolution,[],[f2549,f976])).
fof(f24814,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14))) | s(bool,X21) != s(bool,X23)) ) | $spl63),
  inference(resolution,[],[f2549,f650])).
fof(f24812,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | $spl63),
  inference(resolution,[],[f2549,f649])).
fof(f8840,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | $spl158),
  inference(cnf_transformation,[],[f8840_D])).
fof(f8840_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f24811,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14))) | s(bool,X12) != s(bool,X13)) ) | $spl63),
  inference(resolution,[],[f2549,f648])).
fof(f24810,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | $spl63),
  inference(resolution,[],[f2549,f648])).
fof(f24809,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14))) | s(bool,X6) != s(bool,X8)) ) | $spl63),
  inference(resolution,[],[f2549,f646])).
fof(f24808,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14))) | s(bool,X3) != s(bool,X4)) ) | $spl63),
  inference(resolution,[],[f2549,f644])).
fof(f24807,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | $spl63),
  inference(resolution,[],[f2549,f644])).
fof(f24762,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl34),
  inference(resolution,[],[f1716,f976])).
fof(f24671,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))) | s(bool,X23) != s(bool,X24)) ) | $spl34),
  inference(resolution,[],[f1716,f650])).
fof(f12339,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl170),
  inference(cnf_transformation,[],[f12339_D])).
fof(f12339_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f24670,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl34),
  inference(resolution,[],[f1716,f650])).
fof(f24668,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl34),
  inference(resolution,[],[f1716,f649])).
fof(f24667,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl34),
  inference(resolution,[],[f1716,f648])).
fof(f24666,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))) | s(bool,X8) != s(bool,X9)) ) | $spl34),
  inference(resolution,[],[f1716,f646])).
fof(f24665,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl34),
  inference(resolution,[],[f1716,f646])).
fof(f24664,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl34),
  inference(resolution,[],[f1716,f644])).
fof(f15116,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))) | $spl174),
  inference(cnf_transformation,[],[f15116_D])).
fof(f15116_D,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f2019,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl44),
  inference(cnf_transformation,[],[f2019_D])).
fof(f2019_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f24663,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))) | $spl34),
  inference(resolution,[],[f1716,f909])).
fof(f24662,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl34),
  inference(resolution,[],[f1716,f904])).
fof(f24660,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))),s('type.integer.int',sK14)))) | $spl34),
  inference(resolution,[],[f1716,f901])).
fof(f24610,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f1287,f976])).
fof(f24515,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X21) != s(bool,X23)) ) | $spl5),
  inference(resolution,[],[f1287,f650])).
fof(f24513,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f1287,f649])).
fof(f24646,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl218),
  inference(cnf_transformation,[],[f24646_D])).
fof(f24646_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f24512,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X12) != s(bool,X13)) ) | $spl5),
  inference(resolution,[],[f1287,f648])).
fof(f24511,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f1287,f648])).
fof(f24510,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X6) != s(bool,X8)) ) | $spl5),
  inference(resolution,[],[f1287,f646])).
fof(f24509,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X3) != s(bool,X4)) ) | $spl5),
  inference(resolution,[],[f1287,f644])).
fof(f24508,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f1287,f644])).
fof(f24502,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl5),
  inference(resolution,[],[f1287,f896])).
fof(f23788,plain,(
  ( ! [X187,X188] : (p(s(bool,X187)) | ~p(s(bool,X188)) | s(bool,X188) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f1145,f976])).
fof(f23693,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X22) != s(bool,X24)) ) | $spl5),
  inference(resolution,[],[f1145,f650])).
fof(f23691,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f1145,f649])).
fof(f18062,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl186),
  inference(cnf_transformation,[],[f18062_D])).
fof(f18062_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f23690,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X13) != s(bool,X14)) ) | $spl5),
  inference(resolution,[],[f1145,f648])).
fof(f23689,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f1145,f648])).
fof(f23688,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X7) != s(bool,X9)) ) | $spl5),
  inference(resolution,[],[f1145,f646])).
fof(f23687,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X4) != s(bool,X5)) ) | $spl5),
  inference(resolution,[],[f1145,f644])).
fof(f23686,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f1145,f644])).
fof(f23585,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) | $spl5),
  inference(resolution,[],[f1140,f976])).
fof(f23494,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))))) | s(bool,X23) != s(bool,X24)) ) | $spl5),
  inference(resolution,[],[f1140,f650])).
fof(f11078,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) | $spl166),
  inference(cnf_transformation,[],[f11078_D])).
fof(f11078_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f23493,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) | $spl5),
  inference(resolution,[],[f1140,f650])).
fof(f23491,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) | $spl5),
  inference(resolution,[],[f1140,f649])).
fof(f23490,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) | $spl5),
  inference(resolution,[],[f1140,f648])).
fof(f23489,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))))) | s(bool,X8) != s(bool,X9)) ) | $spl5),
  inference(resolution,[],[f1140,f646])).
fof(f23488,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) | $spl5),
  inference(resolution,[],[f1140,f646])).
fof(f23487,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) | $spl5),
  inference(resolution,[],[f1140,f644])).
fof(f3055,plain,(
  ( ! [X187] : (~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl84),
  inference(cnf_transformation,[],[f3055_D])).
fof(f3055_D,plain,(
  ( ! [X187] : (~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f23436,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl33),
  inference(resolution,[],[f1298,f976])).
fof(f23341,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X21) != s(bool,X23)) ) | $spl33),
  inference(resolution,[],[f1298,f650])).
fof(f23339,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl33),
  inference(resolution,[],[f1298,f649])).
fof(f3056,plain,(
  ( ! [X14] : (s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl86),
  inference(cnf_transformation,[],[f3056_D])).
fof(f3056_D,plain,(
  ( ! [X14] : (s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f23338,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X12) != s(bool,X13)) ) | $spl33),
  inference(resolution,[],[f1298,f648])).
fof(f23337,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl33),
  inference(resolution,[],[f1298,f648])).
fof(f23336,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X6) != s(bool,X8)) ) | $spl33),
  inference(resolution,[],[f1298,f646])).
fof(f23335,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X3) != s(bool,X4)) ) | $spl33),
  inference(resolution,[],[f1298,f644])).
fof(f23334,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl33),
  inference(resolution,[],[f1298,f644])).
fof(f4848,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl110),
  inference(cnf_transformation,[],[f4848_D])).
fof(f4848_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f23284,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl8),
  inference(resolution,[],[f1124,f976])).
fof(f23193,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X23) != s(bool,X24)) ) | $spl8),
  inference(resolution,[],[f1124,f650])).
fof(f2024,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl50),
  inference(cnf_transformation,[],[f2024_D])).
fof(f2024_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f23192,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl8),
  inference(resolution,[],[f1124,f650])).
fof(f23190,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl8),
  inference(resolution,[],[f1124,f649])).
fof(f23189,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl8),
  inference(resolution,[],[f1124,f648])).
fof(f23188,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X8) != s(bool,X9)) ) | $spl8),
  inference(resolution,[],[f1124,f646])).
fof(f23187,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl8),
  inference(resolution,[],[f1124,f646])).
fof(f23186,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl8),
  inference(resolution,[],[f1124,f644])).
fof(f4847,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | $spl108),
  inference(cnf_transformation,[],[f4847_D])).
fof(f4847_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f23183,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | $spl8),
  inference(resolution,[],[f1124,f938])).
fof(f2021,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl47),
  inference(cnf_transformation,[],[f2021_D])).
fof(f2021_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f23181,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl8),
  inference(resolution,[],[f1124,f932])).
fof(f3592,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl98),
  inference(cnf_transformation,[],[f3592_D])).
fof(f3592_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f23006,plain,(
  ( ! [X187,X188] : (p(s(bool,X187)) | ~p(s(bool,X188)) | s(bool,X188) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl2),
  inference(resolution,[],[f991,f976])).
fof(f22911,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X22) != s(bool,X24)) ) | $spl2),
  inference(resolution,[],[f991,f650])).
fof(f22909,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl2),
  inference(resolution,[],[f991,f649])).
fof(f3593,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl100),
  inference(cnf_transformation,[],[f3593_D])).
fof(f3593_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f22908,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X13) != s(bool,X14)) ) | $spl2),
  inference(resolution,[],[f991,f648])).
fof(f22907,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl2),
  inference(resolution,[],[f991,f648])).
fof(f22906,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X7) != s(bool,X9)) ) | $spl2),
  inference(resolution,[],[f991,f646])).
fof(f22905,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X4) != s(bool,X5)) ) | $spl2),
  inference(resolution,[],[f991,f644])).
fof(f22904,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl2),
  inference(resolution,[],[f991,f644])).
fof(f3589,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl95),
  inference(cnf_transformation,[],[f3589_D])).
fof(f3589_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) <=> ~$spl95),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl95])])).
fof(f22902,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl2),
  inference(resolution,[],[f991,f933])).
fof(f3587,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | $spl92),
  inference(cnf_transformation,[],[f3587_D])).
fof(f3587_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f22900,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))),s('type.integer.int',sK14)))) | $spl2),
  inference(resolution,[],[f991,f927])).
fof(f22846,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f989,f976])).
fof(f22751,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X21) != s(bool,X23)) ) | $spl5),
  inference(resolution,[],[f989,f650])).
fof(f22749,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f989,f649])).
fof(f1289,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl28),
  inference(cnf_transformation,[],[f1289_D])).
fof(f1289_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f22748,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X12) != s(bool,X13)) ) | $spl5),
  inference(resolution,[],[f989,f648])).
fof(f22747,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f989,f648])).
fof(f22746,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X6) != s(bool,X8)) ) | $spl5),
  inference(resolution,[],[f989,f646])).
fof(f22745,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X3) != s(bool,X4)) ) | $spl5),
  inference(resolution,[],[f989,f644])).
fof(f22744,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f989,f644])).
fof(f1127,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl12),
  inference(cnf_transformation,[],[f1127_D])).
fof(f1127_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f22703,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f976])).
fof(f22612,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X23) != s(bool,X24)) ) | $spl2),
  inference(resolution,[],[f986,f650])).
fof(f1129,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl16),
  inference(cnf_transformation,[],[f1129_D])).
fof(f1129_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f22611,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f650])).
fof(f22609,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f649])).
fof(f22608,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f648])).
fof(f22607,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X8) != s(bool,X9)) ) | $spl2),
  inference(resolution,[],[f986,f646])).
fof(f22606,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f646])).
fof(f22605,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f644])).
fof(f16374,plain,(
  s('type.integer.int',sK15) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) | $spl179),
  inference(cnf_transformation,[],[f16374_D])).
fof(f16374_D,plain,(
  s('type.integer.int',sK15) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) <=> ~$spl179),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl179])])).
fof(f22050,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | s('type.integer.int',sK15) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) | $spl1),
  inference(forward_demodulation,[],[f22026,f922])).
fof(f22026,plain,(
  s('type.integer.int',sK15) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | $spl1),
  inference(superposition,[],[f2376,f917])).
fof(f2376,plain,(
  s('type.integer.int',sK15) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) | $spl1),
  inference(resolution,[],[f985,f906])).
fof(f22548,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl213),
  inference(cnf_transformation,[],[f22548_D])).
fof(f22548_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) <=> ~$spl213),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl213])])).
fof(f1297,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | $spl32),
  inference(cnf_transformation,[],[f1297_D])).
fof(f1297_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f20853,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl61),
  inference(forward_demodulation,[],[f20709,f962])).
fof(f20709,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl61),
  inference(resolution,[],[f2547,f896])).
fof(f987,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | $spl3),
  inference(cnf_transformation,[],[f987_D])).
fof(f987_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f989,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | $spl5),
  inference(cnf_transformation,[],[f989_D])).
fof(f989_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f19929,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | $spl7),
  inference(resolution,[],[f1121,f901])).
fof(f22546,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | $spl7),
  inference(forward_demodulation,[],[f20048,f922])).
fof(f20048,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | $spl7),
  inference(superposition,[],[f1121,f917])).
fof(f20825,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | s('type.integer.int',sK14) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl61),
  inference(superposition,[],[f2547,f951])).
fof(f22515,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) | $spl210),
  inference(cnf_transformation,[],[f22515_D])).
fof(f22515_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f22474,plain,(
  ( ! [X187,X188] : (p(s(bool,X187)) | ~p(s(bool,X188)) | s(bool,X188) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) | $spl4),
  inference(resolution,[],[f5754,f976])).
fof(f22379,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))))) | s(bool,X22) != s(bool,X24)) ) | $spl4),
  inference(resolution,[],[f5754,f650])).
fof(f22377,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) | $spl4),
  inference(resolution,[],[f5754,f649])).
fof(f22376,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))))) | s(bool,X13) != s(bool,X14)) ) | $spl4),
  inference(resolution,[],[f5754,f648])).
fof(f22375,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) | $spl4),
  inference(resolution,[],[f5754,f648])).
fof(f22374,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))))) | s(bool,X7) != s(bool,X9)) ) | $spl4),
  inference(resolution,[],[f5754,f646])).
fof(f22373,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))))) | s(bool,X4) != s(bool,X5)) ) | $spl4),
  inference(resolution,[],[f5754,f644])).
fof(f22372,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) | $spl4),
  inference(resolution,[],[f5754,f644])).
fof(f3057,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl89),
  inference(cnf_transformation,[],[f3057_D])).
fof(f3057_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) <=> ~$spl89),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl89])])).
fof(f22367,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl4),
  inference(resolution,[],[f5754,f909])).
fof(f22366,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl4),
  inference(resolution,[],[f5754,f904])).
fof(f21465,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) ) | $spl198),
  inference(cnf_transformation,[],[f21465_D])).
fof(f21465_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f22316,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) ) | $spl67),
  inference(resolution,[],[f2554,f976])).
fof(f22221,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15))) | s(bool,X21) != s(bool,X23)) ) | $spl67),
  inference(resolution,[],[f2554,f650])).
fof(f22219,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) ) | $spl67),
  inference(resolution,[],[f2554,f649])).
fof(f21466,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) ) | $spl200),
  inference(cnf_transformation,[],[f21466_D])).
fof(f21466_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f22218,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15))) | s(bool,X12) != s(bool,X13)) ) | $spl67),
  inference(resolution,[],[f2554,f648])).
fof(f22217,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) ) | $spl67),
  inference(resolution,[],[f2554,f648])).
fof(f22216,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15))) | s(bool,X6) != s(bool,X8)) ) | $spl67),
  inference(resolution,[],[f2554,f646])).
fof(f22215,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15))) | s(bool,X3) != s(bool,X4)) ) | $spl67),
  inference(resolution,[],[f2554,f644])).
fof(f22214,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) ) | $spl67),
  inference(resolution,[],[f2554,f644])).
fof(f22196,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) ) | $spl208),
  inference(cnf_transformation,[],[f22196_D])).
fof(f22196_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f22159,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) ) | $spl66),
  inference(resolution,[],[f2553,f976])).
fof(f2553,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl66),
  inference(cnf_transformation,[],[f2553_D])).
fof(f2553_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f22068,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15))) | s(bool,X23) != s(bool,X24)) ) | $spl66),
  inference(resolution,[],[f2553,f650])).
fof(f22067,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) ) | $spl66),
  inference(resolution,[],[f2553,f650])).
fof(f22065,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) ) | $spl66),
  inference(resolution,[],[f2553,f649])).
fof(f22064,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) ) | $spl66),
  inference(resolution,[],[f2553,f648])).
fof(f22063,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15))) | s(bool,X8) != s(bool,X9)) ) | $spl66),
  inference(resolution,[],[f2553,f646])).
fof(f22062,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) ) | $spl66),
  inference(resolution,[],[f2553,f646])).
fof(f22061,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) ) | $spl66),
  inference(resolution,[],[f2553,f644])).
fof(f22195,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))) | $spl206),
  inference(cnf_transformation,[],[f22195_D])).
fof(f22195_D,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f22193,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl204),
  inference(cnf_transformation,[],[f22193_D])).
fof(f22193_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f22060,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))) | $spl66),
  inference(resolution,[],[f2553,f909])).
fof(f22059,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl66),
  inference(resolution,[],[f2553,f904])).
fof(f2554,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl67),
  inference(cnf_transformation,[],[f2554_D])).
fof(f2554_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) <=> ~$spl67),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl67])])).
fof(f21830,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(subsumption_resolution,[],[f21829,f965])).
fof(f21829,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f21828,f959])).
fof(f21828,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(forward_demodulation,[],[f21827,f958])).
fof(f21827,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(forward_demodulation,[],[f21826,f959])).
fof(f21826,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(forward_demodulation,[],[f21784,f922])).
fof(f21784,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(superposition,[],[f985,f956])).
fof(f21825,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl1),
  inference(forward_demodulation,[],[f21783,f922])).
fof(f21783,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl1),
  inference(superposition,[],[f985,f955])).
fof(f21824,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(forward_demodulation,[],[f21782,f922])).
fof(f21782,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(superposition,[],[f985,f954])).
fof(f21823,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl1),
  inference(forward_demodulation,[],[f21781,f922])).
fof(f21781,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl1),
  inference(superposition,[],[f985,f953])).
fof(f21822,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(forward_demodulation,[],[f21780,f922])).
fof(f21780,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(superposition,[],[f985,f952])).
fof(f21821,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f21820,f962])).
fof(f21820,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) | $spl1),
  inference(forward_demodulation,[],[f21779,f922])).
fof(f21779,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) | $spl1),
  inference(superposition,[],[f985,f950])).
fof(f21815,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f21814,f968])).
fof(f21814,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f21813,f961])).
fof(f21813,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl1),
  inference(forward_demodulation,[],[f21775,f922])).
fof(f21775,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl1),
  inference(superposition,[],[f985,f914])).
fof(f2538,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) ) | $spl56),
  inference(cnf_transformation,[],[f2538_D])).
fof(f2538_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f21754,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) ) | $spl1),
  inference(resolution,[],[f985,f976])).
fof(f21659,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15))) | s(bool,X21) != s(bool,X23)) ) | $spl1),
  inference(resolution,[],[f985,f650])).
fof(f21657,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) ) | $spl1),
  inference(resolution,[],[f985,f649])).
fof(f2539,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) ) | $spl58),
  inference(cnf_transformation,[],[f2539_D])).
fof(f2539_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f21656,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15))) | s(bool,X12) != s(bool,X13)) ) | $spl1),
  inference(resolution,[],[f985,f648])).
fof(f21655,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) ) | $spl1),
  inference(resolution,[],[f985,f648])).
fof(f21654,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15))) | s(bool,X6) != s(bool,X8)) ) | $spl1),
  inference(resolution,[],[f985,f646])).
fof(f21653,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15))) | s(bool,X3) != s(bool,X4)) ) | $spl1),
  inference(resolution,[],[f985,f644])).
fof(f21652,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) ) | $spl1),
  inference(resolution,[],[f985,f644])).
fof(f21599,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl4),
  inference(resolution,[],[f988,f976])).
fof(f21508,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X23) != s(bool,X24)) ) | $spl4),
  inference(resolution,[],[f988,f650])).
fof(f21507,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl4),
  inference(resolution,[],[f988,f650])).
fof(f21505,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl4),
  inference(resolution,[],[f988,f649])).
fof(f21504,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl4),
  inference(resolution,[],[f988,f648])).
fof(f21503,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X8) != s(bool,X9)) ) | $spl4),
  inference(resolution,[],[f988,f646])).
fof(f21502,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl4),
  inference(resolution,[],[f988,f646])).
fof(f21501,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl4),
  inference(resolution,[],[f988,f644])).
fof(f21500,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))) | $spl4),
  inference(resolution,[],[f988,f909])).
fof(f21499,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | $spl4),
  inference(resolution,[],[f988,f904])).
fof(f19011,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) | $spl54),
  inference(resolution,[],[f2535,f901])).
fof(f21431,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) ) | $spl67),
  inference(resolution,[],[f2554,f976])).
fof(f21336,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15))) | s(bool,X21) != s(bool,X23)) ) | $spl67),
  inference(resolution,[],[f2554,f650])).
fof(f21334,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) ) | $spl67),
  inference(resolution,[],[f2554,f649])).
fof(f21333,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15))) | s(bool,X12) != s(bool,X13)) ) | $spl67),
  inference(resolution,[],[f2554,f648])).
fof(f21332,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) ) | $spl67),
  inference(resolution,[],[f2554,f648])).
fof(f21331,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15))) | s(bool,X6) != s(bool,X8)) ) | $spl67),
  inference(resolution,[],[f2554,f646])).
fof(f21330,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15))) | s(bool,X3) != s(bool,X4)) ) | $spl67),
  inference(resolution,[],[f2554,f644])).
fof(f21329,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) ) | $spl67),
  inference(resolution,[],[f2554,f644])).
fof(f21245,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | $spl62),
  inference(resolution,[],[f2548,f976])).
fof(f21154,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14))) | s(bool,X23) != s(bool,X24)) ) | $spl62),
  inference(resolution,[],[f2548,f650])).
fof(f21153,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | $spl62),
  inference(resolution,[],[f2548,f650])).
fof(f21151,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | $spl62),
  inference(resolution,[],[f2548,f649])).
fof(f21150,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | $spl62),
  inference(resolution,[],[f2548,f648])).
fof(f21149,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14))) | s(bool,X8) != s(bool,X9)) ) | $spl62),
  inference(resolution,[],[f2548,f646])).
fof(f21148,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | $spl62),
  inference(resolution,[],[f2548,f646])).
fof(f21147,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | $spl62),
  inference(resolution,[],[f2548,f644])).
fof(f21146,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) | $spl62),
  inference(resolution,[],[f2548,f909])).
fof(f21145,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | $spl62),
  inference(resolution,[],[f2548,f904])).
fof(f21047,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | $spl63),
  inference(resolution,[],[f2549,f976])).
fof(f20952,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14))) | s(bool,X21) != s(bool,X23)) ) | $spl63),
  inference(resolution,[],[f2549,f650])).
fof(f20950,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | $spl63),
  inference(resolution,[],[f2549,f649])).
fof(f20949,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14))) | s(bool,X12) != s(bool,X13)) ) | $spl63),
  inference(resolution,[],[f2549,f648])).
fof(f20948,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | $spl63),
  inference(resolution,[],[f2549,f648])).
fof(f20947,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14))) | s(bool,X6) != s(bool,X8)) ) | $spl63),
  inference(resolution,[],[f2549,f646])).
fof(f20946,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14))) | s(bool,X3) != s(bool,X4)) ) | $spl63),
  inference(resolution,[],[f2549,f644])).
fof(f20945,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | $spl63),
  inference(resolution,[],[f2549,f644])).
fof(f20817,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl61),
  inference(resolution,[],[f2547,f976])).
fof(f20722,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X21) != s(bool,X23)) ) | $spl61),
  inference(resolution,[],[f2547,f650])).
fof(f20720,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl61),
  inference(resolution,[],[f2547,f649])).
fof(f20526,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl192),
  inference(cnf_transformation,[],[f20526_D])).
fof(f20526_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f20719,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X12) != s(bool,X13)) ) | $spl61),
  inference(resolution,[],[f2547,f648])).
fof(f20718,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl61),
  inference(resolution,[],[f2547,f648])).
fof(f20717,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X6) != s(bool,X8)) ) | $spl61),
  inference(resolution,[],[f2547,f646])).
fof(f20716,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X3) != s(bool,X4)) ) | $spl61),
  inference(resolution,[],[f2547,f644])).
fof(f20715,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl61),
  inference(resolution,[],[f2547,f644])).
fof(f20682,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl194),
  inference(cnf_transformation,[],[f20682_D])).
fof(f20682_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f20644,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl60),
  inference(resolution,[],[f2546,f976])).
fof(f2546,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | $spl60),
  inference(cnf_transformation,[],[f2546_D])).
fof(f2546_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f20553,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X23) != s(bool,X24)) ) | $spl60),
  inference(resolution,[],[f2546,f650])).
fof(f20552,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl60),
  inference(resolution,[],[f2546,f650])).
fof(f20550,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl60),
  inference(resolution,[],[f2546,f649])).
fof(f20549,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl60),
  inference(resolution,[],[f2546,f648])).
fof(f20548,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X8) != s(bool,X9)) ) | $spl60),
  inference(resolution,[],[f2546,f646])).
fof(f20547,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl60),
  inference(resolution,[],[f2546,f646])).
fof(f20546,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl60),
  inference(resolution,[],[f2546,f644])).
fof(f20545,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) | $spl60),
  inference(resolution,[],[f2546,f909])).
fof(f20544,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | $spl60),
  inference(resolution,[],[f2546,f904])).
fof(f20487,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl61),
  inference(resolution,[],[f2547,f976])).
fof(f20392,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X21) != s(bool,X23)) ) | $spl61),
  inference(resolution,[],[f2547,f650])).
fof(f20390,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl61),
  inference(resolution,[],[f2547,f649])).
fof(f20389,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X12) != s(bool,X13)) ) | $spl61),
  inference(resolution,[],[f2547,f648])).
fof(f20388,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl61),
  inference(resolution,[],[f2547,f648])).
fof(f20387,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X6) != s(bool,X8)) ) | $spl61),
  inference(resolution,[],[f2547,f646])).
fof(f20386,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X3) != s(bool,X4)) ) | $spl61),
  inference(resolution,[],[f2547,f644])).
fof(f20385,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl61),
  inference(resolution,[],[f2547,f644])).
fof(f20236,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl35),
  inference(resolution,[],[f1717,f976])).
fof(f20141,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))) | s(bool,X21) != s(bool,X23)) ) | $spl35),
  inference(resolution,[],[f1717,f650])).
fof(f20139,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl35),
  inference(resolution,[],[f1717,f649])).
fof(f20138,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))) | s(bool,X12) != s(bool,X13)) ) | $spl35),
  inference(resolution,[],[f1717,f648])).
fof(f20137,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl35),
  inference(resolution,[],[f1717,f648])).
fof(f20136,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))) | s(bool,X6) != s(bool,X8)) ) | $spl35),
  inference(resolution,[],[f1717,f646])).
fof(f20135,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))) | s(bool,X3) != s(bool,X4)) ) | $spl35),
  inference(resolution,[],[f1717,f644])).
fof(f20134,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl35),
  inference(resolution,[],[f1717,f644])).
fof(f20040,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl7),
  inference(resolution,[],[f1121,f976])).
fof(f19945,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X21) != s(bool,X23)) ) | $spl7),
  inference(resolution,[],[f1121,f650])).
fof(f19943,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl7),
  inference(resolution,[],[f1121,f649])).
fof(f8102,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl148),
  inference(cnf_transformation,[],[f8102_D])).
fof(f8102_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f19942,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X12) != s(bool,X13)) ) | $spl7),
  inference(resolution,[],[f1121,f648])).
fof(f19941,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl7),
  inference(resolution,[],[f1121,f648])).
fof(f19940,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X6) != s(bool,X8)) ) | $spl7),
  inference(resolution,[],[f1121,f646])).
fof(f19939,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X3) != s(bool,X4)) ) | $spl7),
  inference(resolution,[],[f1121,f644])).
fof(f19938,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl7),
  inference(resolution,[],[f1121,f644])).
fof(f19932,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl7),
  inference(resolution,[],[f1121,f896])).
fof(f19758,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl3),
  inference(resolution,[],[f1579,f976])).
fof(f19667,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X23) != s(bool,X24)) ) | $spl3),
  inference(resolution,[],[f1579,f650])).
fof(f19666,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl3),
  inference(resolution,[],[f1579,f650])).
fof(f19664,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl3),
  inference(resolution,[],[f1579,f649])).
fof(f19663,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl3),
  inference(resolution,[],[f1579,f648])).
fof(f19662,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X8) != s(bool,X9)) ) | $spl3),
  inference(resolution,[],[f1579,f646])).
fof(f19661,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl3),
  inference(resolution,[],[f1579,f646])).
fof(f19660,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl3),
  inference(resolution,[],[f1579,f644])).
fof(f3586,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | $spl93),
  inference(cnf_transformation,[],[f3586_D])).
fof(f3586_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) <=> ~$spl93),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl93])])).
fof(f6197,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | $spl136),
  inference(cnf_transformation,[],[f6197_D])).
fof(f6197_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f19657,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | $spl3),
  inference(resolution,[],[f1579,f938])).
fof(f19655,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl3),
  inference(resolution,[],[f1579,f932])).
fof(f19616,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl32),
  inference(resolution,[],[f1297,f976])).
fof(f19525,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X23) != s(bool,X24)) ) | $spl32),
  inference(resolution,[],[f1297,f650])).
fof(f19524,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl32),
  inference(resolution,[],[f1297,f650])).
fof(f19522,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl32),
  inference(resolution,[],[f1297,f649])).
fof(f19521,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl32),
  inference(resolution,[],[f1297,f648])).
fof(f19520,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X8) != s(bool,X9)) ) | $spl32),
  inference(resolution,[],[f1297,f646])).
fof(f19519,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl32),
  inference(resolution,[],[f1297,f646])).
fof(f19518,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl32),
  inference(resolution,[],[f1297,f644])).
fof(f2547,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | $spl61),
  inference(cnf_transformation,[],[f2547_D])).
fof(f2547_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f19514,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | $spl32),
  inference(resolution,[],[f1297,f901])).
fof(f19475,plain,(
  ( ! [X187,X188] : (p(s(bool,X187)) | ~p(s(bool,X188)) | s(bool,X188) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl9),
  inference(resolution,[],[f1123,f976])).
fof(f19380,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X22) != s(bool,X24)) ) | $spl9),
  inference(resolution,[],[f1123,f650])).
fof(f19378,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl9),
  inference(resolution,[],[f1123,f649])).
fof(f19377,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X13) != s(bool,X14)) ) | $spl9),
  inference(resolution,[],[f1123,f648])).
fof(f19376,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl9),
  inference(resolution,[],[f1123,f648])).
fof(f19375,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X7) != s(bool,X9)) ) | $spl9),
  inference(resolution,[],[f1123,f646])).
fof(f19374,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X4) != s(bool,X5)) ) | $spl9),
  inference(resolution,[],[f1123,f644])).
fof(f19373,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl9),
  inference(resolution,[],[f1123,f644])).
fof(f2020,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl45),
  inference(cnf_transformation,[],[f2020_D])).
fof(f2020_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f19371,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl9),
  inference(resolution,[],[f1123,f933])).
fof(f19369,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) | $spl9),
  inference(resolution,[],[f1123,f927])).
fof(f19328,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl3),
  inference(resolution,[],[f987,f976])).
fof(f19233,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X21) != s(bool,X23)) ) | $spl3),
  inference(resolution,[],[f987,f650])).
fof(f19231,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl3),
  inference(resolution,[],[f987,f649])).
fof(f19230,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X12) != s(bool,X13)) ) | $spl3),
  inference(resolution,[],[f987,f648])).
fof(f19229,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl3),
  inference(resolution,[],[f987,f648])).
fof(f19228,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X6) != s(bool,X8)) ) | $spl3),
  inference(resolution,[],[f987,f646])).
fof(f19227,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X3) != s(bool,X4)) ) | $spl3),
  inference(resolution,[],[f987,f644])).
fof(f19226,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl3),
  inference(resolution,[],[f987,f644])).
fof(f19220,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl3),
  inference(resolution,[],[f987,f896])).
fof(f19217,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | p(s(bool,X1))) ) | $spl3),
  inference(subsumption_resolution,[],[f19207,f851])).
fof(f19207,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | p(s(bool,X1))) ) | $spl3),
  inference(resolution,[],[f6155,f779])).
fof(f19216,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | p(s(bool,X0))) ) | $spl3),
  inference(subsumption_resolution,[],[f19206,f851])).
fof(f19206,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | p(s(bool,X0))) ) | $spl3),
  inference(resolution,[],[f6155,f783])).
fof(f19205,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))))) ) | $spl3),
  inference(subsumption_resolution,[],[f19195,f851])).
fof(f19195,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))))) ) | $spl3),
  inference(resolution,[],[f6154,f779])).
fof(f19204,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))))) ) | $spl3),
  inference(subsumption_resolution,[],[f19194,f851])).
fof(f19194,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))))) ) | $spl3),
  inference(resolution,[],[f6154,f783])).
fof(f18216,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | $spl4),
  inference(forward_demodulation,[],[f18215,f913])).
fof(f18215,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl4),
  inference(forward_demodulation,[],[f18078,f961])).
fof(f18078,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl4),
  inference(resolution,[],[f988,f901])).
fof(f3058,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl88),
  inference(cnf_transformation,[],[f3058_D])).
fof(f3058_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f19006,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | $spl54),
  inference(resolution,[],[f2535,f896])).
fof(f19169,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | $spl54),
  inference(forward_demodulation,[],[f19124,f922])).
fof(f19124,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | $spl54),
  inference(superposition,[],[f2535,f917])).
fof(f1122,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | $spl6),
  inference(cnf_transformation,[],[f1122_D])).
fof(f1122_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f18073,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | $spl4),
  inference(resolution,[],[f988,f901])).
fof(f19152,plain,(
  $false | ($spl1 | $spl6)),
  inference(subsumption_resolution,[],[f2537,f1122])).
fof(f2537,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f2536,f968])).
fof(f2536,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f2373,f961])).
fof(f2373,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))),s('type.integer.int',sK15)))) | $spl1),
  inference(resolution,[],[f985,f899])).
fof(f19113,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl54),
  inference(resolution,[],[f2535,f976])).
fof(f19022,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) | s(bool,X23) != s(bool,X24)) ) | $spl54),
  inference(resolution,[],[f2535,f650])).
fof(f16670,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl182),
  inference(cnf_transformation,[],[f16670_D])).
fof(f16670_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f19021,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl54),
  inference(resolution,[],[f2535,f650])).
fof(f19019,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl54),
  inference(resolution,[],[f2535,f649])).
fof(f19018,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl54),
  inference(resolution,[],[f2535,f648])).
fof(f19017,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) | s(bool,X8) != s(bool,X9)) ) | $spl54),
  inference(resolution,[],[f2535,f646])).
fof(f19016,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl54),
  inference(resolution,[],[f2535,f646])).
fof(f19015,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl54),
  inference(resolution,[],[f2535,f644])).
fof(f10586,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | $spl160),
  inference(cnf_transformation,[],[f10586_D])).
fof(f10586_D,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f3588,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl94),
  inference(cnf_transformation,[],[f3588_D])).
fof(f3588_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f19014,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | $spl54),
  inference(resolution,[],[f2535,f909])).
fof(f19013,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl54),
  inference(resolution,[],[f2535,f904])).
fof(f18997,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl188),
  inference(cnf_transformation,[],[f18997_D])).
fof(f18997_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f18964,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl55),
  inference(resolution,[],[f2534,f976])).
fof(f18869,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) | s(bool,X21) != s(bool,X23)) ) | $spl55),
  inference(resolution,[],[f2534,f650])).
fof(f18867,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl55),
  inference(resolution,[],[f2534,f649])).
fof(f18866,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) | s(bool,X12) != s(bool,X13)) ) | $spl55),
  inference(resolution,[],[f2534,f648])).
fof(f18865,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl55),
  inference(resolution,[],[f2534,f648])).
fof(f18864,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) | s(bool,X6) != s(bool,X8)) ) | $spl55),
  inference(resolution,[],[f2534,f646])).
fof(f18863,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) | s(bool,X3) != s(bool,X4)) ) | $spl55),
  inference(resolution,[],[f2534,f644])).
fof(f18862,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl55),
  inference(resolution,[],[f2534,f644])).
fof(f18805,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl34),
  inference(resolution,[],[f1716,f976])).
fof(f18714,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))) | s(bool,X23) != s(bool,X24)) ) | $spl34),
  inference(resolution,[],[f1716,f650])).
fof(f18713,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl34),
  inference(resolution,[],[f1716,f650])).
fof(f18711,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl34),
  inference(resolution,[],[f1716,f649])).
fof(f18710,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl34),
  inference(resolution,[],[f1716,f648])).
fof(f18709,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))) | s(bool,X8) != s(bool,X9)) ) | $spl34),
  inference(resolution,[],[f1716,f646])).
fof(f18708,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl34),
  inference(resolution,[],[f1716,f646])).
fof(f18707,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl34),
  inference(resolution,[],[f1716,f644])).
fof(f18706,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))) | $spl34),
  inference(resolution,[],[f1716,f909])).
fof(f18705,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl34),
  inference(resolution,[],[f1716,f904])).
fof(f18703,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))),s('type.integer.int',sK14)))) | $spl34),
  inference(resolution,[],[f1716,f901])).
fof(f10587,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl162),
  inference(cnf_transformation,[],[f10587_D])).
fof(f10587_D,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f18662,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl6),
  inference(resolution,[],[f1122,f976])).
fof(f18571,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X23) != s(bool,X24)) ) | $spl6),
  inference(resolution,[],[f1122,f650])).
fof(f18570,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl6),
  inference(resolution,[],[f1122,f650])).
fof(f18568,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl6),
  inference(resolution,[],[f1122,f649])).
fof(f18567,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl6),
  inference(resolution,[],[f1122,f648])).
fof(f18566,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X8) != s(bool,X9)) ) | $spl6),
  inference(resolution,[],[f1122,f646])).
fof(f18565,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl6),
  inference(resolution,[],[f1122,f646])).
fof(f18564,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl6),
  inference(resolution,[],[f1122,f644])).
fof(f18563,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | $spl6),
  inference(resolution,[],[f1122,f909])).
fof(f18562,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | $spl6),
  inference(resolution,[],[f1122,f904])).
fof(f18180,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl4),
  inference(resolution,[],[f988,f976])).
fof(f18089,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X23) != s(bool,X24)) ) | $spl4),
  inference(resolution,[],[f988,f650])).
fof(f18088,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl4),
  inference(resolution,[],[f988,f650])).
fof(f18086,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl4),
  inference(resolution,[],[f988,f649])).
fof(f18085,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl4),
  inference(resolution,[],[f988,f648])).
fof(f18084,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X8) != s(bool,X9)) ) | $spl4),
  inference(resolution,[],[f988,f646])).
fof(f18083,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl4),
  inference(resolution,[],[f988,f646])).
fof(f18082,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl4),
  inference(resolution,[],[f988,f644])).
fof(f18081,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))) | $spl4),
  inference(resolution,[],[f988,f909])).
fof(f18080,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | $spl4),
  inference(resolution,[],[f988,f904])).
fof(f16745,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | $spl2),
  inference(resolution,[],[f986,f901])).
fof(f18019,plain,(
  ( ! [X187,X188] : (p(s(bool,X187)) | ~p(s(bool,X188)) | s(bool,X188) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f1145,f976])).
fof(f17924,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X22) != s(bool,X24)) ) | $spl5),
  inference(resolution,[],[f1145,f650])).
fof(f17922,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f1145,f649])).
fof(f17921,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X13) != s(bool,X14)) ) | $spl5),
  inference(resolution,[],[f1145,f648])).
fof(f17920,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f1145,f648])).
fof(f17919,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X7) != s(bool,X9)) ) | $spl5),
  inference(resolution,[],[f1145,f646])).
fof(f17918,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X4) != s(bool,X5)) ) | $spl5),
  inference(resolution,[],[f1145,f644])).
fof(f17917,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f1145,f644])).
fof(f17728,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) | $spl5),
  inference(resolution,[],[f1140,f976])).
fof(f17637,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))))) | s(bool,X23) != s(bool,X24)) ) | $spl5),
  inference(resolution,[],[f1140,f650])).
fof(f17636,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) | $spl5),
  inference(resolution,[],[f1140,f650])).
fof(f17634,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) | $spl5),
  inference(resolution,[],[f1140,f649])).
fof(f17633,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) | $spl5),
  inference(resolution,[],[f1140,f648])).
fof(f17632,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))))) | s(bool,X8) != s(bool,X9)) ) | $spl5),
  inference(resolution,[],[f1140,f646])).
fof(f17631,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) | $spl5),
  inference(resolution,[],[f1140,f646])).
fof(f17630,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) | $spl5),
  inference(resolution,[],[f1140,f644])).
fof(f17580,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl33),
  inference(resolution,[],[f1298,f976])).
fof(f17485,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X21) != s(bool,X23)) ) | $spl33),
  inference(resolution,[],[f1298,f650])).
fof(f17483,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl33),
  inference(resolution,[],[f1298,f649])).
fof(f17482,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X12) != s(bool,X13)) ) | $spl33),
  inference(resolution,[],[f1298,f648])).
fof(f17481,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl33),
  inference(resolution,[],[f1298,f648])).
fof(f17480,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X6) != s(bool,X8)) ) | $spl33),
  inference(resolution,[],[f1298,f646])).
fof(f17479,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X3) != s(bool,X4)) ) | $spl33),
  inference(resolution,[],[f1298,f644])).
fof(f17478,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl33),
  inference(resolution,[],[f1298,f644])).
fof(f2534,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl55),
  inference(cnf_transformation,[],[f2534_D])).
fof(f2534_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) <=> ~$spl55),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl55])])).
fof(f17472,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl33),
  inference(resolution,[],[f1298,f896])).
fof(f17428,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl8),
  inference(resolution,[],[f1124,f976])).
fof(f17337,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X23) != s(bool,X24)) ) | $spl8),
  inference(resolution,[],[f1124,f650])).
fof(f17336,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl8),
  inference(resolution,[],[f1124,f650])).
fof(f17334,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl8),
  inference(resolution,[],[f1124,f649])).
fof(f17333,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl8),
  inference(resolution,[],[f1124,f648])).
fof(f17332,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X8) != s(bool,X9)) ) | $spl8),
  inference(resolution,[],[f1124,f646])).
fof(f17331,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl8),
  inference(resolution,[],[f1124,f646])).
fof(f17330,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl8),
  inference(resolution,[],[f1124,f644])).
fof(f17327,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | $spl8),
  inference(resolution,[],[f1124,f938])).
fof(f17325,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl8),
  inference(resolution,[],[f1124,f932])).
fof(f17150,plain,(
  ( ! [X187,X188] : (p(s(bool,X187)) | ~p(s(bool,X188)) | s(bool,X188) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl2),
  inference(resolution,[],[f991,f976])).
fof(f17055,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X22) != s(bool,X24)) ) | $spl2),
  inference(resolution,[],[f991,f650])).
fof(f17053,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl2),
  inference(resolution,[],[f991,f649])).
fof(f17052,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X13) != s(bool,X14)) ) | $spl2),
  inference(resolution,[],[f991,f648])).
fof(f17051,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl2),
  inference(resolution,[],[f991,f648])).
fof(f17050,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X7) != s(bool,X9)) ) | $spl2),
  inference(resolution,[],[f991,f646])).
fof(f17049,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X4) != s(bool,X5)) ) | $spl2),
  inference(resolution,[],[f991,f644])).
fof(f17048,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl2),
  inference(resolution,[],[f991,f644])).
fof(f17046,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl2),
  inference(resolution,[],[f991,f933])).
fof(f17044,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))),s('type.integer.int',sK14)))) | $spl2),
  inference(resolution,[],[f991,f927])).
fof(f16990,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f989,f976])).
fof(f16895,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X21) != s(bool,X23)) ) | $spl5),
  inference(resolution,[],[f989,f650])).
fof(f16893,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f989,f649])).
fof(f16892,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X12) != s(bool,X13)) ) | $spl5),
  inference(resolution,[],[f989,f648])).
fof(f16891,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f989,f648])).
fof(f16890,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X6) != s(bool,X8)) ) | $spl5),
  inference(resolution,[],[f989,f646])).
fof(f16889,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X3) != s(bool,X4)) ) | $spl5),
  inference(resolution,[],[f989,f644])).
fof(f16888,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f989,f644])).
fof(f16847,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f976])).
fof(f16756,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X23) != s(bool,X24)) ) | $spl2),
  inference(resolution,[],[f986,f650])).
fof(f16755,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f650])).
fof(f16753,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f649])).
fof(f16752,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f648])).
fof(f16751,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X8) != s(bool,X9)) ) | $spl2),
  inference(resolution,[],[f986,f646])).
fof(f16750,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f646])).
fof(f16749,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f644])).
fof(f15897,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | $spl7),
  inference(resolution,[],[f1121,f901])).
fof(f16692,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | $spl7),
  inference(forward_demodulation,[],[f16016,f922])).
fof(f16016,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | $spl7),
  inference(superposition,[],[f1121,f917])).
fof(f16638,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl54),
  inference(resolution,[],[f2535,f976])).
fof(f16547,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) | s(bool,X23) != s(bool,X24)) ) | $spl54),
  inference(resolution,[],[f2535,f650])).
fof(f16546,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl54),
  inference(resolution,[],[f2535,f650])).
fof(f16544,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl54),
  inference(resolution,[],[f2535,f649])).
fof(f16543,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl54),
  inference(resolution,[],[f2535,f648])).
fof(f16542,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) | s(bool,X8) != s(bool,X9)) ) | $spl54),
  inference(resolution,[],[f2535,f646])).
fof(f16541,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl54),
  inference(resolution,[],[f2535,f646])).
fof(f16540,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) ) | $spl54),
  inference(resolution,[],[f2535,f644])).
fof(f16539,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | $spl54),
  inference(resolution,[],[f2535,f909])).
fof(f16538,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl54),
  inference(resolution,[],[f2535,f904])).
fof(f16519,plain,(
  $false | ($spl32 | $spl55)),
  inference(subsumption_resolution,[],[f16375,f1297])).
fof(f16375,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | $spl55),
  inference(resolution,[],[f2534,f897])).
fof(f897,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f612])).
fof(f2551,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl65),
  inference(cnf_transformation,[],[f2551_D])).
fof(f2551_D,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl65),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl65])])).
fof(f10585,plain,(
  s('type.integer.int',sK15) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | $spl161),
  inference(cnf_transformation,[],[f10585_D])).
fof(f10585_D,plain,(
  s('type.integer.int',sK15) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) <=> ~$spl161),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl161])])).
fof(f16353,plain,(
  s('type.integer.int',sK15) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl1),
  inference(superposition,[],[f2376,f951])).
fof(f16372,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | s('type.integer.int',sK15) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) | $spl1),
  inference(forward_demodulation,[],[f16348,f922])).
fof(f16348,plain,(
  s('type.integer.int',sK15) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | $spl1),
  inference(superposition,[],[f2376,f917])).
fof(f16232,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl35),
  inference(resolution,[],[f1717,f976])).
fof(f16137,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))) | s(bool,X21) != s(bool,X23)) ) | $spl35),
  inference(resolution,[],[f1717,f650])).
fof(f16135,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl35),
  inference(resolution,[],[f1717,f649])).
fof(f16134,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))) | s(bool,X12) != s(bool,X13)) ) | $spl35),
  inference(resolution,[],[f1717,f648])).
fof(f16133,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl35),
  inference(resolution,[],[f1717,f648])).
fof(f16132,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))) | s(bool,X6) != s(bool,X8)) ) | $spl35),
  inference(resolution,[],[f1717,f646])).
fof(f16131,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))) | s(bool,X3) != s(bool,X4)) ) | $spl35),
  inference(resolution,[],[f1717,f644])).
fof(f16130,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl35),
  inference(resolution,[],[f1717,f644])).
fof(f16008,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl7),
  inference(resolution,[],[f1121,f976])).
fof(f15913,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X21) != s(bool,X23)) ) | $spl7),
  inference(resolution,[],[f1121,f650])).
fof(f15911,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl7),
  inference(resolution,[],[f1121,f649])).
fof(f15910,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X12) != s(bool,X13)) ) | $spl7),
  inference(resolution,[],[f1121,f648])).
fof(f15909,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl7),
  inference(resolution,[],[f1121,f648])).
fof(f15908,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X6) != s(bool,X8)) ) | $spl7),
  inference(resolution,[],[f1121,f646])).
fof(f15907,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X3) != s(bool,X4)) ) | $spl7),
  inference(resolution,[],[f1121,f644])).
fof(f15906,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl7),
  inference(resolution,[],[f1121,f644])).
fof(f15900,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl7),
  inference(resolution,[],[f1121,f896])).
fof(f15726,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl3),
  inference(resolution,[],[f1579,f976])).
fof(f15635,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X23) != s(bool,X24)) ) | $spl3),
  inference(resolution,[],[f1579,f650])).
fof(f15634,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl3),
  inference(resolution,[],[f1579,f650])).
fof(f15632,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl3),
  inference(resolution,[],[f1579,f649])).
fof(f15631,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl3),
  inference(resolution,[],[f1579,f648])).
fof(f15630,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X8) != s(bool,X9)) ) | $spl3),
  inference(resolution,[],[f1579,f646])).
fof(f15629,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl3),
  inference(resolution,[],[f1579,f646])).
fof(f15628,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl3),
  inference(resolution,[],[f1579,f644])).
fof(f15625,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | $spl3),
  inference(resolution,[],[f1579,f938])).
fof(f15623,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl3),
  inference(resolution,[],[f1579,f932])).
fof(f15584,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl32),
  inference(resolution,[],[f1297,f976])).
fof(f15493,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X23) != s(bool,X24)) ) | $spl32),
  inference(resolution,[],[f1297,f650])).
fof(f15492,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl32),
  inference(resolution,[],[f1297,f650])).
fof(f15490,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl32),
  inference(resolution,[],[f1297,f649])).
fof(f15489,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl32),
  inference(resolution,[],[f1297,f648])).
fof(f15488,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X8) != s(bool,X9)) ) | $spl32),
  inference(resolution,[],[f1297,f646])).
fof(f15487,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl32),
  inference(resolution,[],[f1297,f646])).
fof(f15486,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl32),
  inference(resolution,[],[f1297,f644])).
fof(f15482,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | $spl32),
  inference(resolution,[],[f1297,f901])).
fof(f15443,plain,(
  ( ! [X187,X188] : (p(s(bool,X187)) | ~p(s(bool,X188)) | s(bool,X188) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl9),
  inference(resolution,[],[f1123,f976])).
fof(f15348,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X22) != s(bool,X24)) ) | $spl9),
  inference(resolution,[],[f1123,f650])).
fof(f15346,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl9),
  inference(resolution,[],[f1123,f649])).
fof(f15345,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X13) != s(bool,X14)) ) | $spl9),
  inference(resolution,[],[f1123,f648])).
fof(f15344,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl9),
  inference(resolution,[],[f1123,f648])).
fof(f15343,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X7) != s(bool,X9)) ) | $spl9),
  inference(resolution,[],[f1123,f646])).
fof(f15342,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X4) != s(bool,X5)) ) | $spl9),
  inference(resolution,[],[f1123,f644])).
fof(f15341,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl9),
  inference(resolution,[],[f1123,f644])).
fof(f15339,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl9),
  inference(resolution,[],[f1123,f933])).
fof(f15337,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) | $spl9),
  inference(resolution,[],[f1123,f927])).
fof(f15192,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl3),
  inference(resolution,[],[f987,f644])).
fof(f15193,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X3) != s(bool,X4)) ) | $spl3),
  inference(resolution,[],[f987,f644])).
fof(f15194,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X6) != s(bool,X8)) ) | $spl3),
  inference(resolution,[],[f987,f646])).
fof(f15195,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl3),
  inference(resolution,[],[f987,f648])).
fof(f15196,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X12) != s(bool,X13)) ) | $spl3),
  inference(resolution,[],[f987,f648])).
fof(f15197,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl3),
  inference(resolution,[],[f987,f649])).
fof(f15199,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X21) != s(bool,X23)) ) | $spl3),
  inference(resolution,[],[f987,f650])).
fof(f15294,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl3),
  inference(resolution,[],[f987,f976])).
fof(f15321,plain,(
  $false | ($spl3 | $spl6)),
  inference(subsumption_resolution,[],[f15187,f1122])).
fof(f15187,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | $spl3),
  inference(resolution,[],[f987,f899])).
fof(f15186,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl3),
  inference(resolution,[],[f987,f896])).
fof(f15183,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | p(s(bool,X1))) ) | $spl3),
  inference(subsumption_resolution,[],[f15173,f851])).
fof(f15173,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | p(s(bool,X1))) ) | $spl3),
  inference(resolution,[],[f6155,f779])).
fof(f15182,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | p(s(bool,X0))) ) | $spl3),
  inference(subsumption_resolution,[],[f15172,f851])).
fof(f15172,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | p(s(bool,X0))) ) | $spl3),
  inference(resolution,[],[f6155,f783])).
fof(f15171,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))))) ) | $spl3),
  inference(subsumption_resolution,[],[f15161,f851])).
fof(f15161,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))))) ) | $spl3),
  inference(resolution,[],[f6154,f779])).
fof(f15170,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))))) ) | $spl3),
  inference(subsumption_resolution,[],[f15160,f851])).
fof(f15160,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))))) ) | $spl3),
  inference(resolution,[],[f6154,f783])).
fof(f986,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | $spl2),
  inference(cnf_transformation,[],[f986_D])).
fof(f986_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f15135,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | $spl6),
  inference(forward_demodulation,[],[f14874,f922])).
fof(f14874,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | $spl6),
  inference(superposition,[],[f1122,f917])).
fof(f15124,plain,(
  $false | ($spl2 | $spl35)),
  inference(subsumption_resolution,[],[f12193,f986])).
fof(f12193,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | $spl35),
  inference(resolution,[],[f1717,f897])).
fof(f15082,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl34),
  inference(resolution,[],[f1716,f976])).
fof(f14991,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))) | s(bool,X23) != s(bool,X24)) ) | $spl34),
  inference(resolution,[],[f1716,f650])).
fof(f14990,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl34),
  inference(resolution,[],[f1716,f650])).
fof(f14988,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl34),
  inference(resolution,[],[f1716,f649])).
fof(f14987,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl34),
  inference(resolution,[],[f1716,f648])).
fof(f14986,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))) | s(bool,X8) != s(bool,X9)) ) | $spl34),
  inference(resolution,[],[f1716,f646])).
fof(f14985,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl34),
  inference(resolution,[],[f1716,f646])).
fof(f14984,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl34),
  inference(resolution,[],[f1716,f644])).
fof(f14983,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))) | $spl34),
  inference(resolution,[],[f1716,f909])).
fof(f14982,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl34),
  inference(resolution,[],[f1716,f904])).
fof(f14980,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))),s('type.integer.int',sK14)))) | $spl34),
  inference(resolution,[],[f1716,f901])).
fof(f14867,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl6),
  inference(resolution,[],[f1122,f976])).
fof(f14776,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X23) != s(bool,X24)) ) | $spl6),
  inference(resolution,[],[f1122,f650])).
fof(f14775,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl6),
  inference(resolution,[],[f1122,f650])).
fof(f14773,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl6),
  inference(resolution,[],[f1122,f649])).
fof(f14772,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl6),
  inference(resolution,[],[f1122,f648])).
fof(f14771,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X8) != s(bool,X9)) ) | $spl6),
  inference(resolution,[],[f1122,f646])).
fof(f14770,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl6),
  inference(resolution,[],[f1122,f646])).
fof(f14769,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl6),
  inference(resolution,[],[f1122,f644])).
fof(f14768,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | $spl6),
  inference(resolution,[],[f1122,f909])).
fof(f14767,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | $spl6),
  inference(resolution,[],[f1122,f904])).
fof(f14363,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl33),
  inference(resolution,[],[f1298,f976])).
fof(f14268,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X21) != s(bool,X23)) ) | $spl33),
  inference(resolution,[],[f1298,f650])).
fof(f14266,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl33),
  inference(resolution,[],[f1298,f649])).
fof(f14265,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X12) != s(bool,X13)) ) | $spl33),
  inference(resolution,[],[f1298,f648])).
fof(f14264,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl33),
  inference(resolution,[],[f1298,f648])).
fof(f14263,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X6) != s(bool,X8)) ) | $spl33),
  inference(resolution,[],[f1298,f646])).
fof(f14262,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X3) != s(bool,X4)) ) | $spl33),
  inference(resolution,[],[f1298,f644])).
fof(f14261,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl33),
  inference(resolution,[],[f1298,f644])).
fof(f14255,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl33),
  inference(resolution,[],[f1298,f896])).
fof(f14189,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl8),
  inference(resolution,[],[f1124,f976])).
fof(f14098,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X23) != s(bool,X24)) ) | $spl8),
  inference(resolution,[],[f1124,f650])).
fof(f14097,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl8),
  inference(resolution,[],[f1124,f650])).
fof(f14095,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl8),
  inference(resolution,[],[f1124,f649])).
fof(f14094,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl8),
  inference(resolution,[],[f1124,f648])).
fof(f14093,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X8) != s(bool,X9)) ) | $spl8),
  inference(resolution,[],[f1124,f646])).
fof(f14092,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl8),
  inference(resolution,[],[f1124,f646])).
fof(f14091,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl8),
  inference(resolution,[],[f1124,f644])).
fof(f14088,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | $spl8),
  inference(resolution,[],[f1124,f938])).
fof(f14086,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl8),
  inference(resolution,[],[f1124,f932])).
fof(f13911,plain,(
  ( ! [X187,X188] : (p(s(bool,X187)) | ~p(s(bool,X188)) | s(bool,X188) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl2),
  inference(resolution,[],[f991,f976])).
fof(f13816,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X22) != s(bool,X24)) ) | $spl2),
  inference(resolution,[],[f991,f650])).
fof(f13814,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl2),
  inference(resolution,[],[f991,f649])).
fof(f13813,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X13) != s(bool,X14)) ) | $spl2),
  inference(resolution,[],[f991,f648])).
fof(f13812,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl2),
  inference(resolution,[],[f991,f648])).
fof(f13811,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X7) != s(bool,X9)) ) | $spl2),
  inference(resolution,[],[f991,f646])).
fof(f13810,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X4) != s(bool,X5)) ) | $spl2),
  inference(resolution,[],[f991,f644])).
fof(f13809,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl2),
  inference(resolution,[],[f991,f644])).
fof(f13807,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl2),
  inference(resolution,[],[f991,f933])).
fof(f13805,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))),s('type.integer.int',sK14)))) | $spl2),
  inference(resolution,[],[f991,f927])).
fof(f13795,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f13794,f968])).
fof(f13794,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f13793,f961])).
fof(f13793,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) | $spl1),
  inference(forward_demodulation,[],[f13745,f922])).
fof(f13745,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) | $spl1),
  inference(superposition,[],[f985,f966])).
fof(f966,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f636])).
fof(f636,plain,(
  ! [X0] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0)))))),
  inference(nnf_transformation,[],[f404])).
fof(f404,plain,(
  ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))) <=> s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X6] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))) <=> s('type.integer.int',X6) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.integer.INT_ABS_LE|split|2','thm.integer.INT_ABS_EQ|48|')).
fof(f13792,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(subsumption_resolution,[],[f13791,f965])).
fof(f13791,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f13790,f959])).
fof(f13790,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(forward_demodulation,[],[f13789,f958])).
fof(f13789,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(forward_demodulation,[],[f13788,f959])).
fof(f13788,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(forward_demodulation,[],[f13744,f922])).
fof(f13744,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(superposition,[],[f985,f956])).
fof(f13787,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl1),
  inference(forward_demodulation,[],[f13743,f922])).
fof(f13743,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl1),
  inference(superposition,[],[f985,f955])).
fof(f13786,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(forward_demodulation,[],[f13742,f922])).
fof(f13742,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(superposition,[],[f985,f954])).
fof(f13785,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl1),
  inference(forward_demodulation,[],[f13741,f922])).
fof(f13741,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl1),
  inference(superposition,[],[f985,f953])).
fof(f13784,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(forward_demodulation,[],[f13740,f922])).
fof(f13740,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(superposition,[],[f985,f952])).
fof(f13783,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f13782,f962])).
fof(f13782,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) | $spl1),
  inference(forward_demodulation,[],[f13739,f922])).
fof(f13739,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) | $spl1),
  inference(superposition,[],[f985,f950])).
fof(f13781,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f13780,f962])).
fof(f13780,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) | $spl1),
  inference(forward_demodulation,[],[f13738,f922])).
fof(f13738,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) | $spl1),
  inference(superposition,[],[f985,f948])).
fof(f948,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f633])).
fof(f633,plain,(
  ! [X0] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))),
  inference(nnf_transformation,[],[f390])).
fof(f390,plain,(
  ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) <=> s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f145])).
fof(f145,axiom,(
  ! [X9] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))) <=> s('type.integer.int',X9) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.integer.INT_ABS_LE|split|2','thm.integer.INT_NEG_EQ|48|')).
fof(f13777,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f13776,f968])).
fof(f13776,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f13775,f961])).
fof(f13775,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl1),
  inference(forward_demodulation,[],[f13735,f922])).
fof(f13735,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl1),
  inference(superposition,[],[f985,f914])).
fof(f13774,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl1),
  inference(forward_demodulation,[],[f13728,f922])).
fof(f13728,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl1),
  inference(superposition,[],[f985,f949])).
fof(f949,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f633])).
fof(f1121,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | $spl7),
  inference(cnf_transformation,[],[f1121_D])).
fof(f1121_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f13727,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl1),
  inference(superposition,[],[f985,f951])).
fof(f13773,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f13722,f922])).
fof(f13722,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | $spl1),
  inference(superposition,[],[f985,f917])).
fof(f13714,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) ) | $spl1),
  inference(resolution,[],[f985,f976])).
fof(f13619,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15))) | s(bool,X21) != s(bool,X23)) ) | $spl1),
  inference(resolution,[],[f985,f650])).
fof(f13617,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) ) | $spl1),
  inference(resolution,[],[f985,f649])).
fof(f13616,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15))) | s(bool,X12) != s(bool,X13)) ) | $spl1),
  inference(resolution,[],[f985,f648])).
fof(f13615,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) ) | $spl1),
  inference(resolution,[],[f985,f648])).
fof(f13614,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15))) | s(bool,X6) != s(bool,X8)) ) | $spl1),
  inference(resolution,[],[f985,f646])).
fof(f13613,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15))) | s(bool,X3) != s(bool,X4)) ) | $spl1),
  inference(resolution,[],[f985,f644])).
fof(f13612,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) ) | $spl1),
  inference(resolution,[],[f985,f644])).
fof(f2533,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl53),
  inference(cnf_transformation,[],[f2533_D])).
fof(f2533_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f2535,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl54),
  inference(cnf_transformation,[],[f2535_D])).
fof(f2535_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f13765,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl1),
  inference(forward_demodulation,[],[f13606,f962])).
fof(f13606,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl1),
  inference(resolution,[],[f985,f896])).
fof(f1296,plain,(
  s('type.integer.int',sK15) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl31),
  inference(cnf_transformation,[],[f1296_D])).
fof(f1296_D,plain,(
  s('type.integer.int',sK15) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f13542,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | s('type.integer.int',sK15) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl4),
  inference(superposition,[],[f988,f951])).
fof(f13535,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl4),
  inference(resolution,[],[f988,f976])).
fof(f13444,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X23) != s(bool,X24)) ) | $spl4),
  inference(resolution,[],[f988,f650])).
fof(f13443,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl4),
  inference(resolution,[],[f988,f650])).
fof(f13441,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl4),
  inference(resolution,[],[f988,f649])).
fof(f13440,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl4),
  inference(resolution,[],[f988,f648])).
fof(f13439,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X8) != s(bool,X9)) ) | $spl4),
  inference(resolution,[],[f988,f646])).
fof(f13438,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl4),
  inference(resolution,[],[f988,f646])).
fof(f13437,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl4),
  inference(resolution,[],[f988,f644])).
fof(f13436,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))) | $spl4),
  inference(resolution,[],[f988,f909])).
fof(f13435,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | $spl4),
  inference(resolution,[],[f988,f904])).
fof(f13571,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | $spl4),
  inference(forward_demodulation,[],[f13570,f913])).
fof(f13570,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl4),
  inference(forward_demodulation,[],[f13433,f961])).
fof(f13433,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl4),
  inference(resolution,[],[f988,f901])).
fof(f13372,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f976])).
fof(f13281,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X23) != s(bool,X24)) ) | $spl2),
  inference(resolution,[],[f986,f650])).
fof(f13280,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f650])).
fof(f13278,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f649])).
fof(f13277,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f648])).
fof(f13276,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X8) != s(bool,X9)) ) | $spl2),
  inference(resolution,[],[f986,f646])).
fof(f13275,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f646])).
fof(f13274,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f644])).
fof(f13270,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | $spl2),
  inference(resolution,[],[f986,f901])).
fof(f11854,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | $spl4),
  inference(forward_demodulation,[],[f11853,f913])).
fof(f11853,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl4),
  inference(forward_demodulation,[],[f11717,f961])).
fof(f11717,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl4),
  inference(resolution,[],[f988,f901])).
fof(f11826,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | s('type.integer.int',sK15) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl4),
  inference(superposition,[],[f988,f951])).
fof(f13223,plain,(
  $false | ($spl1 | $spl4 | $spl10)),
  inference(subsumption_resolution,[],[f13222,f8691])).
fof(f8691,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | ($spl4 | $spl10)),
  inference(forward_demodulation,[],[f988,f1126])).
fof(f13222,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | ($spl1 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f13221,f922])).
fof(f13221,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | ($spl1 | $spl4 | $spl10)),
  inference(subsumption_resolution,[],[f13167,f8691])).
fof(f13167,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | ($spl1 | $spl10)),
  inference(superposition,[],[f4856,f917])).
fof(f4856,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK14)))) | ($spl1 | $spl10)),
  inference(backward_demodulation,[],[f1126,f985])).
fof(f13210,plain,(
  $false | ($spl1 | $spl10)),
  inference(subsumption_resolution,[],[f13052,f888])).
fof(f13052,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ($spl1 | $spl10)),
  inference(resolution,[],[f4856,f900])).
fof(f13003,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | ($spl4 | $spl10)),
  inference(resolution,[],[f8691,f976])).
fof(f12912,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14))) | s(bool,X23) != s(bool,X24)) ) | ($spl4 | $spl10)),
  inference(resolution,[],[f8691,f650])).
fof(f12911,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | ($spl4 | $spl10)),
  inference(resolution,[],[f8691,f650])).
fof(f12909,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | ($spl4 | $spl10)),
  inference(resolution,[],[f8691,f649])).
fof(f12908,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | ($spl4 | $spl10)),
  inference(resolution,[],[f8691,f648])).
fof(f12907,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14))) | s(bool,X8) != s(bool,X9)) ) | ($spl4 | $spl10)),
  inference(resolution,[],[f8691,f646])).
fof(f12906,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | ($spl4 | $spl10)),
  inference(resolution,[],[f8691,f646])).
fof(f12905,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | ($spl4 | $spl10)),
  inference(resolution,[],[f8691,f644])).
fof(f12904,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) | ($spl4 | $spl10)),
  inference(resolution,[],[f8691,f909])).
fof(f12903,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | ($spl4 | $spl10)),
  inference(resolution,[],[f8691,f904])).
fof(f5364,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl124),
  inference(cnf_transformation,[],[f5364_D])).
fof(f5364_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f12854,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | $spl10),
  inference(forward_demodulation,[],[f12853,f1126])).
fof(f12853,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl10),
  inference(forward_demodulation,[],[f12852,f1126])).
fof(f12852,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl10),
  inference(forward_demodulation,[],[f12562,f959])).
fof(f12562,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl10),
  inference(superposition,[],[f956,f1126])).
fof(f12851,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl10),
  inference(forward_demodulation,[],[f12850,f1126])).
fof(f12850,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl10),
  inference(forward_demodulation,[],[f12849,f1126])).
fof(f12849,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl10),
  inference(forward_demodulation,[],[f12561,f959])).
fof(f12561,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl10),
  inference(superposition,[],[f956,f1126])).
fof(f5266,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | $spl122),
  inference(cnf_transformation,[],[f5266_D])).
fof(f5266_D,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f12762,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | $spl10),
  inference(forward_demodulation,[],[f12761,f1126])).
fof(f12761,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | $spl10),
  inference(forward_demodulation,[],[f12482,f922])).
fof(f12482,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | $spl10),
  inference(superposition,[],[f917,f1126])).
fof(f12760,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | $spl10),
  inference(forward_demodulation,[],[f12759,f1126])).
fof(f12759,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))) | $spl10),
  inference(forward_demodulation,[],[f12481,f922])).
fof(f12481,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))) | $spl10),
  inference(superposition,[],[f917,f1126])).
fof(f5267,plain,(
  s('type.integer.int',sK14) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | $spl123),
  inference(cnf_transformation,[],[f5267_D])).
fof(f5267_D,plain,(
  s('type.integer.int',sK14) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) <=> ~$spl123),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl123])])).
fof(f2548,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | $spl62),
  inference(cnf_transformation,[],[f2548_D])).
fof(f2548_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f12758,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | s('type.integer.int',sK14) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | $spl10),
  inference(forward_demodulation,[],[f12757,f1126])).
fof(f12757,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.integer.int',sK14) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | $spl10),
  inference(forward_demodulation,[],[f12480,f922])).
fof(f12480,plain,(
  s('type.integer.int',sK14) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | $spl10),
  inference(superposition,[],[f916,f1126])).
fof(f12756,plain,(
  s('type.integer.int',sK14) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | $spl10),
  inference(forward_demodulation,[],[f12755,f1126])).
fof(f12755,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | s('type.integer.int',sK15) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))) | $spl10),
  inference(forward_demodulation,[],[f12479,f922])).
fof(f12479,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | s('type.integer.int',sK15) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))) | $spl10),
  inference(superposition,[],[f916,f1126])).
fof(f5257,plain,(
  ( ! [X99] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X99),s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X99))))) ) | $spl120),
  inference(cnf_transformation,[],[f5257_D])).
fof(f5257_D,plain,(
  ( ! [X99] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X99),s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X99))))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f12752,plain,(
  ( ! [X101] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X101),s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X101))))) ) | $spl10),
  inference(forward_demodulation,[],[f12751,f1126])).
fof(f12751,plain,(
  ( ! [X101] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X101),s('type.integer.int',sK15)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X101))))) ) | $spl10),
  inference(forward_demodulation,[],[f12472,f922])).
fof(f12472,plain,(
  ( ! [X101] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X101),s('type.integer.int',sK15)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X101))))) ) | $spl10),
  inference(superposition,[],[f910,f1126])).
fof(f5254,plain,(
  ( ! [X98] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',X98)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X98))))) ) | $spl118),
  inference(cnf_transformation,[],[f5254_D])).
fof(f5254_D,plain,(
  ( ! [X98] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',X98)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X98))))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f12750,plain,(
  ( ! [X100] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',X100)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X100))))) ) | $spl10),
  inference(forward_demodulation,[],[f12749,f1126])).
fof(f12749,plain,(
  ( ! [X100] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK15),s('type.integer.int',X100)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X100))))) ) | $spl10),
  inference(forward_demodulation,[],[f12471,f922])).
fof(f12471,plain,(
  ( ! [X100] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK15),s('type.integer.int',X100)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X100))))) ) | $spl10),
  inference(superposition,[],[f910,f1126])).
fof(f12748,plain,(
  ( ! [X99] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X99),s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X99))))) ) | $spl10),
  inference(forward_demodulation,[],[f12747,f1126])).
fof(f12747,plain,(
  ( ! [X99] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X99),s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X99))))) ) | $spl10),
  inference(forward_demodulation,[],[f12470,f922])).
fof(f12470,plain,(
  ( ! [X99] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X99),s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X99))))) ) | $spl10),
  inference(superposition,[],[f910,f1126])).
fof(f12746,plain,(
  ( ! [X98] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',X98)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X98))))) ) | $spl10),
  inference(forward_demodulation,[],[f12745,f1126])).
fof(f12745,plain,(
  ( ! [X98] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',X98)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X98))))) ) | $spl10),
  inference(forward_demodulation,[],[f12469,f922])).
fof(f12469,plain,(
  ( ! [X98] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',X98)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X98)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15))))) ) | $spl10),
  inference(superposition,[],[f910,f1126])).
fof(f5126,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl112),
  inference(cnf_transformation,[],[f5126_D])).
fof(f5126_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f12634,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | $spl10),
  inference(forward_demodulation,[],[f12633,f1126])).
fof(f12633,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | $spl10),
  inference(forward_demodulation,[],[f12632,f921])).
fof(f12632,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | $spl10),
  inference(forward_demodulation,[],[f12631,f919])).
fof(f12631,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl10),
  inference(forward_demodulation,[],[f12378,f922])).
fof(f12378,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl10),
  inference(superposition,[],[f885,f1126])).
fof(f12630,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl10),
  inference(forward_demodulation,[],[f12629,f1126])).
fof(f12629,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl10),
  inference(forward_demodulation,[],[f12628,f922])).
fof(f12628,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | $spl10),
  inference(forward_demodulation,[],[f12627,f921])).
fof(f12627,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | $spl10),
  inference(forward_demodulation,[],[f12377,f919])).
fof(f12377,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | $spl10),
  inference(superposition,[],[f885,f1126])).
fof(f5133,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK44(sK15)))) | $spl116),
  inference(cnf_transformation,[],[f5133_D])).
fof(f5133_D,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK44(sK15)))) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f12626,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK44(sK15)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | $spl10),
  inference(forward_demodulation,[],[f12625,f1126])).
fof(f12625,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK44(sK15)))) | $spl10),
  inference(forward_demodulation,[],[f12376,f922])).
fof(f12376,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK44(sK15)))) | $spl10),
  inference(superposition,[],[f882,f1126])).
fof(f882,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK44(X0)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f608])).
fof(f608,plain,(
  ! [X0] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK44(X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK44])],[f450])).
fof(f450,plain,(
  ! [X0] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | ? [X1] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))),
  inference(ennf_transformation,[],[f341])).
fof(f341,plain,(
  ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) => ? [X1] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))),
  inference(rectify,[],[f144])).
fof(f144,axiom,(
  ! [X17] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X17)))) => ? [X10] : s('type.integer.int',X17) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10))))),
  file('thm.integer.INT_ABS_LE|split|2','thm.integer.NUM_POSINT_EXISTS')).
fof(f5129,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK43(sK15)))))) | $spl114),
  inference(cnf_transformation,[],[f5129_D])).
fof(f5129_D,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK43(sK15)))))) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f12624,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK43(sK15)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl10),
  inference(forward_demodulation,[],[f12623,f1126])).
fof(f12623,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK43(sK15)))))) | $spl10),
  inference(forward_demodulation,[],[f12375,f921])).
fof(f12375,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK43(sK15)))))) | $spl10),
  inference(superposition,[],[f881,f1126])).
fof(f1126,plain,(
  s('type.integer.int',sK14) = s('type.integer.int',sK15) | $spl10),
  inference(cnf_transformation,[],[f1126_D])).
fof(f1126_D,plain,(
  s('type.integer.int',sK14) = s('type.integer.int',sK15) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f11441,plain,(
  s('type.integer.int',sK14) = s('type.integer.int',sK15) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | $spl32),
  inference(resolution,[],[f1297,f904])).
fof(f11442,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | s('type.integer.int',sK14) = s('type.integer.int',sK15) | $spl32),
  inference(resolution,[],[f1297,f909])).
fof(f11292,plain,(
  s('type.integer.int',sK14) = s('type.integer.int',sK15) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | $spl9),
  inference(resolution,[],[f1123,f904])).
fof(f11293,plain,(
  s('type.integer.int',sK14) = s('type.integer.int',sK15) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | $spl9),
  inference(resolution,[],[f1123,f909])).
fof(f11712,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | $spl4),
  inference(resolution,[],[f988,f901])).
fof(f12305,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl35),
  inference(resolution,[],[f1717,f976])).
fof(f12210,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))) | s(bool,X21) != s(bool,X23)) ) | $spl35),
  inference(resolution,[],[f1717,f650])).
fof(f12208,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl35),
  inference(resolution,[],[f1717,f649])).
fof(f12207,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))) | s(bool,X12) != s(bool,X13)) ) | $spl35),
  inference(resolution,[],[f1717,f648])).
fof(f12206,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl35),
  inference(resolution,[],[f1717,f648])).
fof(f12205,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))) | s(bool,X6) != s(bool,X8)) ) | $spl35),
  inference(resolution,[],[f1717,f646])).
fof(f12204,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))) | s(bool,X3) != s(bool,X4)) ) | $spl35),
  inference(resolution,[],[f1717,f644])).
fof(f12203,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) ) | $spl35),
  inference(resolution,[],[f1717,f644])).
fof(f11819,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl4),
  inference(resolution,[],[f988,f976])).
fof(f11728,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X23) != s(bool,X24)) ) | $spl4),
  inference(resolution,[],[f988,f650])).
fof(f11727,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl4),
  inference(resolution,[],[f988,f650])).
fof(f11725,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl4),
  inference(resolution,[],[f988,f649])).
fof(f11724,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl4),
  inference(resolution,[],[f988,f648])).
fof(f11723,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X8) != s(bool,X9)) ) | $spl4),
  inference(resolution,[],[f988,f646])).
fof(f11722,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl4),
  inference(resolution,[],[f988,f646])).
fof(f11721,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl4),
  inference(resolution,[],[f988,f644])).
fof(f11720,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))) | $spl4),
  inference(resolution,[],[f988,f909])).
fof(f11719,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | $spl4),
  inference(resolution,[],[f988,f904])).
fof(f11683,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl3),
  inference(resolution,[],[f1579,f976])).
fof(f11592,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X23) != s(bool,X24)) ) | $spl3),
  inference(resolution,[],[f1579,f650])).
fof(f11591,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl3),
  inference(resolution,[],[f1579,f650])).
fof(f11589,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl3),
  inference(resolution,[],[f1579,f649])).
fof(f11588,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl3),
  inference(resolution,[],[f1579,f648])).
fof(f11587,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X8) != s(bool,X9)) ) | $spl3),
  inference(resolution,[],[f1579,f646])).
fof(f11586,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl3),
  inference(resolution,[],[f1579,f646])).
fof(f11585,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl3),
  inference(resolution,[],[f1579,f644])).
fof(f11582,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | $spl3),
  inference(resolution,[],[f1579,f938])).
fof(f11580,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl3),
  inference(resolution,[],[f1579,f932])).
fof(f11541,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl32),
  inference(resolution,[],[f1297,f976])).
fof(f11450,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X23) != s(bool,X24)) ) | $spl32),
  inference(resolution,[],[f1297,f650])).
fof(f11449,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl32),
  inference(resolution,[],[f1297,f650])).
fof(f11447,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl32),
  inference(resolution,[],[f1297,f649])).
fof(f11446,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl32),
  inference(resolution,[],[f1297,f648])).
fof(f11445,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X8) != s(bool,X9)) ) | $spl32),
  inference(resolution,[],[f1297,f646])).
fof(f11444,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl32),
  inference(resolution,[],[f1297,f646])).
fof(f11443,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl32),
  inference(resolution,[],[f1297,f644])).
fof(f11439,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | $spl32),
  inference(resolution,[],[f1297,f901])).
fof(f11400,plain,(
  ( ! [X187,X188] : (p(s(bool,X187)) | ~p(s(bool,X188)) | s(bool,X188) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl9),
  inference(resolution,[],[f1123,f976])).
fof(f11305,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X22) != s(bool,X24)) ) | $spl9),
  inference(resolution,[],[f1123,f650])).
fof(f11303,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl9),
  inference(resolution,[],[f1123,f649])).
fof(f11302,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X13) != s(bool,X14)) ) | $spl9),
  inference(resolution,[],[f1123,f648])).
fof(f11301,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl9),
  inference(resolution,[],[f1123,f648])).
fof(f11300,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X7) != s(bool,X9)) ) | $spl9),
  inference(resolution,[],[f1123,f646])).
fof(f11299,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X4) != s(bool,X5)) ) | $spl9),
  inference(resolution,[],[f1123,f644])).
fof(f11298,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl9),
  inference(resolution,[],[f1123,f644])).
fof(f11296,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl9),
  inference(resolution,[],[f1123,f933])).
fof(f11294,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) | $spl9),
  inference(resolution,[],[f1123,f927])).
fof(f11253,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl3),
  inference(resolution,[],[f987,f976])).
fof(f11158,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X21) != s(bool,X23)) ) | $spl3),
  inference(resolution,[],[f987,f650])).
fof(f11156,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl3),
  inference(resolution,[],[f987,f649])).
fof(f11155,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X12) != s(bool,X13)) ) | $spl3),
  inference(resolution,[],[f987,f648])).
fof(f11154,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl3),
  inference(resolution,[],[f987,f648])).
fof(f11153,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X6) != s(bool,X8)) ) | $spl3),
  inference(resolution,[],[f987,f646])).
fof(f11152,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X3) != s(bool,X4)) ) | $spl3),
  inference(resolution,[],[f987,f644])).
fof(f11151,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl3),
  inference(resolution,[],[f987,f644])).
fof(f11145,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl3),
  inference(resolution,[],[f987,f896])).
fof(f11142,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | p(s(bool,X1))) ) | $spl3),
  inference(subsumption_resolution,[],[f11132,f851])).
fof(f11132,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | p(s(bool,X1))) ) | $spl3),
  inference(resolution,[],[f6155,f779])).
fof(f11141,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | p(s(bool,X0))) ) | $spl3),
  inference(subsumption_resolution,[],[f11131,f851])).
fof(f11131,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | p(s(bool,X0))) ) | $spl3),
  inference(resolution,[],[f6155,f783])).
fof(f11130,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))))) ) | $spl3),
  inference(subsumption_resolution,[],[f11120,f851])).
fof(f11120,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))))) ) | $spl3),
  inference(resolution,[],[f6154,f779])).
fof(f11129,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))))) ) | $spl3),
  inference(subsumption_resolution,[],[f11119,f851])).
fof(f11119,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))))) ) | $spl3),
  inference(resolution,[],[f6154,f783])).
fof(f10770,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | $spl7),
  inference(resolution,[],[f1121,f901])).
fof(f11094,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | $spl7),
  inference(forward_demodulation,[],[f10889,f922])).
fof(f10889,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | $spl7),
  inference(superposition,[],[f1121,f917])).
fof(f11037,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) | $spl5),
  inference(resolution,[],[f1140,f976])).
fof(f10946,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))))) | s(bool,X23) != s(bool,X24)) ) | $spl5),
  inference(resolution,[],[f1140,f650])).
fof(f10945,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) | $spl5),
  inference(resolution,[],[f1140,f650])).
fof(f10943,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) | $spl5),
  inference(resolution,[],[f1140,f649])).
fof(f10942,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) | $spl5),
  inference(resolution,[],[f1140,f648])).
fof(f10941,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))))) | s(bool,X8) != s(bool,X9)) ) | $spl5),
  inference(resolution,[],[f1140,f646])).
fof(f10940,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) | $spl5),
  inference(resolution,[],[f1140,f646])).
fof(f10939,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) ) | $spl5),
  inference(resolution,[],[f1140,f644])).
fof(f10881,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl7),
  inference(resolution,[],[f1121,f976])).
fof(f10786,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X21) != s(bool,X23)) ) | $spl7),
  inference(resolution,[],[f1121,f650])).
fof(f10784,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl7),
  inference(resolution,[],[f1121,f649])).
fof(f10783,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X12) != s(bool,X13)) ) | $spl7),
  inference(resolution,[],[f1121,f648])).
fof(f10782,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl7),
  inference(resolution,[],[f1121,f648])).
fof(f10781,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X6) != s(bool,X8)) ) | $spl7),
  inference(resolution,[],[f1121,f646])).
fof(f10780,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X3) != s(bool,X4)) ) | $spl7),
  inference(resolution,[],[f1121,f644])).
fof(f10779,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl7),
  inference(resolution,[],[f1121,f644])).
fof(f10773,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl7),
  inference(resolution,[],[f1121,f896])).
fof(f10722,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f989,f976])).
fof(f10627,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X21) != s(bool,X23)) ) | $spl5),
  inference(resolution,[],[f989,f650])).
fof(f10625,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f989,f649])).
fof(f10624,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X12) != s(bool,X13)) ) | $spl5),
  inference(resolution,[],[f989,f648])).
fof(f10623,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f989,f648])).
fof(f10622,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X6) != s(bool,X8)) ) | $spl5),
  inference(resolution,[],[f989,f646])).
fof(f10621,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X3) != s(bool,X4)) ) | $spl5),
  inference(resolution,[],[f989,f644])).
fof(f10620,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f989,f644])).
fof(f10554,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl6),
  inference(resolution,[],[f1122,f976])).
fof(f10463,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X23) != s(bool,X24)) ) | $spl6),
  inference(resolution,[],[f1122,f650])).
fof(f10462,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl6),
  inference(resolution,[],[f1122,f650])).
fof(f10460,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl6),
  inference(resolution,[],[f1122,f649])).
fof(f10459,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl6),
  inference(resolution,[],[f1122,f648])).
fof(f10458,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X8) != s(bool,X9)) ) | $spl6),
  inference(resolution,[],[f1122,f646])).
fof(f10457,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl6),
  inference(resolution,[],[f1122,f646])).
fof(f10456,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl6),
  inference(resolution,[],[f1122,f644])).
fof(f10455,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | $spl6),
  inference(resolution,[],[f1122,f909])).
fof(f10454,plain,(
  s('type.integer.int',sK15) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | $spl6),
  inference(resolution,[],[f1122,f904])).
fof(f9951,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl33),
  inference(resolution,[],[f1298,f976])).
fof(f9856,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X21) != s(bool,X23)) ) | $spl33),
  inference(resolution,[],[f1298,f650])).
fof(f9854,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl33),
  inference(resolution,[],[f1298,f649])).
fof(f9853,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X12) != s(bool,X13)) ) | $spl33),
  inference(resolution,[],[f1298,f648])).
fof(f9852,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl33),
  inference(resolution,[],[f1298,f648])).
fof(f9851,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X6) != s(bool,X8)) ) | $spl33),
  inference(resolution,[],[f1298,f646])).
fof(f9850,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X3) != s(bool,X4)) ) | $spl33),
  inference(resolution,[],[f1298,f644])).
fof(f9849,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl33),
  inference(resolution,[],[f1298,f644])).
fof(f9843,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl33),
  inference(resolution,[],[f1298,f896])).
fof(f9777,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl8),
  inference(resolution,[],[f1124,f976])).
fof(f9686,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X23) != s(bool,X24)) ) | $spl8),
  inference(resolution,[],[f1124,f650])).
fof(f9685,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl8),
  inference(resolution,[],[f1124,f650])).
fof(f9683,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl8),
  inference(resolution,[],[f1124,f649])).
fof(f9682,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl8),
  inference(resolution,[],[f1124,f648])).
fof(f9681,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X8) != s(bool,X9)) ) | $spl8),
  inference(resolution,[],[f1124,f646])).
fof(f9680,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl8),
  inference(resolution,[],[f1124,f646])).
fof(f9679,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl8),
  inference(resolution,[],[f1124,f644])).
fof(f9676,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | $spl8),
  inference(resolution,[],[f1124,f938])).
fof(f9674,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl8),
  inference(resolution,[],[f1124,f932])).
fof(f9546,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f9545,f968])).
fof(f9545,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f9544,f961])).
fof(f9544,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) | $spl1),
  inference(forward_demodulation,[],[f9496,f922])).
fof(f9496,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) | $spl1),
  inference(superposition,[],[f985,f966])).
fof(f9543,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(subsumption_resolution,[],[f9542,f965])).
fof(f9542,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f9541,f959])).
fof(f9541,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(forward_demodulation,[],[f9540,f958])).
fof(f9540,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(forward_demodulation,[],[f9539,f959])).
fof(f9539,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(forward_demodulation,[],[f9495,f922])).
fof(f9495,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(superposition,[],[f985,f956])).
fof(f9538,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl1),
  inference(forward_demodulation,[],[f9494,f922])).
fof(f9494,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl1),
  inference(superposition,[],[f985,f955])).
fof(f9537,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(forward_demodulation,[],[f9493,f922])).
fof(f9493,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(superposition,[],[f985,f954])).
fof(f9536,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl1),
  inference(forward_demodulation,[],[f9492,f922])).
fof(f9492,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl1),
  inference(superposition,[],[f985,f953])).
fof(f9535,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(forward_demodulation,[],[f9491,f922])).
fof(f9491,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(superposition,[],[f985,f952])).
fof(f9534,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f9533,f962])).
fof(f9533,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) | $spl1),
  inference(forward_demodulation,[],[f9490,f922])).
fof(f9490,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) | $spl1),
  inference(superposition,[],[f985,f950])).
fof(f9532,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f9531,f962])).
fof(f9531,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) | $spl1),
  inference(forward_demodulation,[],[f9489,f922])).
fof(f9489,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) | $spl1),
  inference(superposition,[],[f985,f948])).
fof(f9528,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f9527,f968])).
fof(f9527,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f9526,f961])).
fof(f9526,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl1),
  inference(forward_demodulation,[],[f9486,f922])).
fof(f9486,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl1),
  inference(superposition,[],[f985,f914])).
fof(f9525,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl1),
  inference(forward_demodulation,[],[f9479,f922])).
fof(f9479,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl1),
  inference(superposition,[],[f985,f949])).
fof(f9478,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl1),
  inference(superposition,[],[f985,f951])).
fof(f9524,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f9473,f922])).
fof(f9473,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | $spl1),
  inference(superposition,[],[f985,f917])).
fof(f9465,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) ) | $spl1),
  inference(resolution,[],[f985,f976])).
fof(f9370,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15))) | s(bool,X21) != s(bool,X23)) ) | $spl1),
  inference(resolution,[],[f985,f650])).
fof(f9368,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) ) | $spl1),
  inference(resolution,[],[f985,f649])).
fof(f9367,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15))) | s(bool,X12) != s(bool,X13)) ) | $spl1),
  inference(resolution,[],[f985,f648])).
fof(f9366,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) ) | $spl1),
  inference(resolution,[],[f985,f648])).
fof(f9365,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15))) | s(bool,X6) != s(bool,X8)) ) | $spl1),
  inference(resolution,[],[f985,f646])).
fof(f9364,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15))) | s(bool,X3) != s(bool,X4)) ) | $spl1),
  inference(resolution,[],[f985,f644])).
fof(f9363,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) ) | $spl1),
  inference(resolution,[],[f985,f644])).
fof(f9516,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl1),
  inference(forward_demodulation,[],[f9357,f962])).
fof(f9357,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl1),
  inference(resolution,[],[f985,f896])).
fof(f9303,plain,(
  ( ! [X187,X188] : (p(s(bool,X187)) | ~p(s(bool,X188)) | s(bool,X188) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl2),
  inference(resolution,[],[f991,f976])).
fof(f9208,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X22) != s(bool,X24)) ) | $spl2),
  inference(resolution,[],[f991,f650])).
fof(f9206,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl2),
  inference(resolution,[],[f991,f649])).
fof(f9205,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X13) != s(bool,X14)) ) | $spl2),
  inference(resolution,[],[f991,f648])).
fof(f9204,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl2),
  inference(resolution,[],[f991,f648])).
fof(f9203,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X7) != s(bool,X9)) ) | $spl2),
  inference(resolution,[],[f991,f646])).
fof(f9202,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X4) != s(bool,X5)) ) | $spl2),
  inference(resolution,[],[f991,f644])).
fof(f9201,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl2),
  inference(resolution,[],[f991,f644])).
fof(f9199,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl2),
  inference(resolution,[],[f991,f933])).
fof(f9197,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))),s('type.integer.int',sK14)))) | $spl2),
  inference(resolution,[],[f991,f927])).
fof(f9130,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | s('type.integer.int',sK15) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl4),
  inference(superposition,[],[f988,f951])).
fof(f9123,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl4),
  inference(resolution,[],[f988,f976])).
fof(f9032,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X23) != s(bool,X24)) ) | $spl4),
  inference(resolution,[],[f988,f650])).
fof(f9031,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl4),
  inference(resolution,[],[f988,f650])).
fof(f9029,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl4),
  inference(resolution,[],[f988,f649])).
fof(f9028,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl4),
  inference(resolution,[],[f988,f648])).
fof(f9027,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X8) != s(bool,X9)) ) | $spl4),
  inference(resolution,[],[f988,f646])).
fof(f9026,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl4),
  inference(resolution,[],[f988,f646])).
fof(f9025,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl4),
  inference(resolution,[],[f988,f644])).
fof(f9024,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))) | $spl4),
  inference(resolution,[],[f988,f909])).
fof(f9023,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | $spl4),
  inference(resolution,[],[f988,f904])).
fof(f9159,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | $spl4),
  inference(forward_demodulation,[],[f9158,f913])).
fof(f9158,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl4),
  inference(forward_demodulation,[],[f9021,f961])).
fof(f9021,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl4),
  inference(resolution,[],[f988,f901])).
fof(f8960,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f976])).
fof(f8869,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X23) != s(bool,X24)) ) | $spl2),
  inference(resolution,[],[f986,f650])).
fof(f8868,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f650])).
fof(f8866,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f649])).
fof(f8865,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f648])).
fof(f8864,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X8) != s(bool,X9)) ) | $spl2),
  inference(resolution,[],[f986,f646])).
fof(f8863,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f646])).
fof(f8862,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f644])).
fof(f8858,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | $spl2),
  inference(resolution,[],[f986,f901])).
fof(f8802,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | ($spl4 | $spl10)),
  inference(resolution,[],[f8691,f976])).
fof(f8711,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14))) | s(bool,X23) != s(bool,X24)) ) | ($spl4 | $spl10)),
  inference(resolution,[],[f8691,f650])).
fof(f8710,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | ($spl4 | $spl10)),
  inference(resolution,[],[f8691,f650])).
fof(f8708,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | ($spl4 | $spl10)),
  inference(resolution,[],[f8691,f649])).
fof(f8707,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | ($spl4 | $spl10)),
  inference(resolution,[],[f8691,f648])).
fof(f8706,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14))) | s(bool,X8) != s(bool,X9)) ) | ($spl4 | $spl10)),
  inference(resolution,[],[f8691,f646])).
fof(f8705,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | ($spl4 | $spl10)),
  inference(resolution,[],[f8691,f646])).
fof(f8704,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) ) | ($spl4 | $spl10)),
  inference(resolution,[],[f8691,f644])).
fof(f8703,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) | ($spl4 | $spl10)),
  inference(resolution,[],[f8691,f909])).
fof(f8702,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | ($spl4 | $spl10)),
  inference(resolution,[],[f8691,f904])).
fof(f8649,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | $spl10),
  inference(forward_demodulation,[],[f8648,f1126])).
fof(f8648,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl10),
  inference(forward_demodulation,[],[f8647,f1126])).
fof(f8647,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl10),
  inference(forward_demodulation,[],[f8357,f959])).
fof(f8357,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl10),
  inference(superposition,[],[f956,f1126])).
fof(f8646,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl10),
  inference(forward_demodulation,[],[f8645,f1126])).
fof(f8645,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl10),
  inference(forward_demodulation,[],[f8644,f1126])).
fof(f8644,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl10),
  inference(forward_demodulation,[],[f8356,f959])).
fof(f8356,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl10),
  inference(superposition,[],[f956,f1126])).
fof(f8557,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | $spl10),
  inference(forward_demodulation,[],[f8556,f1126])).
fof(f8556,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | $spl10),
  inference(forward_demodulation,[],[f8277,f922])).
fof(f8277,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | $spl10),
  inference(superposition,[],[f917,f1126])).
fof(f8555,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | $spl10),
  inference(forward_demodulation,[],[f8554,f1126])).
fof(f8554,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))) | $spl10),
  inference(forward_demodulation,[],[f8276,f922])).
fof(f8276,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))) | $spl10),
  inference(superposition,[],[f917,f1126])).
fof(f8553,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | s('type.integer.int',sK14) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | $spl10),
  inference(forward_demodulation,[],[f8552,f1126])).
fof(f8552,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.integer.int',sK14) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | $spl10),
  inference(forward_demodulation,[],[f8275,f922])).
fof(f8275,plain,(
  s('type.integer.int',sK14) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | $spl10),
  inference(superposition,[],[f916,f1126])).
fof(f8551,plain,(
  s('type.integer.int',sK14) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | $spl10),
  inference(forward_demodulation,[],[f8550,f1126])).
fof(f8550,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | s('type.integer.int',sK15) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))) | $spl10),
  inference(forward_demodulation,[],[f8274,f922])).
fof(f8274,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | s('type.integer.int',sK15) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))) | $spl10),
  inference(superposition,[],[f916,f1126])).
fof(f8547,plain,(
  ( ! [X101] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X101),s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X101))))) ) | $spl10),
  inference(forward_demodulation,[],[f8546,f1126])).
fof(f8546,plain,(
  ( ! [X101] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X101),s('type.integer.int',sK15)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X101))))) ) | $spl10),
  inference(forward_demodulation,[],[f8267,f922])).
fof(f8267,plain,(
  ( ! [X101] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X101),s('type.integer.int',sK15)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X101))))) ) | $spl10),
  inference(superposition,[],[f910,f1126])).
fof(f8545,plain,(
  ( ! [X100] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',X100)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X100))))) ) | $spl10),
  inference(forward_demodulation,[],[f8544,f1126])).
fof(f8544,plain,(
  ( ! [X100] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK15),s('type.integer.int',X100)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X100))))) ) | $spl10),
  inference(forward_demodulation,[],[f8266,f922])).
fof(f8266,plain,(
  ( ! [X100] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK15),s('type.integer.int',X100)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X100))))) ) | $spl10),
  inference(superposition,[],[f910,f1126])).
fof(f8543,plain,(
  ( ! [X99] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X99),s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X99))))) ) | $spl10),
  inference(forward_demodulation,[],[f8542,f1126])).
fof(f8542,plain,(
  ( ! [X99] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X99),s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X99))))) ) | $spl10),
  inference(forward_demodulation,[],[f8265,f922])).
fof(f8265,plain,(
  ( ! [X99] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X99),s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X99))))) ) | $spl10),
  inference(superposition,[],[f910,f1126])).
fof(f8541,plain,(
  ( ! [X98] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',X98)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X98))))) ) | $spl10),
  inference(forward_demodulation,[],[f8540,f1126])).
fof(f8540,plain,(
  ( ! [X98] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',X98)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X98))))) ) | $spl10),
  inference(forward_demodulation,[],[f8264,f922])).
fof(f8264,plain,(
  ( ! [X98] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',X98)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X98)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15))))) ) | $spl10),
  inference(superposition,[],[f910,f1126])).
fof(f8429,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | $spl10),
  inference(forward_demodulation,[],[f8428,f1126])).
fof(f8428,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | $spl10),
  inference(forward_demodulation,[],[f8427,f921])).
fof(f8427,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | $spl10),
  inference(forward_demodulation,[],[f8426,f919])).
fof(f8426,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl10),
  inference(forward_demodulation,[],[f8173,f922])).
fof(f8173,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl10),
  inference(superposition,[],[f885,f1126])).
fof(f8425,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl10),
  inference(forward_demodulation,[],[f8424,f1126])).
fof(f8424,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl10),
  inference(forward_demodulation,[],[f8423,f922])).
fof(f8423,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | $spl10),
  inference(forward_demodulation,[],[f8422,f921])).
fof(f8422,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | $spl10),
  inference(forward_demodulation,[],[f8172,f919])).
fof(f8172,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | $spl10),
  inference(superposition,[],[f885,f1126])).
fof(f8421,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK44(sK15)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | $spl10),
  inference(forward_demodulation,[],[f8420,f1126])).
fof(f8420,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK44(sK15)))) | $spl10),
  inference(forward_demodulation,[],[f8171,f922])).
fof(f8171,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK44(sK15)))) | $spl10),
  inference(superposition,[],[f882,f1126])).
fof(f8419,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK43(sK15)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl10),
  inference(forward_demodulation,[],[f8418,f1126])).
fof(f8418,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK43(sK15)))))) | $spl10),
  inference(forward_demodulation,[],[f8170,f921])).
fof(f8170,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK43(sK15)))))) | $spl10),
  inference(superposition,[],[f881,f1126])).
fof(f8158,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))))) ) | $spl3),
  inference(subsumption_resolution,[],[f8146,f851])).
fof(f8146,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))))) ) | $spl3),
  inference(resolution,[],[f6154,f779])).
fof(f8155,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))))) ) | $spl3),
  inference(subsumption_resolution,[],[f8145,f851])).
fof(f8145,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))))) ) | $spl3),
  inference(resolution,[],[f6154,f783])).
fof(f8064,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl5),
  inference(resolution,[],[f1139,f976])).
fof(f1139,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | $spl5),
  inference(resolution,[],[f989,f900])).
fof(f7969,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X21) != s(bool,X23)) ) | $spl5),
  inference(resolution,[],[f1139,f650])).
fof(f7967,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl5),
  inference(resolution,[],[f1139,f649])).
fof(f7966,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X12) != s(bool,X13)) ) | $spl5),
  inference(resolution,[],[f1139,f648])).
fof(f7965,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl5),
  inference(resolution,[],[f1139,f648])).
fof(f7964,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X6) != s(bool,X8)) ) | $spl5),
  inference(resolution,[],[f1139,f646])).
fof(f7963,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15))) | s(bool,X3) != s(bool,X4)) ) | $spl5),
  inference(resolution,[],[f1139,f644])).
fof(f7962,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) ) | $spl5),
  inference(resolution,[],[f1139,f644])).
fof(f7956,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl5),
  inference(resolution,[],[f1139,f896])).
fof(f7891,plain,(
  $false | $spl16),
  inference(equality_resolution,[],[f1129])).
fof(f7782,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f7637,f7482])).
fof(f7482,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl20),
  inference(equality_resolution,[],[f1131])).
fof(f7637,plain,(
  ( ! [X339,X340] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X339),s('type.integer.int',X340))))) ) | $spl20),
  inference(resolution,[],[f7482,f903])).
fof(f7781,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f7636,f7482])).
fof(f7636,plain,(
  ( ! [X337,X338] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X337),s('type.integer.int',X338))))) ) | $spl20),
  inference(resolution,[],[f7482,f894])).
fof(f7780,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f7635,f7482])).
fof(f7635,plain,(
  ( ! [X335,X336] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X335),s('type.integer.int',X336))))) ) | $spl20),
  inference(resolution,[],[f7482,f892])).
fof(f7779,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f7630,f7482])).
fof(f7630,plain,(
  ( ! [X329] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X329))))) ) | $spl20),
  inference(resolution,[],[f7482,f885])).
fof(f7778,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f7628,f7482])).
fof(f7628,plain,(
  ( ! [X328] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X328))))))) ) | $spl20),
  inference(resolution,[],[f7482,f885])).
fof(f7777,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f7617,f7482])).
fof(f7617,plain,(
  ( ! [X308,X309] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X308),s('type.integer.int',X309))))) ) | $spl20),
  inference(resolution,[],[f7482,f903])).
fof(f7776,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f7614,f7482])).
fof(f7614,plain,(
  ( ! [X302,X303] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X302),s('type.integer.int',X303))))) ) | $spl20),
  inference(resolution,[],[f7482,f894])).
fof(f7775,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f7613,f7482])).
fof(f7613,plain,(
  ( ! [X300,X301] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X300),s('type.integer.int',X301))))) ) | $spl20),
  inference(resolution,[],[f7482,f892])).
fof(f7612,plain,(
  $false | $spl20),
  inference(resolution,[],[f7482,f888])).
fof(f7629,plain,(
  $false | $spl20),
  inference(resolution,[],[f7482,f923])).
fof(f923,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))))) )),
  inference(cnf_transformation,[],[f375])).
fof(f375,plain,(
  ! [X0] : p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6] : p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))))))),
  file('thm.integer.INT_ABS_LE|split|2','thm.integer.INT_ABS_POS')).
fof(f7631,plain,(
  $false | $spl20),
  inference(resolution,[],[f7482,f887])).
fof(f887,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))))) )),
  inference(cnf_transformation,[],[f348])).
fof(f348,plain,(
  ! [X0] : p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))))))),
  inference(rectify,[],[f158])).
fof(f158,axiom,(
  ! [X10] : p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X10))))))),
  file('thm.integer.INT_ABS_LE|split|2','thm.integer.INT_POS')).
fof(f7633,plain,(
  $false | ($spl2 | $spl20)),
  inference(resolution,[],[f7482,f986])).
fof(f7658,plain,(
  $false | ($spl8 | $spl20)),
  inference(resolution,[],[f7482,f1124])).
fof(f7385,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,X29)) ) | $spl140),
  inference(cnf_transformation,[],[f7385_D])).
fof(f7385_D,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,X29)) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f7380,plain,(
  ( ! [X35,X34] : (s(bool,'F_0') != s(bool,X35) | p(s(bool,X34))) ) | $spl14),
  inference(duplicate_literal_removal,[],[f7324])).
fof(f7324,plain,(
  ( ! [X35,X34] : (s(bool,'F_0') != s(bool,X35) | p(s(bool,X34)) | p(s(bool,X34))) ) | $spl14),
  inference(superposition,[],[f1128,f872])).
fof(f872,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f606])).
fof(f606,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f331])).
fof(f331,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f330])).
fof(f330,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.integer.INT_ABS_LE|split|2','thm.bool.EQ_CLAUSES|split|3')).
fof(f7381,plain,(
  ( ! [X33,X32] : (s(bool,'F_0') != s(bool,X33) | p(s(bool,X32))) ) | $spl14),
  inference(duplicate_literal_removal,[],[f7323])).
fof(f7323,plain,(
  ( ! [X33,X32] : (s(bool,'F_0') != s(bool,X33) | p(s(bool,X32)) | p(s(bool,X32))) ) | $spl14),
  inference(superposition,[],[f1128,f870])).
fof(f870,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f605])).
fof(f605,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f329])).
fof(f329,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f328])).
fof(f328,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.integer.INT_ABS_LE|split|2','thm.bool.EQ_CLAUSES|split|2')).
fof(f7382,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,X29) | p(s(bool,X28))) ) | $spl14),
  inference(duplicate_literal_removal,[],[f7321])).
fof(f7321,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,X29) | p(s(bool,X28)) | p(s(bool,X28))) ) | $spl14),
  inference(superposition,[],[f1128,f848])).
fof(f7278,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl33),
  inference(resolution,[],[f1298,f976])).
fof(f7183,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X21) != s(bool,X23)) ) | $spl33),
  inference(resolution,[],[f1298,f650])).
fof(f7181,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl33),
  inference(resolution,[],[f1298,f649])).
fof(f7180,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X12) != s(bool,X13)) ) | $spl33),
  inference(resolution,[],[f1298,f648])).
fof(f7179,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl33),
  inference(resolution,[],[f1298,f648])).
fof(f7178,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X6) != s(bool,X8)) ) | $spl33),
  inference(resolution,[],[f1298,f646])).
fof(f7177,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X3) != s(bool,X4)) ) | $spl33),
  inference(resolution,[],[f1298,f644])).
fof(f7176,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl33),
  inference(resolution,[],[f1298,f644])).
fof(f7170,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl33),
  inference(resolution,[],[f1298,f896])).
fof(f7139,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl8),
  inference(resolution,[],[f1124,f976])).
fof(f7048,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X23) != s(bool,X24)) ) | $spl8),
  inference(resolution,[],[f1124,f650])).
fof(f7047,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl8),
  inference(resolution,[],[f1124,f650])).
fof(f7045,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl8),
  inference(resolution,[],[f1124,f649])).
fof(f7044,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl8),
  inference(resolution,[],[f1124,f648])).
fof(f7043,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X8) != s(bool,X9)) ) | $spl8),
  inference(resolution,[],[f1124,f646])).
fof(f7042,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl8),
  inference(resolution,[],[f1124,f646])).
fof(f7041,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl8),
  inference(resolution,[],[f1124,f644])).
fof(f7038,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | $spl8),
  inference(resolution,[],[f1124,f938])).
fof(f7036,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl8),
  inference(resolution,[],[f1124,f932])).
fof(f6994,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f6993,f968])).
fof(f6993,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f6992,f961])).
fof(f6992,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) | $spl1),
  inference(forward_demodulation,[],[f6944,f922])).
fof(f6944,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) | $spl1),
  inference(superposition,[],[f985,f966])).
fof(f6991,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(subsumption_resolution,[],[f6990,f965])).
fof(f6990,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f6989,f959])).
fof(f6989,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(forward_demodulation,[],[f6988,f958])).
fof(f6988,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(forward_demodulation,[],[f6987,f959])).
fof(f6987,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(forward_demodulation,[],[f6943,f922])).
fof(f6943,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(superposition,[],[f985,f956])).
fof(f6986,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl1),
  inference(forward_demodulation,[],[f6942,f922])).
fof(f6942,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl1),
  inference(superposition,[],[f985,f955])).
fof(f6985,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(forward_demodulation,[],[f6941,f922])).
fof(f6941,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(superposition,[],[f985,f954])).
fof(f6984,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl1),
  inference(forward_demodulation,[],[f6940,f922])).
fof(f6940,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl1),
  inference(superposition,[],[f985,f953])).
fof(f6983,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(forward_demodulation,[],[f6939,f922])).
fof(f6939,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(superposition,[],[f985,f952])).
fof(f6982,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f6981,f962])).
fof(f6981,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) | $spl1),
  inference(forward_demodulation,[],[f6938,f922])).
fof(f6938,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) | $spl1),
  inference(superposition,[],[f985,f950])).
fof(f6980,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f6979,f962])).
fof(f6979,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) | $spl1),
  inference(forward_demodulation,[],[f6937,f922])).
fof(f6937,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) | $spl1),
  inference(superposition,[],[f985,f948])).
fof(f6976,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f6975,f968])).
fof(f6975,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f6974,f961])).
fof(f6974,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl1),
  inference(forward_demodulation,[],[f6934,f922])).
fof(f6934,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl1),
  inference(superposition,[],[f985,f914])).
fof(f6973,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl1),
  inference(forward_demodulation,[],[f6927,f922])).
fof(f6927,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl1),
  inference(superposition,[],[f985,f949])).
fof(f6926,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl1),
  inference(superposition,[],[f985,f951])).
fof(f6972,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f6921,f922])).
fof(f6921,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | $spl1),
  inference(superposition,[],[f985,f917])).
fof(f6913,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) ) | $spl1),
  inference(resolution,[],[f985,f976])).
fof(f6818,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15))) | s(bool,X21) != s(bool,X23)) ) | $spl1),
  inference(resolution,[],[f985,f650])).
fof(f6816,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) ) | $spl1),
  inference(resolution,[],[f985,f649])).
fof(f6815,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15))) | s(bool,X12) != s(bool,X13)) ) | $spl1),
  inference(resolution,[],[f985,f648])).
fof(f6814,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) ) | $spl1),
  inference(resolution,[],[f985,f648])).
fof(f6813,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15))) | s(bool,X6) != s(bool,X8)) ) | $spl1),
  inference(resolution,[],[f985,f646])).
fof(f6812,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15))) | s(bool,X3) != s(bool,X4)) ) | $spl1),
  inference(resolution,[],[f985,f644])).
fof(f6811,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) ) | $spl1),
  inference(resolution,[],[f985,f644])).
fof(f6964,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl1),
  inference(forward_demodulation,[],[f6805,f962])).
fof(f6805,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl1),
  inference(resolution,[],[f985,f896])).
fof(f6682,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f989,f976])).
fof(f6587,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X21) != s(bool,X23)) ) | $spl5),
  inference(resolution,[],[f989,f650])).
fof(f6585,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f989,f649])).
fof(f6584,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X12) != s(bool,X13)) ) | $spl5),
  inference(resolution,[],[f989,f648])).
fof(f6583,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f989,f648])).
fof(f6582,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X6) != s(bool,X8)) ) | $spl5),
  inference(resolution,[],[f989,f646])).
fof(f6581,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X3) != s(bool,X4)) ) | $spl5),
  inference(resolution,[],[f989,f644])).
fof(f6580,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f989,f644])).
fof(f6506,plain,(
  ( ! [X187,X188] : (p(s(bool,X187)) | ~p(s(bool,X188)) | s(bool,X188) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl2),
  inference(resolution,[],[f991,f976])).
fof(f6411,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X22) != s(bool,X24)) ) | $spl2),
  inference(resolution,[],[f991,f650])).
fof(f6409,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl2),
  inference(resolution,[],[f991,f649])).
fof(f6408,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X13) != s(bool,X14)) ) | $spl2),
  inference(resolution,[],[f991,f648])).
fof(f6407,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl2),
  inference(resolution,[],[f991,f648])).
fof(f6406,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X7) != s(bool,X9)) ) | $spl2),
  inference(resolution,[],[f991,f646])).
fof(f6405,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X4) != s(bool,X5)) ) | $spl2),
  inference(resolution,[],[f991,f644])).
fof(f6404,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl2),
  inference(resolution,[],[f991,f644])).
fof(f6402,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl2),
  inference(resolution,[],[f991,f933])).
fof(f6400,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))),s('type.integer.int',sK14)))) | $spl2),
  inference(resolution,[],[f991,f927])).
fof(f6339,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f976])).
fof(f6248,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X23) != s(bool,X24)) ) | $spl2),
  inference(resolution,[],[f986,f650])).
fof(f6247,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f650])).
fof(f6245,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f649])).
fof(f6244,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f648])).
fof(f6243,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X8) != s(bool,X9)) ) | $spl2),
  inference(resolution,[],[f986,f646])).
fof(f6242,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f646])).
fof(f6241,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f644])).
fof(f6237,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | $spl2),
  inference(resolution,[],[f986,f901])).
fof(f6169,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl3),
  inference(resolution,[],[f1579,f976])).
fof(f6078,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X23) != s(bool,X24)) ) | $spl3),
  inference(resolution,[],[f1579,f650])).
fof(f6077,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl3),
  inference(resolution,[],[f1579,f650])).
fof(f6075,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl3),
  inference(resolution,[],[f1579,f649])).
fof(f6074,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl3),
  inference(resolution,[],[f1579,f648])).
fof(f6073,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X8) != s(bool,X9)) ) | $spl3),
  inference(resolution,[],[f1579,f646])).
fof(f6072,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl3),
  inference(resolution,[],[f1579,f646])).
fof(f6071,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl3),
  inference(resolution,[],[f1579,f644])).
fof(f6068,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | $spl3),
  inference(resolution,[],[f1579,f938])).
fof(f6066,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl3),
  inference(resolution,[],[f1579,f932])).
fof(f6027,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl32),
  inference(resolution,[],[f1297,f976])).
fof(f5936,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X23) != s(bool,X24)) ) | $spl32),
  inference(resolution,[],[f1297,f650])).
fof(f5935,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl32),
  inference(resolution,[],[f1297,f650])).
fof(f5933,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl32),
  inference(resolution,[],[f1297,f649])).
fof(f5932,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl32),
  inference(resolution,[],[f1297,f648])).
fof(f5931,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X8) != s(bool,X9)) ) | $spl32),
  inference(resolution,[],[f1297,f646])).
fof(f5930,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl32),
  inference(resolution,[],[f1297,f646])).
fof(f5929,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl32),
  inference(resolution,[],[f1297,f644])).
fof(f5925,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | $spl32),
  inference(resolution,[],[f1297,f901])).
fof(f5859,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl4),
  inference(resolution,[],[f988,f976])).
fof(f5768,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X23) != s(bool,X24)) ) | $spl4),
  inference(resolution,[],[f988,f650])).
fof(f5767,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl4),
  inference(resolution,[],[f988,f650])).
fof(f5765,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl4),
  inference(resolution,[],[f988,f649])).
fof(f5764,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl4),
  inference(resolution,[],[f988,f648])).
fof(f5763,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X8) != s(bool,X9)) ) | $spl4),
  inference(resolution,[],[f988,f646])).
fof(f5762,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl4),
  inference(resolution,[],[f988,f646])).
fof(f5761,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl4),
  inference(resolution,[],[f988,f644])).
fof(f5760,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))) | $spl4),
  inference(resolution,[],[f988,f909])).
fof(f5759,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | $spl4),
  inference(resolution,[],[f988,f904])).
fof(f5894,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | $spl4),
  inference(forward_demodulation,[],[f5893,f913])).
fof(f5893,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl4),
  inference(forward_demodulation,[],[f5757,f961])).
fof(f5757,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl4),
  inference(resolution,[],[f988,f901])).
fof(f5682,plain,(
  ( ! [X187,X188] : (p(s(bool,X187)) | ~p(s(bool,X188)) | s(bool,X188) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl9),
  inference(resolution,[],[f1123,f976])).
fof(f5587,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X22) != s(bool,X24)) ) | $spl9),
  inference(resolution,[],[f1123,f650])).
fof(f5585,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl9),
  inference(resolution,[],[f1123,f649])).
fof(f5584,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X13) != s(bool,X14)) ) | $spl9),
  inference(resolution,[],[f1123,f648])).
fof(f5583,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl9),
  inference(resolution,[],[f1123,f648])).
fof(f5582,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X7) != s(bool,X9)) ) | $spl9),
  inference(resolution,[],[f1123,f646])).
fof(f5581,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X4) != s(bool,X5)) ) | $spl9),
  inference(resolution,[],[f1123,f644])).
fof(f5580,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl9),
  inference(resolution,[],[f1123,f644])).
fof(f5578,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl9),
  inference(resolution,[],[f1123,f933])).
fof(f5576,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) | $spl9),
  inference(resolution,[],[f1123,f927])).
fof(f5523,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl3),
  inference(resolution,[],[f987,f976])).
fof(f5428,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X21) != s(bool,X23)) ) | $spl3),
  inference(resolution,[],[f987,f650])).
fof(f5426,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl3),
  inference(resolution,[],[f987,f649])).
fof(f5425,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X12) != s(bool,X13)) ) | $spl3),
  inference(resolution,[],[f987,f648])).
fof(f5424,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl3),
  inference(resolution,[],[f987,f648])).
fof(f5423,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X6) != s(bool,X8)) ) | $spl3),
  inference(resolution,[],[f987,f646])).
fof(f5422,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X3) != s(bool,X4)) ) | $spl3),
  inference(resolution,[],[f987,f644])).
fof(f5421,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl3),
  inference(resolution,[],[f987,f644])).
fof(f5415,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl3),
  inference(resolution,[],[f987,f896])).
fof(f5371,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | $spl10),
  inference(forward_demodulation,[],[f5370,f1126])).
fof(f5370,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl10),
  inference(forward_demodulation,[],[f5369,f1126])).
fof(f5369,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl10),
  inference(forward_demodulation,[],[f5063,f959])).
fof(f5063,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl10),
  inference(superposition,[],[f956,f1126])).
fof(f5362,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl10),
  inference(forward_demodulation,[],[f5361,f1126])).
fof(f5361,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl10),
  inference(forward_demodulation,[],[f5360,f1126])).
fof(f5360,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl10),
  inference(forward_demodulation,[],[f5062,f959])).
fof(f5062,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl10),
  inference(superposition,[],[f956,f1126])).
fof(f5273,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | $spl10),
  inference(forward_demodulation,[],[f5272,f1126])).
fof(f5272,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.integer.int',sK14) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | $spl10),
  inference(forward_demodulation,[],[f4983,f922])).
fof(f4983,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | $spl10),
  inference(superposition,[],[f917,f1126])).
fof(f5271,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | $spl10),
  inference(forward_demodulation,[],[f5270,f1126])).
fof(f5270,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))) | $spl10),
  inference(forward_demodulation,[],[f4982,f922])).
fof(f4982,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))) | $spl10),
  inference(superposition,[],[f917,f1126])).
fof(f5269,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | s('type.integer.int',sK14) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | $spl10),
  inference(forward_demodulation,[],[f5268,f1126])).
fof(f5268,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.integer.int',sK14) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | $spl10),
  inference(forward_demodulation,[],[f4981,f922])).
fof(f4981,plain,(
  s('type.integer.int',sK14) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | $spl10),
  inference(superposition,[],[f916,f1126])).
fof(f5265,plain,(
  s('type.integer.int',sK14) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | $spl10),
  inference(forward_demodulation,[],[f5264,f1126])).
fof(f5264,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | s('type.integer.int',sK15) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))) | $spl10),
  inference(forward_demodulation,[],[f4980,f922])).
fof(f4980,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | s('type.integer.int',sK15) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))) | $spl10),
  inference(superposition,[],[f916,f1126])).
fof(f5261,plain,(
  ( ! [X101] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X101),s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X101))))) ) | $spl10),
  inference(forward_demodulation,[],[f5260,f1126])).
fof(f5260,plain,(
  ( ! [X101] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X101),s('type.integer.int',sK15)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X101))))) ) | $spl10),
  inference(forward_demodulation,[],[f4973,f922])).
fof(f4973,plain,(
  ( ! [X101] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X101),s('type.integer.int',sK15)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X101))))) ) | $spl10),
  inference(superposition,[],[f910,f1126])).
fof(f5259,plain,(
  ( ! [X100] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',X100)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X100))))) ) | $spl10),
  inference(forward_demodulation,[],[f5258,f1126])).
fof(f5258,plain,(
  ( ! [X100] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK15),s('type.integer.int',X100)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X100))))) ) | $spl10),
  inference(forward_demodulation,[],[f4972,f922])).
fof(f4972,plain,(
  ( ! [X100] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK15),s('type.integer.int',X100)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X100))))) ) | $spl10),
  inference(superposition,[],[f910,f1126])).
fof(f5256,plain,(
  ( ! [X99] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X99),s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X99))))) ) | $spl10),
  inference(forward_demodulation,[],[f5255,f1126])).
fof(f5255,plain,(
  ( ! [X99] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X99),s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X99))))) ) | $spl10),
  inference(forward_demodulation,[],[f4971,f922])).
fof(f4971,plain,(
  ( ! [X99] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X99),s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X99))))) ) | $spl10),
  inference(superposition,[],[f910,f1126])).
fof(f5253,plain,(
  ( ! [X98] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',X98)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X98))))) ) | $spl10),
  inference(forward_demodulation,[],[f5252,f1126])).
fof(f5252,plain,(
  ( ! [X98] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',X98)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X98))))) ) | $spl10),
  inference(forward_demodulation,[],[f4970,f922])).
fof(f4970,plain,(
  ( ! [X98] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK14),s('type.integer.int',X98)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X98)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15))))) ) | $spl10),
  inference(superposition,[],[f910,f1126])).
fof(f5141,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | $spl10),
  inference(forward_demodulation,[],[f5140,f1126])).
fof(f5140,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | $spl10),
  inference(forward_demodulation,[],[f5139,f921])).
fof(f5139,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | $spl10),
  inference(forward_demodulation,[],[f5138,f919])).
fof(f5138,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl10),
  inference(forward_demodulation,[],[f4879,f922])).
fof(f4879,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | $spl10),
  inference(superposition,[],[f885,f1126])).
fof(f5137,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl10),
  inference(forward_demodulation,[],[f5136,f1126])).
fof(f5136,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl10),
  inference(forward_demodulation,[],[f5135,f922])).
fof(f5135,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | $spl10),
  inference(forward_demodulation,[],[f5134,f921])).
fof(f5134,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | $spl10),
  inference(forward_demodulation,[],[f4878,f919])).
fof(f4878,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | $spl10),
  inference(superposition,[],[f885,f1126])).
fof(f5131,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK44(sK15)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | $spl10),
  inference(forward_demodulation,[],[f5130,f1126])).
fof(f5130,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK44(sK15)))) | $spl10),
  inference(forward_demodulation,[],[f4877,f922])).
fof(f4877,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK44(sK15)))) | $spl10),
  inference(superposition,[],[f882,f1126])).
fof(f5125,plain,(
  s('type.integer.int',sK14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK43(sK15)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl10),
  inference(forward_demodulation,[],[f5124,f1126])).
fof(f5124,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK43(sK15)))))) | $spl10),
  inference(forward_demodulation,[],[f4876,f921])).
fof(f4876,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',sK15) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK43(sK15)))))) | $spl10),
  inference(superposition,[],[f881,f1126])).
fof(f4433,plain,(
  s('type.integer.int',sK14) = s('type.integer.int',sK15) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | $spl2),
  inference(resolution,[],[f991,f904])).
fof(f4434,plain,(
  s('type.integer.int',sK14) = s('type.integer.int',sK15) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | $spl2),
  inference(resolution,[],[f991,f909])).
fof(f4819,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl8),
  inference(resolution,[],[f1124,f976])).
fof(f4728,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X23) != s(bool,X24)) ) | $spl8),
  inference(resolution,[],[f1124,f650])).
fof(f4727,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl8),
  inference(resolution,[],[f1124,f650])).
fof(f4725,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl8),
  inference(resolution,[],[f1124,f649])).
fof(f4724,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl8),
  inference(resolution,[],[f1124,f648])).
fof(f4723,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X8) != s(bool,X9)) ) | $spl8),
  inference(resolution,[],[f1124,f646])).
fof(f4722,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl8),
  inference(resolution,[],[f1124,f646])).
fof(f4721,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl8),
  inference(resolution,[],[f1124,f644])).
fof(f4718,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | $spl8),
  inference(resolution,[],[f1124,f938])).
fof(f4716,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | $spl8),
  inference(resolution,[],[f1124,f932])).
fof(f4541,plain,(
  ( ! [X187,X188] : (p(s(bool,X187)) | ~p(s(bool,X188)) | s(bool,X188) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl2),
  inference(resolution,[],[f991,f976])).
fof(f4446,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X22) != s(bool,X24)) ) | $spl2),
  inference(resolution,[],[f991,f650])).
fof(f4444,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl2),
  inference(resolution,[],[f991,f649])).
fof(f4443,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X13) != s(bool,X14)) ) | $spl2),
  inference(resolution,[],[f991,f648])).
fof(f4442,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl2),
  inference(resolution,[],[f991,f648])).
fof(f4441,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X7) != s(bool,X9)) ) | $spl2),
  inference(resolution,[],[f991,f646])).
fof(f4440,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X4) != s(bool,X5)) ) | $spl2),
  inference(resolution,[],[f991,f644])).
fof(f4439,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl2),
  inference(resolution,[],[f991,f644])).
fof(f4437,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl2),
  inference(resolution,[],[f991,f933])).
fof(f4435,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))),s('type.integer.int',sK14)))) | $spl2),
  inference(resolution,[],[f991,f927])).
fof(f4357,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f989,f976])).
fof(f4262,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X21) != s(bool,X23)) ) | $spl5),
  inference(resolution,[],[f989,f650])).
fof(f4260,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f989,f649])).
fof(f4259,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X12) != s(bool,X13)) ) | $spl5),
  inference(resolution,[],[f989,f648])).
fof(f4258,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f989,f648])).
fof(f4257,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X6) != s(bool,X8)) ) | $spl5),
  inference(resolution,[],[f989,f646])).
fof(f4256,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X3) != s(bool,X4)) ) | $spl5),
  inference(resolution,[],[f989,f644])).
fof(f4255,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f989,f644])).
fof(f4188,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f976])).
fof(f4097,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X23) != s(bool,X24)) ) | $spl2),
  inference(resolution,[],[f986,f650])).
fof(f4096,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f650])).
fof(f4094,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f649])).
fof(f4093,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f648])).
fof(f4092,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X8) != s(bool,X9)) ) | $spl2),
  inference(resolution,[],[f986,f646])).
fof(f4091,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f646])).
fof(f4090,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f644])).
fof(f4089,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | s('type.integer.int',sK14) = s('type.integer.int',sK15) | $spl2),
  inference(resolution,[],[f986,f909])).
fof(f4088,plain,(
  s('type.integer.int',sK14) = s('type.integer.int',sK15) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | $spl2),
  inference(resolution,[],[f986,f904])).
fof(f4086,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | $spl2),
  inference(resolution,[],[f986,f901])).
fof(f4043,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl32),
  inference(resolution,[],[f1297,f976])).
fof(f3952,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X23) != s(bool,X24)) ) | $spl32),
  inference(resolution,[],[f1297,f650])).
fof(f3951,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl32),
  inference(resolution,[],[f1297,f650])).
fof(f3949,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl32),
  inference(resolution,[],[f1297,f649])).
fof(f3948,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl32),
  inference(resolution,[],[f1297,f648])).
fof(f3947,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X8) != s(bool,X9)) ) | $spl32),
  inference(resolution,[],[f1297,f646])).
fof(f3946,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl32),
  inference(resolution,[],[f1297,f646])).
fof(f3945,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl32),
  inference(resolution,[],[f1297,f644])).
fof(f3941,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | $spl32),
  inference(resolution,[],[f1297,f901])).
fof(f3867,plain,(
  ( ! [X187,X188] : (p(s(bool,X187)) | ~p(s(bool,X188)) | s(bool,X188) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl9),
  inference(resolution,[],[f1123,f976])).
fof(f3772,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X22) != s(bool,X24)) ) | $spl9),
  inference(resolution,[],[f1123,f650])).
fof(f3770,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl9),
  inference(resolution,[],[f1123,f649])).
fof(f3769,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X13) != s(bool,X14)) ) | $spl9),
  inference(resolution,[],[f1123,f648])).
fof(f3768,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl9),
  inference(resolution,[],[f1123,f648])).
fof(f3767,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X7) != s(bool,X9)) ) | $spl9),
  inference(resolution,[],[f1123,f646])).
fof(f3766,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X4) != s(bool,X5)) ) | $spl9),
  inference(resolution,[],[f1123,f644])).
fof(f3765,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl9),
  inference(resolution,[],[f1123,f644])).
fof(f3763,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl9),
  inference(resolution,[],[f1123,f933])).
fof(f3761,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) | $spl9),
  inference(resolution,[],[f1123,f927])).
fof(f3720,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl3),
  inference(resolution,[],[f987,f976])).
fof(f3625,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X21) != s(bool,X23)) ) | $spl3),
  inference(resolution,[],[f987,f650])).
fof(f3623,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl3),
  inference(resolution,[],[f987,f649])).
fof(f3622,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X12) != s(bool,X13)) ) | $spl3),
  inference(resolution,[],[f987,f648])).
fof(f3621,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl3),
  inference(resolution,[],[f987,f648])).
fof(f3620,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X6) != s(bool,X8)) ) | $spl3),
  inference(resolution,[],[f987,f646])).
fof(f3619,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X3) != s(bool,X4)) ) | $spl3),
  inference(resolution,[],[f987,f644])).
fof(f3618,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl3),
  inference(resolution,[],[f987,f644])).
fof(f3612,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl3),
  inference(resolution,[],[f987,f896])).
fof(f3555,plain,(
  ( ! [X187,X188] : (p(s(bool,X187)) | ~p(s(bool,X188)) | s(bool,X188) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl2),
  inference(resolution,[],[f991,f976])).
fof(f3460,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X22) != s(bool,X24)) ) | $spl2),
  inference(resolution,[],[f991,f650])).
fof(f3458,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl2),
  inference(resolution,[],[f991,f649])).
fof(f3457,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X13) != s(bool,X14)) ) | $spl2),
  inference(resolution,[],[f991,f648])).
fof(f3456,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl2),
  inference(resolution,[],[f991,f648])).
fof(f3455,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X7) != s(bool,X9)) ) | $spl2),
  inference(resolution,[],[f991,f646])).
fof(f3454,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X4) != s(bool,X5)) ) | $spl2),
  inference(resolution,[],[f991,f644])).
fof(f3453,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl2),
  inference(resolution,[],[f991,f644])).
fof(f3451,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl2),
  inference(resolution,[],[f991,f933])).
fof(f3449,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15))))),s('type.integer.int',sK14)))) | $spl2),
  inference(resolution,[],[f991,f927])).
fof(f3415,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f3414,f968])).
fof(f3414,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f3413,f961])).
fof(f3413,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) | $spl1),
  inference(forward_demodulation,[],[f3365,f922])).
fof(f3365,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) | $spl1),
  inference(superposition,[],[f985,f966])).
fof(f3412,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(subsumption_resolution,[],[f3411,f965])).
fof(f3411,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f3410,f959])).
fof(f3410,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(forward_demodulation,[],[f3409,f958])).
fof(f3409,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(forward_demodulation,[],[f3408,f959])).
fof(f3408,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(forward_demodulation,[],[f3364,f922])).
fof(f3364,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(superposition,[],[f985,f956])).
fof(f3407,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl1),
  inference(forward_demodulation,[],[f3363,f922])).
fof(f3363,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl1),
  inference(superposition,[],[f985,f955])).
fof(f3406,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(forward_demodulation,[],[f3362,f922])).
fof(f3362,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(superposition,[],[f985,f954])).
fof(f3405,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl1),
  inference(forward_demodulation,[],[f3361,f922])).
fof(f3361,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl1),
  inference(superposition,[],[f985,f953])).
fof(f3404,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(forward_demodulation,[],[f3360,f922])).
fof(f3360,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(superposition,[],[f985,f952])).
fof(f3403,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f3402,f962])).
fof(f3402,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) | $spl1),
  inference(forward_demodulation,[],[f3359,f922])).
fof(f3359,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) | $spl1),
  inference(superposition,[],[f985,f950])).
fof(f3401,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f3400,f962])).
fof(f3400,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) | $spl1),
  inference(forward_demodulation,[],[f3358,f922])).
fof(f3358,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) | $spl1),
  inference(superposition,[],[f985,f948])).
fof(f3397,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f3396,f968])).
fof(f3396,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f3395,f961])).
fof(f3395,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl1),
  inference(forward_demodulation,[],[f3355,f922])).
fof(f3355,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl1),
  inference(superposition,[],[f985,f914])).
fof(f3394,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl1),
  inference(forward_demodulation,[],[f3348,f922])).
fof(f3348,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl1),
  inference(superposition,[],[f985,f949])).
fof(f3347,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl1),
  inference(superposition,[],[f985,f951])).
fof(f3393,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f3342,f922])).
fof(f3342,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | $spl1),
  inference(superposition,[],[f985,f917])).
fof(f3334,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) ) | $spl1),
  inference(resolution,[],[f985,f976])).
fof(f3239,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15))) | s(bool,X21) != s(bool,X23)) ) | $spl1),
  inference(resolution,[],[f985,f650])).
fof(f3237,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) ) | $spl1),
  inference(resolution,[],[f985,f649])).
fof(f3236,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15))) | s(bool,X12) != s(bool,X13)) ) | $spl1),
  inference(resolution,[],[f985,f648])).
fof(f3235,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) ) | $spl1),
  inference(resolution,[],[f985,f648])).
fof(f3234,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15))) | s(bool,X6) != s(bool,X8)) ) | $spl1),
  inference(resolution,[],[f985,f646])).
fof(f3233,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15))) | s(bool,X3) != s(bool,X4)) ) | $spl1),
  inference(resolution,[],[f985,f644])).
fof(f3232,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) ) | $spl1),
  inference(resolution,[],[f985,f644])).
fof(f3385,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl1),
  inference(forward_demodulation,[],[f3226,f962])).
fof(f3226,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl1),
  inference(resolution,[],[f985,f896])).
fof(f3165,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f976])).
fof(f3074,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X23) != s(bool,X24)) ) | $spl2),
  inference(resolution,[],[f986,f650])).
fof(f3073,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f650])).
fof(f3071,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f649])).
fof(f3070,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f648])).
fof(f3069,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X8) != s(bool,X9)) ) | $spl2),
  inference(resolution,[],[f986,f646])).
fof(f3068,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f646])).
fof(f3067,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f644])).
fof(f3066,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | s('type.integer.int',sK14) = s('type.integer.int',sK15) | $spl2),
  inference(resolution,[],[f986,f909])).
fof(f3065,plain,(
  s('type.integer.int',sK14) = s('type.integer.int',sK15) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | $spl2),
  inference(resolution,[],[f986,f904])).
fof(f2908,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl33),
  inference(resolution,[],[f1298,f896])).
fof(f2914,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl33),
  inference(resolution,[],[f1298,f644])).
fof(f2915,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X3) != s(bool,X4)) ) | $spl33),
  inference(resolution,[],[f1298,f644])).
fof(f2916,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X6) != s(bool,X8)) ) | $spl33),
  inference(resolution,[],[f1298,f646])).
fof(f2917,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl33),
  inference(resolution,[],[f1298,f648])).
fof(f2918,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X12) != s(bool,X13)) ) | $spl33),
  inference(resolution,[],[f1298,f648])).
fof(f2919,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl33),
  inference(resolution,[],[f1298,f649])).
fof(f2921,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14))) | s(bool,X21) != s(bool,X23)) ) | $spl33),
  inference(resolution,[],[f1298,f650])).
fof(f3016,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) ) | $spl33),
  inference(resolution,[],[f1298,f976])).
fof(f1125,plain,(
  s('type.integer.int',sK14) != s('type.integer.int',sK15) | $spl11),
  inference(cnf_transformation,[],[f1125_D])).
fof(f1125_D,plain,(
  s('type.integer.int',sK14) != s('type.integer.int',sK15) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f2897,plain,(
  s('type.integer.int',sK14) != s('type.integer.int',sK15) | s('type.integer.int',sK15) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl5),
  inference(superposition,[],[f1146,f951])).
fof(f3045,plain,(
  $false | ($spl9 | $spl33)),
  inference(subsumption_resolution,[],[f2910,f1123])).
fof(f2910,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | $spl33),
  inference(resolution,[],[f1298,f903])).
fof(f3044,plain,(
  $false | ($spl9 | $spl33)),
  inference(subsumption_resolution,[],[f2907,f1123])).
fof(f2907,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | $spl33),
  inference(resolution,[],[f1298,f894])).
fof(f3043,plain,(
  $false | ($spl9 | $spl33)),
  inference(subsumption_resolution,[],[f2906,f1123])).
fof(f2906,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | $spl33),
  inference(resolution,[],[f1298,f892])).
fof(f2858,plain,(
  ( ! [X187,X188] : (p(s(bool,X187)) | ~p(s(bool,X188)) | s(bool,X188) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl9),
  inference(resolution,[],[f1123,f976])).
fof(f2763,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X22) != s(bool,X24)) ) | $spl9),
  inference(resolution,[],[f1123,f650])).
fof(f2761,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl9),
  inference(resolution,[],[f1123,f649])).
fof(f2760,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X13) != s(bool,X14)) ) | $spl9),
  inference(resolution,[],[f1123,f648])).
fof(f2759,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl9),
  inference(resolution,[],[f1123,f648])).
fof(f2758,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X7) != s(bool,X9)) ) | $spl9),
  inference(resolution,[],[f1123,f646])).
fof(f2757,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X4) != s(bool,X5)) ) | $spl9),
  inference(resolution,[],[f1123,f644])).
fof(f2756,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl9),
  inference(resolution,[],[f1123,f644])).
fof(f2754,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl9),
  inference(resolution,[],[f1123,f933])).
fof(f2752,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) | $spl9),
  inference(resolution,[],[f1123,f927])).
fof(f2711,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl3),
  inference(resolution,[],[f987,f976])).
fof(f2616,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X21) != s(bool,X23)) ) | $spl3),
  inference(resolution,[],[f987,f650])).
fof(f2614,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl3),
  inference(resolution,[],[f987,f649])).
fof(f2613,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X12) != s(bool,X13)) ) | $spl3),
  inference(resolution,[],[f987,f648])).
fof(f2612,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl3),
  inference(resolution,[],[f987,f648])).
fof(f2611,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X6) != s(bool,X8)) ) | $spl3),
  inference(resolution,[],[f987,f646])).
fof(f2610,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X3) != s(bool,X4)) ) | $spl3),
  inference(resolution,[],[f987,f644])).
fof(f2609,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl3),
  inference(resolution,[],[f987,f644])).
fof(f2603,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl3),
  inference(resolution,[],[f987,f896])).
fof(f2596,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl83),
  inference(cnf_transformation,[],[f2596_D])).
fof(f2596_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) <=> ~$spl83),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl83])])).
fof(f2345,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl5),
  inference(forward_demodulation,[],[f2201,f962])).
fof(f2201,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15)))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl5),
  inference(resolution,[],[f989,f896])).
fof(f2589,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f2588,f968])).
fof(f2588,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f2587,f961])).
fof(f2587,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) | $spl1),
  inference(forward_demodulation,[],[f2511,f922])).
fof(f2511,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) | $spl1),
  inference(superposition,[],[f985,f966])).
fof(f2584,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(subsumption_resolution,[],[f2583,f965])).
fof(f2583,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f2582,f959])).
fof(f2582,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(forward_demodulation,[],[f2581,f958])).
fof(f2581,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(forward_demodulation,[],[f2580,f959])).
fof(f2580,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(forward_demodulation,[],[f2510,f922])).
fof(f2510,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(superposition,[],[f985,f956])).
fof(f2579,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl1),
  inference(forward_demodulation,[],[f2509,f922])).
fof(f2509,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl1),
  inference(superposition,[],[f985,f955])).
fof(f2576,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(forward_demodulation,[],[f2508,f922])).
fof(f2508,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(superposition,[],[f985,f954])).
fof(f2573,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl1),
  inference(forward_demodulation,[],[f2507,f922])).
fof(f2507,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | $spl1),
  inference(superposition,[],[f985,f953])).
fof(f2568,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(forward_demodulation,[],[f2506,f922])).
fof(f2506,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK46('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK45('const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))) | $spl1),
  inference(superposition,[],[f985,f952])).
fof(f2565,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f2564,f962])).
fof(f2564,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) | $spl1),
  inference(forward_demodulation,[],[f2505,f922])).
fof(f2505,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) | $spl1),
  inference(superposition,[],[f985,f950])).
fof(f2563,plain,(
  s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f2562,f962])).
fof(f2562,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) | $spl1),
  inference(forward_demodulation,[],[f2504,f922])).
fof(f2504,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))) | $spl1),
  inference(superposition,[],[f985,f948])).
fof(f2557,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f2556,f968])).
fof(f2556,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f2555,f961])).
fof(f2555,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl1),
  inference(forward_demodulation,[],[f2501,f922])).
fof(f2501,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl1),
  inference(superposition,[],[f985,f914])).
fof(f2552,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl1),
  inference(forward_demodulation,[],[f2494,f922])).
fof(f2494,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl1),
  inference(superposition,[],[f985,f949])).
fof(f2493,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl1),
  inference(superposition,[],[f985,f951])).
fof(f2545,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | $spl1),
  inference(forward_demodulation,[],[f2488,f922])).
fof(f2488,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK14)))) | $spl1),
  inference(superposition,[],[f985,f917])).
fof(f2480,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) ) | $spl1),
  inference(resolution,[],[f985,f976])).
fof(f2385,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15))) | s(bool,X21) != s(bool,X23)) ) | $spl1),
  inference(resolution,[],[f985,f650])).
fof(f2383,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) ) | $spl1),
  inference(resolution,[],[f985,f649])).
fof(f2382,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15))) | s(bool,X12) != s(bool,X13)) ) | $spl1),
  inference(resolution,[],[f985,f648])).
fof(f2381,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) ) | $spl1),
  inference(resolution,[],[f985,f648])).
fof(f2380,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15))) | s(bool,X6) != s(bool,X8)) ) | $spl1),
  inference(resolution,[],[f985,f646])).
fof(f2379,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15))) | s(bool,X3) != s(bool,X4)) ) | $spl1),
  inference(resolution,[],[f985,f644])).
fof(f2378,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) ) | $spl1),
  inference(resolution,[],[f985,f644])).
fof(f2531,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl1),
  inference(forward_demodulation,[],[f2372,f962])).
fof(f2372,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14)))))))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl1),
  inference(resolution,[],[f985,f896])).
fof(f2317,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | s('type.integer.int',sK15) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl5),
  inference(superposition,[],[f989,f951])).
fof(f2309,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f989,f976])).
fof(f2214,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X21) != s(bool,X23)) ) | $spl5),
  inference(resolution,[],[f989,f650])).
fof(f2212,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f989,f649])).
fof(f2211,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X12) != s(bool,X13)) ) | $spl5),
  inference(resolution,[],[f989,f648])).
fof(f2210,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f989,f648])).
fof(f2209,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X6) != s(bool,X8)) ) | $spl5),
  inference(resolution,[],[f989,f646])).
fof(f2208,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X3) != s(bool,X4)) ) | $spl5),
  inference(resolution,[],[f989,f644])).
fof(f2207,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f989,f644])).
fof(f2140,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f976])).
fof(f2049,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X23) != s(bool,X24)) ) | $spl2),
  inference(resolution,[],[f986,f650])).
fof(f2048,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f650])).
fof(f2046,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f649])).
fof(f2045,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f648])).
fof(f2044,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X8) != s(bool,X9)) ) | $spl2),
  inference(resolution,[],[f986,f646])).
fof(f2043,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f646])).
fof(f2042,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f644])).
fof(f2041,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | s('type.integer.int',sK14) = s('type.integer.int',sK15) | $spl2),
  inference(resolution,[],[f986,f909])).
fof(f2040,plain,(
  s('type.integer.int',sK14) = s('type.integer.int',sK15) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | $spl2),
  inference(resolution,[],[f986,f904])).
fof(f2038,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | $spl2),
  inference(resolution,[],[f986,f901])).
fof(f1988,plain,(
  ( ! [X187,X188] : (p(s(bool,X187)) | ~p(s(bool,X188)) | s(bool,X188) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl9),
  inference(resolution,[],[f1123,f976])).
fof(f1893,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X22) != s(bool,X24)) ) | $spl9),
  inference(resolution,[],[f1123,f650])).
fof(f1891,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl9),
  inference(resolution,[],[f1123,f649])).
fof(f1890,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X13) != s(bool,X14)) ) | $spl9),
  inference(resolution,[],[f1123,f648])).
fof(f1889,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl9),
  inference(resolution,[],[f1123,f648])).
fof(f1888,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X7) != s(bool,X9)) ) | $spl9),
  inference(resolution,[],[f1123,f646])).
fof(f1887,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X4) != s(bool,X5)) ) | $spl9),
  inference(resolution,[],[f1123,f644])).
fof(f1886,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl9),
  inference(resolution,[],[f1123,f644])).
fof(f1884,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl9),
  inference(resolution,[],[f1123,f933])).
fof(f1882,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) | $spl9),
  inference(resolution,[],[f1123,f927])).
fof(f1870,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f1738,f987])).
fof(f1738,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | $spl8),
  inference(resolution,[],[f1124,f907])).
fof(f1869,plain,(
  $false | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f1737,f987])).
fof(f1737,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | $spl8),
  inference(resolution,[],[f1124,f905])).
fof(f1689,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl3),
  inference(resolution,[],[f987,f976])).
fof(f1594,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X21) != s(bool,X23)) ) | $spl3),
  inference(resolution,[],[f987,f650])).
fof(f1592,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl3),
  inference(resolution,[],[f987,f649])).
fof(f1591,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X12) != s(bool,X13)) ) | $spl3),
  inference(resolution,[],[f987,f648])).
fof(f1590,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl3),
  inference(resolution,[],[f987,f648])).
fof(f1589,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X6) != s(bool,X8)) ) | $spl3),
  inference(resolution,[],[f987,f646])).
fof(f1588,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X3) != s(bool,X4)) ) | $spl3),
  inference(resolution,[],[f987,f644])).
fof(f1587,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl3),
  inference(resolution,[],[f987,f644])).
fof(f1581,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK14)))))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK15)))))) | $spl3),
  inference(resolution,[],[f987,f896])).
fof(f1571,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f1448,f1133])).
fof(f1448,plain,(
  ( ! [X360,X361] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X360),s('type.integer.int',X361))))) ) | $spl24),
  inference(resolution,[],[f1133,f975])).
fof(f1570,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f1447,f1133])).
fof(f1447,plain,(
  ( ! [X358,X359] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X358),s('type.integer.int',X359))))) ) | $spl24),
  inference(resolution,[],[f1133,f975])).
fof(f1569,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f1439,f1133])).
fof(f1439,plain,(
  ( ! [X343,X342] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X342),s('type.integer.int',X343))))) ) | $spl24),
  inference(resolution,[],[f1133,f902])).
fof(f1568,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f1438,f1133])).
fof(f1438,plain,(
  ( ! [X341,X340] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X340),s('type.integer.int',X341))))) ) | $spl24),
  inference(resolution,[],[f1133,f895])).
fof(f1567,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f1437,f1133])).
fof(f1437,plain,(
  ( ! [X339,X338] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X338),s('type.integer.int',X339))))) ) | $spl24),
  inference(resolution,[],[f1133,f893])).
fof(f1566,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f1421,f1133])).
fof(f1421,plain,(
  ( ! [X312,X311] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X311),s('type.integer.int',X312))))) ) | $spl24),
  inference(resolution,[],[f1133,f902])).
fof(f1565,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f1418,f1133])).
fof(f1418,plain,(
  ( ! [X306,X305] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X305),s('type.integer.int',X306))))) ) | $spl24),
  inference(resolution,[],[f1133,f895])).
fof(f1564,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f1417,f1133])).
fof(f1417,plain,(
  ( ! [X304,X303] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X303),s('type.integer.int',X304))))) ) | $spl24),
  inference(resolution,[],[f1133,f893])).
fof(f1563,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f1413,f1133])).
fof(f1413,plain,(
  ( ! [X298] : (~p(s(bool,X298))) ) | $spl24),
  inference(resolution,[],[f1133,f880])).
fof(f1562,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f1407,f1133])).
fof(f1407,plain,(
  ( ! [X292] : (~p(s(bool,X292))) ) | $spl24),
  inference(resolution,[],[f1133,f846])).
fof(f846,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f590])).
fof(f590,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f589])).
fof(f589,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f441])).
fof(f441,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f320])).
fof(f320,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f319])).
fof(f319,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f163])).
fof(f163,axiom,(
  ! [X12] : ((p(s(bool,X12)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X12)))),
  file('thm.integer.INT_ABS_LE|split|2','thm.bool.IMP_CLAUSES|split|4')).
fof(f1415,plain,(
  $false | $spl24),
  inference(resolution,[],[f1133,f889])).
fof(f1431,plain,(
  $false | ($spl5 | $spl24)),
  inference(resolution,[],[f1133,f989])).
fof(f1460,plain,(
  $false | $spl24),
  inference(resolution,[],[f1133,f965])).
fof(f1258,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK15),s('type.integer.int',sK14)))) | s('type.integer.int',sK15) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl5),
  inference(superposition,[],[f989,f951])).
fof(f1250,plain,(
  ( ! [X187,X186] : (p(s(bool,X186)) | ~p(s(bool,X187)) | s(bool,X187) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f989,f976])).
fof(f1155,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X21) != s(bool,X23)) ) | $spl5),
  inference(resolution,[],[f989,f650])).
fof(f1153,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f989,f649])).
fof(f1152,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X12) != s(bool,X13)) ) | $spl5),
  inference(resolution,[],[f989,f648])).
fof(f1151,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f989,f648])).
fof(f1150,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X6) != s(bool,X8)) ) | $spl5),
  inference(resolution,[],[f989,f646])).
fof(f1149,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14))) | s(bool,X3) != s(bool,X4)) ) | $spl5),
  inference(resolution,[],[f989,f644])).
fof(f1148,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) ) | $spl5),
  inference(resolution,[],[f989,f644])).
fof(f1096,plain,(
  ( ! [X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X186) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f976])).
fof(f1005,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X23) != s(bool,X24)) ) | $spl2),
  inference(resolution,[],[f986,f650])).
fof(f1004,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,X20) | s(bool,X21) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f650])).
fof(f1002,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f649])).
fof(f1001,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f648])).
fof(f1000,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))) | s(bool,X8) != s(bool,X9)) ) | $spl2),
  inference(resolution,[],[f986,f646])).
fof(f999,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f646])).
fof(f998,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) ) | $spl2),
  inference(resolution,[],[f986,f644])).
fof(f997,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | s('type.integer.int',sK14) = s('type.integer.int',sK15) | $spl2),
  inference(resolution,[],[f986,f909])).
fof(f996,plain,(
  s('type.integer.int',sK14) = s('type.integer.int',sK15) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | $spl2),
  inference(resolution,[],[f986,f904])).
fof(f994,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | $spl2),
  inference(resolution,[],[f986,f901])).
fof(f984,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) | $spl0),
  inference(cnf_transformation,[],[f984_D])).
fof(f984_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f988,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | $spl4),
  inference(cnf_transformation,[],[f988_D])).
fof(f988_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f640,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15))))),
  inference(cnf_transformation,[],[f483])).
fof(f483,plain,(
  (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))))) & (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15)))) | (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) & ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14,sK15])],[f482])).
fof(f482,plain,(
  ? [X0,X1] : ((p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) & (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))),s('type.integer.int',X1)))) | (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) & ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))))),
  inference(flattening,[],[f481])).
fof(f481,plain,(
  ? [X0,X1] : ((p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))),s('type.integer.int',X1)))) | (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))) & (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))),s('type.integer.int',X1)))) | (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) & ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))))),
  inference(nnf_transformation,[],[f414])).
fof(f414,plain,(
  ? [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))),s('type.integer.int',X1)))) <~> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(ennf_transformation,[],[f263])).
fof(f263,plain,(
  ~! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X0))))),s('type.integer.int',X1)))) <=> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))))),s('type.integer.int',X7)))) <=> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X7))),s('type.integer.int',X6)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X6),s('type.integer.int',X7))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',X6))))),s('type.integer.int',X7)))) <=> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X7))),s('type.integer.int',X6)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X6),s('type.integer.int',X7))))))),
  file('thm.integer.INT_ABS_LE|split|2','thm.integer.INT_ABS_LE|split|2')).
fof(f641,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK15))),s('type.integer.int',sK14)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15))))),
  inference(cnf_transformation,[],[f483])).
fof(f642,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK14),s('type.integer.int',sK15)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.ABS_1'(s('type.integer.int',sK14))))),s('type.integer.int',sK15))))),
  inference(cnf_transformation,[],[f483])).
% SZS output end Proof for thm.integer
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1014
% Passive clauses: 26859
% Generated clauses: 45294
% Final active clauses: 398
% Final passive clauses: 6054
% Input formulas: 261
% Initial clauses: 336
% 
% Selected by SInE selection: 110
% SInE iterations: 1
% 
% Duplicate literals: 652
% Trivial inequalities: 23
% Fw subsumption resolutions: 2539
% Fw demodulations: 3228
% Bw demodulations: 106
% 
% Simple tautologies: 816
% Equational tautologies: 2
% Forward subsumptions: 8658
% 
% Binary resolution: 19941
% Forward superposition: 10579
% Backward superposition: 2481
% Equality resolution: 4
% 
% Split clauses: 1550
% Split components: 3274
% SAT solver clauses: 1636
% SAT solver unit clauses: 24
% SAT solver binary clauses: 58
% 
% Sat splits: 1550
% Sat splitting refutations: 85
% 
% Memory used [KB]: 11513
% Time elapsed: 0.830 s
% ------------------------------
% ------------------------------
% Success in time 15.068 s
