Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f398,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f395])).
fof(f395,plain,(
  s('type.num.num',sK0) != s('type.num.num',sK0)),
  inference(superposition,[],[f44,f388])).
fof(f388,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.pred_set.MAX_SET_1'(s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0')))))) )),
  inference(forward_demodulation,[],[f387,f50])).
fof(f50,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f21])).
fof(f21,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7] : s('type.num.num',X7) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X7),s('type.num.num','const.num.|48|_0')))),
  file('thm.pred_set.MAX_SET_REWRITES|split|1','thm.arithmetic.MAX_|48||split|0')).
fof(f387,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.pred_set.MAX_SET_1'(s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0')))))) )),
  inference(forward_demodulation,[],[f386,f47])).
fof(f47,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.MAX_SET_1'(s(fun('type.num.num',bool),'const.pred_set.EMPTY_0')))),
  inference(cnf_transformation,[],[f16])).
fof(f16,axiom,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.MAX_SET_1'(s(fun('type.num.num',bool),'const.pred_set.EMPTY_0')))),
  file('thm.pred_set.MAX_SET_REWRITES|split|1','thm.pred_set.MAX_SET_THM|split|0')).
fof(f386,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.MAX_SET_1'(s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))))) = s('type.num.num','const.pred_set.MAX_SET_1'(s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0')))))) )),
  inference(subsumption_resolution,[],[f380,f45])).
fof(f45,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.MAX_SET_REWRITES|split|1',aHL_TRUTH)).
fof(f380,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.MAX_SET_1'(s(fun('type.num.num',bool),'const.pred_set.EMPTY_0'))))) = s('type.num.num','const.pred_set.MAX_SET_1'(s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0')))))) )),
  inference(superposition,[],[f63,f68])).
fof(f68,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(resolution,[],[f54,f49])).
fof(f49,plain,(
  ( ! [X0] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(cnf_transformation,[],[f20])).
fof(f20,plain,(
  ! [X0] : p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X8] : p(s(bool,'const.pred_set.FINITE_1'(s(fun(X8,bool),'const.pred_set.EMPTY_0'))))),
  file('thm.pred_set.MAX_SET_REWRITES|split|1','thm.pred_set.FINITE_EMPTY')).
fof(f54,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f23])).
fof(f23,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.pred_set.MAX_SET_REWRITES|split|1','thm.bool.EQ_CLAUSES|split|0')).
fof(f63,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun('type.num.num',bool),X1)))) | s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.MAX_SET_1'(s(fun('type.num.num',bool),X1))))) = s('type.num.num','const.pred_set.MAX_SET_1'(s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',X0),s(fun('type.num.num',bool),X1)))))) )),
  inference(cnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun('type.num.num',bool),X1)))) | s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.MAX_SET_1'(s(fun('type.num.num',bool),X1))))) = s('type.num.num','const.pred_set.MAX_SET_1'(s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',X0),s(fun('type.num.num',bool),X1))))))),
  inference(ennf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun('type.num.num',bool),X1)))) => s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.MAX_SET_1'(s(fun('type.num.num',bool),X1))))) = s('type.num.num','const.pred_set.MAX_SET_1'(s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',X0),s(fun('type.num.num',bool),X1))))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X6,X11] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun('type.num.num',bool),X11)))) => s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.MAX_SET_1'(s(fun('type.num.num',bool),X11))))) = s('type.num.num','const.pred_set.MAX_SET_1'(s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',X6),s(fun('type.num.num',bool),X11))))))),
  file('thm.pred_set.MAX_SET_REWRITES|split|1','thm.pred_set.MAX_SET_THM|split|1')).
fof(f44,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.pred_set.MAX_SET_1'(s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',sK0),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0')))))),
  inference(cnf_transformation,[],[f37])).
fof(f37,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.pred_set.MAX_SET_1'(s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',sK0),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0')))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f34])).
fof(f34,plain,(
  ? [X0] : s('type.num.num',X0) != s('type.num.num','const.pred_set.MAX_SET_1'(s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0')))))),
  inference(ennf_transformation,[],[f18])).
fof(f18,plain,(
  ~! [X0] : s('type.num.num',X0) = s('type.num.num','const.pred_set.MAX_SET_1'(s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0')))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : s('type.num.num',X6) = s('type.num.num','const.pred_set.MAX_SET_1'(s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',X6),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0')))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : s('type.num.num',X6) = s('type.num.num','const.pred_set.MAX_SET_1'(s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',X6),s(fun('type.num.num',bool),'const.pred_set.EMPTY_0')))))),
  file('thm.pred_set.MAX_SET_REWRITES|split|1','thm.pred_set.MAX_SET_REWRITES|split|1')).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 23
% Passive clauses: 63
% Generated clauses: 333
% Final active clauses: 22
% Final passive clauses: 28
% Input formulas: 16
% Initial clauses: 21
% 
% Duplicate literals: 49
% Trivial inequalities: 11
% Fw subsumption resolutions: 23
% Fw demodulations: 3
% Bw demodulations: 1
% 
% Simple tautologies: 36
% Equational tautologies: 16
% Forward subsumptions: 141
% 
% Binary resolution: 15
% Forward superposition: 77
% Backward superposition: 82
% Self superposition: 25
% Equality factoring: 26
% Equality resolution: 4
% 
% Memory used [KB]: 767
% Time elapsed: 0.017 s
% ------------------------------
% ------------------------------
% Success in time 0.023 s
