Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.pred_set
% (51342)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 119
% Passive clauses: 2229
% Generated clauses: 19543
% Final active clauses: 106
% Final passive clauses: 1858
% Input formulas: 24
% Initial clauses: 48
% 
% Duplicate literals: 901
% Trivial inequalities: 386
% Fw subsumption resolutions: 3371
% Fw demodulations: 2298
% Bw demodulations: 11
% 
% Simple tautologies: 455
% Equational tautologies: 63
% Forward subsumptions: 8316
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 394
% Factoring: 5
% Forward superposition: 5448
% Backward superposition: 6078
% Self superposition: 236
% Equality factoring: 177
% Equality resolution: 29
% 
% Split clauses: 1076
% Split components: 2206
% SAT solver clauses: 1730
% SAT solver unit clauses: 65
% SAT solver binary clauses: 378
% 
% Sat splits: 1076
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 1719
% TWLsolver calls for satisfiability: 141
% 
% Memory used [KB]: 8571
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (51396)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 95
% Passive clauses: 662
% Generated clauses: 12224
% Final active clauses: 80
% Final passive clauses: 499
% Input formulas: 24
% Initial clauses: 48
% 
% Duplicate literals: 901
% Trivial inequalities: 155
% Fw subsumption resolutions: 1334
% Fw demodulations: 998
% Bw demodulations: 74
% Condensations: 1355
% Global subsumptions: 9
% 
% Simple tautologies: 230
% Equational tautologies: 46
% Forward subsumptions: 5410
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 287
% Factoring: 4
% Forward superposition: 3255
% Backward superposition: 2928
% Self superposition: 148
% Equality factoring: 106
% Equality resolution: 19
% 
% Split clauses: 51
% Split components: 111
% SAT solver clauses: 1143
% SAT solver unit clauses: 63
% SAT solver binary clauses: 49
% SAT solver learnt clauses: 18
% 
% Sat splits: 51
% Sat splitting refutations: 9
% 
% TWLsolver clauses: 1107
% TWLsolver calls for satisfiability: 251
% 
% Memory used [KB]: 8571
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f631,plain,(
  $false),
  inference(subsumption_resolution,[],[f630,f135])).
fof(f135,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f90,f87])).
fof(f87,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.CROSS_EMPTY|split|1',aHL_TRUTH)).
fof(f90,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f61])).
fof(f61,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f28])).
fof(f28,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f17])).
fof(f17,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.pred_set.CROSS_EMPTY|split|1','thm.bool.NOT_CLAUSES|split|1')).
fof(f630,plain,(
  p(s(bool,'F_0'))),
  inference(forward_demodulation,[],[f629,f143])).
fof(f143,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f123,f101])).
fof(f101,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f34])).
fof(f34,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X11] : (s(bool,'F_0') = s(bool,X11) <=> ~p(s(bool,X11)))),
  file('thm.pred_set.CROSS_EMPTY|split|1','thm.bool.EQ_CLAUSES|split|3')).
fof(f123,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0] : ! [X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X7] : ! [X12] : ~p(s(bool,'const.bool.IN_2'(s(X7,X12),s(fun(X7,bool),'const.pred_set.EMPTY_0'))))),
  file('thm.pred_set.CROSS_EMPTY|split|1','thm.pred_set.NOT_IN_EMPTY')).
fof(f629,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,'const.pair.FST_1'(s('type.pair.prod'(sK0,sK1),sK3('const.pred_set.EMPTY_0','const.pred_set.CROSS_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),sK2)),'type.pair.prod'(sK0,sK1))))),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))),
  inference(subsumption_resolution,[],[f604,f87])).
fof(f604,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,'const.pair.FST_1'(s('type.pair.prod'(sK0,sK1),sK3('const.pred_set.EMPTY_0','const.pred_set.CROSS_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),sK2)),'type.pair.prod'(sK0,sK1))))),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))),
  inference(superposition,[],[f131,f553])).
fof(f553,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK0,sK1),sK3('const.pred_set.EMPTY_0','const.pred_set.CROSS_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),sK2)),'type.pair.prod'(sK0,sK1))),s(fun('type.pair.prod'(sK0,sK1),bool),'const.pred_set.CROSS_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),sK2)))))),
  inference(unit_resulting_resolution,[],[f135,f545,f177])).
fof(f177,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f174,f135])).
fof(f174,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f126,f94])).
fof(f94,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.CROSS_EMPTY|split|1',aHL_BOOL_CASES)).
fof(f126,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f57])).
fof(f57,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X9,X10] : ((p(s(bool,X9)) => p(s(bool,X10))) => ((p(s(bool,X10)) => p(s(bool,X9))) => s(bool,X9) = s(bool,X10)))),
  file('thm.pred_set.CROSS_EMPTY|split|1','thm.bool.IMP_ANTISYM_AX')).
fof(f545,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK0,sK1),sK3('const.pred_set.EMPTY_0','const.pred_set.CROSS_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),sK2)),'type.pair.prod'(sK0,sK1))),s(fun('type.pair.prod'(sK0,sK1),bool),'const.pred_set.CROSS_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),sK2)))))),
  inference(forward_demodulation,[],[f526,f143])).
fof(f526,plain,(
  s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK0,sK1),sK3('const.pred_set.EMPTY_0','const.pred_set.CROSS_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),sK2)),'type.pair.prod'(sK0,sK1))),s(fun('type.pair.prod'(sK0,sK1),bool),'const.pred_set.EMPTY_0'))) != s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK0,sK1),sK3('const.pred_set.EMPTY_0','const.pred_set.CROSS_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),sK2)),'type.pair.prod'(sK0,sK1))),s(fun('type.pair.prod'(sK0,sK1),bool),'const.pred_set.CROSS_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),sK2)))))),
  inference(unit_resulting_resolution,[],[f86,f125])).
fof(f125,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) )),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f81])).
fof(f81,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f80])).
fof(f80,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X7] : ! [X13,X11] : (s(fun(X7,bool),X11) = s(fun(X7,bool),X13) <=> ! [X12] : s(bool,'const.bool.IN_2'(s(X7,X12),s(fun(X7,bool),X11))) = s(bool,'const.bool.IN_2'(s(X7,X12),s(fun(X7,bool),X13))))),
  file('thm.pred_set.CROSS_EMPTY|split|1','thm.pred_set.EXTENSION')).
fof(f86,plain,(
  s(fun('type.pair.prod'(sK0,sK1),bool),'const.pred_set.EMPTY_0') != s(fun('type.pair.prod'(sK0,sK1),bool),'const.pred_set.CROSS_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),sK2)))),
  inference(cnf_transformation,[],[f60])).
fof(f60,plain,(
  s(fun('type.pair.prod'(sK0,sK1),bool),'const.pred_set.EMPTY_0') != s(fun('type.pair.prod'(sK0,sK1),bool),'const.pred_set.CROSS_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK1,bool),sK2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f55])).
fof(f55,plain,(
  ? [X0,X1,X2] : s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.EMPTY_0') != s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.CROSS_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X1,bool),X2)))),
  inference(ennf_transformation,[],[f27])).
fof(f27,plain,(
  ~! [X0,X1,X2] : s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.CROSS_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X1,bool),X2)))),
  inference(flattening,[],[f26])).
fof(f26,plain,(
  ~! [X0,X1] : ! [X2] : s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.CROSS_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X1,bool),X2)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8] : s(fun('type.pair.prod'(X6,X7),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(X6,X7),bool),'const.pred_set.CROSS_2'(s(fun(X6,bool),'const.pred_set.EMPTY_0'),s(fun(X7,bool),X8)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8] : s(fun('type.pair.prod'(X6,X7),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(X6,X7),bool),'const.pred_set.CROSS_2'(s(fun(X6,bool),'const.pred_set.EMPTY_0'),s(fun(X7,bool),X8)))),
  file('thm.pred_set.CROSS_EMPTY|split|1','thm.pred_set.CROSS_EMPTY|split|1')).
fof(f131,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(X0,X1),X4),s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.CROSS_2'(s(fun(X0,bool),X2),s(fun(X1,bool),X3)))))) | p(s(bool,'const.bool.IN_2'(s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X4))),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(X0,X1),X4),s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.CROSS_2'(s(fun(X0,bool),X2),s(fun(X1,bool),X3)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X4))),s(fun(X0,bool),X2)))) & p(s(bool,'const.bool.IN_2'(s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X4))),s(fun(X1,bool),X3)))))) & (~p(s(bool,'const.bool.IN_2'(s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X4))),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X4))),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(X0,X1),X4),s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.CROSS_2'(s(fun(X0,bool),X2),s(fun(X1,bool),X3))))))))),
  inference(flattening,[],[f84])).
fof(f84,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(X0,X1),X4),s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.CROSS_2'(s(fun(X0,bool),X2),s(fun(X1,bool),X3)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X4))),s(fun(X0,bool),X2)))) & p(s(bool,'const.bool.IN_2'(s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X4))),s(fun(X1,bool),X3)))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X4))),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X4))),s(fun(X1,bool),X3))))) | p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(X0,X1),X4),s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.CROSS_2'(s(fun(X0,bool),X2),s(fun(X1,bool),X3))))))))),
  inference(nnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(X0,X1),X4),s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.CROSS_2'(s(fun(X0,bool),X2),s(fun(X1,bool),X3)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X4))),s(fun(X0,bool),X2)))) & p(s(bool,'const.bool.IN_2'(s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X4))),s(fun(X1,bool),X3))))))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(X0,X1),X4),s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.CROSS_2'(s(fun(X0,bool),X2),s(fun(X1,bool),X3)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X4))),s(fun(X0,bool),X2)))) & p(s(bool,'const.bool.IN_2'(s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X4))),s(fun(X1,bool),X3))))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X7,X14] : ! [X8,X15,X12] : (p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(X7,X14),X12),s(fun('type.pair.prod'(X7,X14),bool),'const.pred_set.CROSS_2'(s(fun(X7,bool),X8),s(fun(X14,bool),X15)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X7,'const.pair.FST_1'(s('type.pair.prod'(X7,X14),X12))),s(fun(X7,bool),X8)))) & p(s(bool,'const.bool.IN_2'(s(X14,'const.pair.SND_1'(s('type.pair.prod'(X7,X14),X12))),s(fun(X14,bool),X15))))))),
  file('thm.pred_set.CROSS_EMPTY|split|1','thm.pred_set.IN_CROSS')).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 24
% Passive clauses: 55
% Generated clauses: 533
% Final active clauses: 21
% Final passive clauses: 29
% Input formulas: 24
% Initial clauses: 48
% 
% Duplicate literals: 44
% Trivial inequalities: 11
% Fw subsumption resolutions: 30
% Fw demodulations: 19
% Bw demodulations: 5
% Condensations: 57
% 
% Simple tautologies: 40
% Equational tautologies: 29
% Forward subsumptions: 232
% 
% Unit resulting resolution: 60
% Factoring: 1
% Forward superposition: 68
% Backward superposition: 138
% Self superposition: 23
% Equality factoring: 26
% Equality resolution: 3
% 
% SAT solver clauses: 57
% SAT solver unit clauses: 15
% SAT solver binary clauses: 14
% 
% Memory used [KB]: 1023
% Time elapsed: 0.016 s
% ------------------------------
% ------------------------------
% Success in time 1.342 s
