Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.ring
% (114334)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 106
% Passive clauses: 1568
% Generated clauses: 17952
% Final active clauses: 103
% Final passive clauses: 1306
% Input formulas: 35
% Initial clauses: 45
% 
% Duplicate literals: 649
% Trivial inequalities: 171
% Fw subsumption resolutions: 1589
% Fw demodulations: 2298
% 
% Simple tautologies: 179
% Equational tautologies: 28
% Forward subsumptions: 11082
% 
% Binary resolution: 4
% Forward superposition: 7690
% Backward superposition: 5498
% Self superposition: 263
% Equality factoring: 65
% Equality resolution: 21
% 
% Split clauses: 188
% Split components: 380
% SAT solver clauses: 542
% SAT solver unit clauses: 43
% SAT solver binary clauses: 139
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 188
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 541
% TWLsolver calls for satisfiability: 115
% 
% Memory used [KB]: 14200
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 294 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ring
% (114361)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 89
% Passive clauses: 576
% Generated clauses: 15344
% Final active clauses: 86
% Final passive clauses: 483
% Input formulas: 35
% Initial clauses: 45
% 
% Duplicate literals: 667
% Trivial inequalities: 172
% Fw subsumption resolutions: 865
% Fw demodulations: 1537
% Bw demodulations: 2
% Condensations: 1918
% 
% Simple tautologies: 166
% Equational tautologies: 27
% Forward subsumptions: 8397
% 
% Binary resolution: 4
% Forward superposition: 5638
% Backward superposition: 3889
% Self superposition: 219
% Equality factoring: 64
% Equality resolution: 19
% 
% Split clauses: 42
% Split components: 88
% SAT solver clauses: 1004
% SAT solver unit clauses: 71
% SAT solver binary clauses: 47
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 42
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 1003
% TWLsolver calls for satisfiability: 275
% 
% Memory used [KB]: 14456
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 287 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.ring
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ring
% SZS output start Proof for thm.ring
fof(f21623,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f115,f166,f164,f162,f160,f158,f156,f114,f165,f113,f163,f112,f161,f111,f159,f110,f157,f155,f173,f199,f202,f200,f756,f757,f912,f758,f913,f911,f2626,f2628,f13866,f13885,f13886,f13888,f13890,f13748,f13892,f13867,f15920,f15943,f15955,f15816,f15956,f15921,f17344,f17367,f17382,f2627,f17235,f13891,f17383,f17345,f13893,f17976,f18007,f18017,f18825,f18861,f18871,f19378,f19415,f19425,f19478,f20893,f20936,f20993,f21047,f21048,f21462,f21506,f21572])).
fof(f21572,plain,(
  $false | ($spl0 | $spl5)),
  inference(subsumption_resolution,[],[f21571,f160])).
fof(f21571,plain,(
  s(sK0,sK2) = s(sK0,sK7) | $spl0),
  inference(forward_demodulation,[],[f21570,f167])).
fof(f167,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,happ(s(fun(X1,X0),'const.combin.K_1'(s(X0,X2))),s(X1,X3)))) )),
  inference(forward_demodulation,[],[f153,f152])).
fof(f152,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.combin.K_2'(s(X1,X2),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.combin.K_2'(s(X1,X2),s(X0,X3)))),
  inference(flattening,[],[f96])).
fof(f96,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.combin.K_2'(s(X1,X2),s(X0,X3)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X17,X6] : ! [X18,X19] : s(X6,X18) = s(X6,'const.combin.K_2'(s(X6,X18),s(X17,X19)))),
  file('thm.ring.ring_literal_|49||49|','thm.combin.K_THM')).
fof(f153,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.combin.K_2'(s(X0,X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),'const.combin.K_1'(s(X0,X2))),s(X1,X3)))) )),
  inference(cnf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.combin.K_2'(s(X0,X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),'const.combin.K_1'(s(X0,X2))),s(X1,X3)))),
  inference(flattening,[],[f98])).
fof(f98,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X0,'const.combin.K_2'(s(X0,X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),'const.combin.K_1'(s(X0,X2))),s(X1,X3)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X6,X17] : ! [X24,X25] : s(X6,'const.combin.K_2'(s(X6,X24),s(X17,X25))) = s(X6,happ(s(fun(X17,X6),'const.combin.K_1'(s(X6,X24))),s(X17,X25)))),
  file('thm.ring.ring_literal_|49||49|','ARITY_1_2_const.combin.K')).
fof(f21570,plain,(
  s(sK0,sK7) = s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | $spl0),
  inference(forward_demodulation,[],[f21569,f167])).
fof(f21569,plain,(
  s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK7))),s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | $spl0),
  inference(forward_demodulation,[],[f21568,f131])).
fof(f131,plain,(
  ( ! [X2,X0,X1] : (s(X0,happ(s(fun(X0,X0),X2),s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))))) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0,X1,X2] : s(X0,happ(s(fun(X0,X0),X2),s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))))) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f64])).
fof(f64,plain,(
  ! [X0] : ! [X1,X2] : s(X0,happ(s(fun(X0,X0),X2),s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))))) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X6] : ! [X20,X21] : s(X6,happ(s(fun(X6,X6),X21),s(X6,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X6),X20))))) = s(X6,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X6),'const.ring.ring_R|49|_fupd_2'(s(fun(X6,X6),X21),s('type.ring.ring'(X6),X20)))))),
  file('thm.ring.ring_literal_|49||49|','thm.ring.ring_39_prove_ltb')).
fof(f21568,plain,(
  s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK7))),s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | $spl0),
  inference(forward_demodulation,[],[f21567,f131])).
fof(f21567,plain,(
  s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK7))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | $spl0),
  inference(forward_demodulation,[],[f21367,f120])).
fof(f120,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0,X1,X2] : s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f42])).
fof(f42,plain,(
  ! [X0] : ! [X1,X2] : s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : ! [X20,X21] : s(X6,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X6),X20))) = s(X6,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X6),'const.ring.ring_R|48|_fupd_2'(s(fun(X6,X6),X21),s('type.ring.ring'(X6),X20)))))),
  file('thm.ring.ring_literal_|49||49|','thm.ring.ring_22_prove_ltb')).
fof(f21367,plain,(
  s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK7))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) | $spl0),
  inference(superposition,[],[f120,f155])).
fof(f21506,plain,(
  $false | ($spl0 | $spl5)),
  inference(subsumption_resolution,[],[f21505,f160])).
fof(f21505,plain,(
  s(sK0,sK2) = s(sK0,sK7) | $spl0),
  inference(forward_demodulation,[],[f21504,f167])).
fof(f21504,plain,(
  s(sK0,sK7) = s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | $spl0),
  inference(forward_demodulation,[],[f21503,f167])).
fof(f21503,plain,(
  s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK7))),s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | $spl0),
  inference(forward_demodulation,[],[f21502,f131])).
fof(f21502,plain,(
  s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK7))),s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | $spl0),
  inference(forward_demodulation,[],[f21501,f131])).
fof(f21501,plain,(
  s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK7))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | $spl0),
  inference(forward_demodulation,[],[f21500,f120])).
fof(f21500,plain,(
  s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK7))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) | $spl0),
  inference(forward_demodulation,[],[f21356,f120])).
fof(f21356,plain,(
  s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) = s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK7))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) | $spl0),
  inference(unit_resulting_resolution,[],[f155,f145])).
fof(f145,plain,(
  ( ! [X2,X0,X1] : (s('type.ring.ring'(X0),X1) != s('type.ring.ring'(X0),X2) | s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X2)))) )),
  inference(cnf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0,X1,X2] : ((s('type.ring.ring'(X0),X1) != s('type.ring.ring'(X0),X2) | (s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X2))) & s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X2))))) & (s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) != s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X2))) | s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) != s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X2))) | s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) != s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X2))) | s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) != s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X2))) | s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) != s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X2))) | s('type.ring.ring'(X0),X1) = s('type.ring.ring'(X0),X2)))),
  inference(flattening,[],[f107])).
fof(f107,plain,(
  ! [X0,X1,X2] : ((s('type.ring.ring'(X0),X1) != s('type.ring.ring'(X0),X2) | (s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X2))) & s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X2))))) & ((s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) != s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X2))) | s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) != s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X2))) | s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) != s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X2))) | s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) != s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X2))) | s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) != s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X2)))) | s('type.ring.ring'(X0),X1) = s('type.ring.ring'(X0),X2)))),
  inference(nnf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0,X1,X2] : (s('type.ring.ring'(X0),X1) = s('type.ring.ring'(X0),X2) <=> (s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X2))) & s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X2)))))),
  inference(flattening,[],[f90])).
fof(f90,plain,(
  ! [X0] : ! [X1,X2] : (s('type.ring.ring'(X0),X1) = s('type.ring.ring'(X0),X2) <=> (s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X2))) & s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X2))) & s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X2)))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X6] : ! [X22,X23] : (s('type.ring.ring'(X6),X22) = s('type.ring.ring'(X6),X23) <=> (s(X6,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X6),X22))) = s(X6,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X6),X23))) & s(X6,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X6),X22))) = s(X6,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X6),X23))) & s(fun(X6,fun(X6,X6)),'const.ring.ring_RP_1'(s('type.ring.ring'(X6),X22))) = s(fun(X6,fun(X6,X6)),'const.ring.ring_RP_1'(s('type.ring.ring'(X6),X23))) & s(fun(X6,fun(X6,X6)),'const.ring.ring_RM_1'(s('type.ring.ring'(X6),X22))) = s(fun(X6,fun(X6,X6)),'const.ring.ring_RM_1'(s('type.ring.ring'(X6),X23))) & s(fun(X6,X6),'const.ring.ring_RN_1'(s('type.ring.ring'(X6),X22))) = s(fun(X6,X6),'const.ring.ring_RN_1'(s('type.ring.ring'(X6),X23)))))),
  file('thm.ring.ring_literal_|49||49|','thm.ring.ring_component_equality')).
fof(f21462,plain,(
  $false | ($spl0 | $spl5)),
  inference(subsumption_resolution,[],[f21461,f160])).
fof(f21461,plain,(
  s(sK0,sK2) = s(sK0,sK7) | $spl0),
  inference(forward_demodulation,[],[f21460,f167])).
fof(f21460,plain,(
  s(sK0,sK7) = s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | $spl0),
  inference(forward_demodulation,[],[f21459,f167])).
fof(f21459,plain,(
  s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK7))),s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | $spl0),
  inference(forward_demodulation,[],[f21458,f131])).
fof(f21458,plain,(
  s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK7))),s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | $spl0),
  inference(forward_demodulation,[],[f21457,f131])).
fof(f21457,plain,(
  s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK7))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | $spl0),
  inference(forward_demodulation,[],[f21456,f120])).
fof(f21456,plain,(
  s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK7))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) | $spl0),
  inference(forward_demodulation,[],[f21361,f120])).
fof(f21361,plain,(
  s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) = s(sK0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK7))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) | $spl0),
  inference(unit_resulting_resolution,[],[f155,f145])).
fof(f21048,plain,(
  $false | ($spl0 | $spl4 | $spl11)),
  inference(subsumption_resolution,[],[f166,f20908])).
fof(f20908,plain,(
  s(fun(sK0,sK0),sK5) = s(fun(sK0,sK0),sK10) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f20907,f15065])).
fof(f15065,plain,(
  ( ! [X4,X5,X3] : (s(fun(X3,X3),X4) = s(fun(X3,X3),'const.ring.ring_RN_1'(s('type.ring.ring'(X3),'const.ring.ring_RN_fupd_2'(s(fun(fun(X3,X3),fun(X3,X3)),'const.combin.K_1'(s(fun(X3,X3),X4))),s('type.ring.ring'(X3),X5)))))) )),
  inference(superposition,[],[f141,f167])).
fof(f141,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),X2),s('type.ring.ring'(X0),X1))))) = s(fun(X0,X0),happ(s(fun(fun(X0,X0),fun(X0,X0)),X2),s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0,X1,X2] : s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),X2),s('type.ring.ring'(X0),X1))))) = s(fun(X0,X0),happ(s(fun(fun(X0,X0),fun(X0,X0)),X2),s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f84])).
fof(f84,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),X2),s('type.ring.ring'(X0),X1))))) = s(fun(X0,X0),happ(s(fun(fun(X0,X0),fun(X0,X0)),X2),s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X6] : ! [X20,X21] : s(fun(X6,X6),'const.ring.ring_RN_1'(s('type.ring.ring'(X6),'const.ring.ring_RN_fupd_2'(s(fun(fun(X6,X6),fun(X6,X6)),X21),s('type.ring.ring'(X6),X20))))) = s(fun(X6,X6),happ(s(fun(fun(X6,X6),fun(X6,X6)),X21),s(fun(X6,X6),'const.ring.ring_RN_1'(s('type.ring.ring'(X6),X20)))))),
  file('thm.ring.ring_literal_|49||49|','thm.ring.ring_42_prove_ltb')).
fof(f20907,plain,(
  s(fun(sK0,sK0),sK10) = s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f20906,f15065])).
fof(f20906,plain,(
  s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f20905,f135])).
fof(f135,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1,X2] : s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f72])).
fof(f72,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6] : ! [X20,X21] : s(fun(X6,X6),'const.ring.ring_RN_1'(s('type.ring.ring'(X6),X20))) = s(fun(X6,X6),'const.ring.ring_RN_1'(s('type.ring.ring'(X6),'const.ring.ring_RM_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),X21),s('type.ring.ring'(X6),X20)))))),
  file('thm.ring.ring_literal_|49||49|','thm.ring.ring_37_prove_ltb')).
fof(f20905,plain,(
  s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f20904,f135])).
fof(f20904,plain,(
  s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f20903,f138])).
fof(f138,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1,X2] : s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f78])).
fof(f78,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6] : ! [X20,X21] : s(fun(X6,X6),'const.ring.ring_RN_1'(s('type.ring.ring'(X6),X20))) = s(fun(X6,X6),'const.ring.ring_RN_1'(s('type.ring.ring'(X6),'const.ring.ring_RP_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),X21),s('type.ring.ring'(X6),X20)))))),
  file('thm.ring.ring_literal_|49||49|','thm.ring.ring_36_prove_ltb')).
fof(f20903,plain,(
  s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f20902,f138])).
fof(f20902,plain,(
  s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f20901,f123])).
fof(f123,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1,X2] : s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6] : ! [X20,X21] : s(fun(X6,X6),'const.ring.ring_RN_1'(s('type.ring.ring'(X6),X20))) = s(fun(X6,X6),'const.ring.ring_RN_1'(s('type.ring.ring'(X6),'const.ring.ring_R|49|_fupd_2'(s(fun(X6,X6),X21),s('type.ring.ring'(X6),X20)))))),
  file('thm.ring.ring_literal_|49||49|','thm.ring.ring_35_prove_ltb')).
fof(f20901,plain,(
  s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f20900,f123])).
fof(f20900,plain,(
  s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f20899,f122])).
fof(f122,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1,X2] : s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : ! [X20,X21] : s(fun(X6,X6),'const.ring.ring_RN_1'(s('type.ring.ring'(X6),X20))) = s(fun(X6,X6),'const.ring.ring_RN_1'(s('type.ring.ring'(X6),'const.ring.ring_R|48|_fupd_2'(s(fun(X6,X6),X21),s('type.ring.ring'(X6),X20)))))),
  file('thm.ring.ring_literal_|49||49|','thm.ring.ring_34_prove_ltb')).
fof(f20899,plain,(
  s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f20799,f122])).
fof(f20799,plain,(
  s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) = s(fun(sK0,sK0),'const.ring.ring_RN_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) | ($spl0 | $spl4)),
  inference(unit_resulting_resolution,[],[f19479,f148])).
fof(f148,plain,(
  ( ! [X2,X0,X1] : (s('type.ring.ring'(X0),X1) != s('type.ring.ring'(X0),X2) | s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,X0),'const.ring.ring_RN_1'(s('type.ring.ring'(X0),X2)))) )),
  inference(cnf_transformation,[],[f108])).
fof(f19479,plain,(
  s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f155,f159])).
fof(f21047,plain,(
  $false | ($spl0 | $spl4 | $spl7)),
  inference(subsumption_resolution,[],[f162,f20919])).
fof(f20919,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK3) = s(fun(sK0,fun(sK0,sK0)),sK8) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f20918,f16040])).
fof(f16040,plain,(
  ( ! [X4,X5,X3] : (s(fun(X3,fun(X3,X3)),X4) = s(fun(X3,fun(X3,X3)),'const.ring.ring_RP_1'(s('type.ring.ring'(X3),'const.ring.ring_RP_fupd_2'(s(fun(fun(X3,fun(X3,X3)),fun(X3,fun(X3,X3))),'const.combin.K_1'(s(fun(X3,fun(X3,X3)),X4))),s('type.ring.ring'(X3),X5)))))) )),
  inference(superposition,[],[f142,f167])).
fof(f142,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1))))) = s(fun(X0,fun(X0,X0)),happ(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0,X1,X2] : s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1))))) = s(fun(X0,fun(X0,X0)),happ(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f86])).
fof(f86,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1))))) = s(fun(X0,fun(X0,X0)),happ(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X6] : ! [X20,X21] : s(fun(X6,fun(X6,X6)),'const.ring.ring_RP_1'(s('type.ring.ring'(X6),'const.ring.ring_RP_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),X21),s('type.ring.ring'(X6),X20))))) = s(fun(X6,fun(X6,X6)),happ(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),X21),s(fun(X6,fun(X6,X6)),'const.ring.ring_RP_1'(s('type.ring.ring'(X6),X20)))))),
  file('thm.ring.ring_literal_|49||49|','thm.ring.ring_4|48|_prove_ltb')).
fof(f20918,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK8) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f20917,f16040])).
fof(f20917,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f20916,f16089])).
fof(f16089,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X3),s('type.ring.ring'(X0),X2))))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X3),s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),X1),s('type.ring.ring'(X0),X2)))))))) )),
  inference(forward_demodulation,[],[f16036,f142])).
fof(f16036,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,fun(X0,X0)),happ(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X3),s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X2))))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X3),s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),X1),s('type.ring.ring'(X0),X2)))))))) )),
  inference(superposition,[],[f142,f136])).
fof(f136,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),X2),s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0,X1,X2] : s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),X2),s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f74])).
fof(f74,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),X2),s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X20,X21] : s(fun(X6,fun(X6,X6)),'const.ring.ring_RP_1'(s('type.ring.ring'(X6),X20))) = s(fun(X6,fun(X6,X6)),'const.ring.ring_RP_1'(s('type.ring.ring'(X6),'const.ring.ring_RN_fupd_2'(s(fun(fun(X6,X6),fun(X6,X6)),X21),s('type.ring.ring'(X6),X20)))))),
  file('thm.ring.ring_literal_|49||49|','thm.ring.ring_29_prove_ltb')).
fof(f20916,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f20915,f16089])).
fof(f20915,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f20914,f16090])).
fof(f16090,plain,(
  ( ! [X6,X4,X7,X5] : (s(fun(X4,fun(X4,X4)),'const.ring.ring_RP_1'(s('type.ring.ring'(X4),'const.ring.ring_RP_fupd_2'(s(fun(fun(X4,fun(X4,X4)),fun(X4,fun(X4,X4))),X7),s('type.ring.ring'(X4),X6))))) = s(fun(X4,fun(X4,X4)),'const.ring.ring_RP_1'(s('type.ring.ring'(X4),'const.ring.ring_RP_fupd_2'(s(fun(fun(X4,fun(X4,X4)),fun(X4,fun(X4,X4))),X7),s('type.ring.ring'(X4),'const.ring.ring_RM_fupd_2'(s(fun(fun(X4,fun(X4,X4)),fun(X4,fun(X4,X4))),X5),s('type.ring.ring'(X4),X6)))))))) )),
  inference(forward_demodulation,[],[f16037,f142])).
fof(f16037,plain,(
  ( ! [X6,X4,X7,X5] : (s(fun(X4,fun(X4,X4)),happ(s(fun(fun(X4,fun(X4,X4)),fun(X4,fun(X4,X4))),X7),s(fun(X4,fun(X4,X4)),'const.ring.ring_RP_1'(s('type.ring.ring'(X4),X6))))) = s(fun(X4,fun(X4,X4)),'const.ring.ring_RP_1'(s('type.ring.ring'(X4),'const.ring.ring_RP_fupd_2'(s(fun(fun(X4,fun(X4,X4)),fun(X4,fun(X4,X4))),X7),s('type.ring.ring'(X4),'const.ring.ring_RM_fupd_2'(s(fun(fun(X4,fun(X4,X4)),fun(X4,fun(X4,X4))),X5),s('type.ring.ring'(X4),X6)))))))) )),
  inference(superposition,[],[f142,f139])).
fof(f139,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0,X1,X2] : s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f80])).
fof(f80,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6] : ! [X20,X21] : s(fun(X6,fun(X6,X6)),'const.ring.ring_RP_1'(s('type.ring.ring'(X6),X20))) = s(fun(X6,fun(X6,X6)),'const.ring.ring_RP_1'(s('type.ring.ring'(X6),'const.ring.ring_RM_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),X21),s('type.ring.ring'(X6),X20)))))),
  file('thm.ring.ring_literal_|49||49|','thm.ring.ring_28_prove_ltb')).
fof(f20914,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f20913,f16090])).
fof(f20913,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f20912,f134])).
fof(f134,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1,X2] : s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X6] : ! [X20,X21] : s(fun(X6,fun(X6,X6)),'const.ring.ring_RP_1'(s('type.ring.ring'(X6),X20))) = s(fun(X6,fun(X6,X6)),'const.ring.ring_RP_1'(s('type.ring.ring'(X6),'const.ring.ring_R|49|_fupd_2'(s(fun(X6,X6),X21),s('type.ring.ring'(X6),X20)))))),
  file('thm.ring.ring_literal_|49||49|','thm.ring.ring_27_prove_ltb')).
fof(f20912,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f20911,f134])).
fof(f20911,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f20910,f130])).
fof(f130,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1,X2] : s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6] : ! [X20,X21] : s(fun(X6,fun(X6,X6)),'const.ring.ring_RP_1'(s('type.ring.ring'(X6),X20))) = s(fun(X6,fun(X6,X6)),'const.ring.ring_RP_1'(s('type.ring.ring'(X6),'const.ring.ring_R|48|_fupd_2'(s(fun(X6,X6),X21),s('type.ring.ring'(X6),X20)))))),
  file('thm.ring.ring_literal_|49||49|','thm.ring.ring_26_prove_ltb')).
fof(f20910,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f20798,f130])).
fof(f20798,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RP_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) | ($spl0 | $spl4)),
  inference(unit_resulting_resolution,[],[f19479,f146])).
fof(f146,plain,(
  ( ! [X2,X0,X1] : (s('type.ring.ring'(X0),X1) != s('type.ring.ring'(X0),X2) | s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X2)))) )),
  inference(cnf_transformation,[],[f108])).
fof(f20993,plain,(
  $false | ($spl0 | $spl3 | $spl4)),
  inference(subsumption_resolution,[],[f20992,f158])).
fof(f20992,plain,(
  s(sK0,sK1) = s(sK0,sK6) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f20991,f167])).
fof(f20991,plain,(
  s(sK0,sK6) = s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f20990,f167])).
fof(f20990,plain,(
  s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) = s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f20804,f127])).
fof(f127,plain,(
  ( ! [X2,X0,X1] : (s(X0,happ(s(fun(X0,X0),X2),s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))))) = s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1,X2] : s(X0,happ(s(fun(X0,X0),X2),s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))))) = s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ! [X0] : ! [X1,X2] : s(X0,happ(s(fun(X0,X0),X2),s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))))) = s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6] : ! [X20,X21] : s(X6,happ(s(fun(X6,X6),X21),s(X6,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X6),X20))))) = s(X6,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X6),'const.ring.ring_R|48|_fupd_2'(s(fun(X6,X6),X21),s('type.ring.ring'(X6),X20)))))),
  file('thm.ring.ring_literal_|49||49|','thm.ring.ring_38_prove_ltb')).
fof(f20804,plain,(
  s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) = s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) | ($spl0 | $spl4)),
  inference(superposition,[],[f127,f19479])).
fof(f20936,plain,(
  $false | ($spl0 | $spl3 | $spl4)),
  inference(subsumption_resolution,[],[f20935,f158])).
fof(f20935,plain,(
  s(sK0,sK1) = s(sK0,sK6) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f20934,f167])).
fof(f20934,plain,(
  s(sK0,sK6) = s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f20933,f167])).
fof(f20933,plain,(
  s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) = s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f20932,f127])).
fof(f20932,plain,(
  s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) = s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f20796,f127])).
fof(f20796,plain,(
  s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) = s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) | ($spl0 | $spl4)),
  inference(unit_resulting_resolution,[],[f19479,f144])).
fof(f144,plain,(
  ( ! [X2,X0,X1] : (s('type.ring.ring'(X0),X1) != s('type.ring.ring'(X0),X2) | s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))) = s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X2)))) )),
  inference(cnf_transformation,[],[f108])).
fof(f20893,plain,(
  $false | ($spl0 | $spl3 | $spl4)),
  inference(subsumption_resolution,[],[f20892,f158])).
fof(f20892,plain,(
  s(sK0,sK1) = s(sK0,sK6) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f20891,f167])).
fof(f20891,plain,(
  s(sK0,sK6) = s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f20890,f167])).
fof(f20890,plain,(
  s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) = s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f20889,f127])).
fof(f20889,plain,(
  s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) = s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f20801,f127])).
fof(f20801,plain,(
  s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) = s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) | ($spl0 | $spl4)),
  inference(unit_resulting_resolution,[],[f19479,f144])).
fof(f19478,plain,(
  $false | ($spl0 | $spl3 | $spl4 | $spl10)),
  inference(subsumption_resolution,[],[f158,f19343])).
fof(f19343,plain,(
  s(sK0,sK1) = s(sK0,sK6) | ($spl0 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f19342,f167])).
fof(f19342,plain,(
  s(sK0,sK6) = s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) | ($spl0 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f19341,f167])).
fof(f19341,plain,(
  s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) = s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) | ($spl0 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f19340,f127])).
fof(f19340,plain,(
  s(sK0,happ(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) = s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) | ($spl0 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f19254,f127])).
fof(f19254,plain,(
  s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) = s(sK0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) | ($spl0 | $spl4 | $spl10)),
  inference(unit_resulting_resolution,[],[f18044,f144])).
fof(f18044,plain,(
  s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f174,f159])).
fof(f174,plain,(
  s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK7))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl10)),
  inference(forward_demodulation,[],[f155,f165])).
fof(f19425,plain,(
  $false | ($spl0 | $spl4 | $spl9 | $spl10)),
  inference(subsumption_resolution,[],[f19424,f164])).
fof(f19424,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK9) | ($spl0 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f19423,f16752])).
fof(f16752,plain,(
  ( ! [X4,X5,X3] : (s(fun(X3,fun(X3,X3)),X4) = s(fun(X3,fun(X3,X3)),'const.ring.ring_RM_1'(s('type.ring.ring'(X3),'const.ring.ring_RM_fupd_2'(s(fun(fun(X3,fun(X3,X3)),fun(X3,fun(X3,X3))),'const.combin.K_1'(s(fun(X3,fun(X3,X3)),X4))),s('type.ring.ring'(X3),X5)))))) )),
  inference(superposition,[],[f143,f167])).
fof(f143,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1))))) = s(fun(X0,fun(X0,X0)),happ(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0,X1,X2] : s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1))))) = s(fun(X0,fun(X0,X0)),happ(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f88])).
fof(f88,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1))))) = s(fun(X0,fun(X0,X0)),happ(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X6] : ! [X20,X21] : s(fun(X6,fun(X6,X6)),'const.ring.ring_RM_1'(s('type.ring.ring'(X6),'const.ring.ring_RM_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),X21),s('type.ring.ring'(X6),X20))))) = s(fun(X6,fun(X6,X6)),happ(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),X21),s(fun(X6,fun(X6,X6)),'const.ring.ring_RM_1'(s('type.ring.ring'(X6),X20)))))),
  file('thm.ring.ring_literal_|49||49|','thm.ring.ring_4|49|_prove_ltb')).
fof(f19423,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK9) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) | ($spl0 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f19422,f16752])).
fof(f19422,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) | ($spl0 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f19421,f16801])).
fof(f16801,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X3),s('type.ring.ring'(X0),X2))))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X3),s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),X1),s('type.ring.ring'(X0),X2)))))))) )),
  inference(forward_demodulation,[],[f16748,f143])).
fof(f16748,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,fun(X0,X0)),happ(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X3),s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X2))))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X3),s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),X1),s('type.ring.ring'(X0),X2)))))))) )),
  inference(superposition,[],[f143,f137])).
fof(f137,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),X2),s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0,X1,X2] : s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),X2),s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f76])).
fof(f76,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),X2),s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : ! [X20,X21] : s(fun(X6,fun(X6,X6)),'const.ring.ring_RM_1'(s('type.ring.ring'(X6),X20))) = s(fun(X6,fun(X6,X6)),'const.ring.ring_RM_1'(s('type.ring.ring'(X6),'const.ring.ring_RN_fupd_2'(s(fun(fun(X6,X6),fun(X6,X6)),X21),s('type.ring.ring'(X6),X20)))))),
  file('thm.ring.ring_literal_|49||49|','thm.ring.ring_33_prove_ltb')).
fof(f19421,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) | ($spl0 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f19420,f16801])).
fof(f19420,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) | ($spl0 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f19419,f140])).
fof(f140,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0,X1,X2] : s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f82])).
fof(f82,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),X2),s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : ! [X20,X21] : s(fun(X6,fun(X6,X6)),'const.ring.ring_RM_1'(s('type.ring.ring'(X6),X20))) = s(fun(X6,fun(X6,X6)),'const.ring.ring_RM_1'(s('type.ring.ring'(X6),'const.ring.ring_RP_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),X21),s('type.ring.ring'(X6),X20)))))),
  file('thm.ring.ring_literal_|49||49|','thm.ring.ring_32_prove_ltb')).
fof(f19419,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f19418,f140])).
fof(f19418,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f19417,f133])).
fof(f133,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1,X2] : s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6] : ! [X20,X21] : s(fun(X6,fun(X6,X6)),'const.ring.ring_RM_1'(s('type.ring.ring'(X6),X20))) = s(fun(X6,fun(X6,X6)),'const.ring.ring_RM_1'(s('type.ring.ring'(X6),'const.ring.ring_R|49|_fupd_2'(s(fun(X6,X6),X21),s('type.ring.ring'(X6),X20)))))),
  file('thm.ring.ring_literal_|49||49|','thm.ring.ring_3|49|_prove_ltb')).
fof(f19417,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f19416,f133])).
fof(f19416,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f19255,f132])).
fof(f132,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f67])).
fof(f67,plain,(
  ! [X0,X1,X2] : s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f66])).
fof(f66,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),X2),s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X20,X21] : s(fun(X6,fun(X6,X6)),'const.ring.ring_RM_1'(s('type.ring.ring'(X6),X20))) = s(fun(X6,fun(X6,X6)),'const.ring.ring_RM_1'(s('type.ring.ring'(X6),'const.ring.ring_R|48|_fupd_2'(s(fun(X6,X6),X21),s('type.ring.ring'(X6),X20)))))),
  file('thm.ring.ring_literal_|49||49|','thm.ring.ring_3|48|_prove_ltb')).
fof(f19255,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) | ($spl0 | $spl4 | $spl10)),
  inference(superposition,[],[f132,f18044])).
fof(f19415,plain,(
  $false | ($spl0 | $spl4 | $spl9 | $spl10)),
  inference(subsumption_resolution,[],[f19414,f164])).
fof(f19414,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK9) | ($spl0 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f19413,f16752])).
fof(f19413,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK9) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) | ($spl0 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f19412,f16752])).
fof(f19412,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) | ($spl0 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f19411,f16801])).
fof(f19411,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) | ($spl0 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f19410,f16801])).
fof(f19410,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) | ($spl0 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f19409,f140])).
fof(f19409,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f19408,f140])).
fof(f19408,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f19407,f133])).
fof(f19407,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f19406,f133])).
fof(f19406,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f19405,f132])).
fof(f19405,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) | ($spl0 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f19245,f132])).
fof(f19245,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) | ($spl0 | $spl4 | $spl10)),
  inference(unit_resulting_resolution,[],[f18044,f147])).
fof(f147,plain,(
  ( ! [X2,X0,X1] : (s('type.ring.ring'(X0),X1) != s('type.ring.ring'(X0),X2) | s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X2)))) )),
  inference(cnf_transformation,[],[f108])).
fof(f19378,plain,(
  $false | ($spl0 | $spl4 | $spl9 | $spl10)),
  inference(subsumption_resolution,[],[f19377,f164])).
fof(f19377,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK9) | ($spl0 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f19376,f16752])).
fof(f19376,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK9) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) | ($spl0 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f19375,f16752])).
fof(f19375,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) | ($spl0 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f19374,f16801])).
fof(f19374,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) | ($spl0 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f19373,f16801])).
fof(f19373,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) | ($spl0 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f19372,f140])).
fof(f19372,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f19371,f140])).
fof(f19371,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f19370,f133])).
fof(f19370,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f19369,f133])).
fof(f19369,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f19368,f132])).
fof(f19368,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) | ($spl0 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f19250,f132])).
fof(f19250,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) | ($spl0 | $spl4 | $spl10)),
  inference(unit_resulting_resolution,[],[f18044,f147])).
fof(f18871,plain,(
  $false | ($spl0 | $spl2 | $spl4 | $spl9 | $spl10)),
  inference(subsumption_resolution,[],[f18870,f164])).
fof(f18870,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK9) | ($spl0 | $spl2 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f18869,f16752])).
fof(f18869,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK9) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) | ($spl0 | $spl2 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f18868,f16752])).
fof(f18868,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) | ($spl0 | $spl2 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f18867,f16801])).
fof(f18867,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) | ($spl0 | $spl2 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f18866,f16801])).
fof(f18866,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) | ($spl0 | $spl2 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f18865,f140])).
fof(f18865,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl2 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f18864,f140])).
fof(f18864,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl2 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f18863,f133])).
fof(f18863,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl2 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f18862,f133])).
fof(f18862,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl2 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f18704,f132])).
fof(f18704,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) | ($spl0 | $spl2 | $spl4 | $spl10)),
  inference(superposition,[],[f132,f18045])).
fof(f18045,plain,(
  s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl2 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f18044,f157])).
fof(f18861,plain,(
  $false | ($spl0 | $spl2 | $spl4 | $spl9 | $spl10)),
  inference(subsumption_resolution,[],[f18860,f164])).
fof(f18860,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK9) | ($spl0 | $spl2 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f18859,f16752])).
fof(f18859,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK9) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) | ($spl0 | $spl2 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f18858,f16752])).
fof(f18858,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) | ($spl0 | $spl2 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f18857,f16801])).
fof(f18857,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) | ($spl0 | $spl2 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f18856,f16801])).
fof(f18856,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) | ($spl0 | $spl2 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f18855,f140])).
fof(f18855,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl2 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f18854,f140])).
fof(f18854,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl2 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f18853,f133])).
fof(f18853,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl2 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f18852,f133])).
fof(f18852,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl2 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f18851,f132])).
fof(f18851,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) | ($spl0 | $spl2 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f18694,f132])).
fof(f18694,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) | ($spl0 | $spl2 | $spl4 | $spl10)),
  inference(unit_resulting_resolution,[],[f18045,f147])).
fof(f18825,plain,(
  $false | ($spl0 | $spl2 | $spl4 | $spl9 | $spl10)),
  inference(subsumption_resolution,[],[f18824,f164])).
fof(f18824,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK9) | ($spl0 | $spl2 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f18823,f16752])).
fof(f18823,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK9) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) | ($spl0 | $spl2 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f18822,f16752])).
fof(f18822,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) | ($spl0 | $spl2 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f18821,f16801])).
fof(f18821,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) | ($spl0 | $spl2 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f18820,f16801])).
fof(f18820,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) | ($spl0 | $spl2 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f18819,f140])).
fof(f18819,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl2 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f18818,f140])).
fof(f18818,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl2 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f18817,f133])).
fof(f18817,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl2 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f18816,f133])).
fof(f18816,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl2 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f18815,f132])).
fof(f18815,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) | ($spl0 | $spl2 | $spl4 | $spl10)),
  inference(forward_demodulation,[],[f18699,f132])).
fof(f18699,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) | ($spl0 | $spl2 | $spl4 | $spl10)),
  inference(unit_resulting_resolution,[],[f18045,f147])).
fof(f18017,plain,(
  $false | ($spl0 | $spl2 | $spl4 | $spl6 | $spl9 | $spl10)),
  inference(subsumption_resolution,[],[f18016,f164])).
fof(f18016,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK9) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f18015,f16752])).
fof(f18015,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK9) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f18014,f16752])).
fof(f18014,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f18013,f16801])).
fof(f18013,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f18012,f16801])).
fof(f18012,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f18011,f140])).
fof(f18011,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f18010,f140])).
fof(f18010,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f18009,f133])).
fof(f18009,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f18008,f133])).
fof(f18008,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f17861,f132])).
fof(f17861,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10)),
  inference(superposition,[],[f132,f177])).
fof(f177,plain,(
  s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f176,f157])).
fof(f176,plain,(
  s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl4 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f175,f159])).
fof(f175,plain,(
  s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK7))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f174,f161])).
fof(f18007,plain,(
  $false | ($spl0 | $spl2 | $spl4 | $spl6 | $spl9 | $spl10)),
  inference(subsumption_resolution,[],[f18006,f164])).
fof(f18006,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK9) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f18005,f16752])).
fof(f18005,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK9) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f18004,f16752])).
fof(f18004,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f18003,f16801])).
fof(f18003,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f18002,f16801])).
fof(f18002,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f18001,f140])).
fof(f18001,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f18000,f140])).
fof(f18000,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f17999,f133])).
fof(f17999,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f17998,f133])).
fof(f17998,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f17997,f132])).
fof(f17997,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f17851,f132])).
fof(f17851,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10)),
  inference(unit_resulting_resolution,[],[f177,f147])).
fof(f17976,plain,(
  $false | ($spl0 | $spl2 | $spl4 | $spl6 | $spl9 | $spl10)),
  inference(subsumption_resolution,[],[f17975,f164])).
fof(f17975,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK9) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f17974,f16752])).
fof(f17974,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK9) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f17973,f16752])).
fof(f17973,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f17972,f16801])).
fof(f17972,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f17971,f16801])).
fof(f17971,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f17970,f140])).
fof(f17970,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f17969,f140])).
fof(f17969,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f17968,f133])).
fof(f17968,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f17967,f133])).
fof(f17967,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f17966,f132])).
fof(f17966,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10)),
  inference(forward_demodulation,[],[f17856,f132])).
fof(f17856,plain,(
  s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) = s(fun(sK0,fun(sK0,sK0)),'const.ring.ring_RM_1'(s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))))) | ($spl0 | $spl2 | $spl4 | $spl6 | $spl10)),
  inference(unit_resulting_resolution,[],[f177,f147])).
fof(f13893,plain,(
  ( ! [X451,X446,X452] : (s(bool,'T_0') = s(bool,happ(s(fun(X451,bool),'const.combin.K_1'(s(bool,X446))),s(X451,X452)))) ) | $spl30),
  inference(cnf_transformation,[],[f13893_D])).
fof(f13893_D,plain,(
  ( ! [X451,X446,X452] : (s(bool,'T_0') = s(bool,happ(s(fun(X451,bool),'const.combin.K_1'(s(bool,X446))),s(X451,X452)))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f17345,plain,(
  ( ! [X142] : (s(bool,'T_0') != s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|48|_fupd_2'(s(fun(bool,bool),'const.combin.K_1'(s(bool,'F_0'))),s('type.ring.ring'(bool),X142)))))) ) | $spl34),
  inference(cnf_transformation,[],[f17345_D])).
fof(f17345_D,plain,(
  ( ! [X142] : (s(bool,'T_0') != s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|48|_fupd_2'(s(fun(bool,bool),'const.combin.K_1'(s(bool,'F_0'))),s('type.ring.ring'(bool),X142)))))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f17383,plain,(
  ( ! [X377,X372,X376,X375] : (s(bool,'T_0') != s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|48|_fupd_2'(s(fun(bool,bool),'const.combin.K_1'(s(bool,'F_0'))),s('type.ring.ring'(bool),X375))))) | s(bool,'T_0') = s(bool,happ(s(fun(X376,bool),'const.combin.K_1'(s(bool,X372))),s(X376,X377)))) ) | ($spl15 | $spl24)),
  inference(subsumption_resolution,[],[f16992,f2900])).
fof(f2900,plain,(
  ( ! [X134,X132,X135,X133,X131] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,X132),s(X134,X135))) | s(bool,'T_0') = s(bool,happ(s(fun(X131,bool),'const.combin.K_1'(s(bool,X132))),s(X131,X133)))) )),
  inference(superposition,[],[f2307,f167])).
fof(f2307,plain,(
  ( ! [X4,X0,X3] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,X0),s(X3,X4))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(forward_demodulation,[],[f2227,f152])).
fof(f2227,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X1,X2))) != s(bool,'const.combin.K_2'(s(bool,X0),s(X3,X4)))) )),
  inference(condensation,[],[f2226])).
fof(f2226,plain,(
  ( ! [X99,X97,X95,X100,X98,X96] : (s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X99,X100))) != s(bool,'const.combin.K_2'(s(bool,X96),s(X97,X98))) | s(bool,'T_0') = s(bool,X95) | s(bool,'T_0') = s(bool,X96)) )),
  inference(duplicate_literal_removal,[],[f2117])).
fof(f2117,plain,(
  ( ! [X99,X97,X95,X100,X98,X96] : (s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X99,X100))) != s(bool,'const.combin.K_2'(s(bool,X96),s(X97,X98))) | s(bool,'T_0') = s(bool,X95) | s(bool,'T_0') = s(bool,X96) | s(bool,'T_0') = s(bool,X95)) )),
  inference(superposition,[],[f1348,f316])).
fof(f316,plain,(
  ( ! [X19,X17,X20,X18] : (s(bool,X20) = s(bool,'const.combin.K_2'(s(bool,X17),s(X18,X19))) | s(bool,'T_0') = s(bool,X17) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f188,f152])).
fof(f188,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f118,f118])).
fof(f118,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.ring.ring_literal_|49||49|',aHL_BOOL_CASES)).
fof(f1348,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X1,X2))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f1330])).
fof(f1330,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X1,X2))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f291])).
fof(f291,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,X24) = s(bool,X27) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X24),s(X25,X26)))) )),
  inference(superposition,[],[f188,f152])).
fof(f16992,plain,(
  ( ! [X377,X372,X374,X376,X373,X375] : (s(bool,'T_0') != s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|48|_fupd_2'(s(fun(bool,bool),'const.combin.K_1'(s(bool,'F_0'))),s('type.ring.ring'(bool),X375))))) | s(bool,'T_0') = s(bool,happ(s(fun(X376,bool),'const.combin.K_1'(s(bool,X372))),s(X376,X377))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X372),s(X373,X374)))) ) | ($spl15 | $spl24)),
  inference(superposition,[],[f2900,f14766])).
fof(f14766,plain,(
  ( ! [X2,X3] : (s(bool,X3) = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|48|_fupd_2'(s(fun(bool,bool),'const.combin.K_1'(s(bool,'F_0'))),s('type.ring.ring'(bool),X2))))) | s(bool,'T_0') = s(bool,X3)) ) | ($spl15 | $spl24)),
  inference(subsumption_resolution,[],[f14636,f202])).
fof(f14636,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X3) = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|48|_fupd_2'(s(fun(bool,bool),'const.combin.K_1'(s(bool,'F_0'))),s('type.ring.ring'(bool),X2))))) | s(bool,'T_0') = s(bool,X3)) ) | $spl24),
  inference(superposition,[],[f14278,f188])).
fof(f14278,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|48|_fupd_2'(s(fun(bool,bool),'const.combin.K_1'(s(bool,'F_0'))),s('type.ring.ring'(bool),X2)))))) ) | $spl24),
  inference(forward_demodulation,[],[f14088,f167])).
fof(f14088,plain,(
  ( ! [X2,X0,X1] : (s(bool,happ(s(fun(X0,bool),'const.combin.K_1'(s(bool,'F_0'))),s(X0,X1))) = s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|48|_fupd_2'(s(fun(bool,bool),'const.combin.K_1'(s(bool,'F_0'))),s('type.ring.ring'(bool),X2)))))) ) | $spl24),
  inference(unit_resulting_resolution,[],[f127,f13899,f13899,f4541])).
fof(f4541,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f4333,f118])).
fof(f4333,plain,(
  ( ! [X112,X115] : (s(bool,'F_0') != s(bool,X112) | s(bool,X112) = s(bool,X115) | s(bool,'T_0') = s(bool,X115)) )),
  inference(superposition,[],[f4199,f152])).
fof(f4199,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,'const.combin.K_2'(s(bool,X1),s(X2,X3))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f4044])).
fof(f4044,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,'const.combin.K_2'(s(bool,X1),s(X2,X3))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1336,f118])).
fof(f1336,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,'const.combin.K_2'(s(bool,X1),s(X2,X3))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f291])).
fof(f13899,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,happ(s(fun(X2,bool),'const.combin.K_1'(s(bool,'F_0'))),s(X2,X3)))) ) | $spl24),
  inference(subsumption_resolution,[],[f13747,f2628])).
fof(f13747,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X0,X1))) | s(bool,'T_0') != s(bool,happ(s(fun(X2,bool),'const.combin.K_1'(s(bool,'F_0'))),s(X2,X3)))) )),
  inference(condensation,[],[f13527])).
fof(f13527,plain,(
  ( ! [X499,X500,X502,X498,X501] : (s(bool,'T_0') != s(bool,happ(s(fun(X500,bool),'const.combin.K_1'(s(bool,'F_0'))),s(X500,X501))) | s(bool,X502) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X498,X499))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X498,X499)))) )),
  inference(superposition,[],[f2541,f13107])).
fof(f13107,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,happ(s(fun(X1,bool),'const.combin.K_1'(s(bool,'F_0'))),s(X1,X2))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(condensation,[],[f13106])).
fof(f13106,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) = s(bool,happ(s(fun(X2,bool),'const.combin.K_1'(s(bool,'F_0'))),s(X2,X3))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f12524])).
fof(f12524,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) = s(bool,happ(s(fun(X2,bool),'const.combin.K_1'(s(bool,'F_0'))),s(X2,X3))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f317,f118])).
fof(f317,plain,(
  ( ! [X24,X23,X21,X22] : (s(bool,X24) = s(bool,happ(s(fun(X21,bool),'const.combin.K_1'(s(bool,X22))),s(X21,X23))) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f188,f167])).
fof(f2541,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X1,X2))) | s(bool,X0) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X1,X2)))) )),
  inference(equality_factoring,[],[f2279])).
fof(f2279,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X1,X2))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(condensation,[],[f2278])).
fof(f2278,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X2,X3))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f2027])).
fof(f2027,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X2,X3))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f316,f118])).
fof(f13891,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X0),s(X1,X2)))) ) | $spl28),
  inference(cnf_transformation,[],[f13891_D])).
fof(f13891_D,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X0),s(X1,X2)))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f17235,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X0),s(X1,X2))) | s(bool,'T_0') != s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|48|_fupd_2'(s(fun(bool,bool),'const.combin.K_1'(s(bool,'F_0'))),s('type.ring.ring'(bool),X3)))))) ) | ($spl15 | $spl24)),
  inference(condensation,[],[f16991])).
fof(f16991,plain,(
  ( ! [X366,X368,X370,X367,X369,X371] : (s(bool,'T_0') != s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|48|_fupd_2'(s(fun(bool,bool),'const.combin.K_1'(s(bool,'F_0'))),s('type.ring.ring'(bool),X369))))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X366),s(X370,X371))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X366),s(X367,X368)))) ) | ($spl15 | $spl24)),
  inference(superposition,[],[f2763,f14766])).
fof(f2763,plain,(
  ( ! [X26,X24,X23,X25,X22] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,X22),s(X23,X24))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X22),s(X25,X26)))) ) | ($spl15 | $spl24)),
  inference(subsumption_resolution,[],[f2682,f202])).
fof(f2682,plain,(
  ( ! [X26,X24,X23,X25,X22] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,X22),s(X23,X24))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X22),s(X25,X26)))) ) | $spl24),
  inference(superposition,[],[f2628,f291])).
fof(f2627,plain,(
  ( ! [X96] : (s(bool,'T_0') = s(bool,X96)) ) | $spl22),
  inference(cnf_transformation,[],[f2627_D])).
fof(f2627_D,plain,(
  ( ! [X96] : (s(bool,'T_0') = s(bool,X96)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f17382,plain,(
  ( ! [X356,X359] : (s(bool,'T_0') != s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|48|_fupd_2'(s(fun(bool,bool),'const.combin.K_1'(s(bool,'F_0'))),s('type.ring.ring'(bool),X359))))) | s(bool,'T_0') = s(bool,X356)) ) | ($spl15 | $spl24)),
  inference(subsumption_resolution,[],[f16989,f2307])).
fof(f16989,plain,(
  ( ! [X356,X358,X357,X359] : (s(bool,'T_0') != s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|48|_fupd_2'(s(fun(bool,bool),'const.combin.K_1'(s(bool,'F_0'))),s('type.ring.ring'(bool),X359))))) | s(bool,'T_0') = s(bool,X356) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X356),s(X357,X358)))) ) | ($spl15 | $spl24)),
  inference(superposition,[],[f2307,f14766])).
fof(f17367,plain,(
  ( ! [X235,X234] : (s(bool,'T_0') != s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|48|_fupd_2'(s(fun(bool,bool),'const.combin.K_1'(s(bool,'F_0'))),s('type.ring.ring'(bool),X235))))) | s(bool,'T_0') = s(bool,X234)) ) | ($spl15 | $spl24)),
  inference(subsumption_resolution,[],[f16972,f13192])).
fof(f13192,plain,(
  ( ! [X4,X0,X3] : (s(bool,'T_0') != s(bool,happ(s(fun(X3,bool),'const.combin.K_1'(s(bool,X0))),s(X3,X4))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(forward_demodulation,[],[f13031,f152])).
fof(f13031,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X1,X2))) != s(bool,happ(s(fun(X3,bool),'const.combin.K_1'(s(bool,X0))),s(X3,X4)))) )),
  inference(condensation,[],[f13030])).
fof(f13030,plain,(
  ( ! [X187,X185,X188,X186,X184,X183] : (s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X187,X188))) != s(bool,happ(s(fun(X184,bool),'const.combin.K_1'(s(bool,X185))),s(X184,X186))) | s(bool,'T_0') = s(bool,X183) | s(bool,'T_0') = s(bool,X185)) )),
  inference(duplicate_literal_removal,[],[f12706])).
fof(f12706,plain,(
  ( ! [X187,X185,X188,X186,X184,X183] : (s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X187,X188))) != s(bool,happ(s(fun(X184,bool),'const.combin.K_1'(s(bool,X185))),s(X184,X186))) | s(bool,'T_0') = s(bool,X183) | s(bool,'T_0') = s(bool,X185) | s(bool,'T_0') = s(bool,X183)) )),
  inference(superposition,[],[f1348,f317])).
fof(f16972,plain,(
  ( ! [X237,X235,X236,X234] : (s(bool,'T_0') != s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|48|_fupd_2'(s(fun(bool,bool),'const.combin.K_1'(s(bool,'F_0'))),s('type.ring.ring'(bool),X235))))) | s(bool,'T_0') = s(bool,happ(s(fun(X236,bool),'const.combin.K_1'(s(bool,X234))),s(X236,X237))) | s(bool,'T_0') = s(bool,X234)) ) | ($spl15 | $spl24)),
  inference(superposition,[],[f7420,f14766])).
fof(f7420,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') != s(bool,X2) | s(bool,'T_0') = s(bool,happ(s(fun(X3,bool),'const.combin.K_1'(s(bool,X2))),s(X3,X4)))) ) | $spl15),
  inference(subsumption_resolution,[],[f7318,f202])).
fof(f7318,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') != s(bool,X2) | s(bool,'T_0') = s(bool,happ(s(fun(X3,bool),'const.combin.K_1'(s(bool,X2))),s(X3,X4))) | s(bool,'T_0') = s(bool,'F_0')) )),
  inference(superposition,[],[f1593,f190])).
fof(f190,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X2),s(X3,X4))) | s(bool,'F_0') = s(bool,X2)) )),
  inference(superposition,[],[f118,f152])).
fof(f1593,plain,(
  ( ! [X103,X101,X105,X102,X104] : (s(bool,X102) != s(bool,'const.combin.K_2'(s(bool,'T_0'),s(X104,X105))) | s(bool,'T_0') = s(bool,happ(s(fun(X101,bool),'const.combin.K_1'(s(bool,X102))),s(X101,X103)))) )),
  inference(superposition,[],[f1348,f167])).
fof(f17344,plain,(
  ( ! [X142,X141] : (s(bool,'T_0') != s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|48|_fupd_2'(s(fun(bool,bool),'const.combin.K_1'(s(bool,'F_0'))),s('type.ring.ring'(bool),X142))))) | s(bool,'T_0') = s(bool,X141)) ) | ($spl15 | $spl24)),
  inference(subsumption_resolution,[],[f16953,f2307])).
fof(f16953,plain,(
  ( ! [X144,X142,X143,X141] : (s(bool,'T_0') != s(bool,'const.ring.ring_R|48|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|48|_fupd_2'(s(fun(bool,bool),'const.combin.K_1'(s(bool,'F_0'))),s('type.ring.ring'(bool),X142))))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X141),s(X143,X144))) | s(bool,'T_0') = s(bool,X141)) ) | ($spl15 | $spl24)),
  inference(superposition,[],[f1533,f14766])).
fof(f1533,plain,(
  ( ! [X76,X77,X75] : (s(bool,'T_0') != s(bool,X75) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X75),s(X76,X77)))) ) | $spl15),
  inference(subsumption_resolution,[],[f1231,f202])).
fof(f1231,plain,(
  ( ! [X76,X77,X75] : (s(bool,'T_0') != s(bool,X75) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X75),s(X76,X77)))) ) | $spl15),
  inference(superposition,[],[f202,f291])).
fof(f15921,plain,(
  ( ! [X142] : (s(bool,'T_0') != s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|49|_fupd_2'(s(fun(bool,bool),'const.combin.K_1'(s(bool,'F_0'))),s('type.ring.ring'(bool),X142)))))) ) | $spl32),
  inference(cnf_transformation,[],[f15921_D])).
fof(f15921_D,plain,(
  ( ! [X142] : (s(bool,'T_0') != s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|49|_fupd_2'(s(fun(bool,bool),'const.combin.K_1'(s(bool,'F_0'))),s('type.ring.ring'(bool),X142)))))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f15956,plain,(
  ( ! [X366,X368,X367,X363] : (s(bool,'T_0') != s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|49|_fupd_2'(s(fun(bool,bool),'const.combin.K_1'(s(bool,'F_0'))),s('type.ring.ring'(bool),X366))))) | s(bool,'T_0') = s(bool,happ(s(fun(X367,bool),'const.combin.K_1'(s(bool,X363))),s(X367,X368)))) ) | ($spl15 | $spl24)),
  inference(subsumption_resolution,[],[f15578,f2900])).
fof(f15578,plain,(
  ( ! [X364,X366,X368,X365,X367,X363] : (s(bool,'T_0') != s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|49|_fupd_2'(s(fun(bool,bool),'const.combin.K_1'(s(bool,'F_0'))),s('type.ring.ring'(bool),X366))))) | s(bool,'T_0') = s(bool,happ(s(fun(X367,bool),'const.combin.K_1'(s(bool,X363))),s(X367,X368))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X363),s(X364,X365)))) ) | ($spl15 | $spl24)),
  inference(superposition,[],[f2900,f14502])).
fof(f14502,plain,(
  ( ! [X2,X3] : (s(bool,X3) = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|49|_fupd_2'(s(fun(bool,bool),'const.combin.K_1'(s(bool,'F_0'))),s('type.ring.ring'(bool),X2))))) | s(bool,'T_0') = s(bool,X3)) ) | ($spl15 | $spl24)),
  inference(subsumption_resolution,[],[f14372,f202])).
fof(f14372,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X3) = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|49|_fupd_2'(s(fun(bool,bool),'const.combin.K_1'(s(bool,'F_0'))),s('type.ring.ring'(bool),X2))))) | s(bool,'T_0') = s(bool,X3)) ) | $spl24),
  inference(superposition,[],[f14277,f188])).
fof(f14277,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|49|_fupd_2'(s(fun(bool,bool),'const.combin.K_1'(s(bool,'F_0'))),s('type.ring.ring'(bool),X2)))))) ) | $spl24),
  inference(forward_demodulation,[],[f14089,f167])).
fof(f14089,plain,(
  ( ! [X2,X0,X1] : (s(bool,happ(s(fun(X0,bool),'const.combin.K_1'(s(bool,'F_0'))),s(X0,X1))) = s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|49|_fupd_2'(s(fun(bool,bool),'const.combin.K_1'(s(bool,'F_0'))),s('type.ring.ring'(bool),X2)))))) ) | $spl24),
  inference(unit_resulting_resolution,[],[f131,f13899,f13899,f4541])).
fof(f15816,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X0),s(X1,X2))) | s(bool,'T_0') != s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|49|_fupd_2'(s(fun(bool,bool),'const.combin.K_1'(s(bool,'F_0'))),s('type.ring.ring'(bool),X3)))))) ) | ($spl15 | $spl24)),
  inference(condensation,[],[f15577])).
fof(f15577,plain,(
  ( ! [X358,X360,X362,X357,X359,X361] : (s(bool,'T_0') != s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|49|_fupd_2'(s(fun(bool,bool),'const.combin.K_1'(s(bool,'F_0'))),s('type.ring.ring'(bool),X360))))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X357),s(X361,X362))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X357),s(X358,X359)))) ) | ($spl15 | $spl24)),
  inference(superposition,[],[f2763,f14502])).
fof(f15955,plain,(
  ( ! [X350,X347] : (s(bool,'T_0') != s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|49|_fupd_2'(s(fun(bool,bool),'const.combin.K_1'(s(bool,'F_0'))),s('type.ring.ring'(bool),X350))))) | s(bool,'T_0') = s(bool,X347)) ) | ($spl15 | $spl24)),
  inference(subsumption_resolution,[],[f15575,f2307])).
fof(f15575,plain,(
  ( ! [X350,X349,X347,X348] : (s(bool,'T_0') != s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|49|_fupd_2'(s(fun(bool,bool),'const.combin.K_1'(s(bool,'F_0'))),s('type.ring.ring'(bool),X350))))) | s(bool,'T_0') = s(bool,X347) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X347),s(X348,X349)))) ) | ($spl15 | $spl24)),
  inference(superposition,[],[f2307,f14502])).
fof(f15943,plain,(
  ( ! [X235,X234] : (s(bool,'T_0') != s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|49|_fupd_2'(s(fun(bool,bool),'const.combin.K_1'(s(bool,'F_0'))),s('type.ring.ring'(bool),X235))))) | s(bool,'T_0') = s(bool,X234)) ) | ($spl15 | $spl24)),
  inference(subsumption_resolution,[],[f15560,f13192])).
fof(f15560,plain,(
  ( ! [X237,X235,X236,X234] : (s(bool,'T_0') != s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|49|_fupd_2'(s(fun(bool,bool),'const.combin.K_1'(s(bool,'F_0'))),s('type.ring.ring'(bool),X235))))) | s(bool,'T_0') = s(bool,happ(s(fun(X236,bool),'const.combin.K_1'(s(bool,X234))),s(X236,X237))) | s(bool,'T_0') = s(bool,X234)) ) | ($spl15 | $spl24)),
  inference(superposition,[],[f7420,f14502])).
fof(f15920,plain,(
  ( ! [X142,X141] : (s(bool,'T_0') != s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|49|_fupd_2'(s(fun(bool,bool),'const.combin.K_1'(s(bool,'F_0'))),s('type.ring.ring'(bool),X142))))) | s(bool,'T_0') = s(bool,X141)) ) | ($spl15 | $spl24)),
  inference(subsumption_resolution,[],[f15541,f2307])).
fof(f15541,plain,(
  ( ! [X144,X142,X143,X141] : (s(bool,'T_0') != s(bool,'const.ring.ring_R|49|_1'(s('type.ring.ring'(bool),'const.ring.ring_R|49|_fupd_2'(s(fun(bool,bool),'const.combin.K_1'(s(bool,'F_0'))),s('type.ring.ring'(bool),X142))))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X141),s(X143,X144))) | s(bool,'T_0') = s(bool,X141)) ) | ($spl15 | $spl24)),
  inference(superposition,[],[f1533,f14502])).
fof(f13867,plain,(
  ( ! [X173,X174] : (s(bool,'T_0') != s(bool,happ(s(fun(X173,bool),'const.combin.K_1'(s(bool,'F_0'))),s(X173,X174)))) ) | $spl26),
  inference(cnf_transformation,[],[f13867_D])).
fof(f13867_D,plain,(
  ( ! [X173,X174] : (s(bool,'T_0') != s(bool,happ(s(fun(X173,bool),'const.combin.K_1'(s(bool,'F_0'))),s(X173,X174)))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f13892,plain,(
  ( ! [X449,X451,X446,X452,X450] : (s(bool,'T_0') != s(bool,happ(s(fun(X449,bool),'const.combin.K_1'(s(bool,'F_0'))),s(X449,X450))) | s(bool,'T_0') = s(bool,happ(s(fun(X451,bool),'const.combin.K_1'(s(bool,X446))),s(X451,X452)))) )),
  inference(subsumption_resolution,[],[f13519,f2900])).
fof(f13519,plain,(
  ( ! [X447,X449,X451,X446,X452,X448,X450] : (s(bool,'T_0') != s(bool,happ(s(fun(X449,bool),'const.combin.K_1'(s(bool,'F_0'))),s(X449,X450))) | s(bool,'T_0') = s(bool,happ(s(fun(X451,bool),'const.combin.K_1'(s(bool,X446))),s(X451,X452))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X446),s(X447,X448)))) )),
  inference(superposition,[],[f2900,f13107])).
fof(f13748,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X0),s(X1,X2))) | s(bool,'T_0') != s(bool,happ(s(fun(X3,bool),'const.combin.K_1'(s(bool,'F_0'))),s(X3,X4)))) ) | ($spl15 | $spl24)),
  inference(condensation,[],[f13518])).
fof(f13518,plain,(
  ( ! [X445,X441,X443,X444,X440,X442,X439] : (s(bool,'T_0') != s(bool,happ(s(fun(X442,bool),'const.combin.K_1'(s(bool,'F_0'))),s(X442,X443))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X439),s(X444,X445))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X439),s(X440,X441)))) ) | ($spl15 | $spl24)),
  inference(superposition,[],[f2763,f13107])).
fof(f13890,plain,(
  ( ! [X430,X431,X427] : (s(bool,'T_0') != s(bool,happ(s(fun(X430,bool),'const.combin.K_1'(s(bool,'F_0'))),s(X430,X431))) | s(bool,'T_0') = s(bool,X427)) )),
  inference(subsumption_resolution,[],[f13516,f2307])).
fof(f13516,plain,(
  ( ! [X428,X430,X429,X431,X427] : (s(bool,'T_0') != s(bool,happ(s(fun(X430,bool),'const.combin.K_1'(s(bool,'F_0'))),s(X430,X431))) | s(bool,'T_0') = s(bool,X427) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X427),s(X428,X429)))) )),
  inference(superposition,[],[f2307,f13107])).
fof(f13888,plain,(
  ( ! [X294,X296,X295] : (s(bool,'T_0') != s(bool,happ(s(fun(X295,bool),'const.combin.K_1'(s(bool,'F_0'))),s(X295,X296))) | s(bool,'T_0') = s(bool,X294)) )),
  inference(forward_demodulation,[],[f13887,f167])).
fof(f13887,plain,(
  ( ! [X294,X296,X298,X295,X297] : (s(bool,happ(s(fun(X295,bool),'const.combin.K_1'(s(bool,'F_0'))),s(X295,X296))) != s(bool,happ(s(fun(X297,bool),'const.combin.K_1'(s(bool,'T_0'))),s(X297,X298))) | s(bool,'T_0') = s(bool,X294)) )),
  inference(subsumption_resolution,[],[f13506,f2307])).
fof(f13506,plain,(
  ( ! [X294,X300,X296,X298,X295,X297,X299] : (s(bool,happ(s(fun(X295,bool),'const.combin.K_1'(s(bool,'F_0'))),s(X295,X296))) != s(bool,happ(s(fun(X297,bool),'const.combin.K_1'(s(bool,'T_0'))),s(X297,X298))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X294),s(X299,X300))) | s(bool,'T_0') = s(bool,X294)) )),
  inference(superposition,[],[f10850,f13107])).
fof(f10850,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X0) != s(bool,happ(s(fun(X3,bool),'const.combin.K_1'(s(bool,'T_0'))),s(X3,X4))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X0),s(X1,X2)))) )),
  inference(condensation,[],[f10849])).
fof(f10849,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (s(bool,X26) != s(bool,happ(s(fun(X27,bool),'const.combin.K_1'(s(bool,'T_0'))),s(X27,X28))) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X26),s(X29,X30)))) )),
  inference(duplicate_literal_removal,[],[f10672])).
fof(f10672,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (s(bool,X26) != s(bool,happ(s(fun(X27,bool),'const.combin.K_1'(s(bool,'T_0'))),s(X27,X28))) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X26),s(X29,X30)))) )),
  inference(superposition,[],[f9971,f291])).
fof(f9971,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,happ(s(fun(X1,bool),'const.combin.K_1'(s(bool,'T_0'))),s(X1,X2))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f9953])).
fof(f9953,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,happ(s(fun(X1,bool),'const.combin.K_1'(s(bool,'T_0'))),s(X1,X2))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f292])).
fof(f292,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31) | s(bool,'T_0') = s(bool,happ(s(fun(X28,bool),'const.combin.K_1'(s(bool,X29))),s(X28,X30)))) )),
  inference(superposition,[],[f188,f167])).
fof(f13886,plain,(
  ( ! [X290,X289,X291] : (s(bool,'T_0') != s(bool,happ(s(fun(X290,bool),'const.combin.K_1'(s(bool,'F_0'))),s(X290,X291))) | s(bool,'T_0') = s(bool,X289)) )),
  inference(forward_demodulation,[],[f13750,f167])).
fof(f13750,plain,(
  ( ! [X292,X290,X293,X289,X291] : (s(bool,happ(s(fun(X290,bool),'const.combin.K_1'(s(bool,'F_0'))),s(X290,X291))) != s(bool,happ(s(fun(X292,bool),'const.combin.K_1'(s(bool,'T_0'))),s(X292,X293))) | s(bool,'T_0') = s(bool,X289)) )),
  inference(duplicate_literal_removal,[],[f13505])).
fof(f13505,plain,(
  ( ! [X292,X290,X293,X289,X291] : (s(bool,happ(s(fun(X290,bool),'const.combin.K_1'(s(bool,'F_0'))),s(X290,X291))) != s(bool,happ(s(fun(X292,bool),'const.combin.K_1'(s(bool,'T_0'))),s(X292,X293))) | s(bool,'T_0') = s(bool,X289) | s(bool,'T_0') = s(bool,X289)) )),
  inference(superposition,[],[f9971,f13107])).
fof(f13885,plain,(
  ( ! [X285,X284,X286] : (s(bool,'T_0') != s(bool,happ(s(fun(X285,bool),'const.combin.K_1'(s(bool,'F_0'))),s(X285,X286))) | s(bool,'T_0') = s(bool,X284)) ) | $spl15),
  inference(subsumption_resolution,[],[f13504,f13192])).
fof(f13504,plain,(
  ( ! [X288,X285,X287,X284,X286] : (s(bool,'T_0') != s(bool,happ(s(fun(X285,bool),'const.combin.K_1'(s(bool,'F_0'))),s(X285,X286))) | s(bool,'T_0') = s(bool,happ(s(fun(X287,bool),'const.combin.K_1'(s(bool,X284))),s(X287,X288))) | s(bool,'T_0') = s(bool,X284)) ) | $spl15),
  inference(superposition,[],[f7420,f13107])).
fof(f13866,plain,(
  ( ! [X173,X174,X172] : (s(bool,'T_0') != s(bool,happ(s(fun(X173,bool),'const.combin.K_1'(s(bool,'F_0'))),s(X173,X174))) | s(bool,'T_0') = s(bool,X172)) ) | $spl15),
  inference(subsumption_resolution,[],[f13485,f2307])).
fof(f13485,plain,(
  ( ! [X175,X173,X176,X174,X172] : (s(bool,'T_0') != s(bool,happ(s(fun(X173,bool),'const.combin.K_1'(s(bool,'F_0'))),s(X173,X174))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X172),s(X175,X176))) | s(bool,'T_0') = s(bool,X172)) ) | $spl15),
  inference(superposition,[],[f1533,f13107])).
fof(f2628,plain,(
  ( ! [X97,X98] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X97,X98)))) ) | $spl24),
  inference(cnf_transformation,[],[f2628_D])).
fof(f2628_D,plain,(
  ( ! [X97,X98] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X97,X98)))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f2626,plain,(
  ( ! [X97,X98,X96] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X97,X98))) | s(bool,'T_0') = s(bool,X96)) ) | $spl15),
  inference(subsumption_resolution,[],[f2470,f2307])).
fof(f2470,plain,(
  ( ! [X99,X97,X100,X98,X96] : (s(bool,'T_0') != s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X97,X98))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X96),s(X99,X100))) | s(bool,'T_0') = s(bool,X96)) ) | $spl15),
  inference(superposition,[],[f1533,f2279])).
fof(f911,plain,(
  ( ! [X12,X17,X16] : (s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X12),s(X16,X17))) | s(bool,'T_0') = s(bool,X12)) ) | $spl16),
  inference(cnf_transformation,[],[f911_D])).
fof(f911_D,plain,(
  ( ! [X12,X17,X16] : (s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X12),s(X16,X17))) | s(bool,'T_0') = s(bool,X12)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f913,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X25,X26)))) ) | $spl20),
  inference(cnf_transformation,[],[f913_D])).
fof(f913_D,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X25,X26)))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f758,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'F_0'),s(X25,X26))) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X24),s(X27,X28))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f240,f118])).
fof(f240,plain,(
  ( ! [X14,X12,X15,X13,X11] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X11),s(X14,X15))) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X11),s(X12,X13)))) )),
  inference(superposition,[],[f190,f152])).
fof(f912,plain,(
  ( ! [X14,X15,X13] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X13),s(X14,X15))) | s(bool,'T_0') = s(bool,X13)) ) | $spl18),
  inference(cnf_transformation,[],[f912_D])).
fof(f912_D,plain,(
  ( ! [X14,X15,X13] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X13),s(X14,X15))) | s(bool,'T_0') = s(bool,X13)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f757,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X19),s(X20,X21))) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X18),s(X22,X23))) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f240,f188])).
fof(f756,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X13),s(X14,X15))) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X12),s(X16,X17))) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f240,f188])).
fof(f200,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl12),
  inference(cnf_transformation,[],[f200_D])).
fof(f200_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f202,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl15),
  inference(cnf_transformation,[],[f202_D])).
fof(f202_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f199,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f118])).
fof(f173,plain,(
  $false | ($spl1 | $spl2 | $spl4 | $spl6 | $spl8 | $spl10)),
  inference(trivial_inequality_removal,[],[f172])).
fof(f172,plain,(
  s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) != s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl1 | $spl2 | $spl4 | $spl6 | $spl8 | $spl10)),
  inference(forward_demodulation,[],[f171,f157])).
fof(f171,plain,(
  s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) != s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl1 | $spl4 | $spl6 | $spl8 | $spl10)),
  inference(forward_demodulation,[],[f170,f159])).
fof(f170,plain,(
  s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) != s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK7))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl1 | $spl6 | $spl8 | $spl10)),
  inference(forward_demodulation,[],[f169,f161])).
fof(f169,plain,(
  s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) != s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK7))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl1 | $spl8 | $spl10)),
  inference(forward_demodulation,[],[f168,f163])).
fof(f168,plain,(
  s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) != s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK7))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | ($spl1 | $spl10)),
  inference(forward_demodulation,[],[f156,f165])).
fof(f155,plain,(
  s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK7))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | $spl0),
  inference(cnf_transformation,[],[f155_D])).
fof(f155_D,plain,(
  s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK7))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f157,plain,(
  s(sK0,sK1) = s(sK0,sK6) | $spl2),
  inference(cnf_transformation,[],[f157_D])).
fof(f157_D,plain,(
  s(sK0,sK1) = s(sK0,sK6) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f110,plain,(
  s(sK0,sK1) = s(sK0,sK6) | s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK7))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0')))))))))))),
  inference(cnf_transformation,[],[f106])).
fof(f106,plain,(
  (s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK7))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | (s(sK0,sK1) = s(sK0,sK6) & s(sK0,sK2) = s(sK0,sK7) & s(fun(sK0,fun(sK0,sK0)),sK3) = s(fun(sK0,fun(sK0,sK0)),sK8) & s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK9) & s(fun(sK0,sK0),sK5) = s(fun(sK0,sK0),sK10))) & (s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) != s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK7))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | s(sK0,sK1) != s(sK0,sK6) | s(sK0,sK2) != s(sK0,sK7) | s(fun(sK0,fun(sK0,sK0)),sK3) != s(fun(sK0,fun(sK0,sK0)),sK8) | s(fun(sK0,fun(sK0,sK0)),sK4) != s(fun(sK0,fun(sK0,sK0)),sK9) | s(fun(sK0,sK0),sK5) != s(fun(sK0,sK0),sK10))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK9,sK10])],[f105])).
fof(f105,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : ((s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X1))),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X2))),s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X3))),s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X4))),s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),'const.combin.K_1'(s(fun(X0,X0),X5))),s('type.ring.ring'(X0),'const.bool.ARB_0'))))))))))) = s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X6))),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X7))),s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X8))),s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X9))),s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),'const.combin.K_1'(s(fun(X0,X0),X10))),s('type.ring.ring'(X0),'const.bool.ARB_0'))))))))))) | (s(X0,X1) = s(X0,X6) & s(X0,X2) = s(X0,X7) & s(fun(X0,fun(X0,X0)),X3) = s(fun(X0,fun(X0,X0)),X8) & s(fun(X0,fun(X0,X0)),X4) = s(fun(X0,fun(X0,X0)),X9) & s(fun(X0,X0),X5) = s(fun(X0,X0),X10))) & (s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X1))),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X2))),s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X3))),s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X4))),s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),'const.combin.K_1'(s(fun(X0,X0),X5))),s('type.ring.ring'(X0),'const.bool.ARB_0'))))))))))) != s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X6))),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X7))),s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X8))),s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X9))),s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),'const.combin.K_1'(s(fun(X0,X0),X10))),s('type.ring.ring'(X0),'const.bool.ARB_0'))))))))))) | s(X0,X1) != s(X0,X6) | s(X0,X2) != s(X0,X7) | s(fun(X0,fun(X0,X0)),X3) != s(fun(X0,fun(X0,X0)),X8) | s(fun(X0,fun(X0,X0)),X4) != s(fun(X0,fun(X0,X0)),X9) | s(fun(X0,X0),X5) != s(fun(X0,X0),X10)))),
  inference(flattening,[],[f104])).
fof(f104,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : ((s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X1))),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X2))),s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X3))),s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X4))),s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),'const.combin.K_1'(s(fun(X0,X0),X5))),s('type.ring.ring'(X0),'const.bool.ARB_0'))))))))))) = s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X6))),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X7))),s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X8))),s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X9))),s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),'const.combin.K_1'(s(fun(X0,X0),X10))),s('type.ring.ring'(X0),'const.bool.ARB_0'))))))))))) | (s(X0,X1) = s(X0,X6) & s(X0,X2) = s(X0,X7) & s(fun(X0,fun(X0,X0)),X3) = s(fun(X0,fun(X0,X0)),X8) & s(fun(X0,fun(X0,X0)),X4) = s(fun(X0,fun(X0,X0)),X9) & s(fun(X0,X0),X5) = s(fun(X0,X0),X10))) & (s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X1))),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X2))),s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X3))),s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X4))),s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),'const.combin.K_1'(s(fun(X0,X0),X5))),s('type.ring.ring'(X0),'const.bool.ARB_0'))))))))))) != s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X6))),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X7))),s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X8))),s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X9))),s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),'const.combin.K_1'(s(fun(X0,X0),X10))),s('type.ring.ring'(X0),'const.bool.ARB_0'))))))))))) | (s(X0,X1) != s(X0,X6) | s(X0,X2) != s(X0,X7) | s(fun(X0,fun(X0,X0)),X3) != s(fun(X0,fun(X0,X0)),X8) | s(fun(X0,fun(X0,X0)),X4) != s(fun(X0,fun(X0,X0)),X9) | s(fun(X0,X0),X5) != s(fun(X0,X0),X10))))),
  inference(nnf_transformation,[],[f102])).
fof(f102,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : (s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X1))),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X2))),s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X3))),s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X4))),s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),'const.combin.K_1'(s(fun(X0,X0),X5))),s('type.ring.ring'(X0),'const.bool.ARB_0'))))))))))) = s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X6))),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X7))),s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X8))),s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X9))),s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),'const.combin.K_1'(s(fun(X0,X0),X10))),s('type.ring.ring'(X0),'const.bool.ARB_0'))))))))))) <~> (s(X0,X1) = s(X0,X6) & s(X0,X2) = s(X0,X7) & s(fun(X0,fun(X0,X0)),X3) = s(fun(X0,fun(X0,X0)),X8) & s(fun(X0,fun(X0,X0)),X4) = s(fun(X0,fun(X0,X0)),X9) & s(fun(X0,X0),X5) = s(fun(X0,X0),X10)))),
  inference(ennf_transformation,[],[f38])).
fof(f38,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : (s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X1))),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X2))),s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X3))),s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X4))),s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),'const.combin.K_1'(s(fun(X0,X0),X5))),s('type.ring.ring'(X0),'const.bool.ARB_0'))))))))))) = s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X6))),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X7))),s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X8))),s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X9))),s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),'const.combin.K_1'(s(fun(X0,X0),X10))),s('type.ring.ring'(X0),'const.bool.ARB_0'))))))))))) <=> (s(X0,X1) = s(X0,X6) & s(X0,X2) = s(X0,X7) & s(fun(X0,fun(X0,X0)),X3) = s(fun(X0,fun(X0,X0)),X8) & s(fun(X0,fun(X0,X0)),X4) = s(fun(X0,fun(X0,X0)),X9) & s(fun(X0,X0),X5) = s(fun(X0,X0),X10)))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ~! [X0] : ! [X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : (s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X1))),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X2))),s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X3))),s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X4))),s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),'const.combin.K_1'(s(fun(X0,X0),X5))),s('type.ring.ring'(X0),'const.bool.ARB_0'))))))))))) = s('type.ring.ring'(X0),'const.ring.ring_R|48|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X6))),s('type.ring.ring'(X0),'const.ring.ring_R|49|_fupd_2'(s(fun(X0,X0),'const.combin.K_1'(s(X0,X7))),s('type.ring.ring'(X0),'const.ring.ring_RP_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X8))),s('type.ring.ring'(X0),'const.ring.ring_RM_fupd_2'(s(fun(fun(X0,fun(X0,X0)),fun(X0,fun(X0,X0))),'const.combin.K_1'(s(fun(X0,fun(X0,X0)),X9))),s('type.ring.ring'(X0),'const.ring.ring_RN_fupd_2'(s(fun(fun(X0,X0),fun(X0,X0)),'const.combin.K_1'(s(fun(X0,X0),X10))),s('type.ring.ring'(X0),'const.bool.ARB_0'))))))))))) <=> (s(X0,X1) = s(X0,X6) & s(X0,X2) = s(X0,X7) & s(fun(X0,fun(X0,X0)),X3) = s(fun(X0,fun(X0,X0)),X8) & s(fun(X0,fun(X0,X0)),X4) = s(fun(X0,fun(X0,X0)),X9) & s(fun(X0,X0),X5) = s(fun(X0,X0),X10)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10,X11,X12,X13,X14,X15,X16] : (s('type.ring.ring'(X6),'const.ring.ring_R|48|_fupd_2'(s(fun(X6,X6),'const.combin.K_1'(s(X6,X7))),s('type.ring.ring'(X6),'const.ring.ring_R|49|_fupd_2'(s(fun(X6,X6),'const.combin.K_1'(s(X6,X8))),s('type.ring.ring'(X6),'const.ring.ring_RP_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),'const.combin.K_1'(s(fun(X6,fun(X6,X6)),X9))),s('type.ring.ring'(X6),'const.ring.ring_RM_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),'const.combin.K_1'(s(fun(X6,fun(X6,X6)),X10))),s('type.ring.ring'(X6),'const.ring.ring_RN_fupd_2'(s(fun(fun(X6,X6),fun(X6,X6)),'const.combin.K_1'(s(fun(X6,X6),X11))),s('type.ring.ring'(X6),'const.bool.ARB_0'))))))))))) = s('type.ring.ring'(X6),'const.ring.ring_R|48|_fupd_2'(s(fun(X6,X6),'const.combin.K_1'(s(X6,X12))),s('type.ring.ring'(X6),'const.ring.ring_R|49|_fupd_2'(s(fun(X6,X6),'const.combin.K_1'(s(X6,X13))),s('type.ring.ring'(X6),'const.ring.ring_RP_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),'const.combin.K_1'(s(fun(X6,fun(X6,X6)),X14))),s('type.ring.ring'(X6),'const.ring.ring_RM_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),'const.combin.K_1'(s(fun(X6,fun(X6,X6)),X15))),s('type.ring.ring'(X6),'const.ring.ring_RN_fupd_2'(s(fun(fun(X6,X6),fun(X6,X6)),'const.combin.K_1'(s(fun(X6,X6),X16))),s('type.ring.ring'(X6),'const.bool.ARB_0'))))))))))) <=> (s(X6,X7) = s(X6,X12) & s(X6,X8) = s(X6,X13) & s(fun(X6,fun(X6,X6)),X9) = s(fun(X6,fun(X6,X6)),X14) & s(fun(X6,fun(X6,X6)),X10) = s(fun(X6,fun(X6,X6)),X15) & s(fun(X6,X6),X11) = s(fun(X6,X6),X16)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10,X11,X12,X13,X14,X15,X16] : (s('type.ring.ring'(X6),'const.ring.ring_R|48|_fupd_2'(s(fun(X6,X6),'const.combin.K_1'(s(X6,X7))),s('type.ring.ring'(X6),'const.ring.ring_R|49|_fupd_2'(s(fun(X6,X6),'const.combin.K_1'(s(X6,X8))),s('type.ring.ring'(X6),'const.ring.ring_RP_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),'const.combin.K_1'(s(fun(X6,fun(X6,X6)),X9))),s('type.ring.ring'(X6),'const.ring.ring_RM_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),'const.combin.K_1'(s(fun(X6,fun(X6,X6)),X10))),s('type.ring.ring'(X6),'const.ring.ring_RN_fupd_2'(s(fun(fun(X6,X6),fun(X6,X6)),'const.combin.K_1'(s(fun(X6,X6),X11))),s('type.ring.ring'(X6),'const.bool.ARB_0'))))))))))) = s('type.ring.ring'(X6),'const.ring.ring_R|48|_fupd_2'(s(fun(X6,X6),'const.combin.K_1'(s(X6,X12))),s('type.ring.ring'(X6),'const.ring.ring_R|49|_fupd_2'(s(fun(X6,X6),'const.combin.K_1'(s(X6,X13))),s('type.ring.ring'(X6),'const.ring.ring_RP_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),'const.combin.K_1'(s(fun(X6,fun(X6,X6)),X14))),s('type.ring.ring'(X6),'const.ring.ring_RM_fupd_2'(s(fun(fun(X6,fun(X6,X6)),fun(X6,fun(X6,X6))),'const.combin.K_1'(s(fun(X6,fun(X6,X6)),X15))),s('type.ring.ring'(X6),'const.ring.ring_RN_fupd_2'(s(fun(fun(X6,X6),fun(X6,X6)),'const.combin.K_1'(s(fun(X6,X6),X16))),s('type.ring.ring'(X6),'const.bool.ARB_0'))))))))))) <=> (s(X6,X7) = s(X6,X12) & s(X6,X8) = s(X6,X13) & s(fun(X6,fun(X6,X6)),X9) = s(fun(X6,fun(X6,X6)),X14) & s(fun(X6,fun(X6,X6)),X10) = s(fun(X6,fun(X6,X6)),X15) & s(fun(X6,X6),X11) = s(fun(X6,X6),X16)))),
  file('thm.ring.ring_literal_|49||49|','thm.ring.ring_literal_|49||49|')).
fof(f159,plain,(
  s(sK0,sK2) = s(sK0,sK7) | $spl4),
  inference(cnf_transformation,[],[f159_D])).
fof(f159_D,plain,(
  s(sK0,sK2) = s(sK0,sK7) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f111,plain,(
  s(sK0,sK2) = s(sK0,sK7) | s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK7))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0')))))))))))),
  inference(cnf_transformation,[],[f106])).
fof(f161,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK3) = s(fun(sK0,fun(sK0,sK0)),sK8) | $spl6),
  inference(cnf_transformation,[],[f161_D])).
fof(f161_D,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK3) = s(fun(sK0,fun(sK0,sK0)),sK8) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f112,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK3) = s(fun(sK0,fun(sK0,sK0)),sK8) | s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK7))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0')))))))))))),
  inference(cnf_transformation,[],[f106])).
fof(f163,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK9) | $spl8),
  inference(cnf_transformation,[],[f163_D])).
fof(f163_D,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK9) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f113,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) = s(fun(sK0,fun(sK0,sK0)),sK9) | s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK7))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0')))))))))))),
  inference(cnf_transformation,[],[f106])).
fof(f165,plain,(
  s(fun(sK0,sK0),sK5) = s(fun(sK0,sK0),sK10) | $spl10),
  inference(cnf_transformation,[],[f165_D])).
fof(f165_D,plain,(
  s(fun(sK0,sK0),sK5) = s(fun(sK0,sK0),sK10) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f114,plain,(
  s(fun(sK0,sK0),sK5) = s(fun(sK0,sK0),sK10) | s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) = s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK7))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0')))))))))))),
  inference(cnf_transformation,[],[f106])).
fof(f156,plain,(
  s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) != s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK7))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) | $spl1),
  inference(cnf_transformation,[],[f156_D])).
fof(f156_D,plain,(
  s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) != s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK7))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f158,plain,(
  s(sK0,sK1) != s(sK0,sK6) | $spl3),
  inference(cnf_transformation,[],[f158_D])).
fof(f158_D,plain,(
  s(sK0,sK1) != s(sK0,sK6) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f160,plain,(
  s(sK0,sK2) != s(sK0,sK7) | $spl5),
  inference(cnf_transformation,[],[f160_D])).
fof(f160_D,plain,(
  s(sK0,sK2) != s(sK0,sK7) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f162,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK3) != s(fun(sK0,fun(sK0,sK0)),sK8) | $spl7),
  inference(cnf_transformation,[],[f162_D])).
fof(f162_D,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK3) != s(fun(sK0,fun(sK0,sK0)),sK8) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f164,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) != s(fun(sK0,fun(sK0,sK0)),sK9) | $spl9),
  inference(cnf_transformation,[],[f164_D])).
fof(f164_D,plain,(
  s(fun(sK0,fun(sK0,sK0)),sK4) != s(fun(sK0,fun(sK0,sK0)),sK9) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f166,plain,(
  s(fun(sK0,sK0),sK5) != s(fun(sK0,sK0),sK10) | $spl11),
  inference(cnf_transformation,[],[f166_D])).
fof(f166_D,plain,(
  s(fun(sK0,sK0),sK5) != s(fun(sK0,sK0),sK10) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f115,plain,(
  s(fun(sK0,sK0),sK5) != s(fun(sK0,sK0),sK10) | s(fun(sK0,fun(sK0,sK0)),sK4) != s(fun(sK0,fun(sK0,sK0)),sK9) | s(fun(sK0,fun(sK0,sK0)),sK3) != s(fun(sK0,fun(sK0,sK0)),sK8) | s(sK0,sK2) != s(sK0,sK7) | s(sK0,sK1) != s(sK0,sK6) | s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK1))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK2))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK3))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK4))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK5))),s('type.ring.ring'(sK0),'const.bool.ARB_0'))))))))))) != s('type.ring.ring'(sK0),'const.ring.ring_R|48|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK6))),s('type.ring.ring'(sK0),'const.ring.ring_R|49|_fupd_2'(s(fun(sK0,sK0),'const.combin.K_1'(s(sK0,sK7))),s('type.ring.ring'(sK0),'const.ring.ring_RP_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK8))),s('type.ring.ring'(sK0),'const.ring.ring_RM_fupd_2'(s(fun(fun(sK0,fun(sK0,sK0)),fun(sK0,fun(sK0,sK0))),'const.combin.K_1'(s(fun(sK0,fun(sK0,sK0)),sK9))),s('type.ring.ring'(sK0),'const.ring.ring_RN_fupd_2'(s(fun(fun(sK0,sK0),fun(sK0,sK0)),'const.combin.K_1'(s(fun(sK0,sK0),sK10))),s('type.ring.ring'(sK0),'const.bool.ARB_0')))))))))))),
  inference(cnf_transformation,[],[f106])).
% SZS output end Proof for thm.ring
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 116
% Passive clauses: 712
% Generated clauses: 21348
% Final active clauses: 99
% Final passive clauses: 578
% Input formulas: 35
% Initial clauses: 45
% 
% Duplicate literals: 840
% Trivial inequalities: 182
% Fw subsumption resolutions: 1040
% Fw demodulations: 3221
% Bw demodulations: 10
% Condensations: 2200
% 
% Simple tautologies: 190
% Equational tautologies: 101
% Forward subsumptions: 12679
% Fw demodulations to eq. taut.: 140
% Bw demodulations to eq. taut.: 2
% 
% Unit resulting resolution: 864
% Forward superposition: 6714
% Backward superposition: 5892
% Self superposition: 230
% Equality factoring: 68
% Equality resolution: 17
% 
% Split clauses: 28
% Split components: 60
% SAT solver clauses: 789
% SAT solver unit clauses: 64
% SAT solver binary clauses: 186
% 
% Sat splits: 28
% Sat splitting refutations: 19
% 
% Memory used [KB]: 14711
% Time elapsed: 0.942 s
% ------------------------------
% ------------------------------
% Success in time 2.25 s
