Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.prelim
% (70382)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 403
% Passive clauses: 4948
% Generated clauses: 30207
% Final active clauses: 198
% Final passive clauses: 3702
% Input formulas: 63
% Initial clauses: 79
% 
% Duplicate literals: 387
% Trivial inequalities: 432
% Fw subsumption resolutions: 897
% Fw demodulations: 7151
% Bw demodulations: 154
% 
% Simple tautologies: 84
% Equational tautologies: 30
% Forward subsumptions: 16679
% Fw demodulations to eq. taut.: 19
% Bw demodulations to eq. taut.: 13
% 
% Binary resolution: 36
% Forward superposition: 11367
% Backward superposition: 9336
% Self superposition: 125
% Equality factoring: 44
% Equality resolution: 34
% 
% Split clauses: 95
% Split components: 190
% SAT solver clauses: 180
% SAT solver unit clauses: 8
% SAT solver binary clauses: 45
% 
% Sat splits: 95
% Sat splitting refutations: 21
% 
% TWLsolver clauses: 180
% TWLsolver calls for satisfiability: 63
% 
% Memory used [KB]: 15351
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.prelim
% (70395)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 345
% Passive clauses: 1479
% Generated clauses: 29783
% Final active clauses: 199
% Final passive clauses: 558
% Input formulas: 63
% Initial clauses: 79
% 
% Duplicate literals: 535
% Trivial inequalities: 428
% Fw subsumption resolutions: 1098
% Fw demodulations: 6257
% Bw demodulations: 580
% Condensations: 510
% Global subsumptions: 12
% 
% Simple tautologies: 83
% Equational tautologies: 33
% Forward subsumptions: 16062
% Fw demodulations to eq. taut.: 10
% Bw demodulations to eq. taut.: 95
% 
% Binary resolution: 36
% Forward superposition: 9693
% Backward superposition: 8509
% Self superposition: 109
% Equality factoring: 50
% Equality resolution: 34
% 
% Split clauses: 67
% Split components: 134
% SAT solver clauses: 3031
% SAT solver unit clauses: 162
% SAT solver binary clauses: 49
% SAT solver learnt clauses: 24
% 
% Sat splits: 67
% Sat splitting refutations: 233
% 
% TWLsolver clauses: 2983
% TWLsolver calls for satisfiability: 413
% 
% Memory used [KB]: 19317
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.prelim
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.prelim
% SZS output start Proof for thm.prelim
fof(f31032,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f382,f387,f385,f692,f707,f705,f1058,f1060,f1068,f1066,f1071,f3208,f3210,f3218,f3216,f3486,f3488,f3496,f3494,f7310,f7312,f7323,f7321,f7349,f7351,f7362,f7360,f9867,f9872,f9759,f9869,f9871,f10060,f10065,f9952,f10062,f10064,f10253,f10258,f10145,f10255,f10257,f10446,f10451,f10338,f10448,f10450,f10590,f10592,f10597,f10601,f10599,f10602,f10604,f10609,f10613,f10611,f10544,f10672,f10594,f10545,f10676,f10674,f10546,f10678,f10606,f10547,f10682,f10680,f10686,f10689,f10771,f10773,f10781,f10782,f10785,f10786,f11187,f11757,f11758,f11759,f11760,f11761,f11762,f11763,f11764,f11765,f11766,f11767,f11768,f11769,f11770,f11771,f11772,f11773,f11774,f11775,f11776,f11777,f11778,f11788,f11792,f11796,f11797,f11798,f11799,f11800,f11801,f11802,f11803,f11804,f11805,f11806,f11807,f11808,f11818,f11820,f11824,f11826,f11832,f11833,f11847,f11850,f11851,f11853,f19811,f19813,f19818,f19822,f19820,f19823,f19825,f19830,f19834,f19832,f19765,f19874,f19815,f19766,f19878,f19876,f19767,f19880,f19827,f19768,f19884,f19882,f19888,f19891,f23062,f23066,f23064,f23069,f23073,f23071,f23113,f23065,f23115,f23117,f23072,f23119,f24475,f24518,f24516,f24477,f24523,f24521,f24506,f24539,f24537,f24507,f24543,f24541,f24508,f24547,f24545,f24509,f24551,f24549,f30956,f30324])).
fof(f30324,plain,(
  $false | ($spl9 | $spl10)),
  inference(unit_resulting_resolution,[],[f148,f30126])).
fof(f30126,plain,(
  ( ! [X2,X0,X3,X1] : (s(X3,X0) = s(X3,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.GREATER_0'),s(X3,X2),s(X3,X1),s(X3,X0)))) ) | ($spl9 | $spl10)),
  inference(forward_demodulation,[],[f30125,f199])).
fof(f199,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f103])).
fof(f103,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : ! [X11,X12] : s(X6,X12) = s(X6,'const.bool.COND_3'(s(bool,'F_0'),s(X6,X11),s(X6,X12)))),
  file('thm.prelim.ordering_case_def|split|2','thm.bool.COND_CLAUSES|split|1')).
fof(f30125,plain,(
  ( ! [X2,X0,X3,X1] : (s(X3,'const.bool.COND_3'(s(bool,'F_0'),s(X3,X2),s(X3,X0))) = s(X3,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.GREATER_0'),s(X3,X2),s(X3,X1),s(X3,X0)))) ) | ($spl9 | $spl10)),
  inference(backward_demodulation,[],[f30123,f9699])).
fof(f9699,plain,(
  ( ! [X2,X0,X3,X1] : (s(X3,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.GREATER_0'),s(X3,X2),s(X3,X1),s(X3,X0))) = s(X3,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(X3,X2),s(X3,X0)))) ) | ($spl9 | $spl10)),
  inference(forward_demodulation,[],[f9676,f199])).
fof(f9676,plain,(
  ( ! [X2,X0,X3,X1] : (s(X3,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.GREATER_0'),s(X3,X2),s(X3,X1),s(X3,X0))) = s(X3,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(X3,X2),s(X3,'const.bool.COND_3'(s(bool,'F_0'),s(X3,X1),s(X3,X0)))))) ) | ($spl9 | $spl10)),
  inference(superposition,[],[f4870,f9612])).
fof(f9612,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,sK4(X0,X1,X2,'const.prelim.GREATER_0',X3))) ) | ($spl9 | $spl10)),
  inference(forward_demodulation,[],[f9552,f1068])).
fof(f9552,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))) = s(bool,sK4(X0,X1,X2,'const.prelim.GREATER_0',X3))) ) | $spl9),
  inference(unit_resulting_resolution,[],[f7299,f4741,f8854])).
fof(f8854,plain,(
  ( ! [X8,X7] : (p(s(bool,X7)) | s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f149,f371])).
fof(f371,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f181,f181])).
fof(f181,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.prelim.ordering_case_def|split|2',aHL_BOOL_CASES)).
fof(f149,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.prelim.ordering_case_def|split|2',aHL_TRUTH)).
fof(f4741,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))) | $spl9),
  inference(backward_demodulation,[],[f4713,f1065])).
fof(f1065,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl9),
  inference(cnf_transformation,[],[f1065_D])).
fof(f1065_D,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0'))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f4713,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.ZERO_0')),
  inference(forward_demodulation,[],[f4712,f733])).
fof(f733,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(forward_demodulation,[],[f717,f199])).
fof(f717,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')))))) )),
  inference(unit_resulting_resolution,[],[f199,f191])).
fof(f191,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) & (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(nnf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X13] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X13))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_distrib|split|18')).
fof(f4712,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(forward_demodulation,[],[f4711,f4424])).
fof(f4424,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f3830,f190])).
fof(f190,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f134])).
fof(f3830,plain,(
  ( ! [X2] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X2)))))) )),
  inference(forward_demodulation,[],[f3719,f167])).
fof(f167,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X13] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))),
  file('thm.prelim.ordering_case_def|split|2','thm.arithmetic.MULT|split|0')).
fof(f3719,plain,(
  ( ! [X2] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X2))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))))) )),
  inference(superposition,[],[f212,f733])).
fof(f212,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X13,X14] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X13),s('type.num.num',X14))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X14)))))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_distrib|split|5')).
fof(f4711,plain,(
  ( ! [X62] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X62)))))) )),
  inference(subsumption_resolution,[],[f4508,f4578])).
fof(f4578,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))),
  inference(forward_demodulation,[],[f4435,f166])).
fof(f166,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X13] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X13),s('type.num.num','const.num.|48|_0')))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_distrib|split|4')).
fof(f4435,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f3830,f3830])).
fof(f4508,plain,(
  ( ! [X62] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X62)))))) )),
  inference(superposition,[],[f190,f3830])).
fof(f7299,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,sK4(X0,X1,X2,'const.prelim.GREATER_0',X3)))) )),
  inference(unit_resulting_resolution,[],[f6978,f234])).
fof(f234,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,sK4(X4,X3,X2,X1,X0))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1)))) )),
  inference(forward_demodulation,[],[f200,f155])).
fof(f155,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f23])).
fof(f23,axiom,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_distrib|split|13')).
fof(f200,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | ~p(s(bool,sK4(X4,X3,X2,X1,X0)))) )),
  inference(cnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,sK4(X4,X3,X2,X1,X0))) | s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) & (s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | p(s(bool,sK4(X4,X3,X2,X1,X0)))) & s(X0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X1),s(X0,X2),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(X0,X2),s(X0,'const.bool.COND_3'(s(bool,sK4(X4,X3,X2,X1,X0)),s(X0,X3),s(X0,X4))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f139])).
fof(f139,plain,(
  ! [X0,X1,X2,X3,X4] : ? [X5] : ((~p(s(bool,X5)) | s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) & (s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | p(s(bool,X5))) & s(X0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X1),s(X0,X2),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(X0,X2),s(X0,'const.bool.COND_3'(s(bool,X5),s(X0,X3),s(X0,X4))))))),
  inference(flattening,[],[f138])).
fof(f138,plain,(
  ! [X0,X1,X2,X3,X4] : ? [X5] : (((~p(s(bool,X5)) | s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) & (s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | p(s(bool,X5)))) & s(X0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X1),s(X0,X2),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(X0,X2),s(X0,'const.bool.COND_3'(s(bool,X5),s(X0,X3),s(X0,X4))))))),
  inference(nnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0,X1,X2,X3,X4] : ? [X5] : ((p(s(bool,X5)) <=> s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) & s(X0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X1),s(X0,X2),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(X0,X2),s(X0,'const.bool.COND_3'(s(bool,X5),s(X0,X3),s(X0,X4))))))),
  inference(flattening,[],[f105])).
fof(f105,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : ? [X5] : ((p(s(bool,X5)) <=> s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) & s(X0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X1),s(X0,X2),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(X0,X2),s(X0,'const.bool.COND_3'(s(bool,X5),s(X0,X3),s(X0,X4))))))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X6] : ! [X10,X7,X8,X9] : ? [X15] : ((p(s(bool,X15)) <=> s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X10))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) & s(X6,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X10),s(X6,X7),s(X6,X8),s(X6,X9))) = s(X6,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X10))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(X6,X7),s(X6,'const.bool.COND_3'(s(bool,X15),s(X6,X8),s(X6,X9))))))),
  file('thm.prelim.ordering_case_def|split|2','thm.prelim.ordering_CASE')).
fof(f6978,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))),
  inference(superposition,[],[f6943,f4715])).
fof(f4715,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(backward_demodulation,[],[f4713,f155])).
fof(f6943,plain,(
  ( ! [X0] : (s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f1943,f4714])).
fof(f4714,plain,(
  s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(backward_demodulation,[],[f4713,f153])).
fof(f153,plain,(
  s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f63])).
fof(f63,axiom,(
  s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  file('thm.prelim.ordering_case_def|split|2','thm.prelim.ordering2num_thm|split|2')).
fof(f1943,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f235,f222])).
fof(f222,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1] : ((s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.num.num',X0) != s('type.num.num',X1) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X13,X14] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X14))) <=> s('type.num.num',X13) = s('type.num.num',X14))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_distrib|split|19')).
fof(f235,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))) )),
  inference(subsumption_resolution,[],[f214,f156])).
fof(f156,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f67])).
fof(f67,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.prelim.ordering_case_def|split|2',aHL_FALSITY)).
fof(f214,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0,X1] : ((s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X13,X14] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X13))) = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X14))) <=> p(s(bool,'F_0')))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_eq|split|4')).
fof(f4870,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X1),s(X0,X2),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(X0,X2),s(X0,'const.bool.COND_3'(s(bool,sK4(X4,X3,X2,X1,X0)),s(X0,X3),s(X0,X4)))))) )),
  inference(backward_demodulation,[],[f4799,f232])).
fof(f232,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X1),s(X0,X2),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s(X0,X2),s(X0,'const.bool.COND_3'(s(bool,sK4(X4,X3,X2,X1,X0)),s(X0,X3),s(X0,X4)))))) )),
  inference(forward_demodulation,[],[f202,f155])).
fof(f202,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X1),s(X0,X2),s(X0,X3),s(X0,X4))) = s(X0,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s(X0,X2),s(X0,'const.bool.COND_3'(s(bool,sK4(X4,X3,X2,X1,X0)),s(X0,X3),s(X0,X4)))))) )),
  inference(cnf_transformation,[],[f140])).
fof(f4799,plain,(
  ( ! [X3] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(backward_demodulation,[],[f4713,f3000])).
fof(f3000,plain,(
  ( ! [X3] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) )),
  inference(forward_demodulation,[],[f2964,f2961])).
fof(f2961,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f208,f155])).
fof(f208,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X13,X14] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X14),s('type.num.num',X13))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X14)))))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_distrib|split|25')).
fof(f2964,plain,(
  ( ! [X3] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3)))))) )),
  inference(superposition,[],[f208,f944])).
fof(f944,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f173,f733])).
fof(f173,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X13] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X13))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13)))))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_distrib|split|14')).
fof(f30123,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f30122,f165])).
fof(f165,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X13] : s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X13),s('type.num.num','const.num.|48|_0')))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_distrib|split|20')).
fof(f30122,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f30075,f205])).
fof(f205,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X13,X14] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X13),s('type.num.num',X14))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X13))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X14)))))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_lt|split|6')).
fof(f30075,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f4806,f22832])).
fof(f22832,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(backward_demodulation,[],[f22761,f4715])).
fof(f22761,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f4737,f4758])).
fof(f4758,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num',X0) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(backward_demodulation,[],[f4713,f1972])).
fof(f1972,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num',X0) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(superposition,[],[f222,f155])).
fof(f4737,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(backward_demodulation,[],[f4713,f944])).
fof(f4806,plain,(
  ( ! [X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(backward_demodulation,[],[f4713,f3047])).
fof(f3047,plain,(
  ( ! [X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(superposition,[],[f209,f153])).
fof(f209,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X13,X14] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X13),s('type.num.num',X14))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X14)))))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_distrib|split|22')).
fof(f148,plain,(
  s(sK0,sK3) != s(sK0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.GREATER_0'),s(sK0,sK1),s(sK0,sK2),s(sK0,sK3)))),
  inference(cnf_transformation,[],[f129])).
fof(f129,plain,(
  s(sK0,sK3) != s(sK0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.GREATER_0'),s(sK0,sK1),s(sK0,sK2),s(sK0,sK3)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f127])).
fof(f127,plain,(
  ? [X0,X1,X2,X3] : s(X0,X3) != s(X0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.GREATER_0'),s(X0,X1),s(X0,X2),s(X0,X3)))),
  inference(ennf_transformation,[],[f66])).
fof(f66,plain,(
  ~! [X0,X1,X2,X3] : s(X0,X3) = s(X0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.GREATER_0'),s(X0,X1),s(X0,X2),s(X0,X3)))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ~! [X0] : ! [X1,X2,X3] : s(X0,X3) = s(X0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.GREATER_0'),s(X0,X1),s(X0,X2),s(X0,X3)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : s(X6,X9) = s(X6,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.GREATER_0'),s(X6,X7),s(X6,X8),s(X6,X9)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : s(X6,X9) = s(X6,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.GREATER_0'),s(X6,X7),s(X6,X8),s(X6,X9)))),
  file('thm.prelim.ordering_case_def|split|2','thm.prelim.ordering_case_def|split|2')).
fof(f30956,plain,(
  $false | ($spl9 | $spl10)),
  inference(trivial_inequality_removal,[],[f30955])).
fof(f30955,plain,(
  s(sK0,sK3) != s(sK0,sK3) | ($spl9 | $spl10)),
  inference(superposition,[],[f148,f30126])).
fof(f24549,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | $spl132),
  inference(cnf_transformation,[],[f24549_D])).
fof(f24549_D,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f24551,plain,(
  s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))) != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | $spl135),
  inference(cnf_transformation,[],[f24551_D])).
fof(f24551_D,plain,(
  s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))) != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl135),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl135])])).
fof(f24509,plain,(
  s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))) != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f4759,f22835])).
fof(f22835,plain,(
  s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(backward_demodulation,[],[f22761,f4740])).
fof(f4740,plain,(
  s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(backward_demodulation,[],[f4713,f1000])).
fof(f1000,plain,(
  s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(superposition,[],[f175,f155])).
fof(f175,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X0))))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X0))))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X13] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X13))))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13)))))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_distrib|split|16')).
fof(f4759,plain,(
  ( ! [X1] : (s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) | s('type.num.num',X1) = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(backward_demodulation,[],[f4713,f1973])).
fof(f1973,plain,(
  ( ! [X1] : (s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) | s('type.num.num',X1) = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(superposition,[],[f222,f153])).
fof(f24545,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) | $spl128),
  inference(cnf_transformation,[],[f24545_D])).
fof(f24545_D,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f24547,plain,(
  s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))) != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) | $spl131),
  inference(cnf_transformation,[],[f24547_D])).
fof(f24547_D,plain,(
  s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))) != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) <=> ~$spl131),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl131])])).
fof(f24508,plain,(
  s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))) != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) | s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f4759,f4739])).
fof(f4739,plain,(
  s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(backward_demodulation,[],[f4713,f999])).
fof(f999,plain,(
  s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(superposition,[],[f175,f153])).
fof(f24541,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | $spl124),
  inference(cnf_transformation,[],[f24541_D])).
fof(f24541_D,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f24543,plain,(
  s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | $spl127),
  inference(cnf_transformation,[],[f24543_D])).
fof(f24543_D,plain,(
  s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl127),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl127])])).
fof(f24507,plain,(
  s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f4759,f22833])).
fof(f22833,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(backward_demodulation,[],[f22761,f4736])).
fof(f4736,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(backward_demodulation,[],[f4713,f943])).
fof(f943,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(superposition,[],[f173,f155])).
fof(f24537,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) | $spl120),
  inference(cnf_transformation,[],[f24537_D])).
fof(f24537_D,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f24539,plain,(
  s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) | $spl123),
  inference(cnf_transformation,[],[f24539_D])).
fof(f24539_D,plain,(
  s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) <=> ~$spl123),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl123])])).
fof(f24506,plain,(
  s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) | s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f4759,f4735])).
fof(f4735,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(backward_demodulation,[],[f4713,f942])).
fof(f942,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(superposition,[],[f173,f153])).
fof(f24521,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl116),
  inference(cnf_transformation,[],[f24521_D])).
fof(f24521_D,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f24523,plain,(
  s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl119),
  inference(cnf_transformation,[],[f24523_D])).
fof(f24523_D,plain,(
  s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl119),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl119])])).
fof(f24477,plain,(
  s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(superposition,[],[f4759,f22835])).
fof(f24516,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl112),
  inference(cnf_transformation,[],[f24516_D])).
fof(f24516_D,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f24518,plain,(
  s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl115),
  inference(cnf_transformation,[],[f24518_D])).
fof(f24518_D,plain,(
  s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl115),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl115])])).
fof(f24475,plain,(
  s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(superposition,[],[f4759,f22833])).
fof(f23119,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) | $spl110),
  inference(cnf_transformation,[],[f23119_D])).
fof(f23119_D,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f23072,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) | $spl107),
  inference(cnf_transformation,[],[f23072_D])).
fof(f23072_D,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) <=> ~$spl107),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl107])])).
fof(f23117,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) | s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f22804,f22761])).
fof(f22804,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) | s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f4758,f4739])).
fof(f23115,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) | $spl108),
  inference(cnf_transformation,[],[f23115_D])).
fof(f23115_D,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f23065,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) | $spl103),
  inference(cnf_transformation,[],[f23065_D])).
fof(f23065_D,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) <=> ~$spl103),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl103])])).
fof(f23113,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) | s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f22802,f22761])).
fof(f22802,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) | s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f4758,f4735])).
fof(f23071,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))) | $spl105),
  inference(cnf_transformation,[],[f23071_D])).
fof(f23071_D,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))) <=> ~$spl105),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl105])])).
fof(f23073,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) | $spl106),
  inference(cnf_transformation,[],[f23073_D])).
fof(f23073_D,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f23069,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) | s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))),
  inference(forward_demodulation,[],[f23068,f4739])).
fof(f23068,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))) | s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(forward_demodulation,[],[f22771,f22761])).
fof(f22771,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))) | s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(superposition,[],[f4758,f4739])).
fof(f23064,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))) | $spl101),
  inference(cnf_transformation,[],[f23064_D])).
fof(f23064_D,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))) <=> ~$spl101),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl101])])).
fof(f23066,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) | $spl102),
  inference(cnf_transformation,[],[f23066_D])).
fof(f23066_D,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f23062,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) | s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))),
  inference(forward_demodulation,[],[f23061,f4735])).
fof(f23061,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))) | s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(forward_demodulation,[],[f22769,f22761])).
fof(f22769,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))) | s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(superposition,[],[f4758,f4735])).
fof(f19891,plain,(
  $false | ($spl9 | $spl76 | $spl90)),
  inference(subsumption_resolution,[],[f19890,f11930])).
fof(f11930,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl9),
  inference(subsumption_resolution,[],[f5338,f4741])).
fof(f5338,plain,(
  s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))) | $spl9),
  inference(superposition,[],[f4741,f181])).
fof(f19890,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl76 | $spl90)),
  inference(backward_demodulation,[],[f19820,f19876])).
fof(f19888,plain,(
  $false | ($spl9 | $spl84 | $spl96)),
  inference(subsumption_resolution,[],[f19887,f11930])).
fof(f19887,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl84 | $spl96)),
  inference(backward_demodulation,[],[f19832,f19882])).
fof(f19882,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl96),
  inference(cnf_transformation,[],[f19882_D])).
fof(f19882_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f19884,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl98),
  inference(cnf_transformation,[],[f19884_D])).
fof(f19884_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f19768,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(superposition,[],[f479,f4740])).
fof(f479,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))) | s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X1)))) )),
  inference(superposition,[],[f171,f181])).
fof(f171,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X0))) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0] : s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X0))) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X13] : s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X13))) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13)))))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_distrib|split|33')).
fof(f19827,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl82),
  inference(cnf_transformation,[],[f19827_D])).
fof(f19827_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f19880,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))) | $spl94),
  inference(cnf_transformation,[],[f19880_D])).
fof(f19880_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f19767,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(superposition,[],[f479,f4739])).
fof(f19876,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl90),
  inference(cnf_transformation,[],[f19876_D])).
fof(f19876_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f19878,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl92),
  inference(cnf_transformation,[],[f19878_D])).
fof(f19878_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f19766,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(superposition,[],[f479,f4736])).
fof(f19815,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl74),
  inference(cnf_transformation,[],[f19815_D])).
fof(f19815_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f19874,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))) | $spl88),
  inference(cnf_transformation,[],[f19874_D])).
fof(f19874_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f19765,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(superposition,[],[f479,f4735])).
fof(f19832,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl84),
  inference(cnf_transformation,[],[f19832_D])).
fof(f19832_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f19834,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl86),
  inference(cnf_transformation,[],[f19834_D])).
fof(f19834_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f19830,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(forward_demodulation,[],[f19829,f4740])).
fof(f19829,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))))),
  inference(forward_demodulation,[],[f19828,f10312])).
fof(f10312,plain,(
  s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(superposition,[],[f171,f4740])).
fof(f19828,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))))),
  inference(forward_demodulation,[],[f19735,f4764])).
fof(f4764,plain,(
  s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(backward_demodulation,[],[f4713,f2253])).
fof(f2253,plain,(
  s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(superposition,[],[f175,f944])).
fof(f19735,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))))),
  inference(superposition,[],[f479,f4740])).
fof(f19825,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))))) | $spl80),
  inference(cnf_transformation,[],[f19825_D])).
fof(f19825_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))))) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f19823,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))))),
  inference(forward_demodulation,[],[f19734,f171])).
fof(f19734,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))))))),
  inference(superposition,[],[f479,f4739])).
fof(f19820,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl76),
  inference(cnf_transformation,[],[f19820_D])).
fof(f19820_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f19822,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl78),
  inference(cnf_transformation,[],[f19822_D])).
fof(f19822_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f19818,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(forward_demodulation,[],[f19817,f4736])).
fof(f19817,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))))),
  inference(forward_demodulation,[],[f19816,f9926])).
fof(f9926,plain,(
  s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(superposition,[],[f171,f4736])).
fof(f19816,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))))),
  inference(forward_demodulation,[],[f19733,f4763])).
fof(f4763,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(backward_demodulation,[],[f4713,f2252])).
fof(f2252,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(superposition,[],[f173,f944])).
fof(f19733,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))))),
  inference(superposition,[],[f479,f4736])).
fof(f19813,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))))) | $spl72),
  inference(cnf_transformation,[],[f19813_D])).
fof(f19813_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))))) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f19811,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))))),
  inference(forward_demodulation,[],[f19732,f171])).
fof(f19732,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))))))),
  inference(superposition,[],[f479,f4735])).
fof(f11853,plain,(
  $false | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f11852,f4741])).
fof(f11852,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f5338,f385])).
fof(f11851,plain,(
  $false | ($spl0 | $spl9 | $spl10)),
  inference(subsumption_resolution,[],[f5337,f385])).
fof(f5337,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ($spl9 | $spl10)),
  inference(superposition,[],[f4741,f1068])).
fof(f11850,plain,(
  $false | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f11849,f4741])).
fof(f11849,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl9)),
  inference(forward_demodulation,[],[f11848,f4713])).
fof(f11848,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f1141,f385])).
fof(f1141,plain,(
  s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl9),
  inference(superposition,[],[f1065,f181])).
fof(f11847,plain,(
  $false | ($spl0 | $spl9 | $spl10)),
  inference(subsumption_resolution,[],[f1140,f385])).
fof(f1140,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ($spl9 | $spl10)),
  inference(superposition,[],[f1065,f1072])).
fof(f1072,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl10),
  inference(backward_demodulation,[],[f1068,f755])).
fof(f755,plain,(
  s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f170,f733])).
fof(f170,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X0))) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0] : s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X0))) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X13] : s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X13))) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13)))))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_distrib|split|32')).
fof(f11833,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f177,f11831])).
fof(f11831,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)) ) | $spl0),
  inference(subsumption_resolution,[],[f11162,f11829])).
fof(f11829,plain,(
  ~p(s(bool,'F_0')) | $spl0),
  inference(global_subsumption,[],[f11178,f148,f149,f150,f151,f152,f153,f154,f155,f228,f156,f158,f159,f160,f161,f162,f163,f164,f165,f166,f167,f168,f169,f170,f171,f172,f173,f229,f175,f176,f177,f178,f179,f180,f181,f230,f231,f191,f190,f195,f194,f198,f199,f232,f233,f234,f203,f204,f205,f206,f207,f208,f209,f210,f211,f212,f213,f235,f219,f218,f221,f220,f223,f222,f225,f224,f226,f371,f375,f383,f382,f460,f461,f462,f477,f478,f479,f545,f568,f620,f683,f682,f684,f685,f691,f692,f696,f697,f711,f712,f733,f782,f786,f785,f817,f896,f901,f904,f909,f942,f943,f944,f999,f1000,f1033,f1058,f1060,f157,f755,f1093,f1163,f1272,f1231,f1227,f1226,f1225,f1223,f1222,f1233,f1234,f1238,f1239,f1240,f1241,f1242,f1243,f1244,f1246,f1248,f1250,f1251,f1290,f1499,f1548,f1637,f1592,f1588,f1587,f1586,f1584,f1583,f1594,f1595,f1599,f1600,f1601,f1602,f1603,f1604,f1605,f1606,f1608,f1610,f1611,f1613,f1614,f1679,f1683,f1710,f1709,f1871,f1870,f1899,f1898,f2004,f2005,f2006,f1943,f2007,f1941,f1937,f1936,f1950,f1951,f1954,f1955,f1956,f1957,f1958,f1959,f1960,f1961,f1962,f1963,f1965,f2021,f1968,f1970,f1971,f1972,f1973,f2036,f2035,f2059,f2310,f2311,f2252,f2253,f2255,f2312,f2288,f2357,f2358,f2360,f2350,f2362,f2384,f2383,f2421,f2422,f2414,f2463,f2464,f2474,f2475,f2467,f2491,f2493,f2494,f2539,f2540,f2541,f2577,f2578,f2580,f2566,f2567,f2581,f2582,f2570,f2655,f2656,f2682,f2683,f2686,f2774,f2789,f2788,f2787,f2797,f2796,f2795,f2844,f2859,f2858,f2857,f2867,f2866,f2865,f2914,f2929,f2928,f2927,f2937,f2936,f2935,f2961,f2962,f3000,f2988,f2989,f3001,f2992,f3046,f3047,f3073,f3074,f3077,f3102,f3101,f3103,f3205,f3204,f3206,f3208,f3210,f3314,f3359,f3360,f3361,f3466,f3387,f3388,f3467,f3468,f3413,f3471,f3473,f3431,f3439,f3486,f3488,f3538,f3539,f3648,f3649,f3567,f3568,f3650,f3651,f3593,f3654,f3656,f3621,f3717,f3718,f3830,f3831,f3747,f3748,f3832,f3833,f3773,f3836,f3838,f3803,f3995,f3853,f3911,f4182,f4186,f4296,f4297,f4298,f4299,f4300,f4301,f4302,f4303,f4304,f4305,f4306,f4307,f4308,f4309,f4310,f4311,f4312,f4313,f4314,f4315,f4316,f4317,f4318,f4319,f4320,f4571,f4424,f4578,f4579,f4580,f4582,f4583,f4584,f4585,f4586,f4587,f4588,f4589,f4590,f4591,f4592,f4602,f4595,f4596,f4604,f4605,f4627,f4628,f4629,f4630,f4631,f4633,f4634,f4635,f4638,f4639,f4642,f4643,f4644,f4645,f4646,f4647,f4484,f4713,f4714,f4860,f4715,f4862,f4717,f4718,f4719,f4720,f4721,f4722,f4723,f4724,f4725,f4726,f4727,f4728,f4729,f4730,f4731,f4732,f4735,f4736,f4737,f4865,f4739,f4740,f4743,f4744,f4748,f4749,f4752,f4753,f4754,f4755,f4756,f4757,f4758,f4759,f4761,f4762,f4763,f4764,f4765,f4766,f4769,f4774,f4775,f4776,f4777,f4780,f4781,f4782,f4783,f4784,f4785,f4786,f4788,f4790,f4791,f4792,f4793,f4794,f4795,f4796,f4797,f4798,f4799,f4870,f4800,f4805,f4806,f4807,f4808,f4816,f4817,f4818,f4819,f4820,f4821,f4822,f4824,f4825,f4826,f4829,f4830,f4831,f4832,f4833,f4834,f4836,f4837,f4838,f4839,f4840,f4841,f4842,f4843,f4844,f4846,f4847,f4848,f4853,f4856,f5035,f5066,f5694,f5818,f5819,f5742,f5744,f5839,f5841,f6310,f6309,f6338,f6350,f6349,f6378,f6568,f6567,f6569,f6609,f6610,f6611,f6641,f6642,f6643,f6759,f6758,f6757,f6822,f6821,f6820,f6878,f6879,f6887,f6886,f6885,f6943,f6949,f6948,f6947,f6978,f6985,f6986,f6987,f7271,f7277,f6990,f6991,f6992,f6993,f6994,f6995,f6996,f6997,f6998,f6999,f7000,f7001,f7002,f7003,f7004,f7005,f7006,f7007,f7008,f7009,f7010,f7011,f7012,f7013,f7014,f7015,f7016,f7017,f7018,f7019,f7020,f7021,f7022,f7023,f7024,f7025,f7026,f7027,f7028,f7029,f7030,f7031,f7032,f7033,f7034,f7035,f7036,f7037,f7038,f7040,f7045,f7046,f7047,f7048,f7049,f7050,f7284,f7053,f7054,f7055,f7056,f7057,f7058,f7059,f7060,f7061,f7062,f7063,f7064,f7065,f7066,f7067,f7068,f7069,f7070,f7071,f7072,f7073,f7074,f7075,f7076,f7077,f7078,f7079,f7080,f7081,f7082,f7083,f7084,f7085,f7086,f7087,f7088,f7089,f7090,f7091,f7092,f7093,f7094,f7095,f7096,f7097,f7098,f7099,f7100,f7102,f7107,f7108,f7109,f7110,f7111,f7112,f7113,f7114,f7115,f7116,f7117,f7118,f7119,f7120,f7121,f7122,f7123,f7124,f7125,f7126,f7127,f7128,f7129,f7130,f7131,f7132,f7133,f7134,f7135,f7136,f7137,f7138,f7139,f7140,f7141,f7142,f7143,f7144,f7145,f7146,f7147,f7148,f7149,f7150,f7151,f7152,f7153,f7154,f7155,f7156,f7157,f7158,f7159,f7160,f7162,f7167,f7168,f7169,f7170,f7171,f7172,f7173,f7174,f7219,f7223,f7227,f7256,f7305,f7304,f7306,f7299,f7310,f7312,f7317,f7318,f7319,f7349,f7351,f7356,f7357,f7358,f7591,f7596,f7597,f7598,f7601,f7640,f7644,f7648,f7677,f7713,f7712,f7714,f7922,f7927,f7928,f7929,f7935,f7974,f7978,f7982,f8011,f8156,f8155,f8189,f8198,f8197,f8231,f8312,f8311,f8345,f8354,f8353,f8387,f8466,f8465,f8464,f8503,f8502,f8535,f8554,f8559,f8560,f8561,f8562,f8564,f8565,f8569,f8570,f8572,f8574,f8580,f8586,f8587,f8588,f8589,f8590,f8637,f8648,f8649,f8650,f8651,f8652,f8654,f8655,f8656,f8657,f8658,f8660,f8661,f8662,f8663,f8664,f8665,f8666,f8717,f8749,f8758,f8759,f8760,f8767,f8772,f8774,f8853,f8854,f8860,f8861,f8862,f9293,f9294,f9295,f9296,f9297,f9298,f9299,f9300,f9301,f9302,f9303,f9304,f8912,f8913,f8914,f8921,f8926,f8928,f9732,f9733,f9734,f9735,f9736,f9737,f9738,f9739,f9740,f9741,f9742,f9743,f9744,f9745,f9746,f9747,f9748,f9749,f9750,f9867,f9872,f9754,f9755,f9756,f9757,f9758,f9759,f9858,f9859,f9860,f9925,f9926,f9927,f9928,f9929,f9930,f9931,f9932,f9933,f9934,f9935,f9936,f9937,f9938,f9939,f9940,f9941,f9942,f9943,f10060,f10065,f9947,f9948,f9949,f9950,f9951,f9952,f10051,f10052,f10053,f10118,f10119,f10120,f10121,f10122,f10123,f10124,f10125,f10126,f10127,f10128,f10129,f10130,f10131,f10132,f10133,f10134,f10135,f10136,f10253,f10258,f10140,f10141,f10142,f10143,f10144,f10145,f10244,f10245,f10246,f10311,f10312,f10313,f10314,f10315,f10316,f10317,f10318,f10319,f10320,f10321,f10322,f10323,f10324,f10325,f10326,f10327,f10328,f10329,f10446,f10451,f10333,f10334,f10335,f10336,f10337,f10338,f10437,f10438,f10439,f10590,f10597,f10602,f10609,f10624,f10632,f10648,f10544,f10545,f10546,f10547,f10564,f10565,f10566])).
fof(f10566,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X29),s(bool,X30),s(bool,X31),s(bool,'T_0'))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X29))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X30),s(bool,'const.bool.COND_3'(s(bool,sK4('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X28)))),X31,X30,X29,bool)),s(bool,X31),s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X28)))) )),
  inference(superposition,[],[f4870,f462])).
fof(f10565,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X25),s(bool,'T_0'),s(bool,X26),s(bool,X27))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X25))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,sK4(X27,X26,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X24)))),X25,bool)),s(bool,X26),s(bool,X27))))) | s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X24)))) )),
  inference(superposition,[],[f4870,f462])).
fof(f10564,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X21),s(bool,X22),s(bool,'T_0'),s(bool,X23))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X21))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X22),s(bool,'const.bool.COND_3'(s(bool,sK4(X23,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X20)))),X22,X21,bool)),s(bool,'T_0'),s(bool,X23))))) | s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X20)))) )),
  inference(superposition,[],[f4870,f462])).
fof(f10648,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X16),s('type.num.num',X17))))))) | s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X16),s('type.num.num',X17)))))) )),
  inference(forward_demodulation,[],[f10647,f210])).
fof(f10647,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X16),s('type.num.num',X17))))) | s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X16))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X17)))))))) )),
  inference(forward_demodulation,[],[f10646,f170])).
fof(f10646,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X16),s('type.num.num',X17))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X16))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X17)))))))) )),
  inference(forward_demodulation,[],[f10528,f170])).
fof(f10528,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X16),s('type.num.num',X17))))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X16))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X17)))))))) )),
  inference(superposition,[],[f462,f210])).
fof(f10632,plain,(
  ( ! [X12,X11] : (s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X11),s('type.num.num',X12))))))) | s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X11),s('type.num.num',X12)))))) )),
  inference(forward_demodulation,[],[f10631,f211])).
fof(f10631,plain,(
  ( ! [X12,X11] : (s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X11),s('type.num.num',X12))))) | s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X11))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X12)))))))) )),
  inference(forward_demodulation,[],[f10630,f170])).
fof(f10630,plain,(
  ( ! [X12,X11] : (s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X11),s('type.num.num',X12))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X11))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X12)))))))) )),
  inference(forward_demodulation,[],[f10522,f170])).
fof(f10522,plain,(
  ( ! [X12,X11] : (s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X11),s('type.num.num',X12))))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X11))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X12)))))))) )),
  inference(superposition,[],[f462,f211])).
fof(f10624,plain,(
  ( ! [X8,X7] : (s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X7),s('type.num.num',X8))))))) | s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X7),s('type.num.num',X8)))))) )),
  inference(forward_demodulation,[],[f10623,f212])).
fof(f10623,plain,(
  ( ! [X8,X7] : (s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X7),s('type.num.num',X8))))) | s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8)))))))) )),
  inference(forward_demodulation,[],[f10622,f170])).
fof(f10622,plain,(
  ( ! [X8,X7] : (s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X7),s('type.num.num',X8))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8)))))))) )),
  inference(forward_demodulation,[],[f10519,f170])).
fof(f10519,plain,(
  ( ! [X8,X7] : (s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X7),s('type.num.num',X8))))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8)))))))) )),
  inference(superposition,[],[f462,f212])).
fof(f10439,plain,(
  ( ! [X76,X74,X75] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X74),s('type.num.num',X75),s('type.num.num',X76),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X74))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X75),s('type.num.num','const.bool.COND_3'(s(bool,sK4('const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))),X76,X75,X74,'type.num.num')),s('type.num.num',X76),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(superposition,[],[f4870,f4740])).
fof(f10438,plain,(
  ( ! [X72,X71,X73] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X71),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X72),s('type.num.num',X73))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X71))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X73,X72,'const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))),X71,'type.num.num')),s('type.num.num',X72),s('type.num.num',X73)))))) )),
  inference(superposition,[],[f4870,f4740])).
fof(f10437,plain,(
  ( ! [X70,X68,X69] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X68),s('type.num.num',X69),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X70))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X68))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X69),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X70,'const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))),X69,X68,'type.num.num')),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X70)))))) )),
  inference(superposition,[],[f4870,f4740])).
fof(f10337,plain,(
  s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.|48|_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f4721,f4740])).
fof(f10336,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(superposition,[],[f4720,f4740])).
fof(f10335,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f4719,f4740])).
fof(f10334,plain,(
  ( ! [X17] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X17))),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X17),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(superposition,[],[f4605,f4740])).
fof(f10333,plain,(
  ( ! [X16] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X16))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X16)))))) )),
  inference(superposition,[],[f4605,f4740])).
fof(f10329,plain,(
  ( ! [X14] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X14))) != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num',X14) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f222,f4740])).
fof(f10328,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X13),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13))),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f212,f4740])).
fof(f10327,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X12))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X12)))))) )),
  inference(superposition,[],[f212,f4740])).
fof(f10326,plain,(
  ( ! [X11] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X11),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X11))),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f211,f4740])).
fof(f10325,plain,(
  ( ! [X10] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X10))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10)))))) )),
  inference(superposition,[],[f211,f4740])).
fof(f10324,plain,(
  ( ! [X9] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X9),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X9))),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f210,f4740])).
fof(f10323,plain,(
  ( ! [X8] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X8))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8)))))) )),
  inference(superposition,[],[f210,f4740])).
fof(f10322,plain,(
  ( ! [X7] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f209,f4740])).
fof(f10321,plain,(
  ( ! [X6] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X6))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6)))))) )),
  inference(superposition,[],[f209,f4740])).
fof(f10320,plain,(
  ( ! [X5] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X5))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X5))),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f208,f4740])).
fof(f10319,plain,(
  ( ! [X4] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X4)))))) )),
  inference(superposition,[],[f208,f4740])).
fof(f10318,plain,(
  ( ! [X3] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3))),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f204,f4740])).
fof(f10317,plain,(
  ( ! [X2] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X2))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))))) )),
  inference(superposition,[],[f204,f4740])).
fof(f10316,plain,(
  ( ! [X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X1))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f203,f4740])).
fof(f10315,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f203,f4740])).
fof(f10314,plain,(
  s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))))),
  inference(superposition,[],[f175,f4740])).
fof(f10313,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))))),
  inference(superposition,[],[f173,f4740])).
fof(f10311,plain,(
  s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(superposition,[],[f170,f4740])).
fof(f10246,plain,(
  ( ! [X76,X74,X75] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X74),s('type.num.num',X75),s('type.num.num',X76),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X74))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X75),s('type.num.num','const.bool.COND_3'(s(bool,sK4('const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))),X76,X75,X74,'type.num.num')),s('type.num.num',X76),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))))) )),
  inference(superposition,[],[f4870,f4739])).
fof(f10245,plain,(
  ( ! [X72,X71,X73] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X71),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))),s('type.num.num',X72),s('type.num.num',X73))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X71))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X73,X72,'const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))),X71,'type.num.num')),s('type.num.num',X72),s('type.num.num',X73)))))) )),
  inference(superposition,[],[f4870,f4739])).
fof(f10244,plain,(
  ( ! [X70,X68,X69] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X68),s('type.num.num',X69),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))),s('type.num.num',X70))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X68))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X69),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X70,'const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))),X69,X68,'type.num.num')),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))),s('type.num.num',X70)))))) )),
  inference(superposition,[],[f4870,f4739])).
fof(f10144,plain,(
  s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.|48|_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))),s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f4721,f4739])).
fof(f10143,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))),
  inference(superposition,[],[f4720,f4739])).
fof(f10142,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))),s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f4719,f4739])).
fof(f10141,plain,(
  ( ! [X17] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X17))),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X17),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(superposition,[],[f4605,f4739])).
fof(f10140,plain,(
  ( ! [X16] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X16))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X16)))))) )),
  inference(superposition,[],[f4605,f4739])).
fof(f10136,plain,(
  ( ! [X14] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X14))) != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) | s('type.num.num',X14) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f222,f4739])).
fof(f10135,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X13),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13))),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))) )),
  inference(superposition,[],[f212,f4739])).
fof(f10134,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X12))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X12)))))) )),
  inference(superposition,[],[f212,f4739])).
fof(f10133,plain,(
  ( ! [X11] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X11),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X11))),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))) )),
  inference(superposition,[],[f211,f4739])).
fof(f10132,plain,(
  ( ! [X10] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X10))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10)))))) )),
  inference(superposition,[],[f211,f4739])).
fof(f10131,plain,(
  ( ! [X9] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X9),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X9))),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))) )),
  inference(superposition,[],[f210,f4739])).
fof(f10130,plain,(
  ( ! [X8] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X8))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8)))))) )),
  inference(superposition,[],[f210,f4739])).
fof(f10129,plain,(
  ( ! [X7] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))) )),
  inference(superposition,[],[f209,f4739])).
fof(f10128,plain,(
  ( ! [X6] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X6))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6)))))) )),
  inference(superposition,[],[f209,f4739])).
fof(f10127,plain,(
  ( ! [X5] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X5))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X5))),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))) )),
  inference(superposition,[],[f208,f4739])).
fof(f10126,plain,(
  ( ! [X4] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X4)))))) )),
  inference(superposition,[],[f208,f4739])).
fof(f10125,plain,(
  ( ! [X3] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3))),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))) )),
  inference(superposition,[],[f204,f4739])).
fof(f10124,plain,(
  ( ! [X2] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X2))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))))) )),
  inference(superposition,[],[f204,f4739])).
fof(f10123,plain,(
  ( ! [X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X1))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))) )),
  inference(superposition,[],[f203,f4739])).
fof(f10122,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f203,f4739])).
fof(f10121,plain,(
  s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))))))),
  inference(superposition,[],[f175,f4739])).
fof(f10120,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))))))),
  inference(superposition,[],[f173,f4739])).
fof(f10119,plain,(
  s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))),
  inference(superposition,[],[f171,f4739])).
fof(f10118,plain,(
  s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))),
  inference(superposition,[],[f170,f4739])).
fof(f10053,plain,(
  ( ! [X76,X74,X75] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X74),s('type.num.num',X75),s('type.num.num',X76),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X74))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X75),s('type.num.num','const.bool.COND_3'(s(bool,sK4('const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))),X76,X75,X74,'type.num.num')),s('type.num.num',X76),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(superposition,[],[f4870,f4736])).
fof(f10052,plain,(
  ( ! [X72,X71,X73] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X71),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X72),s('type.num.num',X73))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X71))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X73,X72,'const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))),X71,'type.num.num')),s('type.num.num',X72),s('type.num.num',X73)))))) )),
  inference(superposition,[],[f4870,f4736])).
fof(f10051,plain,(
  ( ! [X70,X68,X69] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X68),s('type.num.num',X69),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X70))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X68))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X69),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X70,'const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))),X69,X68,'type.num.num')),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X70)))))) )),
  inference(superposition,[],[f4870,f4736])).
fof(f9951,plain,(
  s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.|48|_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f4721,f4736])).
fof(f9950,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(superposition,[],[f4720,f4736])).
fof(f9949,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f4719,f4736])).
fof(f9948,plain,(
  ( ! [X17] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X17))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X17),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(superposition,[],[f4605,f4736])).
fof(f9947,plain,(
  ( ! [X16] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X16))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X16)))))) )),
  inference(superposition,[],[f4605,f4736])).
fof(f9943,plain,(
  ( ! [X14] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X14))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num',X14) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f222,f4736])).
fof(f9942,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X13),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f212,f4736])).
fof(f9941,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X12))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X12)))))) )),
  inference(superposition,[],[f212,f4736])).
fof(f9940,plain,(
  ( ! [X11] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X11),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X11))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f211,f4736])).
fof(f9939,plain,(
  ( ! [X10] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X10))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10)))))) )),
  inference(superposition,[],[f211,f4736])).
fof(f9938,plain,(
  ( ! [X9] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X9),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X9))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f210,f4736])).
fof(f9937,plain,(
  ( ! [X8] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X8))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8)))))) )),
  inference(superposition,[],[f210,f4736])).
fof(f9936,plain,(
  ( ! [X7] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f209,f4736])).
fof(f9935,plain,(
  ( ! [X6] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X6))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6)))))) )),
  inference(superposition,[],[f209,f4736])).
fof(f9934,plain,(
  ( ! [X5] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X5))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X5))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f208,f4736])).
fof(f9933,plain,(
  ( ! [X4] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X4)))))) )),
  inference(superposition,[],[f208,f4736])).
fof(f9932,plain,(
  ( ! [X3] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f204,f4736])).
fof(f9931,plain,(
  ( ! [X2] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X2))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))))) )),
  inference(superposition,[],[f204,f4736])).
fof(f9930,plain,(
  ( ! [X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X1))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f203,f4736])).
fof(f9929,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f203,f4736])).
fof(f9928,plain,(
  s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))))),
  inference(superposition,[],[f175,f4736])).
fof(f9927,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))))),
  inference(superposition,[],[f173,f4736])).
fof(f9925,plain,(
  s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(superposition,[],[f170,f4736])).
fof(f9860,plain,(
  ( ! [X76,X74,X75] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X74),s('type.num.num',X75),s('type.num.num',X76),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X74))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X75),s('type.num.num','const.bool.COND_3'(s(bool,sK4('const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))),X76,X75,X74,'type.num.num')),s('type.num.num',X76),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))))) )),
  inference(superposition,[],[f4870,f4735])).
fof(f9859,plain,(
  ( ! [X72,X71,X73] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X71),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))),s('type.num.num',X72),s('type.num.num',X73))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X71))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X73,X72,'const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))),X71,'type.num.num')),s('type.num.num',X72),s('type.num.num',X73)))))) )),
  inference(superposition,[],[f4870,f4735])).
fof(f9858,plain,(
  ( ! [X70,X68,X69] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X68),s('type.num.num',X69),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))),s('type.num.num',X70))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X68))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X69),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X70,'const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))),X69,X68,'type.num.num')),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))),s('type.num.num',X70)))))) )),
  inference(superposition,[],[f4870,f4735])).
fof(f9758,plain,(
  s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.|48|_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))),s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f4721,f4735])).
fof(f9757,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))),
  inference(superposition,[],[f4720,f4735])).
fof(f9756,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))),s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f4719,f4735])).
fof(f9755,plain,(
  ( ! [X17] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X17))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X17),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(superposition,[],[f4605,f4735])).
fof(f9754,plain,(
  ( ! [X16] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X16))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X16)))))) )),
  inference(superposition,[],[f4605,f4735])).
fof(f9750,plain,(
  ( ! [X14] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X14))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) | s('type.num.num',X14) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f222,f4735])).
fof(f9749,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X13),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))) )),
  inference(superposition,[],[f212,f4735])).
fof(f9748,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X12))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X12)))))) )),
  inference(superposition,[],[f212,f4735])).
fof(f9747,plain,(
  ( ! [X11] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X11),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X11))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))) )),
  inference(superposition,[],[f211,f4735])).
fof(f9746,plain,(
  ( ! [X10] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X10))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10)))))) )),
  inference(superposition,[],[f211,f4735])).
fof(f9745,plain,(
  ( ! [X9] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X9),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X9))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))) )),
  inference(superposition,[],[f210,f4735])).
fof(f9744,plain,(
  ( ! [X8] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X8))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8)))))) )),
  inference(superposition,[],[f210,f4735])).
fof(f9743,plain,(
  ( ! [X7] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))) )),
  inference(superposition,[],[f209,f4735])).
fof(f9742,plain,(
  ( ! [X6] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X6))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6)))))) )),
  inference(superposition,[],[f209,f4735])).
fof(f9741,plain,(
  ( ! [X5] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X5))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X5))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))) )),
  inference(superposition,[],[f208,f4735])).
fof(f9740,plain,(
  ( ! [X4] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X4)))))) )),
  inference(superposition,[],[f208,f4735])).
fof(f9739,plain,(
  ( ! [X3] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))) )),
  inference(superposition,[],[f204,f4735])).
fof(f9738,plain,(
  ( ! [X2] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X2))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))))) )),
  inference(superposition,[],[f204,f4735])).
fof(f9737,plain,(
  ( ! [X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X1))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))) )),
  inference(superposition,[],[f203,f4735])).
fof(f9736,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f203,f4735])).
fof(f9735,plain,(
  s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))))))),
  inference(superposition,[],[f175,f4735])).
fof(f9734,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))))))),
  inference(superposition,[],[f173,f4735])).
fof(f9733,plain,(
  s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))),
  inference(superposition,[],[f171,f4735])).
fof(f9732,plain,(
  s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))),
  inference(superposition,[],[f170,f4735])).
fof(f8928,plain,(
  ( ! [X196,X194,X195,X193] : (s(X195,f(s(bool,'T_0'))) != s(X195,g(s(bool,'T_0'))) | s(fun(bool,X195),X193) = s(fun(bool,X195),X194) | s(bool,X196) = s(bool,sK5(X193,X194,X195,bool)) | s(bool,'T_0') = s(bool,X196)) )),
  inference(superposition,[],[f226,f371])).
fof(f8926,plain,(
  ( ! [X187,X185,X186,X184,X183] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X186),s(bool,X185),s(bool,X184),s(bool,X183))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X186))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X185),s(bool,'T_0'))) | s(bool,X187) = s(bool,'const.bool.COND_3'(s(bool,sK4(X183,X184,X185,X186,bool)),s(bool,X184),s(bool,X183))) | s(bool,'T_0') = s(bool,X187)) )),
  inference(superposition,[],[f4870,f371])).
fof(f8921,plain,(
  ( ! [X167,X165,X166,X168] : (s(bool,'T_0') != s(bool,f(s(X165,sK5(X166,X167,bool,X165)))) | s(fun(X165,bool),X166) = s(fun(X165,bool),X167) | s(bool,X168) = s(bool,g(s(X165,sK5(X166,X167,bool,X165)))) | s(bool,'T_0') = s(bool,X168)) )),
  inference(superposition,[],[f226,f371])).
fof(f8914,plain,(
  ( ! [X134,X136,X135,X133,X137] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X134),s(bool,X135),s(bool,X136),s(bool,'T_0'))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X134))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X135),s(bool,'const.bool.COND_3'(s(bool,sK4(X133,X136,X135,X134,bool)),s(bool,X136),s(bool,'T_0'))))) | s(bool,X133) = s(bool,X137) | s(bool,'T_0') = s(bool,X137)) )),
  inference(superposition,[],[f4870,f371])).
fof(f8913,plain,(
  ( ! [X132,X130,X128,X131,X129] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X129),s(bool,'T_0'),s(bool,X130),s(bool,X131))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X129))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,sK4(X131,X130,X128,X129,bool)),s(bool,X130),s(bool,X131))))) | s(bool,X128) = s(bool,X132) | s(bool,'T_0') = s(bool,X132)) )),
  inference(superposition,[],[f4870,f371])).
fof(f8912,plain,(
  ( ! [X127,X125,X123,X126,X124] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X124),s(bool,X125),s(bool,'T_0'),s(bool,X126))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X124))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X125),s(bool,'const.bool.COND_3'(s(bool,sK4(X126,X123,X125,X124,bool)),s(bool,'T_0'),s(bool,X126))))) | s(bool,X123) = s(bool,X127) | s(bool,'T_0') = s(bool,X127)) )),
  inference(superposition,[],[f4870,f371])).
fof(f9304,plain,(
  ( ! [X109,X107,X110,X108,X106] : (s(bool,X110) = s(bool,sK4(X106,X107,X108,'const.prelim.GREATER_0',X109)) | s(bool,'T_0') = s(bool,X110)) )),
  inference(subsumption_resolution,[],[f8907,f149])).
fof(f8907,plain,(
  ( ! [X109,X107,X110,X108,X106] : (~p(s(bool,'T_0')) | s(bool,X110) = s(bool,sK4(X106,X107,X108,'const.prelim.GREATER_0',X109)) | s(bool,'T_0') = s(bool,X110)) )),
  inference(superposition,[],[f7299,f371])).
fof(f9303,plain,(
  ( ! [X94,X92,X90,X95,X93,X91] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X93))) | s(bool,X95) = s(bool,sK4(X90,X91,X92,X93,X94)) | s(bool,'T_0') = s(bool,X95)) )),
  inference(subsumption_resolution,[],[f8905,f149])).
fof(f8905,plain,(
  ( ! [X94,X92,X90,X95,X93,X91] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X93))) | s(bool,X95) = s(bool,sK4(X90,X91,X92,X93,X94)) | s(bool,'T_0') = s(bool,X95)) )),
  inference(superposition,[],[f234,f371])).
fof(f9302,plain,(
  ( ! [X88,X87,X85,X89,X86,X84] : (s(X88,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X87),s(X88,X86),s(X88,X85),s(X88,X84))) = s(X88,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X87))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(X88,X86),s(X88,X85))) | s(bool,X89) = s(bool,sK4(X84,X85,X86,X87,X88)) | s(bool,'T_0') = s(bool,X89)) )),
  inference(forward_demodulation,[],[f8904,f198])).
fof(f8904,plain,(
  ( ! [X88,X87,X85,X89,X86,X84] : (s(X88,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X87),s(X88,X86),s(X88,X85),s(X88,X84))) = s(X88,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X87))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(X88,X86),s(X88,'const.bool.COND_3'(s(bool,'T_0'),s(X88,X85),s(X88,X84))))) | s(bool,X89) = s(bool,sK4(X84,X85,X86,X87,X88)) | s(bool,'T_0') = s(bool,X89)) )),
  inference(superposition,[],[f4870,f371])).
fof(f9301,plain,(
  ( ! [X74] : (s(bool,X74) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f8898,f149])).
fof(f8898,plain,(
  ( ! [X74] : (~p(s(bool,'T_0')) | s(bool,X74) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f785,f371])).
fof(f9300,plain,(
  ( ! [X73] : (s(bool,X73) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X73)) )),
  inference(subsumption_resolution,[],[f8897,f149])).
fof(f8897,plain,(
  ( ! [X73] : (~p(s(bool,'T_0')) | s(bool,X73) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X73)) )),
  inference(superposition,[],[f786,f371])).
fof(f9299,plain,(
  ( ! [X72,X71] : (s(bool,X72) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X71))))) | s(bool,'T_0') = s(bool,X72)) )),
  inference(subsumption_resolution,[],[f8896,f149])).
fof(f8896,plain,(
  ( ! [X72,X71] : (~p(s(bool,'T_0')) | s(bool,X72) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X71))))) | s(bool,'T_0') = s(bool,X72)) )),
  inference(superposition,[],[f5035,f371])).
fof(f9298,plain,(
  ( ! [X70,X69] : (s(bool,X70) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X69))))) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f8895,f149])).
fof(f8895,plain,(
  ( ! [X70,X69] : (~p(s(bool,'T_0')) | s(bool,X70) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X69))))) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f5066,f371])).
fof(f9297,plain,(
  ( ! [X68,X67] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X67) | s(bool,X68) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X67))) | s(bool,'T_0') = s(bool,X68)) )),
  inference(subsumption_resolution,[],[f8894,f149])).
fof(f8894,plain,(
  ( ! [X68,X67] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X67) | s(bool,X68) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X67))) | s(bool,'T_0') = s(bool,X68)) )),
  inference(superposition,[],[f194,f371])).
fof(f9296,plain,(
  ( ! [X62,X63] : (s(bool,X63) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X62))),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X63)) )),
  inference(subsumption_resolution,[],[f8892,f149])).
fof(f8892,plain,(
  ( ! [X62,X63] : (~p(s(bool,'T_0')) | s(bool,X63) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X62))),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X63)) )),
  inference(superposition,[],[f4785,f371])).
fof(f9295,plain,(
  ( ! [X61,X60] : (s(bool,X61) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X60))),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X61)) )),
  inference(subsumption_resolution,[],[f8891,f149])).
fof(f8891,plain,(
  ( ! [X61,X60] : (~p(s(bool,'T_0')) | s(bool,X61) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X60))),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X61)) )),
  inference(superposition,[],[f4786,f371])).
fof(f9294,plain,(
  ( ! [X45,X43,X41,X46,X44,X42] : (s(X42,X43) = s(X42,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X41),s(X42,X43),s(X42,X44),s(X42,X45))) | s(bool,X46) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X41))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X46)) )),
  inference(forward_demodulation,[],[f8885,f198])).
fof(f8885,plain,(
  ( ! [X45,X43,X41,X46,X44,X42] : (s(X42,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X41),s(X42,X43),s(X42,X44),s(X42,X45))) = s(X42,'const.bool.COND_3'(s(bool,'T_0'),s(X42,X43),s(X42,'const.bool.COND_3'(s(bool,sK4(X45,X44,X43,X41,X42)),s(X42,X44),s(X42,X45))))) | s(bool,X46) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X41))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f4870,f371])).
fof(f9293,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X33),s('type.num.num',X32)))) | s(bool,X34) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X32))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X33))))) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f8882,f149])).
fof(f8882,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X33),s('type.num.num',X32)))) | s(bool,X34) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X32))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X33))))) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f224,f371])).
fof(f8862,plain,(
  ( ! [X37,X35,X33,X36,X34] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X34),s(bool,X35),s(bool,X36),s(bool,X33))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X34))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X35),s(bool,'const.bool.COND_3'(s(bool,sK4('T_0',X36,X35,X34,bool)),s(bool,X36),s(bool,X33))))) | s(bool,X33) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f4870,f371])).
fof(f8861,plain,(
  ( ! [X30,X28,X31,X29,X32] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X29),s(bool,X28),s(bool,X30),s(bool,X31))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X29))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X28),s(bool,'const.bool.COND_3'(s(bool,sK4(X31,X30,'T_0',X29,bool)),s(bool,X30),s(bool,X31))))) | s(bool,X28) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f4870,f371])).
fof(f8860,plain,(
  ( ! [X26,X24,X23,X27,X25] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X24),s(bool,X25),s(bool,X23),s(bool,X26))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X24))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X25),s(bool,'const.bool.COND_3'(s(bool,sK4(X26,'T_0',X25,X24,bool)),s(bool,X23),s(bool,X26))))) | s(bool,X23) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f4870,f371])).
fof(f8853,plain,(
  ( ! [X4,X2,X5,X3,X1] : (s(X2,X3) = s(X2,'const.bool.COND_3'(s(bool,X1),s(X2,X3),s(X2,X4))) | s(bool,X1) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f198,f371])).
fof(f8774,plain,(
  ( ! [X216,X215,X217,X214] : (s(X216,f(s(bool,X217))) != s(X216,g(s(bool,X217))) | s(fun(bool,X216),X214) = s(fun(bool,X216),X215) | s(bool,'T_0') = s(bool,X217) | s(bool,'T_0') = s(bool,sK5(X214,X215,X216,bool))) )),
  inference(superposition,[],[f226,f371])).
fof(f8772,plain,(
  ( ! [X208,X206,X204,X207,X205] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X207),s(bool,X206),s(bool,X205),s(bool,X204))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X207))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X206),s(bool,X208))) | s(bool,'T_0') = s(bool,X208) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,sK4(X204,X205,X206,X207,bool)),s(bool,X205),s(bool,X204)))) )),
  inference(superposition,[],[f4870,f371])).
fof(f8767,plain,(
  ( ! [X189,X187,X188,X186] : (s(bool,X189) != s(bool,f(s(X186,sK5(X187,X188,bool,X186)))) | s(fun(X186,bool),X187) = s(fun(X186,bool),X188) | s(bool,'T_0') = s(bool,X189) | s(bool,'T_0') = s(bool,g(s(X186,sK5(X187,X188,bool,X186))))) )),
  inference(superposition,[],[f226,f371])).
fof(f8760,plain,(
  ( ! [X156,X154,X157,X155,X153] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X155),s(bool,X156),s(bool,X157),s(bool,X154))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X155))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X156),s(bool,'const.bool.COND_3'(s(bool,sK4(X153,X157,X156,X155,bool)),s(bool,X157),s(bool,X154))))) | s(bool,'T_0') = s(bool,X154) | s(bool,'T_0') = s(bool,X153)) )),
  inference(superposition,[],[f4870,f371])).
fof(f8759,plain,(
  ( ! [X152,X151,X149,X150,X148] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X150),s(bool,X149),s(bool,X151),s(bool,X152))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X150))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X149),s(bool,'const.bool.COND_3'(s(bool,sK4(X152,X151,X148,X150,bool)),s(bool,X151),s(bool,X152))))) | s(bool,'T_0') = s(bool,X149) | s(bool,'T_0') = s(bool,X148)) )),
  inference(superposition,[],[f4870,f371])).
fof(f8758,plain,(
  ( ! [X146,X144,X147,X145,X143] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X145),s(bool,X146),s(bool,X144),s(bool,X147))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X145))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X146),s(bool,'const.bool.COND_3'(s(bool,sK4(X147,X143,X146,X145,bool)),s(bool,X144),s(bool,X147))))) | s(bool,'T_0') = s(bool,X144) | s(bool,'T_0') = s(bool,X143)) )),
  inference(superposition,[],[f4870,f371])).
fof(f8749,plain,(
  ( ! [X103,X107,X105,X108,X106,X104] : (s(X107,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X106),s(X107,X105),s(X107,X104),s(X107,X103))) = s(X107,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X106))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(X107,X105),s(X107,'const.bool.COND_3'(s(bool,X108),s(X107,X104),s(X107,X103))))) | s(bool,'T_0') = s(bool,X108) | s(bool,'T_0') = s(bool,sK4(X103,X104,X105,X106,X107))) )),
  inference(superposition,[],[f4870,f371])).
fof(f8717,plain,(
  ( ! [X47,X52,X50,X48,X51,X49] : (s(X49,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X47),s(X49,X50),s(X49,X51),s(X49,X52))) = s(X49,'const.bool.COND_3'(s(bool,X48),s(X49,X50),s(X49,'const.bool.COND_3'(s(bool,sK4(X52,X51,X50,X47,X49)),s(X49,X51),s(X49,X52))))) | s(bool,'T_0') = s(bool,X48) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X47))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f4870,f371])).
fof(f8666,plain,(
  ( ! [X76,X74,X75,X73] : (s(bool,'T_0') = s(bool,X74) | s(bool,X76) = s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.EQUAL_0'),s(bool,X73),s(bool,X74),s(bool,X75))) | s(bool,'T_0') = s(bool,X76)) )),
  inference(superposition,[],[f371,f7277])).
fof(f8665,plain,(
  ( ! [X70,X72,X71,X69] : (s(bool,'T_0') = s(bool,X69) | s(bool,X72) = s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.LESS_0'),s(bool,X69),s(bool,X70),s(bool,X71))) | s(bool,'T_0') = s(bool,X72)) )),
  inference(superposition,[],[f371,f7271])).
fof(f8664,plain,(
  ( ! [X68,X66,X64,X67,X65] : (s(bool,'T_0') = s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X64),s(bool,X65),s(bool,X67),s(bool,X66))) | s(bool,X68) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X64))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X65),s(bool,'const.bool.COND_3'(s(bool,sK4(X66,X67,X65,X64,bool)),s(bool,X67),s(bool,X66))))) | s(bool,'T_0') = s(bool,X68)) )),
  inference(superposition,[],[f371,f4870])).
fof(f8663,plain,(
  ( ! [X61,X62,X63] : (s(bool,'T_0') = s(bool,X62) | s(bool,X63) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X61),s(bool,X62))) | s(bool,'T_0') = s(bool,X63)) )),
  inference(superposition,[],[f371,f199])).
fof(f8662,plain,(
  ( ! [X59,X60,X58] : (s(bool,'T_0') = s(bool,X58) | s(bool,X60) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X58),s(bool,X59))) | s(bool,'T_0') = s(bool,X60)) )),
  inference(superposition,[],[f371,f198])).
fof(f8661,plain,(
  ( ! [X52,X53] : (s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X52))) | s(bool,X53) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X52))))) | s(bool,'T_0') = s(bool,X53)) )),
  inference(superposition,[],[f371,f171])).
fof(f8660,plain,(
  ( ! [X50,X51] : (s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X50))) | s(bool,X51) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X50))))) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f371,f170])).
fof(f8658,plain,(
  ( ! [X47,X48,X46] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X47),s('type.num.num',X46))) | s(bool,X48) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X46))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X47))))) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f371,f203])).
fof(f8657,plain,(
  ( ! [X43,X44] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X43),s('type.num.num','const.num.|48|_0'))) | s(bool,X44) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X43))),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f371,f4721])).
fof(f8656,plain,(
  ( ! [X41,X42,X40] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X40),s('type.num.num',X41))) | s(bool,X42) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X40))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X41))))) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f371,f204])).
fof(f8655,plain,(
  ( ! [X37,X38] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X37))) | s(bool,X38) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X37))),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f371,f4719])).
fof(f8654,plain,(
  ( ! [X35,X36,X34] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X35),s('type.num.num',X34))) | s(bool,X36) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X34))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X35))))) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f371,f208])).
fof(f8652,plain,(
  ( ! [X30,X31,X29] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X29),s('type.num.num',X30))) | s(bool,X31) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X29))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X30))))) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f371,f206])).
fof(f8651,plain,(
  ( ! [X28,X26,X27] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X26),s('type.num.num',X27))) | s(bool,X28) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X26))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X27))))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f371,f205])).
fof(f8650,plain,(
  ( ! [X24,X23,X25] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X23),s('type.num.num',X24))) | s(bool,X25) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X23))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X24))))) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f371,f207])).
fof(f8649,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X20),s('type.num.num',X21))) | s(bool,X22) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X20))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X21))))) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f371,f209])).
fof(f8648,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X16))) | s(bool,X17) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X16))))) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f371,f4720])).
fof(f8637,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f371,f371])).
fof(f8590,plain,(
  ( ! [X99,X97,X100,X98] : (s(bool,X98) = s(bool,X100) | s(bool,'T_0') = s(bool,X100) | s(bool,'T_0') = s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.EQUAL_0'),s(bool,X97),s(bool,X98),s(bool,X99)))) )),
  inference(superposition,[],[f371,f7277])).
fof(f8589,plain,(
  ( ! [X94,X95,X93,X96] : (s(bool,X93) = s(bool,X96) | s(bool,'T_0') = s(bool,X96) | s(bool,'T_0') = s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.LESS_0'),s(bool,X93),s(bool,X94),s(bool,X95)))) )),
  inference(superposition,[],[f371,f7271])).
fof(f8588,plain,(
  ( ! [X92,X90,X88,X91,X89] : (s(bool,X92) = s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X88),s(bool,X89),s(bool,X91),s(bool,X90))) | s(bool,'T_0') = s(bool,X92) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X88))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X89),s(bool,'const.bool.COND_3'(s(bool,sK4(X90,X91,X89,X88,bool)),s(bool,X91),s(bool,X90)))))) )),
  inference(superposition,[],[f371,f4870])).
fof(f8587,plain,(
  ( ! [X87,X85,X86] : (s(bool,X86) = s(bool,X87) | s(bool,'T_0') = s(bool,X87) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X85),s(bool,X86)))) )),
  inference(superposition,[],[f371,f199])).
fof(f8586,plain,(
  ( ! [X83,X84,X82] : (s(bool,X82) = s(bool,X84) | s(bool,'T_0') = s(bool,X84) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X82),s(bool,X83)))) )),
  inference(superposition,[],[f371,f198])).
fof(f8580,plain,(
  ( ! [X72,X71] : (s(bool,X72) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X71))) | s(bool,'T_0') = s(bool,X72) | s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X71)))))) )),
  inference(superposition,[],[f371,f171])).
fof(f8574,plain,(
  ( ! [X64,X65] : (s(bool,X65) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X64))) | s(bool,'T_0') = s(bool,X65) | s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X64)))))) )),
  inference(superposition,[],[f371,f170])).
fof(f8572,plain,(
  ( ! [X61,X62,X60] : (s(bool,X62) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X61),s('type.num.num',X60))) | s(bool,'T_0') = s(bool,X62) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X60))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X61)))))) )),
  inference(superposition,[],[f371,f203])).
fof(f8570,plain,(
  ( ! [X57,X56] : (s(bool,X57) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X56),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X57) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X56))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f371,f4721])).
fof(f8569,plain,(
  ( ! [X54,X55,X53] : (s(bool,X55) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X53),s('type.num.num',X54))) | s(bool,'T_0') = s(bool,X55) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X53))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X54)))))) )),
  inference(superposition,[],[f371,f204])).
fof(f8565,plain,(
  ( ! [X47,X48] : (s(bool,X48) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X47))) | s(bool,'T_0') = s(bool,X48) | s(bool,'T_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X47))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f371,f4719])).
fof(f8564,plain,(
  ( ! [X45,X46,X44] : (s(bool,X46) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X45),s('type.num.num',X44))) | s(bool,'T_0') = s(bool,X46) | s(bool,'T_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X44))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X45)))))) )),
  inference(superposition,[],[f371,f208])).
fof(f8562,plain,(
  ( ! [X39,X41,X40] : (s(bool,X41) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X39),s('type.num.num',X40))) | s(bool,'T_0') = s(bool,X41) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X39))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X40)))))) )),
  inference(superposition,[],[f371,f206])).
fof(f8561,plain,(
  ( ! [X37,X38,X36] : (s(bool,X38) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X36),s('type.num.num',X37))) | s(bool,'T_0') = s(bool,X38) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X36))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X37)))))) )),
  inference(superposition,[],[f371,f205])).
fof(f8560,plain,(
  ( ! [X35,X33,X34] : (s(bool,X35) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X33),s('type.num.num',X34))) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X33))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X34)))))) )),
  inference(superposition,[],[f371,f207])).
fof(f8559,plain,(
  ( ! [X30,X31,X32] : (s(bool,X32) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X30),s('type.num.num',X31))) | s(bool,'T_0') = s(bool,X32) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X30))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X31)))))) )),
  inference(superposition,[],[f371,f209])).
fof(f8554,plain,(
  ( ! [X23,X22] : (s(bool,X23) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X22))) | s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X22)))))) )),
  inference(superposition,[],[f371,f4720])).
fof(f8535,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))) )),
  inference(forward_demodulation,[],[f8501,f4737])).
fof(f8501,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))) )),
  inference(unit_resulting_resolution,[],[f2005,f222])).
fof(f8502,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))) )),
  inference(unit_resulting_resolution,[],[f2005,f220])).
fof(f8503,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))) )),
  inference(unit_resulting_resolution,[],[f2005,f218])).
fof(f8464,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))))) )),
  inference(unit_resulting_resolution,[],[f2004,f222])).
fof(f8465,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))))) )),
  inference(unit_resulting_resolution,[],[f2004,f220])).
fof(f8466,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))))) )),
  inference(unit_resulting_resolution,[],[f2004,f218])).
fof(f8387,plain,(
  ( ! [X0] : (s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))) )),
  inference(forward_demodulation,[],[f8352,f4714])).
fof(f8352,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))) )),
  inference(unit_resulting_resolution,[],[f1588,f222])).
fof(f8353,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))) )),
  inference(unit_resulting_resolution,[],[f1588,f220])).
fof(f8354,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))) )),
  inference(unit_resulting_resolution,[],[f1588,f218])).
fof(f8345,plain,(
  ( ! [X0] : (s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))))) )),
  inference(forward_demodulation,[],[f8310,f4714])).
fof(f8310,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))))) )),
  inference(unit_resulting_resolution,[],[f1587,f222])).
fof(f8311,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))))) )),
  inference(unit_resulting_resolution,[],[f1587,f220])).
fof(f8312,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))))) )),
  inference(unit_resulting_resolution,[],[f1587,f218])).
fof(f8231,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))) )),
  inference(forward_demodulation,[],[f8196,f4715])).
fof(f8196,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))) )),
  inference(unit_resulting_resolution,[],[f1227,f222])).
fof(f8197,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))) )),
  inference(unit_resulting_resolution,[],[f1227,f220])).
fof(f8198,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))) )),
  inference(unit_resulting_resolution,[],[f1227,f218])).
fof(f8189,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))))) )),
  inference(forward_demodulation,[],[f8154,f4715])).
fof(f8154,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))))) )),
  inference(unit_resulting_resolution,[],[f1226,f222])).
fof(f8155,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))))) )),
  inference(unit_resulting_resolution,[],[f1226,f220])).
fof(f8156,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))))) )),
  inference(unit_resulting_resolution,[],[f1226,f218])).
fof(f8011,plain,(
  ( ! [X344,X343,X345] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X344))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.EQUAL_0'),s('type.num.num',X343),s('type.num.num',X344),s('type.num.num',X345)))) )),
  inference(superposition,[],[f4722,f7277])).
fof(f7982,plain,(
  ( ! [X241,X239,X240,X238] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X239))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X241))) | s('type.num.num',X241) = s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.EQUAL_0'),s('type.num.num',X238),s('type.num.num',X239),s('type.num.num',X240)))) )),
  inference(superposition,[],[f222,f7277])).
fof(f7978,plain,(
  ( ! [X222,X225,X223,X224] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X223))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X225))) | s('type.num.num',X225) = s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.EQUAL_0'),s('type.num.num',X222),s('type.num.num',X223),s('type.num.num',X224)))) )),
  inference(superposition,[],[f220,f7277])).
fof(f7974,plain,(
  ( ! [X208,X206,X209,X207] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X207))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X209))) | s('type.num.num',X209) = s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.EQUAL_0'),s('type.num.num',X206),s('type.num.num',X207),s('type.num.num',X208)))) )),
  inference(superposition,[],[f218,f7277])).
fof(f7935,plain,(
  ( ! [X70,X68,X66,X71,X69,X67,X65] : (s(X68,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X66),s(X68,X69),s(X68,X70),s(X68,X71))) = s(X68,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X66))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(X68,X69),s(X68,'const.bool.COND_3'(s(bool,sK4(X71,X70,X69,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.EQUAL_0'),s('type.prelim.ordering',X65),s('type.prelim.ordering',X66),s('type.prelim.ordering',X67)),X68)),s(X68,X70),s(X68,X71)))))) )),
  inference(superposition,[],[f4870,f7277])).
fof(f7929,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (s(X34,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X38),s(X34,X39),s(X34,X40),s(X34,X36))) = s(X34,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X38))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(X34,X39),s(X34,'const.bool.COND_3'(s(bool,sK4('const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.EQUAL_0'),s(X34,X35),s(X34,X36),s(X34,X37)),X40,X39,X38,X34)),s(X34,X40),s(X34,X36)))))) )),
  inference(superposition,[],[f4870,f7277])).
fof(f7928,plain,(
  ( ! [X30,X28,X33,X31,X29,X27,X32] : (s(X27,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X31),s(X27,X29),s(X27,X32),s(X27,X33))) = s(X27,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X31))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(X27,X29),s(X27,'const.bool.COND_3'(s(bool,sK4(X33,X32,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.EQUAL_0'),s(X27,X28),s(X27,X29),s(X27,X30)),X31,X27)),s(X27,X32),s(X27,X33)))))) )),
  inference(superposition,[],[f4870,f7277])).
fof(f7927,plain,(
  ( ! [X26,X24,X23,X21,X25,X22,X20] : (s(X20,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X24),s(X20,X25),s(X20,X22),s(X20,X26))) = s(X20,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X24))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(X20,X25),s(X20,'const.bool.COND_3'(s(bool,sK4(X26,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.EQUAL_0'),s(X20,X21),s(X20,X22),s(X20,X23)),X25,X24,X20)),s(X20,X22),s(X20,X26)))))) )),
  inference(superposition,[],[f4870,f7277])).
fof(f7922,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.EQUAL_0'),s(bool,X0),s(bool,X1),s(bool,X2)))) )),
  inference(superposition,[],[f7277,f181])).
fof(f7714,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))),
  inference(forward_demodulation,[],[f7711,f4737])).
fof(f7711,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))),
  inference(unit_resulting_resolution,[],[f7306,f222])).
fof(f7712,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))),
  inference(unit_resulting_resolution,[],[f7306,f220])).
fof(f7713,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))),
  inference(unit_resulting_resolution,[],[f7306,f218])).
fof(f7677,plain,(
  ( ! [X325,X327,X326] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X325))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.LESS_0'),s('type.num.num',X325),s('type.num.num',X326),s('type.num.num',X327)))) )),
  inference(superposition,[],[f4722,f7271])).
fof(f7648,plain,(
  ( ! [X222,X220,X223,X221] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X220))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X223))) | s('type.num.num',X223) = s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.LESS_0'),s('type.num.num',X220),s('type.num.num',X221),s('type.num.num',X222)))) )),
  inference(superposition,[],[f222,f7271])).
fof(f7644,plain,(
  ( ! [X206,X204,X207,X205] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X204))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X207))) | s('type.num.num',X207) = s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.LESS_0'),s('type.num.num',X204),s('type.num.num',X205),s('type.num.num',X206)))) )),
  inference(superposition,[],[f220,f7271])).
fof(f7640,plain,(
  ( ! [X191,X189,X190,X188] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X188))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X191))) | s('type.num.num',X191) = s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.LESS_0'),s('type.num.num',X188),s('type.num.num',X189),s('type.num.num',X190)))) )),
  inference(superposition,[],[f218,f7271])).
fof(f7601,plain,(
  ( ! [X47,X52,X50,X48,X53,X51,X49] : (s(X50,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X47),s(X50,X51),s(X50,X52),s(X50,X53))) = s(X50,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X47))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(X50,X51),s(X50,'const.bool.COND_3'(s(bool,sK4(X53,X52,X51,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.LESS_0'),s('type.prelim.ordering',X47),s('type.prelim.ordering',X48),s('type.prelim.ordering',X49)),X50)),s(X50,X52),s(X50,X53)))))) )),
  inference(superposition,[],[f4870,f7271])).
fof(f7598,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (s(X34,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X38),s(X34,X39),s(X34,X40),s(X34,X35))) = s(X34,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X38))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(X34,X39),s(X34,'const.bool.COND_3'(s(bool,sK4('const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.LESS_0'),s(X34,X35),s(X34,X36),s(X34,X37)),X40,X39,X38,X34)),s(X34,X40),s(X34,X35)))))) )),
  inference(superposition,[],[f4870,f7271])).
fof(f7597,plain,(
  ( ! [X30,X28,X33,X31,X29,X27,X32] : (s(X27,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X31),s(X27,X28),s(X27,X32),s(X27,X33))) = s(X27,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X31))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(X27,X28),s(X27,'const.bool.COND_3'(s(bool,sK4(X33,X32,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.LESS_0'),s(X27,X28),s(X27,X29),s(X27,X30)),X31,X27)),s(X27,X32),s(X27,X33)))))) )),
  inference(superposition,[],[f4870,f7271])).
fof(f7596,plain,(
  ( ! [X26,X24,X23,X21,X25,X22,X20] : (s(X20,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X24),s(X20,X25),s(X20,X21),s(X20,X26))) = s(X20,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X24))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(X20,X25),s(X20,'const.bool.COND_3'(s(bool,sK4(X26,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.LESS_0'),s(X20,X21),s(X20,X22),s(X20,X23)),X25,X24,X20)),s(X20,X21),s(X20,X26)))))) )),
  inference(superposition,[],[f4870,f7271])).
fof(f7591,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.LESS_0'),s(bool,X0),s(bool,X1),s(bool,X2)))) )),
  inference(superposition,[],[f7271,f181])).
fof(f7358,plain,(
  ( ! [X12,X10,X11] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X10),s(bool,X11),s(bool,X12),s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X10))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X11),s(bool,'const.bool.COND_3'(s(bool,sK4('const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X12,X11,X10,bool)),s(bool,X12),s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(superposition,[],[f4870,f4727])).
fof(f7357,plain,(
  ( ! [X8,X7,X9] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X7),s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X8),s(bool,X9))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X7))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,'const.bool.COND_3'(s(bool,sK4(X9,X8,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X7,bool)),s(bool,X8),s(bool,X9)))))) )),
  inference(superposition,[],[f4870,f4727])).
fof(f7356,plain,(
  ( ! [X6,X4,X5] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X4),s(bool,X5),s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X6))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X4))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X5),s(bool,'const.bool.COND_3'(s(bool,sK4(X6,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X5,X4,bool)),s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X6)))))) )),
  inference(superposition,[],[f4870,f4727])).
fof(f7319,plain,(
  ( ! [X12,X10,X11] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X10),s(bool,X11),s(bool,X12),s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X10))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X11),s(bool,'const.bool.COND_3'(s(bool,sK4('const.arithmetic.EVEN_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))),X12,X11,X10,bool)),s(bool,X12),s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(superposition,[],[f4870,f4726])).
fof(f7318,plain,(
  ( ! [X8,X7,X9] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X7),s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X8),s(bool,X9))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X7))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,'const.bool.COND_3'(s(bool,sK4(X9,X8,'const.arithmetic.EVEN_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))),X7,bool)),s(bool,X8),s(bool,X9)))))) )),
  inference(superposition,[],[f4870,f4726])).
fof(f7317,plain,(
  ( ! [X6,X4,X5] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X4),s(bool,X5),s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X6))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X4))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X5),s(bool,'const.bool.COND_3'(s(bool,sK4(X6,'const.arithmetic.EVEN_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))),X5,X4,bool)),s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X6)))))) )),
  inference(superposition,[],[f4870,f4726])).
fof(f7306,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))),
  inference(forward_demodulation,[],[f7303,f4737])).
fof(f7303,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))),
  inference(unit_resulting_resolution,[],[f6978,f222])).
fof(f7304,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))),
  inference(unit_resulting_resolution,[],[f6978,f220])).
fof(f7305,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))),
  inference(unit_resulting_resolution,[],[f6978,f218])).
fof(f7256,plain,(
  ( ! [X379,X380,X378,X381] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X378),s('type.num.num',X379),s('type.num.num',X381),s('type.num.num',X380))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X378))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X379),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X380,X381,X379,X378,'type.num.num')),s('type.num.num',X381),s('type.num.num',X380)))))) )),
  inference(superposition,[],[f4722,f4870])).
fof(f7227,plain,(
  ( ! [X246,X244,X248,X247,X245] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X248))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X244),s('type.num.num',X245),s('type.num.num',X247),s('type.num.num',X246))))) | s('type.num.num',X248) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X244))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X245),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X246,X247,X245,X244,'type.num.num')),s('type.num.num',X247),s('type.num.num',X246)))))) )),
  inference(superposition,[],[f222,f4870])).
fof(f7223,plain,(
  ( ! [X227,X225,X228,X226,X224] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X228))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X224),s('type.num.num',X225),s('type.num.num',X227),s('type.num.num',X226))))) | s('type.num.num',X228) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X224))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X225),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X226,X227,X225,X224,'type.num.num')),s('type.num.num',X227),s('type.num.num',X226)))))) )),
  inference(superposition,[],[f220,f4870])).
fof(f7219,plain,(
  ( ! [X208,X206,X204,X207,X205] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X208))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X204),s('type.num.num',X205),s('type.num.num',X207),s('type.num.num',X206))))) | s('type.num.num',X208) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X204))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X205),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X206,X207,X205,X204,'type.num.num')),s('type.num.num',X207),s('type.num.num',X206)))))) )),
  inference(superposition,[],[f218,f4870])).
fof(f7174,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X0),s(bool,X1),s(bool,X3),s(bool,X2))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X1),s(bool,'const.bool.COND_3'(s(bool,sK4(X2,X3,X1,X0,bool)),s(bool,X3),s(bool,X2)))))) )),
  inference(superposition,[],[f4870,f181])).
fof(f7173,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X3),s(bool,X2),s(bool,X1),s(bool,X0))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X3))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X2),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,sK4(X0,X1,X2,X3,bool)),s(bool,X1),s(bool,X0)))) )),
  inference(superposition,[],[f4870,f181])).
fof(f7172,plain,(
  ( ! [X241,X239,X237,X235,X240,X238,X236,X234] : (s(X234,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X239),s(X234,X240),s(X234,X241),s(X234,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X235),s(X234,X236),s(X234,X238),s(X234,X237))))) = s(X234,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X239))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(X234,X240),s(X234,'const.bool.COND_3'(s(bool,sK4('const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X235))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(X234,X236),s(X234,'const.bool.COND_3'(s(bool,sK4(X237,X238,X236,X235,X234)),s(X234,X238),s(X234,X237)))),X241,X240,X239,X234)),s(X234,X241),s(X234,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X235),s(X234,X236),s(X234,X238),s(X234,X237)))))))) )),
  inference(superposition,[],[f4870,f4870])).
fof(f7171,plain,(
  ( ! [X231,X229,X233,X230,X228,X232] : (s(X228,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X231),s(X228,X232),s(X228,X233),s(X228,X230))) = s(X228,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X231))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(X228,X232),s(X228,'const.bool.COND_3'(s(bool,sK4('const.bool.COND_3'(s(bool,'F_0'),s(X228,X229),s(X228,X230)),X233,X232,X231,X228)),s(X228,X233),s(X228,X230)))))) )),
  inference(superposition,[],[f4870,f199])).
fof(f7170,plain,(
  ( ! [X222,X227,X225,X223,X226,X224] : (s(X222,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X225),s(X222,X226),s(X222,X227),s(X222,X223))) = s(X222,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X225))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(X222,X226),s(X222,'const.bool.COND_3'(s(bool,sK4('const.bool.COND_3'(s(bool,'T_0'),s(X222,X223),s(X222,X224)),X227,X226,X225,X222)),s(X222,X227),s(X222,X223)))))) )),
  inference(superposition,[],[f4870,f198])).
fof(f7169,plain,(
  ( ! [X218,X216,X220,X215,X219,X217,X221] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X219),s(bool,X220),s(bool,X221),s(bool,'T_0'))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X219))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X220),s(bool,'const.bool.COND_3'(s(bool,sK4(sK4(X215,X216,X217,'const.prelim.EQUAL_0',X218),X221,X220,X219,bool)),s(bool,X221),s(bool,'T_0')))))) )),
  inference(superposition,[],[f4870,f2494])).
fof(f7168,plain,(
  ( ! [X212,X213,X211,X214] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X212),s(bool,X213),s(bool,X214),s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X211))))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X212))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X213),s(bool,'const.bool.COND_3'(s(bool,sK4('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X211)))),X214,X213,X212,bool)),s(bool,X214),s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X211)))))))) )),
  inference(superposition,[],[f4870,f171])).
fof(f7167,plain,(
  ( ! [X210,X208,X209] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X208),s(bool,X209),s(bool,X210),s(bool,'T_0'))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X208))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X209),s(bool,'const.bool.COND_3'(s(bool,sK4('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X210,X209,X208,bool)),s(bool,X210),s(bool,'T_0')))))) )),
  inference(superposition,[],[f4870,f383])).
fof(f7162,plain,(
  ( ! [X194,X192,X195,X193] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X193),s(bool,X194),s(bool,X195),s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X192))))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X193))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X194),s(bool,'const.bool.COND_3'(s(bool,sK4('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X192)))),X195,X194,X193,bool)),s(bool,X195),s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X192)))))))) )),
  inference(superposition,[],[f4870,f170])).
fof(f7160,plain,(
  ( ! [X187,X185,X188,X186,X184] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X186),s(bool,X187),s(bool,X188),s(bool,'const.arithmetic.<=_2'(s('type.num.num',X185),s('type.num.num',X184))))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X186))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X187),s(bool,'const.bool.COND_3'(s(bool,sK4('const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X184))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X185)))),X188,X187,X186,bool)),s(bool,X188),s(bool,'const.arithmetic.<=_2'(s('type.num.num',X185),s('type.num.num',X184)))))))) )),
  inference(superposition,[],[f4870,f203])).
fof(f7159,plain,(
  ( ! [X182,X180,X183,X181] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X181),s(bool,X182),s(bool,X183),s(bool,'T_0'))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X181))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X182),s(bool,'const.bool.COND_3'(s(bool,sK4('const.arithmetic.>=_2'(s('type.num.num',X180),s('type.num.num','const.num.|48|_0')),X183,X182,X181,bool)),s(bool,X183),s(bool,'T_0')))))) )),
  inference(superposition,[],[f4870,f161])).
fof(f7158,plain,(
  ( ! [X177,X178,X176,X179] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X177),s(bool,X178),s(bool,X179),s(bool,'const.arithmetic.<=_2'(s('type.num.num',X176),s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X177))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X178),s(bool,'const.bool.COND_3'(s(bool,sK4('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X176))),s('type.num.num','const.num.|48|_0')),X179,X178,X177,bool)),s(bool,X179),s(bool,'const.arithmetic.<=_2'(s('type.num.num',X176),s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f4870,f4721])).
fof(f7157,plain,(
  ( ! [X175,X173,X171,X174,X172] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X173),s(bool,X174),s(bool,X175),s(bool,'const.arithmetic.<=_2'(s('type.num.num',X171),s('type.num.num',X172))))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X173))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X174),s(bool,'const.bool.COND_3'(s(bool,sK4('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X171))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X172)))),X175,X174,X173,bool)),s(bool,X175),s(bool,'const.arithmetic.<=_2'(s('type.num.num',X171),s('type.num.num',X172)))))))) )),
  inference(superposition,[],[f4870,f204])).
fof(f7156,plain,(
  ( ! [X167,X169,X170,X168] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X168),s(bool,X169),s(bool,X170),s(bool,'T_0'))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X168))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X169),s(bool,'const.bool.COND_3'(s(bool,sK4('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X167)),X170,X169,X168,bool)),s(bool,X170),s(bool,'T_0')))))) )),
  inference(superposition,[],[f4870,f164])).
fof(f7155,plain,(
  ( ! [X165,X166,X164] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X164),s(bool,X165),s(bool,X166),s(bool,'T_0'))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X164))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X165),s(bool,'const.bool.COND_3'(s(bool,sK4('const.arithmetic.>_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))),s('type.num.num','const.num.|48|_0')),X166,X165,X164,bool)),s(bool,X166),s(bool,'T_0')))))) )),
  inference(superposition,[],[f4870,f2358])).
fof(f7154,plain,(
  ( ! [X163,X161,X162] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X161),s(bool,X162),s(bool,X163),s(bool,'T_0'))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X161))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X162),s(bool,'const.bool.COND_3'(s(bool,sK4('const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')),X163,X162,X161,bool)),s(bool,X163),s(bool,'T_0')))))) )),
  inference(superposition,[],[f4870,f2357])).
fof(f7153,plain,(
  ( ! [X158,X159,X157,X160] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X158),s(bool,X159),s(bool,X160),s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X157))))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X158))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X159),s(bool,'const.bool.COND_3'(s(bool,sK4('const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X157))),s('type.num.num','const.num.|48|_0')),X160,X159,X158,bool)),s(bool,X160),s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X157)))))))) )),
  inference(superposition,[],[f4870,f4719])).
fof(f7152,plain,(
  ( ! [X156,X154,X152,X155,X153] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X154),s(bool,X155),s(bool,X156),s(bool,'const.prim_rec.<_2'(s('type.num.num',X153),s('type.num.num',X152))))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X154))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X155),s(bool,'const.bool.COND_3'(s(bool,sK4('const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X152))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X153)))),X156,X155,X154,bool)),s(bool,X156),s(bool,'const.prim_rec.<_2'(s('type.num.num',X153),s('type.num.num',X152)))))))) )),
  inference(superposition,[],[f4870,f208])).
fof(f7151,plain,(
  ( ! [X151,X149,X150,X148] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X149),s(bool,X150),s(bool,X151),s(bool,'F_0'))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X149))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X150),s(bool,'const.bool.COND_3'(s(bool,sK4('const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X148)),X151,X150,X149,bool)),s(bool,X151),s(bool,'F_0')))))) )),
  inference(superposition,[],[f4870,f168])).
fof(f7150,plain,(
  ( ! [X146,X144,X147,X145,X143] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X145),s(bool,X146),s(bool,X147),s(bool,'const.prim_rec.<_2'(s('type.num.num',X143),s('type.num.num',X144))))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X145))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X146),s(bool,'const.bool.COND_3'(s(bool,sK4('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X143))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X144)))),X147,X146,X145,bool)),s(bool,X147),s(bool,'const.prim_rec.<_2'(s('type.num.num',X143),s('type.num.num',X144)))))))) )),
  inference(superposition,[],[f4870,f206])).
fof(f7149,plain,(
  ( ! [X142,X140,X138,X141,X139] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X140),s(bool,X141),s(bool,X142),s(bool,'const.prim_rec.<_2'(s('type.num.num',X138),s('type.num.num',X139))))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X140))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X141),s(bool,'const.bool.COND_3'(s(bool,sK4('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X138))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X139)))),X142,X141,X140,bool)),s(bool,X142),s(bool,'const.prim_rec.<_2'(s('type.num.num',X138),s('type.num.num',X139)))))))) )),
  inference(superposition,[],[f4870,f205])).
fof(f7148,plain,(
  ( ! [X134,X136,X135,X133,X137] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X135),s(bool,X136),s(bool,X137),s(bool,'const.prim_rec.<_2'(s('type.num.num',X133),s('type.num.num',X134))))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X135))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X136),s(bool,'const.bool.COND_3'(s(bool,sK4('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X133))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X134)))),X137,X136,X135,bool)),s(bool,X137),s(bool,'const.prim_rec.<_2'(s('type.num.num',X133),s('type.num.num',X134)))))))) )),
  inference(superposition,[],[f4870,f207])).
fof(f7147,plain,(
  ( ! [X132,X130,X128,X131,X129] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X130),s(bool,X131),s(bool,X132),s(bool,'const.prim_rec.<_2'(s('type.num.num',X128),s('type.num.num',X129))))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X130))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X131),s(bool,'const.bool.COND_3'(s(bool,sK4('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X128))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X129)))),X132,X131,X130,bool)),s(bool,X132),s(bool,'const.prim_rec.<_2'(s('type.num.num',X128),s('type.num.num',X129)))))))) )),
  inference(superposition,[],[f4870,f209])).
fof(f7146,plain,(
  ( ! [X127,X125,X126] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X125),s(bool,X126),s(bool,X127),s(bool,'T_0'))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X125))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X126),s(bool,'const.bool.COND_3'(s(bool,sK4('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))),X127,X126,X125,bool)),s(bool,X127),s(bool,'T_0')))))) )),
  inference(superposition,[],[f4870,f2422])).
fof(f7145,plain,(
  ( ! [X123,X124,X122] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X122),s(bool,X123),s(bool,X124),s(bool,'T_0'))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X122))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X123),s(bool,'const.bool.COND_3'(s(bool,sK4('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X124,X123,X122,bool)),s(bool,X124),s(bool,'T_0')))))) )),
  inference(superposition,[],[f4870,f2421])).
fof(f7144,plain,(
  ( ! [X121,X118,X120,X119] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X119),s(bool,X120),s(bool,X121),s(bool,'T_0'))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X119))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X120),s(bool,'const.bool.COND_3'(s(bool,sK4('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X118)))),X121,X120,X119,bool)),s(bool,X121),s(bool,'T_0')))))) )),
  inference(superposition,[],[f4870,f4717])).
fof(f7143,plain,(
  ( ! [X116,X114,X117,X115] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X115),s(bool,X116),s(bool,X117),s(bool,'T_0'))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X115))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X116),s(bool,'const.bool.COND_3'(s(bool,sK4('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X114)))),X117,X116,X115,bool)),s(bool,X117),s(bool,'T_0')))))) )),
  inference(superposition,[],[f4870,f4718])).
fof(f7142,plain,(
  ( ! [X111,X112,X110,X113] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X111),s(bool,X112),s(bool,X113),s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X110))))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X111))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X112),s(bool,'const.bool.COND_3'(s(bool,sK4('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X110)))),X113,X112,X111,bool)),s(bool,X113),s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X110)))))))) )),
  inference(superposition,[],[f4870,f4720])).
fof(f7141,plain,(
  ( ! [X109,X107,X108,X106] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X107),s(bool,X108),s(bool,X109),s(bool,'F_0'))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X107))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X108),s(bool,'const.bool.COND_3'(s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',X106),s('type.num.num','const.num.|48|_0')),X109,X108,X107,bool)),s(bool,X109),s(bool,'F_0')))))) )),
  inference(superposition,[],[f4870,f165])).
fof(f7140,plain,(
  ( ! [X103,X105,X102,X104] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X103),s(bool,X104),s(bool,X105),s(bool,X102))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X103))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X104),s(bool,'const.bool.COND_3'(s(bool,sK4('F_0',X105,X104,X103,bool)),s(bool,X105),s(bool,X102))))) | s(bool,'T_0') = s(bool,X102)) )),
  inference(superposition,[],[f4870,f181])).
fof(f7139,plain,(
  ( ! [X101,X99,X100,X98] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X99),s(bool,X100),s(bool,X101),s(bool,'F_0'))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X99))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X100),s(bool,'const.bool.COND_3'(s(bool,sK4(X98,X101,X100,X99,bool)),s(bool,X101),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X98)) )),
  inference(superposition,[],[f4870,f181])).
fof(f7138,plain,(
  ( ! [X97,X95,X96] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X95),s('type.num.num',X96),s('type.num.num',X97),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X95))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X96),s('type.num.num','const.bool.COND_3'(s(bool,sK4('const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.EQUAL_0')),X97,X96,X95,'type.num.num')),s('type.num.num',X97),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f4870,f228])).
fof(f7137,plain,(
  ( ! [X94,X92,X93] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X92),s('type.num.num',X93),s('type.num.num',X94),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X92))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X93),s('type.num.num','const.bool.COND_3'(s(bool,sK4('const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.LESS_0')),X94,X93,X92,'type.num.num')),s('type.num.num',X94),s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f4870,f152])).
fof(f7136,plain,(
  ( ! [X90,X88,X91,X89] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X89),s('type.num.num',X90),s('type.num.num',X91),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X88))))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X89))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X90),s('type.num.num','const.bool.COND_3'(s(bool,sK4('const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X88)))),X91,X90,X89,'type.num.num')),s('type.num.num',X91),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X88)))))))))) )),
  inference(superposition,[],[f4870,f175])).
fof(f7135,plain,(
  ( ! [X87,X85,X86] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X85),s('type.num.num',X86),s('type.num.num',X87),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X85))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X86),s('type.num.num','const.bool.COND_3'(s(bool,sK4('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')),X87,X86,X85,'type.num.num')),s('type.num.num',X87),s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f4870,f151])).
fof(f7134,plain,(
  ( ! [X83,X81,X84,X82] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X82),s('type.num.num',X83),s('type.num.num',X84),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X81))))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X82))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X83),s('type.num.num','const.bool.COND_3'(s(bool,sK4('const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X81)))),X84,X83,X82,'type.num.num')),s('type.num.num',X84),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X81)))))))))) )),
  inference(superposition,[],[f4870,f173])).
fof(f7133,plain,(
  ( ! [X80,X78,X79] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X78),s('type.num.num',X79),s('type.num.num',X80),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X78))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X79),s('type.num.num','const.bool.COND_3'(s(bool,sK4('const.arithmetic.ZERO_0',X80,X79,X78,'type.num.num')),s('type.num.num',X80),s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f4870,f4713])).
fof(f7132,plain,(
  ( ! [X76,X74,X77,X75,X73] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X75),s('type.num.num',X76),s('type.num.num',X77),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X73),s('type.num.num',X74))))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X75))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X76),s('type.num.num','const.bool.COND_3'(s(bool,sK4('const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X73))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X74)))),X77,X76,X75,'type.num.num')),s('type.num.num',X77),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X73),s('type.num.num',X74)))))))))) )),
  inference(superposition,[],[f4870,f210])).
fof(f7131,plain,(
  ( ! [X70,X72,X71] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X70),s('type.num.num',X71),s('type.num.num',X72),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X70))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X71),s('type.num.num','const.bool.COND_3'(s(bool,sK4('const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X72,X71,X70,'type.num.num')),s('type.num.num',X72),s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f4870,f1163])).
fof(f7130,plain,(
  ( ! [X68,X69,X67] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X67),s('type.num.num',X68),s('type.num.num',X69),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X67))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X68),s('type.num.num','const.bool.COND_3'(s(bool,sK4('const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))),X69,X68,X67,'type.num.num')),s('type.num.num',X69),s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f4870,f1093])).
fof(f7129,plain,(
  ( ! [X66,X64,X65,X63] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X64),s('type.num.num',X65),s('type.num.num',X66),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X64))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X65),s('type.num.num','const.bool.COND_3'(s(bool,sK4('const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X63)))))),X66,X65,X64,'type.num.num')),s('type.num.num',X66),s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f4870,f177])).
fof(f7128,plain,(
  ( ! [X61,X59,X62,X60] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X60),s('type.num.num',X61),s('type.num.num',X62),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X60))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X61),s('type.num.num','const.bool.COND_3'(s(bool,sK4('const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X59)))))),X62,X61,X60,'type.num.num')),s('type.num.num',X62),s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f4870,f176])).
fof(f7127,plain,(
  ( ! [X57,X58,X56,X55] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X56),s('type.num.num',X57),s('type.num.num',X58),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X56))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X57),s('type.num.num','const.bool.COND_3'(s(bool,sK4('const.arithmetic.EXP_2'(s('type.num.num',X55),s('type.num.num','const.num.|48|_0')),X58,X57,X56,'type.num.num')),s('type.num.num',X58),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f4870,f229])).
fof(f7126,plain,(
  ( ! [X54,X52,X50,X53,X51] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X52),s('type.num.num',X53),s('type.num.num',X54),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X50),s('type.num.num',X51))))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X52))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X53),s('type.num.num','const.bool.COND_3'(s(bool,sK4('const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X50))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X51)))),X54,X53,X52,'type.num.num')),s('type.num.num',X54),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X50),s('type.num.num',X51)))))))))) )),
  inference(superposition,[],[f4870,f211])).
fof(f7125,plain,(
  ( ! [X47,X48,X46,X49] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X47),s('type.num.num',X48),s('type.num.num',X49),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X47))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X48),s('type.num.num','const.bool.COND_3'(s(bool,sK4('const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X46)),X49,X48,X47,'type.num.num')),s('type.num.num',X49),s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f4870,f162])).
fof(f7124,plain,(
  ( ! [X45,X43,X44,X42] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X43),s('type.num.num',X44),s('type.num.num',X45),s('type.num.num',X42))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X43))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X44),s('type.num.num','const.bool.COND_3'(s(bool,sK4('const.arithmetic.-_2'(s('type.num.num',X42),s('type.num.num','const.num.|48|_0')),X45,X44,X43,'type.num.num')),s('type.num.num',X45),s('type.num.num',X42)))))) )),
  inference(superposition,[],[f4870,f163])).
fof(f7123,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X39),s('type.num.num',X40),s('type.num.num',X41),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X37),s('type.num.num',X38))))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X39))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X40),s('type.num.num','const.bool.COND_3'(s(bool,sK4('const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X37))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X38)))),X41,X40,X39,'type.num.num')),s('type.num.num',X41),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X37),s('type.num.num',X38)))))))))) )),
  inference(superposition,[],[f4870,f212])).
fof(f7122,plain,(
  ( ! [X35,X33,X36,X34] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X34),s('type.num.num',X35),s('type.num.num',X36),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X34))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X35),s('type.num.num','const.bool.COND_3'(s(bool,sK4('const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X33)),X36,X35,X34,'type.num.num')),s('type.num.num',X36),s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f4870,f167])).
fof(f7121,plain,(
  ( ! [X30,X31,X29,X32] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X30),s('type.num.num',X31),s('type.num.num',X32),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X30))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X31),s('type.num.num','const.bool.COND_3'(s(bool,sK4('const.arithmetic.*_2'(s('type.num.num',X29),s('type.num.num','const.num.|48|_0')),X32,X31,X30,'type.num.num')),s('type.num.num',X32),s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f4870,f166])).
fof(f7120,plain,(
  ( ! [X28,X26,X27,X25] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X26),s('type.num.num',X27),s('type.num.num',X28),s('type.num.num',X25))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X26))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X27),s('type.num.num','const.bool.COND_3'(s(bool,sK4('const.numeral.iZ_1'(s('type.num.num',X25)),X28,X27,X26,'type.num.num')),s('type.num.num',X28),s('type.num.num',X25)))))) )),
  inference(superposition,[],[f4870,f5694])).
fof(f7119,plain,(
  ( ! [X24,X23,X22] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X22),s('type.num.num',X23),s('type.num.num',X24),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X22))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X23),s('type.num.num','const.bool.COND_3'(s(bool,sK4('const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X24,X23,X22,'type.num.num')),s('type.num.num',X24),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f4870,f4737])).
fof(f7118,plain,(
  ( ! [X21,X19,X20] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X19),s('type.num.num',X20),s('type.num.num',X21),s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X19))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X20),s('type.num.num','const.bool.COND_3'(s(bool,sK4('const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))),X21,X20,X19,'type.num.num')),s('type.num.num',X21),s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))) )),
  inference(superposition,[],[f4870,f4714])).
fof(f7117,plain,(
  ( ! [X17,X18,X16] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X16),s('type.num.num',X17),s('type.num.num',X18),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X16))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X17),s('type.num.num','const.bool.COND_3'(s(bool,sK4('const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),X18,X17,X16,'type.num.num')),s('type.num.num',X18),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f4870,f4715])).
fof(f7116,plain,(
  ( ! [X14,X15,X13] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X13),s('type.num.num',X14),s('type.num.num',X15),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X13))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X14),s('type.num.num','const.bool.COND_3'(s(bool,sK4('const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')),X15,X14,X13,'type.num.num')),s('type.num.num',X15),s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f4870,f4578])).
fof(f7115,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X10),s('type.num.num',X11),s('type.num.num',X12),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X9))))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X10))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X11),s('type.num.num','const.bool.COND_3'(s(bool,sK4('const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X9)))),X12,X11,X10,'type.num.num')),s('type.num.num',X12),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X9)))))))))) )),
  inference(superposition,[],[f4870,f4605])).
fof(f7114,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X5),s('type.num.num',X6),s('type.num.num',X7),s('type.num.num',X4))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X5))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X6),s('type.num.num','const.bool.COND_3'(s(bool,sK4('const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')),X7,X6,X5,'type.num.num')),s('type.num.num',X7),s('type.num.num',X4)))))) )),
  inference(superposition,[],[f4870,f160])).
fof(f7113,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X1),s('type.num.num',X2),s('type.num.num',X3),s('type.num.num',X0))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X2),s('type.num.num','const.bool.COND_3'(s(bool,sK4('const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)),X3,X2,X1,'type.num.num')),s('type.num.num',X3),s('type.num.num',X0)))))) )),
  inference(superposition,[],[f4870,f158])).
fof(f7112,plain,(
  ( ! [X241,X239,X237,X235,X240,X238,X236,X234] : (s(X234,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X239),s(X234,X240),s(X234,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X235),s(X234,X236),s(X234,X238),s(X234,X237))),s(X234,X241))) = s(X234,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X239))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(X234,X240),s(X234,'const.bool.COND_3'(s(bool,sK4(X241,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X235))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(X234,X236),s(X234,'const.bool.COND_3'(s(bool,sK4(X237,X238,X236,X235,X234)),s(X234,X238),s(X234,X237)))),X240,X239,X234)),s(X234,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X235),s(X234,X236),s(X234,X238),s(X234,X237))),s(X234,X241)))))) )),
  inference(superposition,[],[f4870,f4870])).
fof(f7111,plain,(
  ( ! [X231,X229,X233,X230,X228,X232] : (s(X228,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X231),s(X228,X232),s(X228,X230),s(X228,X233))) = s(X228,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X231))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(X228,X232),s(X228,'const.bool.COND_3'(s(bool,sK4(X233,'const.bool.COND_3'(s(bool,'F_0'),s(X228,X229),s(X228,X230)),X232,X231,X228)),s(X228,X230),s(X228,X233)))))) )),
  inference(superposition,[],[f4870,f199])).
fof(f7110,plain,(
  ( ! [X222,X227,X225,X223,X226,X224] : (s(X222,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X225),s(X222,X226),s(X222,X223),s(X222,X227))) = s(X222,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X225))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(X222,X226),s(X222,'const.bool.COND_3'(s(bool,sK4(X227,'const.bool.COND_3'(s(bool,'T_0'),s(X222,X223),s(X222,X224)),X226,X225,X222)),s(X222,X223),s(X222,X227)))))) )),
  inference(superposition,[],[f4870,f198])).
fof(f7109,plain,(
  ( ! [X218,X216,X220,X215,X219,X217,X221] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X219),s(bool,X220),s(bool,'T_0'),s(bool,X221))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X219))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X220),s(bool,'const.bool.COND_3'(s(bool,sK4(X221,sK4(X215,X216,X217,'const.prelim.EQUAL_0',X218),X220,X219,bool)),s(bool,'T_0'),s(bool,X221)))))) )),
  inference(superposition,[],[f4870,f2494])).
fof(f7108,plain,(
  ( ! [X212,X213,X211,X214] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X212),s(bool,X213),s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X211))),s(bool,X214))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X212))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X213),s(bool,'const.bool.COND_3'(s(bool,sK4(X214,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X211)))),X213,X212,bool)),s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X211))),s(bool,X214)))))) )),
  inference(superposition,[],[f4870,f171])).
fof(f7107,plain,(
  ( ! [X210,X208,X209] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X208),s(bool,X209),s(bool,'T_0'),s(bool,X210))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X208))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X209),s(bool,'const.bool.COND_3'(s(bool,sK4(X210,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X209,X208,bool)),s(bool,'T_0'),s(bool,X210)))))) )),
  inference(superposition,[],[f4870,f383])).
fof(f7102,plain,(
  ( ! [X194,X192,X195,X193] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X193),s(bool,X194),s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X192))),s(bool,X195))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X193))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X194),s(bool,'const.bool.COND_3'(s(bool,sK4(X195,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X192)))),X194,X193,bool)),s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X192))),s(bool,X195)))))) )),
  inference(superposition,[],[f4870,f170])).
fof(f7100,plain,(
  ( ! [X187,X185,X188,X186,X184] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X186),s(bool,X187),s(bool,'const.arithmetic.<=_2'(s('type.num.num',X185),s('type.num.num',X184))),s(bool,X188))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X186))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X187),s(bool,'const.bool.COND_3'(s(bool,sK4(X188,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X184))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X185)))),X187,X186,bool)),s(bool,'const.arithmetic.<=_2'(s('type.num.num',X185),s('type.num.num',X184))),s(bool,X188)))))) )),
  inference(superposition,[],[f4870,f203])).
fof(f7099,plain,(
  ( ! [X182,X180,X183,X181] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X181),s(bool,X182),s(bool,'T_0'),s(bool,X183))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X181))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X182),s(bool,'const.bool.COND_3'(s(bool,sK4(X183,'const.arithmetic.>=_2'(s('type.num.num',X180),s('type.num.num','const.num.|48|_0')),X182,X181,bool)),s(bool,'T_0'),s(bool,X183)))))) )),
  inference(superposition,[],[f4870,f161])).
fof(f7098,plain,(
  ( ! [X177,X178,X176,X179] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X177),s(bool,X178),s(bool,'const.arithmetic.<=_2'(s('type.num.num',X176),s('type.num.num','const.num.|48|_0'))),s(bool,X179))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X177))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X178),s(bool,'const.bool.COND_3'(s(bool,sK4(X179,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X176))),s('type.num.num','const.num.|48|_0')),X178,X177,bool)),s(bool,'const.arithmetic.<=_2'(s('type.num.num',X176),s('type.num.num','const.num.|48|_0'))),s(bool,X179)))))) )),
  inference(superposition,[],[f4870,f4721])).
fof(f7097,plain,(
  ( ! [X175,X173,X171,X174,X172] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X173),s(bool,X174),s(bool,'const.arithmetic.<=_2'(s('type.num.num',X171),s('type.num.num',X172))),s(bool,X175))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X173))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X174),s(bool,'const.bool.COND_3'(s(bool,sK4(X175,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X171))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X172)))),X174,X173,bool)),s(bool,'const.arithmetic.<=_2'(s('type.num.num',X171),s('type.num.num',X172))),s(bool,X175)))))) )),
  inference(superposition,[],[f4870,f204])).
fof(f7096,plain,(
  ( ! [X167,X169,X170,X168] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X168),s(bool,X169),s(bool,'T_0'),s(bool,X170))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X168))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X169),s(bool,'const.bool.COND_3'(s(bool,sK4(X170,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X167)),X169,X168,bool)),s(bool,'T_0'),s(bool,X170)))))) )),
  inference(superposition,[],[f4870,f164])).
fof(f7095,plain,(
  ( ! [X165,X166,X164] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X164),s(bool,X165),s(bool,'T_0'),s(bool,X166))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X164))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X165),s(bool,'const.bool.COND_3'(s(bool,sK4(X166,'const.arithmetic.>_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))),s('type.num.num','const.num.|48|_0')),X165,X164,bool)),s(bool,'T_0'),s(bool,X166)))))) )),
  inference(superposition,[],[f4870,f2358])).
fof(f7094,plain,(
  ( ! [X163,X161,X162] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X161),s(bool,X162),s(bool,'T_0'),s(bool,X163))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X161))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X162),s(bool,'const.bool.COND_3'(s(bool,sK4(X163,'const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')),X162,X161,bool)),s(bool,'T_0'),s(bool,X163)))))) )),
  inference(superposition,[],[f4870,f2357])).
fof(f7093,plain,(
  ( ! [X158,X159,X157,X160] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X158),s(bool,X159),s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X157))),s(bool,X160))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X158))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X159),s(bool,'const.bool.COND_3'(s(bool,sK4(X160,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X157))),s('type.num.num','const.num.|48|_0')),X159,X158,bool)),s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X157))),s(bool,X160)))))) )),
  inference(superposition,[],[f4870,f4719])).
fof(f7092,plain,(
  ( ! [X156,X154,X152,X155,X153] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X154),s(bool,X155),s(bool,'const.prim_rec.<_2'(s('type.num.num',X153),s('type.num.num',X152))),s(bool,X156))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X154))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X155),s(bool,'const.bool.COND_3'(s(bool,sK4(X156,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X152))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X153)))),X155,X154,bool)),s(bool,'const.prim_rec.<_2'(s('type.num.num',X153),s('type.num.num',X152))),s(bool,X156)))))) )),
  inference(superposition,[],[f4870,f208])).
fof(f7091,plain,(
  ( ! [X151,X149,X150,X148] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X149),s(bool,X150),s(bool,'F_0'),s(bool,X151))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X149))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X150),s(bool,'const.bool.COND_3'(s(bool,sK4(X151,'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X148)),X150,X149,bool)),s(bool,'F_0'),s(bool,X151)))))) )),
  inference(superposition,[],[f4870,f168])).
fof(f7090,plain,(
  ( ! [X146,X144,X147,X145,X143] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X145),s(bool,X146),s(bool,'const.prim_rec.<_2'(s('type.num.num',X143),s('type.num.num',X144))),s(bool,X147))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X145))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X146),s(bool,'const.bool.COND_3'(s(bool,sK4(X147,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X143))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X144)))),X146,X145,bool)),s(bool,'const.prim_rec.<_2'(s('type.num.num',X143),s('type.num.num',X144))),s(bool,X147)))))) )),
  inference(superposition,[],[f4870,f206])).
fof(f7089,plain,(
  ( ! [X142,X140,X138,X141,X139] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X140),s(bool,X141),s(bool,'const.prim_rec.<_2'(s('type.num.num',X138),s('type.num.num',X139))),s(bool,X142))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X140))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X141),s(bool,'const.bool.COND_3'(s(bool,sK4(X142,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X138))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X139)))),X141,X140,bool)),s(bool,'const.prim_rec.<_2'(s('type.num.num',X138),s('type.num.num',X139))),s(bool,X142)))))) )),
  inference(superposition,[],[f4870,f205])).
fof(f7088,plain,(
  ( ! [X134,X136,X135,X133,X137] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X135),s(bool,X136),s(bool,'const.prim_rec.<_2'(s('type.num.num',X133),s('type.num.num',X134))),s(bool,X137))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X135))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X136),s(bool,'const.bool.COND_3'(s(bool,sK4(X137,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X133))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X134)))),X136,X135,bool)),s(bool,'const.prim_rec.<_2'(s('type.num.num',X133),s('type.num.num',X134))),s(bool,X137)))))) )),
  inference(superposition,[],[f4870,f207])).
fof(f7087,plain,(
  ( ! [X132,X130,X128,X131,X129] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X130),s(bool,X131),s(bool,'const.prim_rec.<_2'(s('type.num.num',X128),s('type.num.num',X129))),s(bool,X132))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X130))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X131),s(bool,'const.bool.COND_3'(s(bool,sK4(X132,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X128))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X129)))),X131,X130,bool)),s(bool,'const.prim_rec.<_2'(s('type.num.num',X128),s('type.num.num',X129))),s(bool,X132)))))) )),
  inference(superposition,[],[f4870,f209])).
fof(f7086,plain,(
  ( ! [X127,X125,X126] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X125),s(bool,X126),s(bool,'T_0'),s(bool,X127))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X125))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X126),s(bool,'const.bool.COND_3'(s(bool,sK4(X127,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))),X126,X125,bool)),s(bool,'T_0'),s(bool,X127)))))) )),
  inference(superposition,[],[f4870,f2422])).
fof(f7085,plain,(
  ( ! [X123,X124,X122] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X122),s(bool,X123),s(bool,'T_0'),s(bool,X124))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X122))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X123),s(bool,'const.bool.COND_3'(s(bool,sK4(X124,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X123,X122,bool)),s(bool,'T_0'),s(bool,X124)))))) )),
  inference(superposition,[],[f4870,f2421])).
fof(f7084,plain,(
  ( ! [X121,X118,X120,X119] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X119),s(bool,X120),s(bool,'T_0'),s(bool,X121))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X119))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X120),s(bool,'const.bool.COND_3'(s(bool,sK4(X121,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X118)))),X120,X119,bool)),s(bool,'T_0'),s(bool,X121)))))) )),
  inference(superposition,[],[f4870,f4717])).
fof(f7083,plain,(
  ( ! [X116,X114,X117,X115] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X115),s(bool,X116),s(bool,'T_0'),s(bool,X117))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X115))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X116),s(bool,'const.bool.COND_3'(s(bool,sK4(X117,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X114)))),X116,X115,bool)),s(bool,'T_0'),s(bool,X117)))))) )),
  inference(superposition,[],[f4870,f4718])).
fof(f7082,plain,(
  ( ! [X111,X112,X110,X113] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X111),s(bool,X112),s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X110))),s(bool,X113))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X111))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X112),s(bool,'const.bool.COND_3'(s(bool,sK4(X113,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X110)))),X112,X111,bool)),s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X110))),s(bool,X113)))))) )),
  inference(superposition,[],[f4870,f4720])).
fof(f7081,plain,(
  ( ! [X109,X107,X108,X106] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X107),s(bool,X108),s(bool,'F_0'),s(bool,X109))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X107))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X108),s(bool,'const.bool.COND_3'(s(bool,sK4(X109,'const.prim_rec.<_2'(s('type.num.num',X106),s('type.num.num','const.num.|48|_0')),X108,X107,bool)),s(bool,'F_0'),s(bool,X109)))))) )),
  inference(superposition,[],[f4870,f165])).
fof(f7080,plain,(
  ( ! [X103,X105,X102,X104] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X103),s(bool,X104),s(bool,X102),s(bool,X105))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X103))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X104),s(bool,'const.bool.COND_3'(s(bool,sK4(X105,'F_0',X104,X103,bool)),s(bool,X102),s(bool,X105))))) | s(bool,'T_0') = s(bool,X102)) )),
  inference(superposition,[],[f4870,f181])).
fof(f7079,plain,(
  ( ! [X101,X99,X100,X98] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X99),s(bool,X100),s(bool,'F_0'),s(bool,X101))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X99))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X100),s(bool,'const.bool.COND_3'(s(bool,sK4(X101,X98,X100,X99,bool)),s(bool,'F_0'),s(bool,X101))))) | s(bool,'T_0') = s(bool,X98)) )),
  inference(superposition,[],[f4870,f181])).
fof(f7078,plain,(
  ( ! [X97,X95,X96] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X95),s('type.num.num',X96),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X97))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X95))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X96),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X97,'const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.EQUAL_0')),X96,X95,'type.num.num')),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X97)))))) )),
  inference(superposition,[],[f4870,f228])).
fof(f7077,plain,(
  ( ! [X94,X92,X93] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X92),s('type.num.num',X93),s('type.num.num','const.num.|48|_0'),s('type.num.num',X94))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X92))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X93),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X94,'const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.LESS_0')),X93,X92,'type.num.num')),s('type.num.num','const.num.|48|_0'),s('type.num.num',X94)))))) )),
  inference(superposition,[],[f4870,f152])).
fof(f7076,plain,(
  ( ! [X90,X88,X91,X89] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X89),s('type.num.num',X90),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X88))))),s('type.num.num',X91))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X89))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X90),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X91,'const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X88)))),X90,X89,'type.num.num')),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X88))))),s('type.num.num',X91)))))) )),
  inference(superposition,[],[f4870,f175])).
fof(f7075,plain,(
  ( ! [X87,X85,X86] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X85),s('type.num.num',X86),s('type.num.num','const.num.|48|_0'),s('type.num.num',X87))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X85))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X86),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X87,'const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')),X86,X85,'type.num.num')),s('type.num.num','const.num.|48|_0'),s('type.num.num',X87)))))) )),
  inference(superposition,[],[f4870,f151])).
fof(f7074,plain,(
  ( ! [X83,X81,X84,X82] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X82),s('type.num.num',X83),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X81))))),s('type.num.num',X84))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X82))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X83),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X84,'const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X81)))),X83,X82,'type.num.num')),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X81))))),s('type.num.num',X84)))))) )),
  inference(superposition,[],[f4870,f173])).
fof(f7073,plain,(
  ( ! [X80,X78,X79] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X78),s('type.num.num',X79),s('type.num.num','const.num.|48|_0'),s('type.num.num',X80))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X78))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X79),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X80,'const.arithmetic.ZERO_0',X79,X78,'type.num.num')),s('type.num.num','const.num.|48|_0'),s('type.num.num',X80)))))) )),
  inference(superposition,[],[f4870,f4713])).
fof(f7072,plain,(
  ( ! [X76,X74,X77,X75,X73] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X75),s('type.num.num',X76),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X73),s('type.num.num',X74))))),s('type.num.num',X77))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X75))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X76),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X77,'const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X73))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X74)))),X76,X75,'type.num.num')),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X73),s('type.num.num',X74))))),s('type.num.num',X77)))))) )),
  inference(superposition,[],[f4870,f210])).
fof(f7071,plain,(
  ( ! [X70,X72,X71] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X70),s('type.num.num',X71),s('type.num.num','const.num.|48|_0'),s('type.num.num',X72))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X70))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X71),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X72,'const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X71,X70,'type.num.num')),s('type.num.num','const.num.|48|_0'),s('type.num.num',X72)))))) )),
  inference(superposition,[],[f4870,f1163])).
fof(f7070,plain,(
  ( ! [X68,X69,X67] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X67),s('type.num.num',X68),s('type.num.num','const.num.|48|_0'),s('type.num.num',X69))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X67))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X68),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X69,'const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))),X68,X67,'type.num.num')),s('type.num.num','const.num.|48|_0'),s('type.num.num',X69)))))) )),
  inference(superposition,[],[f4870,f1093])).
fof(f7069,plain,(
  ( ! [X66,X64,X65,X63] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X64),s('type.num.num',X65),s('type.num.num','const.num.|48|_0'),s('type.num.num',X66))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X64))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X65),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X66,'const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X63)))))),X65,X64,'type.num.num')),s('type.num.num','const.num.|48|_0'),s('type.num.num',X66)))))) )),
  inference(superposition,[],[f4870,f177])).
fof(f7068,plain,(
  ( ! [X61,X59,X62,X60] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X60),s('type.num.num',X61),s('type.num.num','const.num.|48|_0'),s('type.num.num',X62))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X60))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X61),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X62,'const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X59)))))),X61,X60,'type.num.num')),s('type.num.num','const.num.|48|_0'),s('type.num.num',X62)))))) )),
  inference(superposition,[],[f4870,f176])).
fof(f7067,plain,(
  ( ! [X57,X58,X56,X55] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X56),s('type.num.num',X57),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X58))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X56))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X57),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X58,'const.arithmetic.EXP_2'(s('type.num.num',X55),s('type.num.num','const.num.|48|_0')),X57,X56,'type.num.num')),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X58)))))) )),
  inference(superposition,[],[f4870,f229])).
fof(f7066,plain,(
  ( ! [X54,X52,X50,X53,X51] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X52),s('type.num.num',X53),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X50),s('type.num.num',X51))))),s('type.num.num',X54))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X52))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X53),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X54,'const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X50))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X51)))),X53,X52,'type.num.num')),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X50),s('type.num.num',X51))))),s('type.num.num',X54)))))) )),
  inference(superposition,[],[f4870,f211])).
fof(f7065,plain,(
  ( ! [X47,X48,X46,X49] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X47),s('type.num.num',X48),s('type.num.num','const.num.|48|_0'),s('type.num.num',X49))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X47))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X48),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X49,'const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X46)),X48,X47,'type.num.num')),s('type.num.num','const.num.|48|_0'),s('type.num.num',X49)))))) )),
  inference(superposition,[],[f4870,f162])).
fof(f7064,plain,(
  ( ! [X45,X43,X44,X42] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X43),s('type.num.num',X44),s('type.num.num',X42),s('type.num.num',X45))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X43))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X44),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X45,'const.arithmetic.-_2'(s('type.num.num',X42),s('type.num.num','const.num.|48|_0')),X44,X43,'type.num.num')),s('type.num.num',X42),s('type.num.num',X45)))))) )),
  inference(superposition,[],[f4870,f163])).
fof(f7063,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X39),s('type.num.num',X40),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X37),s('type.num.num',X38))))),s('type.num.num',X41))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X39))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X40),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X41,'const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X37))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X38)))),X40,X39,'type.num.num')),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X37),s('type.num.num',X38))))),s('type.num.num',X41)))))) )),
  inference(superposition,[],[f4870,f212])).
fof(f7062,plain,(
  ( ! [X35,X33,X36,X34] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X34),s('type.num.num',X35),s('type.num.num','const.num.|48|_0'),s('type.num.num',X36))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X34))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X35),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X36,'const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X33)),X35,X34,'type.num.num')),s('type.num.num','const.num.|48|_0'),s('type.num.num',X36)))))) )),
  inference(superposition,[],[f4870,f167])).
fof(f7061,plain,(
  ( ! [X30,X31,X29,X32] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X30),s('type.num.num',X31),s('type.num.num','const.num.|48|_0'),s('type.num.num',X32))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X30))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X31),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X32,'const.arithmetic.*_2'(s('type.num.num',X29),s('type.num.num','const.num.|48|_0')),X31,X30,'type.num.num')),s('type.num.num','const.num.|48|_0'),s('type.num.num',X32)))))) )),
  inference(superposition,[],[f4870,f166])).
fof(f7060,plain,(
  ( ! [X28,X26,X27,X25] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X26),s('type.num.num',X27),s('type.num.num',X25),s('type.num.num',X28))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X26))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X27),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X28,'const.numeral.iZ_1'(s('type.num.num',X25)),X27,X26,'type.num.num')),s('type.num.num',X25),s('type.num.num',X28)))))) )),
  inference(superposition,[],[f4870,f5694])).
fof(f7059,plain,(
  ( ! [X24,X23,X22] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X22),s('type.num.num',X23),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X24))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X22))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X23),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X24,'const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X23,X22,'type.num.num')),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X24)))))) )),
  inference(superposition,[],[f4870,f4737])).
fof(f7058,plain,(
  ( ! [X21,X19,X20] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X19),s('type.num.num',X20),s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))),s('type.num.num',X21))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X19))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X20),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X21,'const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))),X20,X19,'type.num.num')),s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))),s('type.num.num',X21)))))) )),
  inference(superposition,[],[f4870,f4714])).
fof(f7057,plain,(
  ( ! [X17,X18,X16] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X16),s('type.num.num',X17),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X18))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X16))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X17),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X18,'const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),X17,X16,'type.num.num')),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X18)))))) )),
  inference(superposition,[],[f4870,f4715])).
fof(f7056,plain,(
  ( ! [X14,X15,X13] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X13),s('type.num.num',X14),s('type.num.num','const.num.|48|_0'),s('type.num.num',X15))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X13))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X14),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X15,'const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')),X14,X13,'type.num.num')),s('type.num.num','const.num.|48|_0'),s('type.num.num',X15)))))) )),
  inference(superposition,[],[f4870,f4578])).
fof(f7055,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X10),s('type.num.num',X11),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X9))))),s('type.num.num',X12))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X10))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X11),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X12,'const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X9)))),X11,X10,'type.num.num')),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X9))))),s('type.num.num',X12)))))) )),
  inference(superposition,[],[f4870,f4605])).
fof(f7054,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X5),s('type.num.num',X6),s('type.num.num',X4),s('type.num.num',X7))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X5))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X6),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X7,'const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')),X6,X5,'type.num.num')),s('type.num.num',X4),s('type.num.num',X7)))))) )),
  inference(superposition,[],[f4870,f160])).
fof(f7053,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X1),s('type.num.num',X2),s('type.num.num',X0),s('type.num.num',X3))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X2),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X3,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)),X2,X1,'type.num.num')),s('type.num.num',X0),s('type.num.num',X3)))))) )),
  inference(superposition,[],[f4870,f158])).
fof(f7284,plain,(
  ( ! [X6,X4,X8,X7,X5] : (s(X8,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X7),s(X8,X6),s(X8,X5),s(X8,X4))) = s(X8,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X7))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(X8,X6),s(X8,X4))) | s(bool,'T_0') = s(bool,sK4(X4,X5,X6,X7,X8))) )),
  inference(forward_demodulation,[],[f7052,f199])).
fof(f7052,plain,(
  ( ! [X6,X4,X8,X7,X5] : (s(X8,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X7),s(X8,X6),s(X8,X5),s(X8,X4))) = s(X8,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X7))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(X8,X6),s(X8,'const.bool.COND_3'(s(bool,'F_0'),s(X8,X5),s(X8,X4))))) | s(bool,'T_0') = s(bool,sK4(X4,X5,X6,X7,X8))) )),
  inference(superposition,[],[f4870,f181])).
fof(f7050,plain,(
  ( ! [X241,X239,X237,X235,X240,X238,X236,X234] : (s(X234,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X239),s(X234,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X235),s(X234,X236),s(X234,X238),s(X234,X237))),s(X234,X240),s(X234,X241))) = s(X234,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X239))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(X234,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X235),s(X234,X236),s(X234,X238),s(X234,X237))),s(X234,'const.bool.COND_3'(s(bool,sK4(X241,X240,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X235))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(X234,X236),s(X234,'const.bool.COND_3'(s(bool,sK4(X237,X238,X236,X235,X234)),s(X234,X238),s(X234,X237)))),X239,X234)),s(X234,X240),s(X234,X241)))))) )),
  inference(superposition,[],[f4870,f4870])).
fof(f7049,plain,(
  ( ! [X231,X229,X233,X230,X228,X232] : (s(X228,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X231),s(X228,X230),s(X228,X232),s(X228,X233))) = s(X228,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X231))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(X228,X230),s(X228,'const.bool.COND_3'(s(bool,sK4(X233,X232,'const.bool.COND_3'(s(bool,'F_0'),s(X228,X229),s(X228,X230)),X231,X228)),s(X228,X232),s(X228,X233)))))) )),
  inference(superposition,[],[f4870,f199])).
fof(f7048,plain,(
  ( ! [X222,X227,X225,X223,X226,X224] : (s(X222,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X225),s(X222,X223),s(X222,X226),s(X222,X227))) = s(X222,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X225))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(X222,X223),s(X222,'const.bool.COND_3'(s(bool,sK4(X227,X226,'const.bool.COND_3'(s(bool,'T_0'),s(X222,X223),s(X222,X224)),X225,X222)),s(X222,X226),s(X222,X227)))))) )),
  inference(superposition,[],[f4870,f198])).
fof(f7047,plain,(
  ( ! [X218,X216,X220,X215,X219,X217,X221] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X219),s(bool,'T_0'),s(bool,X220),s(bool,X221))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X219))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,sK4(X221,X220,sK4(X215,X216,X217,'const.prelim.EQUAL_0',X218),X219,bool)),s(bool,X220),s(bool,X221)))))) )),
  inference(superposition,[],[f4870,f2494])).
fof(f7046,plain,(
  ( ! [X212,X213,X211,X214] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X212),s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X211))),s(bool,X213),s(bool,X214))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X212))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X211))),s(bool,'const.bool.COND_3'(s(bool,sK4(X214,X213,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X211)))),X212,bool)),s(bool,X213),s(bool,X214)))))) )),
  inference(superposition,[],[f4870,f171])).
fof(f7045,plain,(
  ( ! [X210,X208,X209] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X208),s(bool,'T_0'),s(bool,X209),s(bool,X210))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X208))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,sK4(X210,X209,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X208,bool)),s(bool,X209),s(bool,X210)))))) )),
  inference(superposition,[],[f4870,f383])).
fof(f7040,plain,(
  ( ! [X194,X192,X195,X193] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X193),s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X192))),s(bool,X194),s(bool,X195))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X193))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X192))),s(bool,'const.bool.COND_3'(s(bool,sK4(X195,X194,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X192)))),X193,bool)),s(bool,X194),s(bool,X195)))))) )),
  inference(superposition,[],[f4870,f170])).
fof(f7038,plain,(
  ( ! [X187,X185,X188,X186,X184] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X186),s(bool,'const.arithmetic.<=_2'(s('type.num.num',X185),s('type.num.num',X184))),s(bool,X187),s(bool,X188))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X186))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,'const.arithmetic.<=_2'(s('type.num.num',X185),s('type.num.num',X184))),s(bool,'const.bool.COND_3'(s(bool,sK4(X188,X187,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X184))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X185)))),X186,bool)),s(bool,X187),s(bool,X188)))))) )),
  inference(superposition,[],[f4870,f203])).
fof(f7037,plain,(
  ( ! [X182,X180,X183,X181] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X181),s(bool,'T_0'),s(bool,X182),s(bool,X183))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X181))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,sK4(X183,X182,'const.arithmetic.>=_2'(s('type.num.num',X180),s('type.num.num','const.num.|48|_0')),X181,bool)),s(bool,X182),s(bool,X183)))))) )),
  inference(superposition,[],[f4870,f161])).
fof(f7036,plain,(
  ( ! [X177,X178,X176,X179] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X177),s(bool,'const.arithmetic.<=_2'(s('type.num.num',X176),s('type.num.num','const.num.|48|_0'))),s(bool,X178),s(bool,X179))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X177))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,'const.arithmetic.<=_2'(s('type.num.num',X176),s('type.num.num','const.num.|48|_0'))),s(bool,'const.bool.COND_3'(s(bool,sK4(X179,X178,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X176))),s('type.num.num','const.num.|48|_0')),X177,bool)),s(bool,X178),s(bool,X179)))))) )),
  inference(superposition,[],[f4870,f4721])).
fof(f7035,plain,(
  ( ! [X175,X173,X171,X174,X172] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X173),s(bool,'const.arithmetic.<=_2'(s('type.num.num',X171),s('type.num.num',X172))),s(bool,X174),s(bool,X175))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X173))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,'const.arithmetic.<=_2'(s('type.num.num',X171),s('type.num.num',X172))),s(bool,'const.bool.COND_3'(s(bool,sK4(X175,X174,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X171))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X172)))),X173,bool)),s(bool,X174),s(bool,X175)))))) )),
  inference(superposition,[],[f4870,f204])).
fof(f7034,plain,(
  ( ! [X167,X169,X170,X168] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X168),s(bool,'T_0'),s(bool,X169),s(bool,X170))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X168))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,sK4(X170,X169,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X167)),X168,bool)),s(bool,X169),s(bool,X170)))))) )),
  inference(superposition,[],[f4870,f164])).
fof(f7033,plain,(
  ( ! [X165,X166,X164] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X164),s(bool,'T_0'),s(bool,X165),s(bool,X166))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X164))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,sK4(X166,X165,'const.arithmetic.>_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))),s('type.num.num','const.num.|48|_0')),X164,bool)),s(bool,X165),s(bool,X166)))))) )),
  inference(superposition,[],[f4870,f2358])).
fof(f7032,plain,(
  ( ! [X163,X161,X162] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X161),s(bool,'T_0'),s(bool,X162),s(bool,X163))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X161))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,sK4(X163,X162,'const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')),X161,bool)),s(bool,X162),s(bool,X163)))))) )),
  inference(superposition,[],[f4870,f2357])).
fof(f7031,plain,(
  ( ! [X158,X159,X157,X160] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X158),s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X157))),s(bool,X159),s(bool,X160))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X158))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X157))),s(bool,'const.bool.COND_3'(s(bool,sK4(X160,X159,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X157))),s('type.num.num','const.num.|48|_0')),X158,bool)),s(bool,X159),s(bool,X160)))))) )),
  inference(superposition,[],[f4870,f4719])).
fof(f7030,plain,(
  ( ! [X156,X154,X152,X155,X153] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X154),s(bool,'const.prim_rec.<_2'(s('type.num.num',X153),s('type.num.num',X152))),s(bool,X155),s(bool,X156))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X154))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,'const.prim_rec.<_2'(s('type.num.num',X153),s('type.num.num',X152))),s(bool,'const.bool.COND_3'(s(bool,sK4(X156,X155,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X152))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X153)))),X154,bool)),s(bool,X155),s(bool,X156)))))) )),
  inference(superposition,[],[f4870,f208])).
fof(f7029,plain,(
  ( ! [X151,X149,X150,X148] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X149),s(bool,'F_0'),s(bool,X150),s(bool,X151))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X149))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,'F_0'),s(bool,'const.bool.COND_3'(s(bool,sK4(X151,X150,'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X148)),X149,bool)),s(bool,X150),s(bool,X151)))))) )),
  inference(superposition,[],[f4870,f168])).
fof(f7028,plain,(
  ( ! [X146,X144,X147,X145,X143] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X145),s(bool,'const.prim_rec.<_2'(s('type.num.num',X143),s('type.num.num',X144))),s(bool,X146),s(bool,X147))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X145))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,'const.prim_rec.<_2'(s('type.num.num',X143),s('type.num.num',X144))),s(bool,'const.bool.COND_3'(s(bool,sK4(X147,X146,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X143))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X144)))),X145,bool)),s(bool,X146),s(bool,X147)))))) )),
  inference(superposition,[],[f4870,f206])).
fof(f7027,plain,(
  ( ! [X142,X140,X138,X141,X139] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X140),s(bool,'const.prim_rec.<_2'(s('type.num.num',X138),s('type.num.num',X139))),s(bool,X141),s(bool,X142))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X140))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,'const.prim_rec.<_2'(s('type.num.num',X138),s('type.num.num',X139))),s(bool,'const.bool.COND_3'(s(bool,sK4(X142,X141,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X138))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X139)))),X140,bool)),s(bool,X141),s(bool,X142)))))) )),
  inference(superposition,[],[f4870,f205])).
fof(f7026,plain,(
  ( ! [X134,X136,X135,X133,X137] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X135),s(bool,'const.prim_rec.<_2'(s('type.num.num',X133),s('type.num.num',X134))),s(bool,X136),s(bool,X137))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X135))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,'const.prim_rec.<_2'(s('type.num.num',X133),s('type.num.num',X134))),s(bool,'const.bool.COND_3'(s(bool,sK4(X137,X136,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X133))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X134)))),X135,bool)),s(bool,X136),s(bool,X137)))))) )),
  inference(superposition,[],[f4870,f207])).
fof(f7025,plain,(
  ( ! [X132,X130,X128,X131,X129] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X130),s(bool,'const.prim_rec.<_2'(s('type.num.num',X128),s('type.num.num',X129))),s(bool,X131),s(bool,X132))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X130))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,'const.prim_rec.<_2'(s('type.num.num',X128),s('type.num.num',X129))),s(bool,'const.bool.COND_3'(s(bool,sK4(X132,X131,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X128))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X129)))),X130,bool)),s(bool,X131),s(bool,X132)))))) )),
  inference(superposition,[],[f4870,f209])).
fof(f7024,plain,(
  ( ! [X127,X125,X126] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X125),s(bool,'T_0'),s(bool,X126),s(bool,X127))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X125))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,sK4(X127,X126,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))),X125,bool)),s(bool,X126),s(bool,X127)))))) )),
  inference(superposition,[],[f4870,f2422])).
fof(f7023,plain,(
  ( ! [X123,X124,X122] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X122),s(bool,'T_0'),s(bool,X123),s(bool,X124))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X122))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,sK4(X124,X123,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X122,bool)),s(bool,X123),s(bool,X124)))))) )),
  inference(superposition,[],[f4870,f2421])).
fof(f7022,plain,(
  ( ! [X121,X118,X120,X119] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X119),s(bool,'T_0'),s(bool,X120),s(bool,X121))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X119))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,sK4(X121,X120,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X118)))),X119,bool)),s(bool,X120),s(bool,X121)))))) )),
  inference(superposition,[],[f4870,f4717])).
fof(f7021,plain,(
  ( ! [X116,X114,X117,X115] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X115),s(bool,'T_0'),s(bool,X116),s(bool,X117))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X115))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,sK4(X117,X116,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X114)))),X115,bool)),s(bool,X116),s(bool,X117)))))) )),
  inference(superposition,[],[f4870,f4718])).
fof(f7020,plain,(
  ( ! [X111,X112,X110,X113] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X111),s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X110))),s(bool,X112),s(bool,X113))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X111))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X110))),s(bool,'const.bool.COND_3'(s(bool,sK4(X113,X112,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X110)))),X111,bool)),s(bool,X112),s(bool,X113)))))) )),
  inference(superposition,[],[f4870,f4720])).
fof(f7019,plain,(
  ( ! [X109,X107,X108,X106] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X107),s(bool,'F_0'),s(bool,X108),s(bool,X109))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X107))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,'F_0'),s(bool,'const.bool.COND_3'(s(bool,sK4(X109,X108,'const.prim_rec.<_2'(s('type.num.num',X106),s('type.num.num','const.num.|48|_0')),X107,bool)),s(bool,X108),s(bool,X109)))))) )),
  inference(superposition,[],[f4870,f165])).
fof(f7018,plain,(
  ( ! [X103,X105,X102,X104] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X103),s(bool,X102),s(bool,X104),s(bool,X105))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X103))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,X102),s(bool,'const.bool.COND_3'(s(bool,sK4(X105,X104,'F_0',X103,bool)),s(bool,X104),s(bool,X105))))) | s(bool,'T_0') = s(bool,X102)) )),
  inference(superposition,[],[f4870,f181])).
fof(f7017,plain,(
  ( ! [X101,X99,X100,X98] : (s(bool,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X99),s(bool,'F_0'),s(bool,X100),s(bool,X101))) = s(bool,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X99))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(bool,'F_0'),s(bool,'const.bool.COND_3'(s(bool,sK4(X101,X100,X98,X99,bool)),s(bool,X100),s(bool,X101))))) | s(bool,'T_0') = s(bool,X98)) )),
  inference(superposition,[],[f4870,f181])).
fof(f7016,plain,(
  ( ! [X97,X95,X96] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X95),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X96),s('type.num.num',X97))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X95))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X97,X96,'const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.EQUAL_0')),X95,'type.num.num')),s('type.num.num',X96),s('type.num.num',X97)))))) )),
  inference(superposition,[],[f4870,f228])).
fof(f7015,plain,(
  ( ! [X94,X92,X93] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X92),s('type.num.num','const.num.|48|_0'),s('type.num.num',X93),s('type.num.num',X94))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X92))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X94,X93,'const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.LESS_0')),X92,'type.num.num')),s('type.num.num',X93),s('type.num.num',X94)))))) )),
  inference(superposition,[],[f4870,f152])).
fof(f7014,plain,(
  ( ! [X90,X88,X91,X89] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X89),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X88))))),s('type.num.num',X90),s('type.num.num',X91))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X89))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X88))))),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X91,X90,'const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X88)))),X89,'type.num.num')),s('type.num.num',X90),s('type.num.num',X91)))))) )),
  inference(superposition,[],[f4870,f175])).
fof(f7013,plain,(
  ( ! [X87,X85,X86] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X85),s('type.num.num','const.num.|48|_0'),s('type.num.num',X86),s('type.num.num',X87))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X85))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X87,X86,'const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')),X85,'type.num.num')),s('type.num.num',X86),s('type.num.num',X87)))))) )),
  inference(superposition,[],[f4870,f151])).
fof(f7012,plain,(
  ( ! [X83,X81,X84,X82] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X82),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X81))))),s('type.num.num',X83),s('type.num.num',X84))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X82))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X81))))),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X84,X83,'const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X81)))),X82,'type.num.num')),s('type.num.num',X83),s('type.num.num',X84)))))) )),
  inference(superposition,[],[f4870,f173])).
fof(f7011,plain,(
  ( ! [X80,X78,X79] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X78),s('type.num.num','const.num.|48|_0'),s('type.num.num',X79),s('type.num.num',X80))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X78))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X80,X79,'const.arithmetic.ZERO_0',X78,'type.num.num')),s('type.num.num',X79),s('type.num.num',X80)))))) )),
  inference(superposition,[],[f4870,f4713])).
fof(f7010,plain,(
  ( ! [X76,X74,X77,X75,X73] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X75),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X73),s('type.num.num',X74))))),s('type.num.num',X76),s('type.num.num',X77))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X75))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X73),s('type.num.num',X74))))),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X77,X76,'const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X73))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X74)))),X75,'type.num.num')),s('type.num.num',X76),s('type.num.num',X77)))))) )),
  inference(superposition,[],[f4870,f210])).
fof(f7009,plain,(
  ( ! [X70,X72,X71] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X70),s('type.num.num','const.num.|48|_0'),s('type.num.num',X71),s('type.num.num',X72))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X70))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X72,X71,'const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X70,'type.num.num')),s('type.num.num',X71),s('type.num.num',X72)))))) )),
  inference(superposition,[],[f4870,f1163])).
fof(f7008,plain,(
  ( ! [X68,X69,X67] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X67),s('type.num.num','const.num.|48|_0'),s('type.num.num',X68),s('type.num.num',X69))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X67))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X69,X68,'const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))),X67,'type.num.num')),s('type.num.num',X68),s('type.num.num',X69)))))) )),
  inference(superposition,[],[f4870,f1093])).
fof(f7007,plain,(
  ( ! [X66,X64,X65,X63] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X64),s('type.num.num','const.num.|48|_0'),s('type.num.num',X65),s('type.num.num',X66))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X64))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X66,X65,'const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X63)))))),X64,'type.num.num')),s('type.num.num',X65),s('type.num.num',X66)))))) )),
  inference(superposition,[],[f4870,f177])).
fof(f7006,plain,(
  ( ! [X61,X59,X62,X60] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X60),s('type.num.num','const.num.|48|_0'),s('type.num.num',X61),s('type.num.num',X62))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X60))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X62,X61,'const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X59)))))),X60,'type.num.num')),s('type.num.num',X61),s('type.num.num',X62)))))) )),
  inference(superposition,[],[f4870,f176])).
fof(f7005,plain,(
  ( ! [X57,X58,X56,X55] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X56),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X57),s('type.num.num',X58))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X56))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X58,X57,'const.arithmetic.EXP_2'(s('type.num.num',X55),s('type.num.num','const.num.|48|_0')),X56,'type.num.num')),s('type.num.num',X57),s('type.num.num',X58)))))) )),
  inference(superposition,[],[f4870,f229])).
fof(f7004,plain,(
  ( ! [X54,X52,X50,X53,X51] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X52),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X50),s('type.num.num',X51))))),s('type.num.num',X53),s('type.num.num',X54))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X52))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X50),s('type.num.num',X51))))),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X54,X53,'const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X50))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X51)))),X52,'type.num.num')),s('type.num.num',X53),s('type.num.num',X54)))))) )),
  inference(superposition,[],[f4870,f211])).
fof(f7003,plain,(
  ( ! [X47,X48,X46,X49] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X47),s('type.num.num','const.num.|48|_0'),s('type.num.num',X48),s('type.num.num',X49))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X47))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X49,X48,'const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X46)),X47,'type.num.num')),s('type.num.num',X48),s('type.num.num',X49)))))) )),
  inference(superposition,[],[f4870,f162])).
fof(f7002,plain,(
  ( ! [X45,X43,X44,X42] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X43),s('type.num.num',X42),s('type.num.num',X44),s('type.num.num',X45))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X43))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X42),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X45,X44,'const.arithmetic.-_2'(s('type.num.num',X42),s('type.num.num','const.num.|48|_0')),X43,'type.num.num')),s('type.num.num',X44),s('type.num.num',X45)))))) )),
  inference(superposition,[],[f4870,f163])).
fof(f7001,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X39),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X37),s('type.num.num',X38))))),s('type.num.num',X40),s('type.num.num',X41))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X39))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X37),s('type.num.num',X38))))),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X41,X40,'const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X37))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X38)))),X39,'type.num.num')),s('type.num.num',X40),s('type.num.num',X41)))))) )),
  inference(superposition,[],[f4870,f212])).
fof(f7000,plain,(
  ( ! [X35,X33,X36,X34] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X34),s('type.num.num','const.num.|48|_0'),s('type.num.num',X35),s('type.num.num',X36))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X34))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X36,X35,'const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X33)),X34,'type.num.num')),s('type.num.num',X35),s('type.num.num',X36)))))) )),
  inference(superposition,[],[f4870,f167])).
fof(f6999,plain,(
  ( ! [X30,X31,X29,X32] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X30),s('type.num.num','const.num.|48|_0'),s('type.num.num',X31),s('type.num.num',X32))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X30))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X32,X31,'const.arithmetic.*_2'(s('type.num.num',X29),s('type.num.num','const.num.|48|_0')),X30,'type.num.num')),s('type.num.num',X31),s('type.num.num',X32)))))) )),
  inference(superposition,[],[f4870,f166])).
fof(f6998,plain,(
  ( ! [X28,X26,X27,X25] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X26),s('type.num.num',X25),s('type.num.num',X27),s('type.num.num',X28))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X26))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X25),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X28,X27,'const.numeral.iZ_1'(s('type.num.num',X25)),X26,'type.num.num')),s('type.num.num',X27),s('type.num.num',X28)))))) )),
  inference(superposition,[],[f4870,f5694])).
fof(f6997,plain,(
  ( ! [X24,X23,X22] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X22),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X23),s('type.num.num',X24))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X22))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X24,X23,'const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X22,'type.num.num')),s('type.num.num',X23),s('type.num.num',X24)))))) )),
  inference(superposition,[],[f4870,f4737])).
fof(f6996,plain,(
  ( ! [X21,X19,X20] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X19),s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))),s('type.num.num',X20),s('type.num.num',X21))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X19))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X21,X20,'const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))),X19,'type.num.num')),s('type.num.num',X20),s('type.num.num',X21)))))) )),
  inference(superposition,[],[f4870,f4714])).
fof(f6995,plain,(
  ( ! [X17,X18,X16] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X16),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X17),s('type.num.num',X18))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X16))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X18,X17,'const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),X16,'type.num.num')),s('type.num.num',X17),s('type.num.num',X18)))))) )),
  inference(superposition,[],[f4870,f4715])).
fof(f6994,plain,(
  ( ! [X14,X15,X13] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X13),s('type.num.num','const.num.|48|_0'),s('type.num.num',X14),s('type.num.num',X15))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X13))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X15,X14,'const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')),X13,'type.num.num')),s('type.num.num',X14),s('type.num.num',X15)))))) )),
  inference(superposition,[],[f4870,f4578])).
fof(f6993,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X10),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X9))))),s('type.num.num',X11),s('type.num.num',X12))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X10))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X9))))),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X12,X11,'const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X9)))),X10,'type.num.num')),s('type.num.num',X11),s('type.num.num',X12)))))) )),
  inference(superposition,[],[f4870,f4605])).
fof(f6992,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X5),s('type.num.num',X4),s('type.num.num',X6),s('type.num.num',X7))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X5))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X4),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X7,X6,'const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')),X5,'type.num.num')),s('type.num.num',X6),s('type.num.num',X7)))))) )),
  inference(superposition,[],[f4870,f160])).
fof(f6991,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X1),s('type.num.num',X0),s('type.num.num',X2),s('type.num.num',X3))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X0),s('type.num.num','const.bool.COND_3'(s(bool,sK4(X3,X2,'const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)),X1,'type.num.num')),s('type.num.num',X2),s('type.num.num',X3)))))) )),
  inference(superposition,[],[f4870,f158])).
fof(f6990,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X1,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X0),s(X1,X2),s(X1,X3),s(X1,X4))) = s(X1,'const.bool.COND_3'(s(bool,'F_0'),s(X1,X2),s(X1,'const.bool.COND_3'(s(bool,sK4(X4,X3,X2,X0,X1)),s(X1,X3),s(X1,X4))))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f4870,f181])).
fof(f7277,plain,(
  ( ! [X6,X4,X7,X5] : (s(X4,X6) = s(X4,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.EQUAL_0'),s(X4,X5),s(X4,X6),s(X4,X7)))) )),
  inference(forward_demodulation,[],[f7276,f199])).
fof(f7276,plain,(
  ( ! [X6,X4,X7,X5] : (s(X4,'const.bool.COND_3'(s(bool,'F_0'),s(X4,X5),s(X4,X6))) = s(X4,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.EQUAL_0'),s(X4,X5),s(X4,X6),s(X4,X7)))) )),
  inference(forward_demodulation,[],[f7275,f165])).
fof(f7275,plain,(
  ( ! [X6,X4,X7,X5] : (s(X4,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.EQUAL_0'),s(X4,X5),s(X4,X6),s(X4,X7))) = s(X4,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))),s(X4,X5),s(X4,X6)))) )),
  inference(forward_demodulation,[],[f7274,f207])).
fof(f7274,plain,(
  ( ! [X6,X4,X7,X5] : (s(X4,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.EQUAL_0'),s(X4,X5),s(X4,X6),s(X4,X7))) = s(X4,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(X4,X5),s(X4,X6)))) )),
  inference(forward_demodulation,[],[f7273,f4800])).
fof(f7273,plain,(
  ( ! [X6,X4,X7,X5] : (s(X4,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.EQUAL_0'),s(X4,X5),s(X4,X6),s(X4,X7))) = s(X4,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(X4,X5),s(X4,X6)))) )),
  inference(forward_demodulation,[],[f7272,f198])).
fof(f7272,plain,(
  ( ! [X6,X4,X7,X5] : (s(X4,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.EQUAL_0'),s(X4,X5),s(X4,X6),s(X4,X7))) = s(X4,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(X4,X5),s(X4,'const.bool.COND_3'(s(bool,'T_0'),s(X4,X6),s(X4,X7)))))) )),
  inference(forward_demodulation,[],[f6989,f2494])).
fof(f6989,plain,(
  ( ! [X6,X4,X7,X5] : (s(X4,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.EQUAL_0'),s(X4,X5),s(X4,X6),s(X4,X7))) = s(X4,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(X4,X5),s(X4,'const.bool.COND_3'(s(bool,sK4(X7,X6,X5,'const.prelim.EQUAL_0',X4)),s(X4,X6),s(X4,X7)))))) )),
  inference(superposition,[],[f4870,f228])).
fof(f7271,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X1) = s(X0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.LESS_0'),s(X0,X1),s(X0,X2),s(X0,X3)))) )),
  inference(forward_demodulation,[],[f7270,f198])).
fof(f7270,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.LESS_0'),s(X0,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,'const.bool.COND_3'(s(bool,sK4(X3,X2,X1,'const.prelim.LESS_0',X0)),s(X0,X2),s(X0,X3)))))) )),
  inference(forward_demodulation,[],[f6988,f4718])).
fof(f6988,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.LESS_0'),s(X0,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(X0,X1),s(X0,'const.bool.COND_3'(s(bool,sK4(X3,X2,X1,'const.prelim.LESS_0',X0)),s(X0,X2),s(X0,X3)))))) )),
  inference(superposition,[],[f4870,f152])).
fof(f6987,plain,(
  ( ! [X14,X12,X19,X17,X15,X13,X18,X16] : (s(X16,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X12),s('type.prelim.ordering',X13),s('type.prelim.ordering',X15),s('type.prelim.ordering',X14))),s(X16,X17),s(X16,X18),s(X16,X19))) = s(X16,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X12),s('type.prelim.ordering',X13),s('type.prelim.ordering',X15),s('type.prelim.ordering',X14))))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(X16,X17),s(X16,'const.bool.COND_3'(s(bool,sK4(X19,X18,X17,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X12))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.prelim.ordering',X13),s('type.prelim.ordering','const.bool.COND_3'(s(bool,sK4(X14,X15,X13,X12,'type.prelim.ordering')),s('type.prelim.ordering',X15),s('type.prelim.ordering',X14)))),X16)),s(X16,X18),s(X16,X19)))))) )),
  inference(superposition,[],[f4870,f4870])).
fof(f6986,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (s(X8,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X7),s(X8,X9),s(X8,X10),s(X8,X11))) = s(X8,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X7))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(X8,X9),s(X8,'const.bool.COND_3'(s(bool,sK4(X11,X10,X9,'const.bool.COND_3'(s(bool,'F_0'),s('type.prelim.ordering',X6),s('type.prelim.ordering',X7)),X8)),s(X8,X10),s(X8,X11)))))) )),
  inference(superposition,[],[f4870,f199])).
fof(f6985,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X2,'const.prelim.ordering_CASE_4'(s('type.prelim.ordering',X0),s(X2,X3),s(X2,X4),s(X2,X5))) = s(X2,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s(X2,X3),s(X2,'const.bool.COND_3'(s(bool,sK4(X5,X4,X3,'const.bool.COND_3'(s(bool,'T_0'),s('type.prelim.ordering',X0),s('type.prelim.ordering',X1)),X2)),s(X2,X4),s(X2,X5)))))) )),
  inference(superposition,[],[f4870,f198])).
fof(f6947,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))) )),
  inference(unit_resulting_resolution,[],[f6943,f222])).
fof(f6948,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))) )),
  inference(unit_resulting_resolution,[],[f6943,f220])).
fof(f6949,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))) )),
  inference(unit_resulting_resolution,[],[f6943,f218])).
fof(f6885,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))) )),
  inference(unit_resulting_resolution,[],[f1943,f222])).
fof(f6886,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))) )),
  inference(unit_resulting_resolution,[],[f1943,f220])).
fof(f6887,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))) )),
  inference(unit_resulting_resolution,[],[f1943,f218])).
fof(f6879,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK5(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK5(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f226,f181])).
fof(f6878,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK5(X0,X1,X2,bool))) )),
  inference(superposition,[],[f226,f181])).
fof(f6820,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))) )),
  inference(unit_resulting_resolution,[],[f1586,f222])).
fof(f6821,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))) )),
  inference(unit_resulting_resolution,[],[f1586,f220])).
fof(f6822,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))) )),
  inference(unit_resulting_resolution,[],[f1586,f218])).
fof(f6757,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))) )),
  inference(unit_resulting_resolution,[],[f1225,f222])).
fof(f6758,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))) )),
  inference(unit_resulting_resolution,[],[f1225,f220])).
fof(f6759,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))) )),
  inference(unit_resulting_resolution,[],[f1225,f218])).
fof(f6643,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f4605,f4737])).
fof(f6642,plain,(
  ( ! [X2] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))),s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f4605,f4714])).
fof(f6641,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f4605,f4715])).
fof(f6611,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X3)))))) )),
  inference(superposition,[],[f4605,f4737])).
fof(f6610,plain,(
  ( ! [X2] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X2)))))) )),
  inference(superposition,[],[f4605,f4714])).
fof(f6609,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X1)))))) )),
  inference(superposition,[],[f4605,f4715])).
fof(f6569,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))))),
  inference(forward_demodulation,[],[f6566,f4578])).
fof(f6566,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))))),
  inference(unit_resulting_resolution,[],[f4602,f222])).
fof(f6567,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))))),
  inference(unit_resulting_resolution,[],[f4602,f220])).
fof(f6568,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))))),
  inference(unit_resulting_resolution,[],[f4602,f218])).
fof(f6378,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))) )),
  inference(forward_demodulation,[],[f6348,f4578])).
fof(f6348,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))) )),
  inference(unit_resulting_resolution,[],[f4580,f222])).
fof(f6349,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))) )),
  inference(unit_resulting_resolution,[],[f4580,f220])).
fof(f6350,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))) )),
  inference(unit_resulting_resolution,[],[f4580,f218])).
fof(f6338,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))))) )),
  inference(forward_demodulation,[],[f6308,f4578])).
fof(f6308,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))))) )),
  inference(unit_resulting_resolution,[],[f4579,f222])).
fof(f6309,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))))) )),
  inference(unit_resulting_resolution,[],[f4579,f220])).
fof(f6310,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))))) )),
  inference(unit_resulting_resolution,[],[f4579,f218])).
fof(f5841,plain,(
  ( ! [X54] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X54))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X54),s('type.num.num','const.num.|48|_0'))))) )),
  inference(forward_demodulation,[],[f5840,f5694])).
fof(f5840,plain,(
  ( ! [X54] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X54),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num',X54)))))) )),
  inference(forward_demodulation,[],[f5762,f4782])).
fof(f5762,plain,(
  ( ! [X54] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X54)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num',X54)))))) )),
  inference(superposition,[],[f195,f4604])).
fof(f5839,plain,(
  ( ! [X53] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X53),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X53)))) )),
  inference(forward_demodulation,[],[f5838,f5694])).
fof(f5838,plain,(
  ( ! [X53] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X53),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num',X53)))))) )),
  inference(forward_demodulation,[],[f5761,f4782])).
fof(f5761,plain,(
  ( ! [X53] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X53)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num',X53)))))) )),
  inference(superposition,[],[f194,f4604])).
fof(f5744,plain,(
  ( ! [X36] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X36))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numeral.iZ_1'(s('type.num.num',X36)))) )),
  inference(superposition,[],[f4722,f4604])).
fof(f5742,plain,(
  ( ! [X33,X32] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X32))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X33))) | s('type.num.num',X33) = s('type.num.num','const.numeral.iZ_1'(s('type.num.num',X32)))) )),
  inference(superposition,[],[f222,f4604])).
fof(f5819,plain,(
  ( ! [X14,X15] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X14),s('type.num.num',X15))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X14))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X15)))))) )),
  inference(backward_demodulation,[],[f5694,f4853])).
fof(f5818,plain,(
  ( ! [X59,X60,X58] : (s('type.num.num',X60) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X58),s('type.num.num',X59))) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X60))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X58))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X59)))))) )),
  inference(backward_demodulation,[],[f5694,f4320])).
fof(f5694,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.numeral.iZ_1'(s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f4604,f222])).
fof(f5066,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))) )),
  inference(unit_resulting_resolution,[],[f4725,f194])).
fof(f5035,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))) )),
  inference(unit_resulting_resolution,[],[f4724,f194])).
fof(f4856,plain,(
  ( ! [X6] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6))) | s('type.num.num',X6) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(backward_demodulation,[],[f4713,f4583])).
fof(f4853,plain,(
  ( ! [X14,X15] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X14),s('type.num.num',X15))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X14))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X15)))))) )),
  inference(backward_demodulation,[],[f4713,f4305])).
fof(f4848,plain,(
  ( ! [X50,X51] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X50),s('type.num.num',X51))),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X50),s('type.num.num',X51)))))) )),
  inference(backward_demodulation,[],[f4713,f3838])).
fof(f4847,plain,(
  ( ! [X48,X49] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X48),s('type.num.num',X49))),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X48),s('type.num.num',X49)))))) )),
  inference(backward_demodulation,[],[f4713,f3836])).
fof(f4846,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(backward_demodulation,[],[f4713,f3833])).
fof(f4844,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X3))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X3)))))) )),
  inference(backward_demodulation,[],[f4713,f3831])).
fof(f4843,plain,(
  ( ! [X45,X44] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X44))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X45))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X44),s('type.num.num',X45)))))))) )),
  inference(backward_demodulation,[],[f4713,f3773])).
fof(f4842,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))) )),
  inference(backward_demodulation,[],[f4713,f3748])).
fof(f4841,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(backward_demodulation,[],[f4713,f3747])).
fof(f4840,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(backward_demodulation,[],[f4713,f3718])).
fof(f4839,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(backward_demodulation,[],[f4713,f3717])).
fof(f4838,plain,(
  ( ! [X50,X51] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X50),s('type.num.num',X51))),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X50),s('type.num.num',X51)))))) )),
  inference(backward_demodulation,[],[f4713,f3656])).
fof(f4837,plain,(
  ( ! [X48,X49] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X48),s('type.num.num',X49))),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X48),s('type.num.num',X49)))))) )),
  inference(backward_demodulation,[],[f4713,f3654])).
fof(f4836,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(backward_demodulation,[],[f4713,f3651])).
fof(f4834,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X3))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X3)))))) )),
  inference(backward_demodulation,[],[f4713,f3649])).
fof(f4833,plain,(
  ( ! [X45,X44] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X44))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X45))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X44),s('type.num.num',X45)))))))) )),
  inference(backward_demodulation,[],[f4713,f3593])).
fof(f4832,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))) )),
  inference(backward_demodulation,[],[f4713,f3568])).
fof(f4831,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(backward_demodulation,[],[f4713,f3567])).
fof(f4830,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(backward_demodulation,[],[f4713,f3539])).
fof(f4829,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(backward_demodulation,[],[f4713,f3538])).
fof(f4826,plain,(
  ( ! [X50,X51] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X50),s('type.num.num',X51))),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X50),s('type.num.num',X51)))))) )),
  inference(backward_demodulation,[],[f4713,f3473])).
fof(f4825,plain,(
  ( ! [X48,X49] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X48),s('type.num.num',X49))),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X48),s('type.num.num',X49)))))) )),
  inference(backward_demodulation,[],[f4713,f3471])).
fof(f4824,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(backward_demodulation,[],[f4713,f3468])).
fof(f4822,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X3))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X3)))))) )),
  inference(backward_demodulation,[],[f4713,f3466])).
fof(f4821,plain,(
  ( ! [X45,X44] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X44))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X45))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X44),s('type.num.num',X45)))))))) )),
  inference(backward_demodulation,[],[f4713,f3413])).
fof(f4820,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))) )),
  inference(backward_demodulation,[],[f4713,f3388])).
fof(f4819,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(backward_demodulation,[],[f4713,f3387])).
fof(f4818,plain,(
  ( ! [X2] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X2))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))))) )),
  inference(backward_demodulation,[],[f4713,f3361])).
fof(f4817,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(backward_demodulation,[],[f4713,f3360])).
fof(f4816,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(backward_demodulation,[],[f4713,f3359])).
fof(f4808,plain,(
  ( ! [X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))) )),
  inference(backward_demodulation,[],[f4713,f3074])).
fof(f4807,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(backward_demodulation,[],[f4713,f3073])).
fof(f4805,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(backward_demodulation,[],[f4713,f3046])).
fof(f4800,plain,(
  ( ! [X3] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X3))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X3)))) )),
  inference(backward_demodulation,[],[f4713,f3001])).
fof(f4798,plain,(
  ( ! [X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X1))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))) )),
  inference(backward_demodulation,[],[f4713,f2989])).
fof(f4797,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(backward_demodulation,[],[f4713,f2988])).
fof(f4796,plain,(
  ( ! [X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(backward_demodulation,[],[f4713,f2962])).
fof(f4795,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(backward_demodulation,[],[f4713,f2961])).
fof(f4794,plain,(
  ( ! [X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))) )),
  inference(backward_demodulation,[],[f4713,f2683])).
fof(f4793,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(backward_demodulation,[],[f4713,f2682])).
fof(f4792,plain,(
  ( ! [X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(backward_demodulation,[],[f4713,f2656])).
fof(f4791,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(backward_demodulation,[],[f4713,f2655])).
fof(f4790,plain,(
  ( ! [X3] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X3))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X3)))) )),
  inference(backward_demodulation,[],[f4713,f2582])).
fof(f4788,plain,(
  ( ! [X3] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(backward_demodulation,[],[f4713,f2580])).
fof(f4786,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0'))))) )),
  inference(backward_demodulation,[],[f4713,f2578])).
fof(f4785,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0'))))) )),
  inference(backward_demodulation,[],[f4713,f2577])).
fof(f4784,plain,(
  ( ! [X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X1))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))) )),
  inference(backward_demodulation,[],[f4713,f2567])).
fof(f4783,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(backward_demodulation,[],[f4713,f2566])).
fof(f4782,plain,(
  ( ! [X2] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num','const.num.|48|_0'))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))))) )),
  inference(backward_demodulation,[],[f4713,f2541])).
fof(f4781,plain,(
  ( ! [X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(backward_demodulation,[],[f4713,f2540])).
fof(f4780,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(backward_demodulation,[],[f4713,f2539])).
fof(f4777,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.num.|48|_0')))) )),
  inference(backward_demodulation,[],[f4713,f2467])).
fof(f4776,plain,(
  s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))),s('type.num.num','const.num.|48|_0')))),
  inference(backward_demodulation,[],[f4713,f2464])).
fof(f4775,plain,(
  s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))),
  inference(backward_demodulation,[],[f4713,f2463])).
fof(f4774,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(backward_demodulation,[],[f4713,f2414])).
fof(f4769,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))) | s(bool,'T_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.num.|48|_0')))) )),
  inference(backward_demodulation,[],[f4713,f2350])).
fof(f4766,plain,(
  ( ! [X10] : (s('type.num.num',X10) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(backward_demodulation,[],[f4713,f2288])).
fof(f4765,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(backward_demodulation,[],[f4713,f2255])).
fof(f4762,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))) )),
  inference(backward_demodulation,[],[f4713,f2036])).
fof(f4761,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))) )),
  inference(backward_demodulation,[],[f4713,f2035])).
fof(f4757,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))) )),
  inference(backward_demodulation,[],[f4713,f1899])).
fof(f4756,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))) )),
  inference(backward_demodulation,[],[f4713,f1898])).
fof(f4755,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))) )),
  inference(backward_demodulation,[],[f4713,f1871])).
fof(f4754,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))) )),
  inference(backward_demodulation,[],[f4713,f1870])).
fof(f4753,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))) )),
  inference(backward_demodulation,[],[f4713,f1710])).
fof(f4752,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))) )),
  inference(backward_demodulation,[],[f4713,f1709])).
fof(f4749,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))) )),
  inference(backward_demodulation,[],[f4713,f1584])).
fof(f4748,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))) )),
  inference(backward_demodulation,[],[f4713,f1583])).
fof(f4744,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))) )),
  inference(backward_demodulation,[],[f4713,f1223])).
fof(f4743,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))) )),
  inference(backward_demodulation,[],[f4713,f1222])).
fof(f4865,plain,(
  ( ! [X15,X16] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X16))) | s('type.num.num',X16) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X15),s('type.num.num','const.num.|48|_0')))) )),
  inference(backward_demodulation,[],[f4737,f1959])).
fof(f4732,plain,(
  ( ! [X10,X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X9),s('type.num.num',X10)))) )),
  inference(backward_demodulation,[],[f4713,f697])).
fof(f4731,plain,(
  ( ! [X8,X7] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X7),s('type.num.num',X8)))) )),
  inference(backward_demodulation,[],[f4713,f696])).
fof(f4730,plain,(
  ( ! [X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X5))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X5),s('type.num.num','const.num.|48|_0')))) )),
  inference(backward_demodulation,[],[f4713,f691])).
fof(f4729,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')))) )),
  inference(backward_demodulation,[],[f4713,f685])).
fof(f4728,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(backward_demodulation,[],[f4713,f684])).
fof(f4727,plain,(
  s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(backward_demodulation,[],[f4713,f478])).
fof(f4726,plain,(
  s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))),
  inference(backward_demodulation,[],[f4713,f477])).
fof(f4725,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))) )),
  inference(backward_demodulation,[],[f4713,f231])).
fof(f4724,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))) )),
  inference(backward_demodulation,[],[f4713,f230])).
fof(f4723,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(backward_demodulation,[],[f4713,f191])).
fof(f4722,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(backward_demodulation,[],[f4713,f190])).
fof(f4721,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))) )),
  inference(backward_demodulation,[],[f4713,f180])).
fof(f4720,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(backward_demodulation,[],[f4713,f179])).
fof(f4719,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))) )),
  inference(backward_demodulation,[],[f4713,f178])).
fof(f4718,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))) )),
  inference(backward_demodulation,[],[f4713,f172])).
fof(f4717,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))) )),
  inference(backward_demodulation,[],[f4713,f169])).
fof(f4862,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))),
  inference(backward_demodulation,[],[f4715,f2787])).
fof(f4860,plain,(
  s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))),
  inference(backward_demodulation,[],[f4714,f2857])).
fof(f4484,plain,(
  ( ! [X37,X36] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X37))) | s('type.num.num',X37) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X36)))) )),
  inference(superposition,[],[f222,f3830])).
fof(f4647,plain,(
  ( ! [X57,X56,X55] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X57),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X55),s('type.num.num',X56))))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X57),s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X55),s('type.num.num',X56)))))))))) )),
  inference(forward_demodulation,[],[f4626,f212])).
fof(f4626,plain,(
  ( ! [X57,X56,X55] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X57),s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X55),s('type.num.num',X56))))))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X57))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X55),s('type.num.num',X56)))))))) )),
  inference(backward_demodulation,[],[f4605,f4319])).
fof(f4646,plain,(
  ( ! [X54,X52,X53] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X52),s('type.num.num',X53))),s('type.num.num',X54))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X52),s('type.num.num',X53))))),s('type.num.num',X54)))))) )),
  inference(forward_demodulation,[],[f4625,f212])).
fof(f4625,plain,(
  ( ! [X54,X52,X53] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X52),s('type.num.num',X53))))),s('type.num.num',X54))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X52),s('type.num.num',X53))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X54)))))) )),
  inference(backward_demodulation,[],[f4605,f4318])).
fof(f4645,plain,(
  ( ! [X50,X51,X49] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X51),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X49),s('type.num.num',X50))))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X51),s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X49),s('type.num.num',X50)))))))))) )),
  inference(forward_demodulation,[],[f4624,f211])).
fof(f4624,plain,(
  ( ! [X50,X51,X49] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X51),s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X49),s('type.num.num',X50))))))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X51))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X49),s('type.num.num',X50)))))))) )),
  inference(backward_demodulation,[],[f4605,f4317])).
fof(f4644,plain,(
  ( ! [X47,X48,X46] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X46),s('type.num.num',X47))),s('type.num.num',X48))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X46),s('type.num.num',X47))))),s('type.num.num',X48)))))) )),
  inference(forward_demodulation,[],[f4623,f211])).
fof(f4623,plain,(
  ( ! [X47,X48,X46] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X46),s('type.num.num',X47))))),s('type.num.num',X48))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X46),s('type.num.num',X47))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X48)))))) )),
  inference(backward_demodulation,[],[f4605,f4316])).
fof(f4643,plain,(
  ( ! [X45,X43,X44] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X45),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X43),s('type.num.num',X44))))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X45),s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X43),s('type.num.num',X44)))))))))) )),
  inference(forward_demodulation,[],[f4622,f210])).
fof(f4622,plain,(
  ( ! [X45,X43,X44] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X45),s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X43),s('type.num.num',X44))))))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X45))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X43),s('type.num.num',X44)))))))) )),
  inference(backward_demodulation,[],[f4605,f4315])).
fof(f4642,plain,(
  ( ! [X41,X42,X40] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))),s('type.num.num',X42))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))))),s('type.num.num',X42)))))) )),
  inference(forward_demodulation,[],[f4621,f210])).
fof(f4621,plain,(
  ( ! [X41,X42,X40] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))))),s('type.num.num',X42))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X42)))))) )),
  inference(backward_demodulation,[],[f4605,f4314])).
fof(f4639,plain,(
  ( ! [X33,X31,X32] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X31),s('type.num.num',X32))),s('type.num.num',X33))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X31),s('type.num.num',X32))))),s('type.num.num',X33)))) )),
  inference(forward_demodulation,[],[f4618,f208])).
fof(f4618,plain,(
  ( ! [X33,X31,X32] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X31),s('type.num.num',X32))))),s('type.num.num',X33))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X33))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X31),s('type.num.num',X32)))))))) )),
  inference(backward_demodulation,[],[f4605,f4311])).
fof(f4638,plain,(
  ( ! [X30,X28,X29] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X30),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X28),s('type.num.num',X29))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X30),s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X28),s('type.num.num',X29)))))))) )),
  inference(forward_demodulation,[],[f4617,f208])).
fof(f4617,plain,(
  ( ! [X30,X28,X29] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X30),s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X28),s('type.num.num',X29))))))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X28),s('type.num.num',X29))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X30)))))) )),
  inference(backward_demodulation,[],[f4605,f4310])).
fof(f4635,plain,(
  ( ! [X21,X19,X20] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',X20))),s('type.num.num',X21))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',X20))))),s('type.num.num',X21)))) )),
  inference(forward_demodulation,[],[f4614,f203])).
fof(f4614,plain,(
  ( ! [X21,X19,X20] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',X20))))),s('type.num.num',X21))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X21))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',X20)))))))) )),
  inference(backward_demodulation,[],[f4605,f4307])).
fof(f4634,plain,(
  ( ! [X17,X18,X16] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X18),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X16),s('type.num.num',X17))))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X18),s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X16),s('type.num.num',X17)))))))) )),
  inference(forward_demodulation,[],[f4613,f203])).
fof(f4613,plain,(
  ( ! [X17,X18,X16] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X18),s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X16),s('type.num.num',X17))))))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X16),s('type.num.num',X17))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X18)))))) )),
  inference(backward_demodulation,[],[f4605,f4306])).
fof(f4633,plain,(
  ( ! [X12,X13] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num',X13))),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num',X13))))),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(forward_demodulation,[],[f4612,f180])).
fof(f4612,plain,(
  ( ! [X12,X13] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num',X13))))),s('type.num.num','const.num.|48|_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num',X13))))),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(backward_demodulation,[],[f4605,f4304])).
fof(f4631,plain,(
  ( ! [X8,X9] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X9))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X9)))))))) )),
  inference(forward_demodulation,[],[f4610,f178])).
fof(f4610,plain,(
  ( ! [X8,X9] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X9))))))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X9))))),s('type.num.num','const.num.|48|_0')))) )),
  inference(backward_demodulation,[],[f4605,f4302])).
fof(f4630,plain,(
  ( ! [X6,X7] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7))))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7)))))))))) )),
  inference(forward_demodulation,[],[f4609,f175])).
fof(f4609,plain,(
  ( ! [X6,X7] : (s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7))))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7)))))))))) )),
  inference(backward_demodulation,[],[f4605,f4301])).
fof(f4629,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X5))))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X5)))))))))) )),
  inference(forward_demodulation,[],[f4608,f173])).
fof(f4608,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X5))))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X5)))))))))) )),
  inference(backward_demodulation,[],[f4605,f4300])).
fof(f4628,plain,(
  ( ! [X2,X3] : (s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))))) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3)))))))) )),
  inference(forward_demodulation,[],[f4607,f171])).
fof(f4607,plain,(
  ( ! [X2,X3] : (s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))))))) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3)))))))) )),
  inference(backward_demodulation,[],[f4605,f4299])).
fof(f4627,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))))))) )),
  inference(forward_demodulation,[],[f4606,f170])).
fof(f4606,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))))) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))))))) )),
  inference(backward_demodulation,[],[f4605,f4298])).
fof(f4605,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(backward_demodulation,[],[f4604,f213])).
fof(f4604,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num',X0)))))) )),
  inference(forward_demodulation,[],[f4600,f158])).
fof(f4600,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(backward_demodulation,[],[f4578,f4296])).
fof(f4596,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))),
  inference(backward_demodulation,[],[f4578,f2929])).
fof(f4595,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))),
  inference(backward_demodulation,[],[f4578,f2928])).
fof(f4602,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))),
  inference(forward_demodulation,[],[f4594,f4578])).
fof(f4594,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))),
  inference(backward_demodulation,[],[f4578,f2927])).
fof(f4592,plain,(
  ( ! [X29] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X29))) | s('type.num.num',X29) = s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.LESS_0')))) )),
  inference(backward_demodulation,[],[f4578,f1968])).
fof(f4591,plain,(
  ( ! [X25] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X25))) | s('type.num.num',X25) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(backward_demodulation,[],[f4578,f1965])).
fof(f4590,plain,(
  ( ! [X22] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X22))) | s('type.num.num',X22) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(backward_demodulation,[],[f4578,f1963])).
fof(f4589,plain,(
  ( ! [X21] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X21))) | s('type.num.num',X21) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))) )),
  inference(backward_demodulation,[],[f4578,f1962])).
fof(f4588,plain,(
  ( ! [X19,X20] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X20))) | s('type.num.num',X20) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X19)))))))) )),
  inference(backward_demodulation,[],[f4578,f1961])).
fof(f4587,plain,(
  ( ! [X17,X18] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X18))) | s('type.num.num',X18) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X17)))))))) )),
  inference(backward_demodulation,[],[f4578,f1960])).
fof(f4586,plain,(
  ( ! [X12,X11] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X12))) | s('type.num.num',X12) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X11)))) )),
  inference(backward_demodulation,[],[f4578,f1957])).
fof(f4585,plain,(
  ( ! [X10,X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))) | s('type.num.num',X10) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X9)))) )),
  inference(backward_demodulation,[],[f4578,f1956])).
fof(f4584,plain,(
  ( ! [X8,X7] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8))) | s('type.num.num',X8) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X7),s('type.num.num','const.num.|48|_0')))) )),
  inference(backward_demodulation,[],[f4578,f1955])).
fof(f4583,plain,(
  ( ! [X6] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6))) | s('type.num.num',X6) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(backward_demodulation,[],[f4578,f1954])).
fof(f4582,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))),
  inference(backward_demodulation,[],[f4578,f1941])).
fof(f4580,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))) )),
  inference(backward_demodulation,[],[f4578,f1937])).
fof(f4579,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))) )),
  inference(backward_demodulation,[],[f4578,f1936])).
fof(f4571,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0'))))) )),
  inference(forward_demodulation,[],[f4425,f2541])).
fof(f4425,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0))))))))) )),
  inference(unit_resulting_resolution,[],[f3830,f195])).
fof(f4320,plain,(
  ( ! [X59,X60,X58] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X60))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X58))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X59))))) | s('type.num.num',X60) = s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X58),s('type.num.num',X59)))))) )),
  inference(superposition,[],[f222,f213])).
fof(f4319,plain,(
  ( ! [X57,X56,X55] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X57),s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X55),s('type.num.num',X56))))))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X57))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X55))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X56)))))))) )),
  inference(superposition,[],[f212,f213])).
fof(f4318,plain,(
  ( ! [X54,X52,X53] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X52),s('type.num.num',X53))))),s('type.num.num',X54))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X52))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X53))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X54)))))) )),
  inference(superposition,[],[f212,f213])).
fof(f4317,plain,(
  ( ! [X50,X51,X49] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X51),s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X49),s('type.num.num',X50))))))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X51))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X49))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X50)))))))) )),
  inference(superposition,[],[f211,f213])).
fof(f4316,plain,(
  ( ! [X47,X48,X46] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X46),s('type.num.num',X47))))),s('type.num.num',X48))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X46))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X47))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X48)))))) )),
  inference(superposition,[],[f211,f213])).
fof(f4315,plain,(
  ( ! [X45,X43,X44] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X45),s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X43),s('type.num.num',X44))))))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X45))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X43))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X44)))))))) )),
  inference(superposition,[],[f210,f213])).
fof(f4314,plain,(
  ( ! [X41,X42,X40] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))))),s('type.num.num',X42))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X40))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X41))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X42)))))) )),
  inference(superposition,[],[f210,f213])).
fof(f4313,plain,(
  ( ! [X39,X37,X38] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X39),s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X37),s('type.num.num',X38))))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X39))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X37))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X38)))))))) )),
  inference(superposition,[],[f209,f213])).
fof(f4312,plain,(
  ( ! [X35,X36,X34] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X34),s('type.num.num',X35))))),s('type.num.num',X36))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X34))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X35))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X36)))))) )),
  inference(superposition,[],[f209,f213])).
fof(f4311,plain,(
  ( ! [X33,X31,X32] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X31),s('type.num.num',X32))))),s('type.num.num',X33))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X33))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X31))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X32)))))))) )),
  inference(superposition,[],[f208,f213])).
fof(f4310,plain,(
  ( ! [X30,X28,X29] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X30),s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X28),s('type.num.num',X29))))))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X28))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X29))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X30)))))) )),
  inference(superposition,[],[f208,f213])).
fof(f4309,plain,(
  ( ! [X26,X27,X25] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X27),s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X25),s('type.num.num',X26))))))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X27))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X25))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X26)))))))) )),
  inference(superposition,[],[f204,f213])).
fof(f4308,plain,(
  ( ! [X24,X23,X22] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X22),s('type.num.num',X23))))),s('type.num.num',X24))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X22))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X23))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X24)))))) )),
  inference(superposition,[],[f204,f213])).
fof(f4307,plain,(
  ( ! [X21,X19,X20] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',X20))))),s('type.num.num',X21))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X21))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X19))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X20)))))))) )),
  inference(superposition,[],[f203,f213])).
fof(f4306,plain,(
  ( ! [X17,X18,X16] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X18),s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X16),s('type.num.num',X17))))))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X16))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X17))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X18)))))) )),
  inference(superposition,[],[f203,f213])).
fof(f4305,plain,(
  ( ! [X14,X15] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X14))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X15))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X14),s('type.num.num',X15)))))) )),
  inference(superposition,[],[f190,f213])).
fof(f4304,plain,(
  ( ! [X12,X13] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num',X13))))),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X12))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13))))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f180,f213])).
fof(f4303,plain,(
  ( ! [X10,X11] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',X11))))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X11)))))))) )),
  inference(superposition,[],[f179,f213])).
fof(f4302,plain,(
  ( ! [X8,X9] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X9))))))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X9))))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f178,f213])).
fof(f4301,plain,(
  ( ! [X6,X7] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7))))))))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7)))))))) )),
  inference(superposition,[],[f175,f213])).
fof(f4300,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X5))))))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X4))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X5)))))))) )),
  inference(superposition,[],[f173,f213])).
fof(f4299,plain,(
  ( ! [X2,X3] : (s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))))))) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3)))))))) )),
  inference(superposition,[],[f171,f213])).
fof(f4298,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))))) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))) )),
  inference(superposition,[],[f170,f213])).
fof(f4297,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f213,f160])).
fof(f4296,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f213,f158])).
fof(f4186,plain,(
  ( ! [X74,X75] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X75))) | s('type.num.num',X75) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X74)))) )),
  inference(superposition,[],[f220,f3853])).
fof(f4182,plain,(
  ( ! [X66,X67] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X67))) | s('type.num.num',X67) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X66)))) )),
  inference(superposition,[],[f218,f3853])).
fof(f3911,plain,(
  ( ! [X37,X36] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X37))) | s('type.num.num',X37) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X36)))) )),
  inference(superposition,[],[f222,f3648])).
fof(f3853,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f3648,f190])).
fof(f3995,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0'))))) )),
  inference(forward_demodulation,[],[f3854,f2541])).
fof(f3854,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0))))))))) )),
  inference(unit_resulting_resolution,[],[f3648,f195])).
fof(f3803,plain,(
  ( ! [X132,X130,X131] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X132))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X130),s('type.num.num',X131))))))) | s('type.num.num',X132) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X130))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X131)))))) )),
  inference(superposition,[],[f222,f212])).
fof(f3838,plain,(
  ( ! [X50,X51] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X50),s('type.num.num',X51))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X50),s('type.num.num',X51))),s('type.num.num','const.arithmetic.ZERO_0'))))) )),
  inference(forward_demodulation,[],[f3837,f212])).
fof(f3837,plain,(
  ( ! [X50,X51] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X50),s('type.num.num',X51))),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X50))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X51)))))) )),
  inference(forward_demodulation,[],[f3776,f2541])).
fof(f3776,plain,(
  ( ! [X50,X51] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X50),s('type.num.num',X51)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X50))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X51)))))) )),
  inference(superposition,[],[f195,f212])).
fof(f3836,plain,(
  ( ! [X48,X49] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X48),s('type.num.num',X49))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X48),s('type.num.num',X49))),s('type.num.num','const.arithmetic.ZERO_0'))))) )),
  inference(forward_demodulation,[],[f3835,f212])).
fof(f3835,plain,(
  ( ! [X48,X49] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X48),s('type.num.num',X49))),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X48))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X49)))))) )),
  inference(forward_demodulation,[],[f3775,f2541])).
fof(f3775,plain,(
  ( ! [X48,X49] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X48),s('type.num.num',X49)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X48))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X49)))))) )),
  inference(superposition,[],[f194,f212])).
fof(f3773,plain,(
  ( ! [X45,X44] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X44),s('type.num.num',X45))))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X44))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X45)))))) )),
  inference(superposition,[],[f190,f212])).
fof(f3833,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(forward_demodulation,[],[f3750,f3747])).
fof(f3750,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f212,f944])).
fof(f3832,plain,(
  ( ! [X2] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.ZERO_0')))))) )),
  inference(forward_demodulation,[],[f3749,f166])).
fof(f3749,plain,(
  ( ! [X2] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f212,f733])).
fof(f3748,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))) )),
  inference(superposition,[],[f212,f153])).
fof(f3747,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f212,f155])).
fof(f3831,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X3))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X3)))))) )),
  inference(forward_demodulation,[],[f3720,f3717])).
fof(f3720,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X3))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3)))))) )),
  inference(superposition,[],[f212,f944])).
fof(f3718,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(superposition,[],[f212,f153])).
fof(f3717,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f212,f155])).
fof(f3621,plain,(
  ( ! [X125,X126,X124] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X126))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X124),s('type.num.num',X125))))))) | s('type.num.num',X126) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X124))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X125)))))) )),
  inference(superposition,[],[f222,f211])).
fof(f3656,plain,(
  ( ! [X50,X51] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X50),s('type.num.num',X51))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X50),s('type.num.num',X51))),s('type.num.num','const.arithmetic.ZERO_0'))))) )),
  inference(forward_demodulation,[],[f3655,f211])).
fof(f3655,plain,(
  ( ! [X50,X51] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X50),s('type.num.num',X51))),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X50))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X51)))))) )),
  inference(forward_demodulation,[],[f3596,f2541])).
fof(f3596,plain,(
  ( ! [X50,X51] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X50),s('type.num.num',X51)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X50))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X51)))))) )),
  inference(superposition,[],[f195,f211])).
fof(f3654,plain,(
  ( ! [X48,X49] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X48),s('type.num.num',X49))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X48),s('type.num.num',X49))),s('type.num.num','const.arithmetic.ZERO_0'))))) )),
  inference(forward_demodulation,[],[f3653,f211])).
fof(f3653,plain,(
  ( ! [X48,X49] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X48),s('type.num.num',X49))),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X48))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X49)))))) )),
  inference(forward_demodulation,[],[f3595,f2541])).
fof(f3595,plain,(
  ( ! [X48,X49] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X48),s('type.num.num',X49)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X48))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X49)))))) )),
  inference(superposition,[],[f194,f211])).
fof(f3593,plain,(
  ( ! [X45,X44] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X44),s('type.num.num',X45))))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X44))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X45)))))) )),
  inference(superposition,[],[f190,f211])).
fof(f3651,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(forward_demodulation,[],[f3570,f3567])).
fof(f3570,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f211,f944])).
fof(f3650,plain,(
  ( ! [X2] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.ZERO_0')))))) )),
  inference(forward_demodulation,[],[f3569,f163])).
fof(f3569,plain,(
  ( ! [X2] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f211,f733])).
fof(f3568,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))) )),
  inference(superposition,[],[f211,f153])).
fof(f3567,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f211,f155])).
fof(f3649,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X3))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X3)))))) )),
  inference(forward_demodulation,[],[f3541,f3538])).
fof(f3541,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X3))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3)))))) )),
  inference(superposition,[],[f211,f944])).
fof(f3648,plain,(
  ( ! [X2] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X2)))))) )),
  inference(forward_demodulation,[],[f3540,f162])).
fof(f3540,plain,(
  ( ! [X2] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X2))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))))) )),
  inference(superposition,[],[f211,f733])).
fof(f3539,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(superposition,[],[f211,f153])).
fof(f3538,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f211,f155])).
fof(f3439,plain,(
  ( ! [X118,X120,X119] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X120))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X118),s('type.num.num',X119))))))) | s('type.num.num',X120) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X118))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X119)))))) )),
  inference(superposition,[],[f222,f210])).
fof(f3431,plain,(
  ( ! [X94,X95,X96] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X96))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X94),s('type.num.num',X95))))))) | s('type.num.num',X96) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X94))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X95)))))) )),
  inference(superposition,[],[f218,f210])).
fof(f3473,plain,(
  ( ! [X50,X51] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X50),s('type.num.num',X51))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X50),s('type.num.num',X51))),s('type.num.num','const.arithmetic.ZERO_0'))))) )),
  inference(forward_demodulation,[],[f3472,f210])).
fof(f3472,plain,(
  ( ! [X50,X51] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X50),s('type.num.num',X51))),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X50))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X51)))))) )),
  inference(forward_demodulation,[],[f3416,f2541])).
fof(f3416,plain,(
  ( ! [X50,X51] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X50),s('type.num.num',X51)))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X50))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X51)))))) )),
  inference(superposition,[],[f195,f210])).
fof(f3471,plain,(
  ( ! [X48,X49] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X48),s('type.num.num',X49))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X48),s('type.num.num',X49))),s('type.num.num','const.arithmetic.ZERO_0'))))) )),
  inference(forward_demodulation,[],[f3470,f210])).
fof(f3470,plain,(
  ( ! [X48,X49] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X48),s('type.num.num',X49))),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X48))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X49)))))) )),
  inference(forward_demodulation,[],[f3415,f2541])).
fof(f3415,plain,(
  ( ! [X48,X49] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X48),s('type.num.num',X49)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X48))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X49)))))) )),
  inference(superposition,[],[f194,f210])).
fof(f3413,plain,(
  ( ! [X45,X44] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X44),s('type.num.num',X45))))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X44))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X45)))))) )),
  inference(superposition,[],[f190,f210])).
fof(f3468,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(forward_demodulation,[],[f3390,f3387])).
fof(f3390,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f210,f944])).
fof(f3467,plain,(
  ( ! [X2] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.ZERO_0')))))) )),
  inference(forward_demodulation,[],[f3389,f229])).
fof(f3389,plain,(
  ( ! [X2] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f210,f733])).
fof(f3388,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))) )),
  inference(superposition,[],[f210,f153])).
fof(f3387,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f210,f155])).
fof(f3466,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X3))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X3)))))) )),
  inference(forward_demodulation,[],[f3362,f3359])).
fof(f3362,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X3))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3)))))) )),
  inference(superposition,[],[f210,f944])).
fof(f3361,plain,(
  ( ! [X2] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X2))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))))) )),
  inference(superposition,[],[f210,f733])).
fof(f3360,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(superposition,[],[f210,f153])).
fof(f3359,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f210,f155])).
fof(f3314,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) )),
  inference(subsumption_resolution,[],[f3267,f156])).
fof(f3267,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) )),
  inference(superposition,[],[f225,f181])).
fof(f3206,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(forward_demodulation,[],[f3203,f155])).
fof(f3203,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(unit_resulting_resolution,[],[f2384,f222])).
fof(f3204,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(unit_resulting_resolution,[],[f2384,f220])).
fof(f3205,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(unit_resulting_resolution,[],[f2384,f218])).
fof(f3103,plain,(
  s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(forward_demodulation,[],[f3100,f153])).
fof(f3100,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(unit_resulting_resolution,[],[f2383,f222])).
fof(f3101,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(unit_resulting_resolution,[],[f2383,f220])).
fof(f3102,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(unit_resulting_resolution,[],[f2383,f218])).
fof(f3077,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',X3))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3)))))) )),
  inference(superposition,[],[f209,f181])).
fof(f3074,plain,(
  ( ! [X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))) )),
  inference(superposition,[],[f209,f153])).
fof(f3073,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f209,f155])).
fof(f3046,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X0))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f209,f155])).
fof(f2992,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',X2))) | s(bool,'T_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3)))))) )),
  inference(superposition,[],[f208,f181])).
fof(f3001,plain,(
  ( ! [X3] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X3))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X3)))) )),
  inference(forward_demodulation,[],[f2991,f2988])).
fof(f2991,plain,(
  ( ! [X3] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X3))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f208,f944])).
fof(f2989,plain,(
  ( ! [X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X1))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))) )),
  inference(superposition,[],[f208,f153])).
fof(f2988,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X0))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f208,f155])).
fof(f2962,plain,(
  ( ! [X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(superposition,[],[f208,f153])).
fof(f2935,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f2007,f222])).
fof(f2936,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f2007,f220])).
fof(f2937,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f2007,f218])).
fof(f2927,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))),
  inference(unit_resulting_resolution,[],[f1941,f222])).
fof(f2928,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))),
  inference(unit_resulting_resolution,[],[f1941,f220])).
fof(f2929,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))),
  inference(unit_resulting_resolution,[],[f1941,f218])).
fof(f2914,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',X3))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X2))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X3)))))) )),
  inference(superposition,[],[f207,f181])).
fof(f2865,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f1637,f222])).
fof(f2866,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f1637,f220])).
fof(f2867,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f1637,f218])).
fof(f2857,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))),
  inference(unit_resulting_resolution,[],[f1592,f222])).
fof(f2858,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))),
  inference(unit_resulting_resolution,[],[f1592,f220])).
fof(f2859,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))),
  inference(unit_resulting_resolution,[],[f1592,f218])).
fof(f2844,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',X3))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X2))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X3)))))) )),
  inference(superposition,[],[f206,f181])).
fof(f2795,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f1272,f222])).
fof(f2796,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f1272,f220])).
fof(f2797,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f1272,f218])).
fof(f2787,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))),
  inference(unit_resulting_resolution,[],[f1231,f222])).
fof(f2788,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))),
  inference(unit_resulting_resolution,[],[f1231,f220])).
fof(f2789,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))),
  inference(unit_resulting_resolution,[],[f1231,f218])).
fof(f2774,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',X3))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X2))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X3)))))) )),
  inference(superposition,[],[f205,f181])).
fof(f2686,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X3))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3)))))) )),
  inference(superposition,[],[f204,f181])).
fof(f2683,plain,(
  ( ! [X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))) )),
  inference(superposition,[],[f204,f153])).
fof(f2682,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f204,f155])).
fof(f2656,plain,(
  ( ! [X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(superposition,[],[f204,f153])).
fof(f2655,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X0))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f204,f155])).
fof(f2570,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num',X2))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3)))))) )),
  inference(superposition,[],[f203,f181])).
fof(f2582,plain,(
  ( ! [X3] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X3))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X3)))) )),
  inference(forward_demodulation,[],[f2569,f2566])).
fof(f2569,plain,(
  ( ! [X3] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X3))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f203,f944])).
fof(f2581,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X2)))) )),
  inference(forward_demodulation,[],[f2568,f161])).
fof(f2568,plain,(
  ( ! [X2] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X2))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f203,f733])).
fof(f2567,plain,(
  ( ! [X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X1))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))) )),
  inference(superposition,[],[f203,f153])).
fof(f2566,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X0))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f203,f155])).
fof(f2580,plain,(
  ( ! [X3] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) )),
  inference(forward_demodulation,[],[f2542,f2539])).
fof(f2542,plain,(
  ( ! [X3] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3)))))) )),
  inference(superposition,[],[f203,f944])).
fof(f2578,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0'))))) )),
  inference(backward_demodulation,[],[f2541,f909])).
fof(f2577,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0'))))) )),
  inference(backward_demodulation,[],[f2541,f817])).
fof(f2541,plain,(
  ( ! [X2] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))))) )),
  inference(superposition,[],[f203,f733])).
fof(f2540,plain,(
  ( ! [X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(superposition,[],[f203,f153])).
fof(f2539,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f203,f155])).
fof(f2494,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,sK4(X0,X1,X2,'const.prelim.EQUAL_0',X3))) )),
  inference(unit_resulting_resolution,[],[f2491,f375])).
fof(f2493,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X3))) | s(bool,'T_0') = s(bool,sK4(X0,X1,X2,X3,X4))) )),
  inference(subsumption_resolution,[],[f2492,f156])).
fof(f2492,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'F_0')) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X3))) | s(bool,'T_0') = s(bool,sK4(X0,X1,X2,X3,X4))) )),
  inference(superposition,[],[f233,f181])).
fof(f2491,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,sK4(X0,X1,X2,'const.prelim.EQUAL_0',X3)))) )),
  inference(unit_resulting_resolution,[],[f228,f233])).
fof(f2467,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.ZERO_0'))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f180,f181])).
fof(f2475,plain,(
  s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(forward_demodulation,[],[f2466,f2463])).
fof(f2466,plain,(
  s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f180,f944])).
fof(f2474,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(forward_demodulation,[],[f2465,f164])).
fof(f2465,plain,(
  s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f180,f733])).
fof(f2464,plain,(
  s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))),s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f180,f153])).
fof(f2463,plain,(
  s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f180,f155])).
fof(f2414,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X1))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(superposition,[],[f179,f181])).
fof(f2422,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))),
  inference(forward_demodulation,[],[f2411,f169])).
fof(f2411,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f179,f153])).
fof(f2421,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f2410,f172])).
fof(f2410,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f179,f155])).
fof(f2383,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(unit_resulting_resolution,[],[f2362,f220])).
fof(f2384,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(unit_resulting_resolution,[],[f2362,f218])).
fof(f2362,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(unit_resulting_resolution,[],[f2312,f191])).
fof(f2350,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X1))) | s(bool,'T_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f178,f181])).
fof(f2360,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(forward_demodulation,[],[f2349,f2357])).
fof(f2349,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f178,f944])).
fof(f2358,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))),s('type.num.num','const.num.|48|_0')))),
  inference(forward_demodulation,[],[f2347,f169])).
fof(f2347,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))),s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f178,f153])).
fof(f2357,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))),
  inference(forward_demodulation,[],[f2346,f172])).
fof(f2346,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f178,f155])).
fof(f2288,plain,(
  ( ! [X10] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num',X10) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) )),
  inference(superposition,[],[f222,f944])).
fof(f2312,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(subsumption_resolution,[],[f2279,f786])).
fof(f2279,plain,(
  p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f195,f944])).
fof(f2255,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(superposition,[],[f222,f944])).
fof(f2311,plain,(
  s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(forward_demodulation,[],[f2251,f478])).
fof(f2251,plain,(
  s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f171,f944])).
fof(f2310,plain,(
  s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(forward_demodulation,[],[f2250,f461])).
fof(f2250,plain,(
  s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f170,f944])).
fof(f2059,plain,(
  ( ! [X0] : (s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))) )),
  inference(forward_demodulation,[],[f2034,f153])).
fof(f2034,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))) )),
  inference(unit_resulting_resolution,[],[f1584,f222])).
fof(f2035,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))) )),
  inference(unit_resulting_resolution,[],[f1584,f220])).
fof(f2036,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))) )),
  inference(unit_resulting_resolution,[],[f1584,f218])).
fof(f1971,plain,(
  ( ! [X35,X36,X34] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X35))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X36))) | s('type.num.num',X36) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X34),s('type.num.num',X35)))) )),
  inference(superposition,[],[f222,f199])).
fof(f1970,plain,(
  ( ! [X33,X31,X32] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X31))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X33))) | s('type.num.num',X33) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X31),s('type.num.num',X32)))) )),
  inference(superposition,[],[f222,f198])).
fof(f1968,plain,(
  ( ! [X29] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X29))) | s('type.num.num',X29) = s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.LESS_0')))) )),
  inference(superposition,[],[f222,f152])).
fof(f2021,plain,(
  ( ! [X28] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X28))) | s('type.num.num',X28) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(forward_demodulation,[],[f1967,f733])).
fof(f1967,plain,(
  ( ! [X28] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X28))) | s('type.num.num',X28) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(superposition,[],[f222,f1290])).
fof(f1965,plain,(
  ( ! [X25] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X25))) | s('type.num.num',X25) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f222,f151])).
fof(f1963,plain,(
  ( ! [X22] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X22))) | s('type.num.num',X22) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f222,f1163])).
fof(f1962,plain,(
  ( ! [X21] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X21))) | s('type.num.num',X21) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))) )),
  inference(superposition,[],[f222,f1093])).
fof(f1961,plain,(
  ( ! [X19,X20] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X20))) | s('type.num.num',X20) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X19)))))))) )),
  inference(superposition,[],[f222,f177])).
fof(f1960,plain,(
  ( ! [X17,X18] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X18))) | s('type.num.num',X18) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X17)))))))) )),
  inference(superposition,[],[f222,f176])).
fof(f1959,plain,(
  ( ! [X15,X16] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X16))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num',X16) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X15),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f222,f229])).
fof(f1958,plain,(
  ( ! [X14,X13] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X14))) | s('type.num.num',X14) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X13),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f222,f163])).
fof(f1957,plain,(
  ( ! [X12,X11] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X12))) | s('type.num.num',X12) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X11)))) )),
  inference(superposition,[],[f222,f162])).
fof(f1956,plain,(
  ( ! [X10,X9] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))) | s('type.num.num',X10) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X9)))) )),
  inference(superposition,[],[f222,f167])).
fof(f1955,plain,(
  ( ! [X8,X7] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8))) | s('type.num.num',X8) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X7),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f222,f166])).
fof(f1954,plain,(
  ( ! [X6] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6))) | s('type.num.num',X6) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(superposition,[],[f222,f733])).
fof(f1951,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3))) | s('type.num.num',X3) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f222,f160])).
fof(f1950,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) | s('type.num.num',X1) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(superposition,[],[f222,f158])).
fof(f1936,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))) )),
  inference(unit_resulting_resolution,[],[f682,f222])).
fof(f1937,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))) )),
  inference(unit_resulting_resolution,[],[f683,f222])).
fof(f1941,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))),
  inference(unit_resulting_resolution,[],[f711,f222])).
fof(f2007,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f1942,f228])).
fof(f1942,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.EQUAL_0')))))),
  inference(unit_resulting_resolution,[],[f901,f222])).
fof(f2006,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))) )),
  inference(forward_demodulation,[],[f1944,f155])).
fof(f1944,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))) )),
  inference(unit_resulting_resolution,[],[f1222,f222])).
fof(f2005,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))) )),
  inference(forward_demodulation,[],[f1945,f155])).
fof(f1945,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))) )),
  inference(unit_resulting_resolution,[],[f1223,f222])).
fof(f2004,plain,(
  ( ! [X0] : (s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))) )),
  inference(forward_demodulation,[],[f1946,f153])).
fof(f1946,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))) )),
  inference(unit_resulting_resolution,[],[f1583,f222])).
fof(f1898,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))) )),
  inference(unit_resulting_resolution,[],[f1583,f220])).
fof(f1899,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))) )),
  inference(unit_resulting_resolution,[],[f1583,f218])).
fof(f1870,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))) )),
  inference(unit_resulting_resolution,[],[f1223,f220])).
fof(f1871,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))) )),
  inference(unit_resulting_resolution,[],[f1223,f218])).
fof(f1709,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))) )),
  inference(unit_resulting_resolution,[],[f1222,f220])).
fof(f1710,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))) )),
  inference(unit_resulting_resolution,[],[f1222,f218])).
fof(f1683,plain,(
  ( ! [X4] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X4))) | s('type.num.num',X4) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f220,f1163])).
fof(f1679,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) | s('type.num.num',X0) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f218,f1163])).
fof(f1614,plain,(
  ( ! [X35,X33,X34] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X34))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X35))) | s('type.num.num',X35) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X33),s('type.num.num',X34)))) )),
  inference(superposition,[],[f220,f199])).
fof(f1613,plain,(
  ( ! [X30,X31,X32] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X30))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X32))) | s('type.num.num',X32) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X30),s('type.num.num',X31)))) )),
  inference(superposition,[],[f220,f198])).
fof(f1611,plain,(
  ( ! [X28] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X28))) | s('type.num.num',X28) = s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.LESS_0')))) )),
  inference(superposition,[],[f220,f152])).
fof(f1610,plain,(
  ( ! [X27] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X27))) | s('type.num.num',X27) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(superposition,[],[f220,f1290])).
fof(f1608,plain,(
  ( ! [X24] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X24))) | s('type.num.num',X24) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f220,f151])).
fof(f1606,plain,(
  ( ! [X21] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X21))) | s('type.num.num',X21) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))) )),
  inference(superposition,[],[f220,f1093])).
fof(f1605,plain,(
  ( ! [X19,X20] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X20))) | s('type.num.num',X20) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X19)))))))) )),
  inference(superposition,[],[f220,f177])).
fof(f1604,plain,(
  ( ! [X17,X18] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X18))) | s('type.num.num',X18) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X17)))))))) )),
  inference(superposition,[],[f220,f176])).
fof(f1603,plain,(
  ( ! [X15,X16] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X16))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num',X16) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X15),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f220,f229])).
fof(f1602,plain,(
  ( ! [X14,X13] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X13))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X14))) | s('type.num.num',X14) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X13),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f220,f163])).
fof(f1601,plain,(
  ( ! [X12,X11] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X12))) | s('type.num.num',X12) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X11)))) )),
  inference(superposition,[],[f220,f162])).
fof(f1600,plain,(
  ( ! [X10,X9] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X10))) | s('type.num.num',X10) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X9)))) )),
  inference(superposition,[],[f220,f167])).
fof(f1599,plain,(
  ( ! [X8,X7] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X8))) | s('type.num.num',X8) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X7),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f220,f166])).
fof(f1595,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X2))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X3))) | s('type.num.num',X3) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f220,f160])).
fof(f1594,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))) | s('type.num.num',X1) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(superposition,[],[f220,f158])).
fof(f1583,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f230,f220])).
fof(f1584,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f231,f220])).
fof(f1586,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f235,f220])).
fof(f1587,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))) )),
  inference(unit_resulting_resolution,[],[f682,f220])).
fof(f1588,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))) )),
  inference(unit_resulting_resolution,[],[f683,f220])).
fof(f1592,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))),
  inference(unit_resulting_resolution,[],[f711,f220])).
fof(f1637,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f1593,f228])).
fof(f1593,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.EQUAL_0')))))),
  inference(unit_resulting_resolution,[],[f901,f220])).
fof(f1548,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) | s('type.num.num',X0) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))) )),
  inference(superposition,[],[f218,f1093])).
fof(f1499,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) | s('type.num.num',X0) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(superposition,[],[f218,f1290])).
fof(f1290,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(unit_resulting_resolution,[],[f1033,f190])).
fof(f1251,plain,(
  ( ! [X33,X31,X32] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X32))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X33))) | s('type.num.num',X33) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X31),s('type.num.num',X32)))) )),
  inference(superposition,[],[f218,f199])).
fof(f1250,plain,(
  ( ! [X30,X28,X29] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X28))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X30))) | s('type.num.num',X30) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X28),s('type.num.num',X29)))) )),
  inference(superposition,[],[f218,f198])).
fof(f1248,plain,(
  ( ! [X26] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X26))) | s('type.num.num',X26) = s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.LESS_0')))) )),
  inference(superposition,[],[f218,f152])).
fof(f1246,plain,(
  ( ! [X23] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X23))) | s('type.num.num',X23) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f218,f151])).
fof(f1244,plain,(
  ( ! [X19,X20] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X20))) | s('type.num.num',X20) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X19)))))))) )),
  inference(superposition,[],[f218,f177])).
fof(f1243,plain,(
  ( ! [X17,X18] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X18))) | s('type.num.num',X18) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X17)))))))) )),
  inference(superposition,[],[f218,f176])).
fof(f1242,plain,(
  ( ! [X15,X16] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X16))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num',X16) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X15),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f218,f229])).
fof(f1241,plain,(
  ( ! [X14,X13] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X13))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X14))) | s('type.num.num',X14) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X13),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f218,f163])).
fof(f1240,plain,(
  ( ! [X12,X11] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X12))) | s('type.num.num',X12) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X11)))) )),
  inference(superposition,[],[f218,f162])).
fof(f1239,plain,(
  ( ! [X10,X9] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X10))) | s('type.num.num',X10) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X9)))) )),
  inference(superposition,[],[f218,f167])).
fof(f1238,plain,(
  ( ! [X8,X7] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X8))) | s('type.num.num',X8) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X7),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f218,f166])).
fof(f1234,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X2))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X3))) | s('type.num.num',X3) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f218,f160])).
fof(f1233,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))) | s('type.num.num',X1) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(superposition,[],[f218,f158])).
fof(f1222,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f230,f218])).
fof(f1223,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f231,f218])).
fof(f1225,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f235,f218])).
fof(f1226,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))) )),
  inference(unit_resulting_resolution,[],[f682,f218])).
fof(f1227,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))) )),
  inference(unit_resulting_resolution,[],[f683,f218])).
fof(f1231,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))),
  inference(unit_resulting_resolution,[],[f711,f218])).
fof(f1272,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f1232,f228])).
fof(f1232,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.EQUAL_0')))))),
  inference(unit_resulting_resolution,[],[f901,f218])).
fof(f1163,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f177,f155])).
fof(f1093,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))),
  inference(superposition,[],[f176,f153])).
fof(f157,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(cnf_transformation,[],[f68])).
fof(f68,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(flattening,[],[f44])).
fof(f44,axiom,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_distrib|split|34')).
fof(f1033,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(forward_demodulation,[],[f1001,f151])).
fof(f1001,plain,(
  s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f175,f733])).
fof(f909,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))) )),
  inference(unit_resulting_resolution,[],[f683,f194])).
fof(f904,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(subsumption_resolution,[],[f861,f156])).
fof(f861,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(superposition,[],[f195,f181])).
fof(f901,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.EQUAL_0')))),
  inference(subsumption_resolution,[],[f857,f786])).
fof(f857,plain,(
  p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.EQUAL_0')))),
  inference(superposition,[],[f195,f228])).
fof(f896,plain,(
  ( ! [X6] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f854,f786])).
fof(f854,plain,(
  ( ! [X6] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f195,f229])).
fof(f817,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))))) )),
  inference(unit_resulting_resolution,[],[f682,f194])).
fof(f785,plain,(
  ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))),
  inference(unit_resulting_resolution,[],[f711,f194])).
fof(f786,plain,(
  ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(unit_resulting_resolution,[],[f712,f194])).
fof(f782,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(forward_demodulation,[],[f754,f383])).
fof(f754,plain,(
  s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0'))) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f171,f733])).
fof(f712,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),
  inference(subsumption_resolution,[],[f699,f231])).
fof(f699,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f190,f155])).
fof(f711,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))),
  inference(subsumption_resolution,[],[f698,f230])).
fof(f698,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f190,f153])).
fof(f697,plain,(
  ( ! [X10,X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X9),s('type.num.num',X10)))) )),
  inference(superposition,[],[f190,f199])).
fof(f696,plain,(
  ( ! [X8,X7] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X7),s('type.num.num',X8)))) )),
  inference(superposition,[],[f190,f198])).
fof(f691,plain,(
  ( ! [X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X5))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X5),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f190,f163])).
fof(f685,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f190,f160])).
fof(f684,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(superposition,[],[f190,f158])).
fof(f682,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f230,f190])).
fof(f683,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f231,f190])).
fof(f620,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,X1))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f199,f181])).
fof(f568,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.bool.COND_3'(s(bool,X0),s(X1,X3),s(X1,X2))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f199,f181])).
fof(f545,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f198,f181])).
fof(f478,plain,(
  s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f171,f155])).
fof(f477,plain,(
  s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))),
  inference(superposition,[],[f171,f153])).
fof(f462,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))) | s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X1)))) )),
  inference(superposition,[],[f170,f181])).
fof(f461,plain,(
  s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f170,f155])).
fof(f460,plain,(
  s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))),
  inference(superposition,[],[f170,f153])).
fof(f383,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),
  inference(subsumption_resolution,[],[f380,f156])).
fof(f380,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f150,f181])).
fof(f375,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f156,f181])).
fof(f226,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK5(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK5(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK5(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK5(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f128])).
fof(f128,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f125])).
fof(f125,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.prelim.ordering_case_def|split|2',aHL_EXT)).
fof(f224,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))))),
  inference(nnf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) <=> ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f123])).
fof(f123,plain,(
  ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) <=> ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X13,X14] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X13))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X14)))))) <=> ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X14),s('type.num.num',X13)))))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_lt|split|5')).
fof(f225,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f146])).
fof(f223,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num',X1) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f145])).
fof(f220,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0,X1] : ((s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.num.num',X0) != s('type.num.num',X1) | s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X13,X14] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X13))) = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X14))) <=> s('type.num.num',X13) = s('type.num.num',X14))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_eq|split|7')).
fof(f221,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num',X1) | s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f144])).
fof(f218,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1] : ((s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.num.num',X0) != s('type.num.num',X1) | s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X13,X14] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X13))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X14))) <=> s('type.num.num',X13) = s('type.num.num',X14))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_eq|split|6')).
fof(f219,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num',X1) | s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f143])).
fof(f213,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))))))) )),
  inference(cnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X13,X14] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X14))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X13),s('type.num.num',X14)))))))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_distrib|split|2')).
fof(f211,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X13,X14] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X13),s('type.num.num',X14))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X14)))))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_distrib|split|8')).
fof(f210,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X13,X14] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X13),s('type.num.num',X14))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X14)))))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_distrib|split|12')).
fof(f207,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X13,X14] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X13),s('type.num.num',X14))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X13))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X14)))))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_lt|split|3')).
fof(f206,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X13,X14] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X13),s('type.num.num',X14))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X13))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X14)))))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_lt|split|4')).
fof(f204,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0,X1] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X13,X14] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X13),s('type.num.num',X14))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X14)))))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_distrib|split|28')).
fof(f203,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0,X1] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X13,X14] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X14),s('type.num.num',X13))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X14)))))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_distrib|split|31')).
fof(f233,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,sK4(X4,X3,X2,X1,X0))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1)))) )),
  inference(forward_demodulation,[],[f201,f155])).
fof(f201,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,sK4(X4,X3,X2,X1,X0))) | s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) )),
  inference(cnf_transformation,[],[f140])).
fof(f198,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f101])).
fof(f101,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X11,X12] : s(X6,X11) = s(X6,'const.bool.COND_3'(s(bool,'T_0'),s(X6,X11),s(X6,X12)))),
  file('thm.prelim.ordering_case_def|split|2','thm.bool.COND_CLAUSES|split|0')).
fof(f194,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0] : ((~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))))),
  inference(nnf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X0))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X13] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X13))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_distrib|split|30')).
fof(f195,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f136])).
fof(f231,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))) )),
  inference(subsumption_resolution,[],[f186,f156])).
fof(f186,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0] : ((s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))),
  inference(nnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X13] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X13))) <=> p(s(bool,'F_0')))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_eq|split|1')).
fof(f230,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))) )),
  inference(subsumption_resolution,[],[f182,f156])).
fof(f182,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0] : ((s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))),
  inference(nnf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X13] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X13))) <=> p(s(bool,'F_0')))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_eq|split|3')).
fof(f180,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X13] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X13),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13))),s('type.num.num','const.num.|48|_0')))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_distrib|split|27')).
fof(f179,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0] : s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X13] : s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X13))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13)))))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_distrib|split|21')).
fof(f178,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0] : s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X13] : s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X13))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13))),s('type.num.num','const.num.|48|_0')))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_distrib|split|24')).
fof(f176,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))) )),
  inference(cnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X13] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X13)))))))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_distrib|split|10')).
fof(f229,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f174,f155])).
fof(f174,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X13] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X13),s('type.num.num','const.num.|48|_0')))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_distrib|split|11')).
fof(f172,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X13] : s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X13)))))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_lt|split|0')).
fof(f169,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X13] : s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X13)))))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_lt|split|1')).
fof(f168,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X13] : s(bool,'F_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_distrib|split|23')).
fof(f164,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X13] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_distrib|split|26')).
fof(f163,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X13] : s('type.num.num',X13) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X13),s('type.num.num','const.num.|48|_0')))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_distrib|split|7')).
fof(f162,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X13] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_distrib|split|6')).
fof(f161,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X13] : s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X13),s('type.num.num','const.num.|48|_0')))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_distrib|split|29')).
fof(f160,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X13] : s('type.num.num',X13) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X13),s('type.num.num','const.num.|48|_0')))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_distrib|split|1')).
fof(f159,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(cnf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X13] : s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X13),s('type.num.num','const.arithmetic.ZERO_0')))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_lt|split|2')).
fof(f158,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X13] : s('type.num.num',X13) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X13)))),
  file('thm.prelim.ordering_case_def|split|2','thm.arithmetic.ADD|split|0')).
fof(f228,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.EQUAL_0')))),
  inference(backward_demodulation,[],[f155,f154])).
fof(f154,plain,(
  s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.EQUAL_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f62])).
fof(f62,axiom,(
  s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.EQUAL_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  file('thm.prelim.ordering_case_def|split|2','thm.prelim.ordering2num_thm|split|1')).
fof(f152,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.LESS_0')))),
  inference(cnf_transformation,[],[f61])).
fof(f61,axiom,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.LESS_0')))),
  file('thm.prelim.ordering_case_def|split|2','thm.prelim.ordering2num_thm|split|0')).
fof(f151,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))),
  inference(cnf_transformation,[],[f25])).
fof(f25,axiom,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))),
  file('thm.prelim.ordering_case_def|split|2','thm.prim_rec.PRE|split|0')).
fof(f150,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(cnf_transformation,[],[f45])).
fof(f45,axiom,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_distrib|split|35')).
fof(f11178,plain,(
  ( ! [X1] : (~p(s(bool,'F_0')) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering',X1)))) ) | $spl0),
  inference(backward_demodulation,[],[f385,f234])).
fof(f11162,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)) ) | $spl0),
  inference(backward_demodulation,[],[f385,f195])).
fof(f177,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))) )),
  inference(cnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X13] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X13)))))))),
  file('thm.prelim.ordering_case_def|split|2','thm.numeral.numeral_distrib|split|9')).
fof(f11832,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f176,f11831])).
fof(f11826,plain,(
  $false | ($spl0 | $spl41 | $spl48 | $spl62)),
  inference(subsumption_resolution,[],[f11825,f149])).
fof(f11825,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl41 | $spl48 | $spl62)),
  inference(forward_demodulation,[],[f11164,f10688])).
fof(f10688,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl48 | $spl62)),
  inference(backward_demodulation,[],[f10599,f10674])).
fof(f11164,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl41)),
  inference(backward_demodulation,[],[f385,f10476])).
fof(f10476,plain,(
  ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl41),
  inference(unit_resulting_resolution,[],[f10448,f194])).
fof(f11824,plain,(
  $false | ($spl0 | $spl37 | $spl48 | $spl62)),
  inference(subsumption_resolution,[],[f11823,f149])).
fof(f11823,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl37 | $spl48 | $spl62)),
  inference(forward_demodulation,[],[f11163,f10688])).
fof(f11163,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl37)),
  inference(backward_demodulation,[],[f385,f10283])).
fof(f10283,plain,(
  ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))) | $spl37),
  inference(unit_resulting_resolution,[],[f10255,f194])).
fof(f11820,plain,(
  $false | ($spl0 | $spl29 | $spl48 | $spl62)),
  inference(subsumption_resolution,[],[f11819,f149])).
fof(f11819,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl29 | $spl48 | $spl62)),
  inference(forward_demodulation,[],[f11153,f10688])).
fof(f11153,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl29)),
  inference(backward_demodulation,[],[f385,f10111])).
fof(f10111,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.|48|_0')))) | $spl29),
  inference(forward_demodulation,[],[f10110,f9758])).
fof(f10110,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))),s('type.num.num','const.num.|48|_0')))) | $spl29),
  inference(forward_demodulation,[],[f10101,f4782])).
fof(f10101,plain,(
  ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))))) | $spl29),
  inference(unit_resulting_resolution,[],[f9904,f194])).
fof(f9904,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))) | $spl29),
  inference(forward_demodulation,[],[f9901,f4578])).
fof(f9901,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))) | $spl29),
  inference(unit_resulting_resolution,[],[f9869,f222])).
fof(f11818,plain,(
  $false | ($spl0 | $spl33 | $spl48 | $spl62)),
  inference(subsumption_resolution,[],[f11817,f149])).
fof(f11817,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl33 | $spl48 | $spl62)),
  inference(forward_demodulation,[],[f11149,f10688])).
fof(f11149,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl33)),
  inference(backward_demodulation,[],[f385,f10503])).
fof(f10503,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.|48|_0')))) | $spl33),
  inference(forward_demodulation,[],[f10502,f9951])).
fof(f10502,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.|48|_0')))) | $spl33),
  inference(forward_demodulation,[],[f10487,f4782])).
fof(f10487,plain,(
  ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) | $spl33),
  inference(unit_resulting_resolution,[],[f10097,f194])).
fof(f10097,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl33),
  inference(forward_demodulation,[],[f10094,f4578])).
fof(f10094,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl33),
  inference(unit_resulting_resolution,[],[f10062,f222])).
fof(f11808,plain,(
  $false | ($spl0 | $spl48 | $spl62)),
  inference(subsumption_resolution,[],[f11753,f4578])).
fof(f11753,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl48 | $spl62)),
  inference(backward_demodulation,[],[f11571,f4602])).
fof(f11571,plain,(
  ( ! [X53] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X53)))) ) | ($spl0 | $spl48 | $spl62)),
  inference(subsumption_resolution,[],[f11570,f149])).
fof(f11570,plain,(
  ( ! [X53] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X53)))) ) | ($spl0 | $spl48 | $spl62)),
  inference(forward_demodulation,[],[f11118,f10688])).
fof(f11118,plain,(
  ( ! [X53] : (~p(s(bool,'F_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X53)))) ) | $spl0),
  inference(backward_demodulation,[],[f385,f5839])).
fof(f11807,plain,(
  $false | ($spl0 | $spl29 | $spl48 | $spl62)),
  inference(subsumption_resolution,[],[f11741,f4578])).
fof(f11741,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl29 | $spl48 | $spl62)),
  inference(backward_demodulation,[],[f11571,f10108])).
fof(f10108,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))))) | $spl29),
  inference(forward_demodulation,[],[f10105,f4578])).
fof(f10105,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))))) | $spl29),
  inference(unit_resulting_resolution,[],[f9904,f222])).
fof(f11806,plain,(
  $false | ($spl0 | $spl48 | $spl62)),
  inference(subsumption_resolution,[],[f11724,f11795])).
fof(f11795,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num',X1)) ) | ($spl0 | $spl48 | $spl62)),
  inference(subsumption_resolution,[],[f11577,f11571])).
fof(f11577,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) ) | ($spl0 | $spl48 | $spl62)),
  inference(backward_demodulation,[],[f11571,f222])).
fof(f11724,plain,(
  s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl48 | $spl62)),
  inference(backward_demodulation,[],[f11571,f8345])).
fof(f11805,plain,(
  $false | ($spl0 | $spl48 | $spl62)),
  inference(subsumption_resolution,[],[f11718,f4578])).
fof(f11718,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl48 | $spl62)),
  inference(backward_demodulation,[],[f11571,f4579])).
fof(f11804,plain,(
  $false | ($spl0 | $spl48 | $spl62)),
  inference(subsumption_resolution,[],[f11701,f11795])).
fof(f11701,plain,(
  s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl48 | $spl62)),
  inference(backward_demodulation,[],[f11571,f8387])).
fof(f11803,plain,(
  $false | ($spl0 | $spl48 | $spl62)),
  inference(subsumption_resolution,[],[f11698,f11795])).
fof(f11698,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl48 | $spl62)),
  inference(backward_demodulation,[],[f11571,f8231])).
fof(f11802,plain,(
  $false | ($spl0 | $spl48 | $spl62)),
  inference(subsumption_resolution,[],[f11692,f11795])).
fof(f11692,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))) ) | ($spl0 | $spl48 | $spl62)),
  inference(backward_demodulation,[],[f11571,f6887])).
fof(f11801,plain,(
  $false | ($spl0 | $spl48 | $spl62)),
  inference(subsumption_resolution,[],[f11691,f11795])).
fof(f11691,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))) ) | ($spl0 | $spl48 | $spl62)),
  inference(backward_demodulation,[],[f11571,f6886])).
fof(f11800,plain,(
  $false | ($spl0 | $spl48 | $spl62)),
  inference(subsumption_resolution,[],[f11689,f4578])).
fof(f11689,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl48 | $spl62)),
  inference(backward_demodulation,[],[f11571,f4580])).
fof(f11799,plain,(
  $false | ($spl0 | $spl48 | $spl62)),
  inference(subsumption_resolution,[],[f11676,f4578])).
fof(f11676,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl48 | $spl62)),
  inference(backward_demodulation,[],[f11571,f6569])).
fof(f11798,plain,(
  $false | ($spl0 | $spl48 | $spl62)),
  inference(subsumption_resolution,[],[f11671,f4578])).
fof(f11671,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl48 | $spl62)),
  inference(backward_demodulation,[],[f11571,f6338])).
fof(f11797,plain,(
  $false | ($spl0 | $spl48 | $spl62)),
  inference(subsumption_resolution,[],[f11667,f4578])).
fof(f11667,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl48 | $spl62)),
  inference(backward_demodulation,[],[f11571,f6378])).
fof(f11796,plain,(
  $false | ($spl0 | $spl48 | $spl62)),
  inference(subsumption_resolution,[],[f11666,f11571])).
fof(f11666,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))))) ) | ($spl0 | $spl48 | $spl62)),
  inference(backward_demodulation,[],[f11571,f6885])).
fof(f11792,plain,(
  $false | ($spl0 | $spl48 | $spl62)),
  inference(subsumption_resolution,[],[f11791,f896])).
fof(f11791,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl48 | $spl62)),
  inference(forward_demodulation,[],[f11790,f11571])).
fof(f11790,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) ) | ($spl0 | $spl48 | $spl62)),
  inference(forward_demodulation,[],[f11574,f11571])).
fof(f11574,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num',X1)))))) ) | ($spl0 | $spl48 | $spl62)),
  inference(backward_demodulation,[],[f11571,f210])).
fof(f11788,plain,(
  $false | ($spl0 | $spl48 | $spl62)),
  inference(subsumption_resolution,[],[f11787,f712])).
fof(f11787,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl48 | $spl62)),
  inference(forward_demodulation,[],[f11572,f11571])).
fof(f11572,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) ) | ($spl0 | $spl48 | $spl62)),
  inference(backward_demodulation,[],[f11571,f173])).
fof(f11778,plain,(
  $false | ($spl0 | $spl48 | $spl62)),
  inference(trivial_inequality_removal,[],[f11668])).
fof(f11668,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl48 | $spl62)),
  inference(backward_demodulation,[],[f11571,f6349])).
fof(f11777,plain,(
  $false | ($spl0 | $spl48 | $spl62)),
  inference(trivial_inequality_removal,[],[f11669])).
fof(f11669,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl48 | $spl62)),
  inference(backward_demodulation,[],[f11571,f6350])).
fof(f11776,plain,(
  $false | ($spl0 | $spl48 | $spl62)),
  inference(trivial_inequality_removal,[],[f11672])).
fof(f11672,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl48 | $spl62)),
  inference(backward_demodulation,[],[f11571,f6309])).
fof(f11775,plain,(
  $false | ($spl0 | $spl48 | $spl62)),
  inference(trivial_inequality_removal,[],[f11673])).
fof(f11673,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl48 | $spl62)),
  inference(backward_demodulation,[],[f11571,f6310])).
fof(f11774,plain,(
  $false | ($spl0 | $spl48 | $spl62)),
  inference(trivial_inequality_removal,[],[f11677])).
fof(f11677,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl48 | $spl62)),
  inference(backward_demodulation,[],[f11571,f6567])).
fof(f11773,plain,(
  $false | ($spl0 | $spl48 | $spl62)),
  inference(trivial_inequality_removal,[],[f11678])).
fof(f11678,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl48 | $spl62)),
  inference(backward_demodulation,[],[f11571,f6568])).
fof(f11772,plain,(
  $false | ($spl0 | $spl48 | $spl62)),
  inference(trivial_inequality_removal,[],[f11684])).
fof(f11684,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl48 | $spl62)),
  inference(backward_demodulation,[],[f11571,f1227])).
fof(f11771,plain,(
  $false | ($spl0 | $spl48 | $spl62)),
  inference(trivial_inequality_removal,[],[f11686])).
fof(f11686,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl48 | $spl62)),
  inference(backward_demodulation,[],[f11571,f1588])).
fof(f11770,plain,(
  $false | ($spl0 | $spl48 | $spl62)),
  inference(trivial_inequality_removal,[],[f11696])).
fof(f11696,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl48 | $spl62)),
  inference(backward_demodulation,[],[f11571,f8197])).
fof(f11769,plain,(
  $false | ($spl0 | $spl48 | $spl62)),
  inference(trivial_inequality_removal,[],[f11697])).
fof(f11697,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl48 | $spl62)),
  inference(backward_demodulation,[],[f11571,f8198])).
fof(f11768,plain,(
  $false | ($spl0 | $spl48 | $spl62)),
  inference(trivial_inequality_removal,[],[f11699])).
fof(f11699,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl48 | $spl62)),
  inference(backward_demodulation,[],[f11571,f8353])).
fof(f11767,plain,(
  $false | ($spl0 | $spl48 | $spl62)),
  inference(trivial_inequality_removal,[],[f11700])).
fof(f11700,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl48 | $spl62)),
  inference(backward_demodulation,[],[f11571,f8354])).
fof(f11766,plain,(
  $false | ($spl0 | $spl48 | $spl62)),
  inference(trivial_inequality_removal,[],[f11714])).
fof(f11714,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl48 | $spl62)),
  inference(backward_demodulation,[],[f11571,f1226])).
fof(f11765,plain,(
  $false | ($spl0 | $spl48 | $spl62)),
  inference(trivial_inequality_removal,[],[f11716])).
fof(f11716,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl48 | $spl62)),
  inference(backward_demodulation,[],[f11571,f1587])).
fof(f11764,plain,(
  $false | ($spl0 | $spl48 | $spl62)),
  inference(trivial_inequality_removal,[],[f11720])).
fof(f11720,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl48 | $spl62)),
  inference(backward_demodulation,[],[f11571,f8155])).
fof(f11763,plain,(
  $false | ($spl0 | $spl48 | $spl62)),
  inference(trivial_inequality_removal,[],[f11721])).
fof(f11721,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl48 | $spl62)),
  inference(backward_demodulation,[],[f11571,f8156])).
fof(f11762,plain,(
  $false | ($spl0 | $spl48 | $spl62)),
  inference(trivial_inequality_removal,[],[f11722])).
fof(f11722,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl48 | $spl62)),
  inference(backward_demodulation,[],[f11571,f8311])).
fof(f11761,plain,(
  $false | ($spl0 | $spl48 | $spl62)),
  inference(trivial_inequality_removal,[],[f11723])).
fof(f11723,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl48 | $spl62)),
  inference(backward_demodulation,[],[f11571,f8312])).
fof(f11760,plain,(
  $false | ($spl0 | $spl29 | $spl48 | $spl62)),
  inference(trivial_inequality_removal,[],[f11742])).
fof(f11742,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl29 | $spl48 | $spl62)),
  inference(backward_demodulation,[],[f11571,f10106])).
fof(f10106,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))))) | $spl29),
  inference(unit_resulting_resolution,[],[f9904,f220])).
fof(f11759,plain,(
  $false | ($spl0 | $spl29 | $spl48 | $spl62)),
  inference(trivial_inequality_removal,[],[f11743])).
fof(f11743,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl29 | $spl48 | $spl62)),
  inference(backward_demodulation,[],[f11571,f10107])).
fof(f10107,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))))) | $spl29),
  inference(unit_resulting_resolution,[],[f9904,f218])).
fof(f11758,plain,(
  $false | ($spl0 | $spl48 | $spl62)),
  inference(trivial_inequality_removal,[],[f11751])).
fof(f11751,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl48 | $spl62)),
  inference(backward_demodulation,[],[f11571,f4595])).
fof(f11757,plain,(
  $false | ($spl0 | $spl48 | $spl62)),
  inference(trivial_inequality_removal,[],[f11752])).
fof(f11752,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl48 | $spl62)),
  inference(backward_demodulation,[],[f11571,f4596])).
fof(f11187,plain,(
  $false | ($spl0 | $spl48 | $spl62)),
  inference(subsumption_resolution,[],[f11186,f149])).
fof(f11186,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl48 | $spl62)),
  inference(forward_demodulation,[],[f11185,f10688])).
fof(f11185,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl48 | $spl62)),
  inference(forward_demodulation,[],[f11184,f385])).
fof(f11184,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))) ) | ($spl0 | $spl48 | $spl62)),
  inference(subsumption_resolution,[],[f11183,f149])).
fof(f11183,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))) ) | ($spl0 | $spl48 | $spl62)),
  inference(forward_demodulation,[],[f10927,f10688])).
fof(f10927,plain,(
  ( ! [X0,X1] : (~p(s(bool,'F_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))) ) | $spl0),
  inference(backward_demodulation,[],[f385,f224])).
fof(f10786,plain,(
  $false | ($spl9 | $spl10 | $spl48 | $spl62)),
  inference(subsumption_resolution,[],[f1140,f10688])).
fof(f10785,plain,(
  $false | ($spl9 | $spl48 | $spl62)),
  inference(subsumption_resolution,[],[f10784,f4741])).
fof(f10784,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))) | ($spl9 | $spl48 | $spl62)),
  inference(forward_demodulation,[],[f10783,f4713])).
fof(f10783,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | ($spl9 | $spl48 | $spl62)),
  inference(subsumption_resolution,[],[f1141,f10688])).
fof(f10782,plain,(
  $false | ($spl9 | $spl10 | $spl48 | $spl62)),
  inference(subsumption_resolution,[],[f5337,f10688])).
fof(f10781,plain,(
  $false | ($spl9 | $spl48 | $spl62)),
  inference(subsumption_resolution,[],[f10780,f4741])).
fof(f10780,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))) | ($spl9 | $spl48 | $spl62)),
  inference(subsumption_resolution,[],[f5338,f10688])).
fof(f10773,plain,(
  $false | ($spl9 | $spl10 | $spl48 | $spl62)),
  inference(subsumption_resolution,[],[f10699,f4741])).
fof(f10699,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))) | ($spl10 | $spl48 | $spl62)),
  inference(backward_demodulation,[],[f10688,f1068])).
fof(f10771,plain,(
  $false | ($spl48 | $spl62)),
  inference(subsumption_resolution,[],[f10691,f149])).
fof(f10691,plain,(
  ~p(s(bool,'T_0')) | ($spl48 | $spl62)),
  inference(backward_demodulation,[],[f10688,f156])).
fof(f10689,plain,(
  $false | ($spl3 | $spl48 | $spl62)),
  inference(subsumption_resolution,[],[f10688,f387])).
fof(f10686,plain,(
  $false | ($spl3 | $spl56 | $spl68)),
  inference(subsumption_resolution,[],[f10685,f387])).
fof(f10685,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl56 | $spl68)),
  inference(backward_demodulation,[],[f10611,f10680])).
fof(f10680,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl68),
  inference(cnf_transformation,[],[f10680_D])).
fof(f10680_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f10682,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl70),
  inference(cnf_transformation,[],[f10682_D])).
fof(f10682_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f10547,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(superposition,[],[f462,f4740])).
fof(f10606,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl54),
  inference(cnf_transformation,[],[f10606_D])).
fof(f10606_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f10678,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))) | $spl66),
  inference(cnf_transformation,[],[f10678_D])).
fof(f10678_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f10546,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(superposition,[],[f462,f4739])).
fof(f10674,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl62),
  inference(cnf_transformation,[],[f10674_D])).
fof(f10674_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f10676,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl64),
  inference(cnf_transformation,[],[f10676_D])).
fof(f10676_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f10545,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(superposition,[],[f462,f4736])).
fof(f10594,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl46),
  inference(cnf_transformation,[],[f10594_D])).
fof(f10594_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f10672,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))) | $spl60),
  inference(cnf_transformation,[],[f10672_D])).
fof(f10672_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f10544,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(superposition,[],[f462,f4735])).
fof(f10611,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl56),
  inference(cnf_transformation,[],[f10611_D])).
fof(f10611_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f10613,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl58),
  inference(cnf_transformation,[],[f10613_D])).
fof(f10613_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f10609,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(forward_demodulation,[],[f10608,f4740])).
fof(f10608,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))))),
  inference(forward_demodulation,[],[f10607,f10311])).
fof(f10607,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))))),
  inference(forward_demodulation,[],[f10515,f4764])).
fof(f10515,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))))),
  inference(superposition,[],[f462,f4740])).
fof(f10604,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))))) | $spl52),
  inference(cnf_transformation,[],[f10604_D])).
fof(f10604_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f10602,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))))),
  inference(forward_demodulation,[],[f10514,f170])).
fof(f10514,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))))))),
  inference(superposition,[],[f462,f4739])).
fof(f10599,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl48),
  inference(cnf_transformation,[],[f10599_D])).
fof(f10599_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f10601,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl50),
  inference(cnf_transformation,[],[f10601_D])).
fof(f10601_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f10597,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(forward_demodulation,[],[f10596,f4736])).
fof(f10596,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))))),
  inference(forward_demodulation,[],[f10595,f9925])).
fof(f10595,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))))),
  inference(forward_demodulation,[],[f10513,f4763])).
fof(f10513,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))))),
  inference(superposition,[],[f462,f4736])).
fof(f10592,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))))) | $spl44),
  inference(cnf_transformation,[],[f10592_D])).
fof(f10592_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f10590,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))))))),
  inference(forward_demodulation,[],[f10512,f170])).
fof(f10512,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))))))),
  inference(superposition,[],[f462,f4735])).
fof(f10450,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | $spl42),
  inference(cnf_transformation,[],[f10450_D])).
fof(f10450_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f10448,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl41),
  inference(cnf_transformation,[],[f10448_D])).
fof(f10448_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f10338,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f4722,f4740])).
fof(f10451,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f10332,f152])).
fof(f10332,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.LESS_0'))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f4592,f4740])).
fof(f10446,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f10331,f151])).
fof(f10331,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f4591,f4740])).
fof(f10257,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) | $spl38),
  inference(cnf_transformation,[],[f10257_D])).
fof(f10257_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f10255,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) | $spl37),
  inference(cnf_transformation,[],[f10255_D])).
fof(f10255_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f10145,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f4722,f4739])).
fof(f10258,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))),
  inference(forward_demodulation,[],[f10139,f152])).
fof(f10139,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) | s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.LESS_0'))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f4592,f4739])).
fof(f10253,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))),
  inference(forward_demodulation,[],[f10138,f151])).
fof(f10138,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) | s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f4591,f4739])).
fof(f10064,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | $spl34),
  inference(cnf_transformation,[],[f10064_D])).
fof(f10064_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f10062,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl33),
  inference(cnf_transformation,[],[f10062_D])).
fof(f10062_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f9952,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f4722,f4736])).
fof(f10065,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f9946,f152])).
fof(f9946,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.LESS_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f4592,f4736])).
fof(f10060,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f9945,f151])).
fof(f9945,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f4591,f4736])).
fof(f9871,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) | $spl30),
  inference(cnf_transformation,[],[f9871_D])).
fof(f9871_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f9869,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) | $spl29),
  inference(cnf_transformation,[],[f9869_D])).
fof(f9869_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f9759,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f4722,f4735])).
fof(f9872,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))),
  inference(forward_demodulation,[],[f9753,f152])).
fof(f9753,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) | s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.LESS_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f4592,f4735])).
fof(f9867,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))),
  inference(forward_demodulation,[],[f9752,f151])).
fof(f9752,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) | s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f4591,f4735])).
fof(f7360,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl24),
  inference(cnf_transformation,[],[f7360_D])).
fof(f7360_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f7362,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | $spl26),
  inference(cnf_transformation,[],[f7362_D])).
fof(f7362_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f7351,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f181,f4727])).
fof(f7349,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f4727,f181])).
fof(f7321,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) | $spl20),
  inference(cnf_transformation,[],[f7321_D])).
fof(f7321_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f7323,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) | $spl22),
  inference(cnf_transformation,[],[f7323_D])).
fof(f7323_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f7312,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))),
  inference(superposition,[],[f181,f4726])).
fof(f7310,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))),
  inference(superposition,[],[f4726,f181])).
fof(f3494,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl16),
  inference(cnf_transformation,[],[f3494_D])).
fof(f3494_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f3496,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl18),
  inference(cnf_transformation,[],[f3496_D])).
fof(f3496_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f3488,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f181,f461])).
fof(f3486,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f461,f181])).
fof(f3216,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) | $spl12),
  inference(cnf_transformation,[],[f3216_D])).
fof(f3216_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0'))))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f3218,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl14),
  inference(cnf_transformation,[],[f3218_D])).
fof(f3218_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f3210,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))),
  inference(superposition,[],[f181,f460])).
fof(f3208,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prelim.ordering2num_1'(s('type.prelim.ordering','const.prelim.GREATER_0')))))),
  inference(superposition,[],[f460,f181])).
fof(f1071,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f1070,f149])).
fof(f1070,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(backward_demodulation,[],[f1069,f157])).
fof(f1069,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))) | $spl8),
  inference(backward_demodulation,[],[f1066,f755])).
fof(f1066,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl8),
  inference(cnf_transformation,[],[f1066_D])).
fof(f1066_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0'))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1068,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))) | $spl10),
  inference(cnf_transformation,[],[f1068_D])).
fof(f1068_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1060,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f181,f755])).
fof(f1058,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f755,f181])).
fof(f705,plain,(
  ( ! [X6] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))) ) | $spl4),
  inference(cnf_transformation,[],[f705_D])).
fof(f705_D,plain,(
  ( ! [X6] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f707,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl7),
  inference(cnf_transformation,[],[f707_D])).
fof(f707_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f692,plain,(
  ( ! [X6] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f190,f229])).
fof(f385,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl0),
  inference(cnf_transformation,[],[f385_D])).
fof(f385_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f387,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl3),
  inference(cnf_transformation,[],[f387_D])).
fof(f387_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f382,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f181])).
% SZS output end Proof for thm.prelim
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 465
% Passive clauses: 2775
% Generated clauses: 31710
% Final active clauses: 281
% Final passive clauses: 1199
% Input formulas: 63
% Initial clauses: 79
% 
% Duplicate literals: 244
% Trivial inequalities: 233
% Fw subsumption resolutions: 432
% Fw demodulations: 5856
% Bw demodulations: 1198
% Condensations: 296
% Global subsumptions: 2
% 
% Simple tautologies: 49
% Equational tautologies: 56
% Forward subsumptions: 21133
% Fw demodulations to eq. taut.: 523
% Bw demodulations to eq. taut.: 31
% 
% Unit resulting resolution: 1988
% Forward superposition: 10918
% Backward superposition: 9287
% Self superposition: 113
% Equality factoring: 36
% Equality resolution: 30
% 
% Split clauses: 50
% Split components: 100
% SAT solver clauses: 2940
% SAT solver unit clauses: 1392
% SAT solver binary clauses: 962
% 
% Sat splits: 50
% Sat splitting refutations: 61
% 
% Memory used [KB]: 20468
% Time elapsed: 0.898 s
% ------------------------------
% ------------------------------
% Success in time 2.218 s
