Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.ring
% (86108)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 161
% Passive clauses: 3300
% Generated clauses: 23871
% Final active clauses: 116
% Final passive clauses: 2918
% Input formulas: 68
% Initial clauses: 80
% 
% Duplicate literals: 758
% Trivial inequalities: 256
% Fw subsumption resolutions: 2997
% Fw demodulations: 4916
% Bw demodulations: 39
% 
% Simple tautologies: 255
% Equational tautologies: 76
% Forward subsumptions: 10478
% 
% Binary resolution: 484
% Factoring: 1
% Forward superposition: 7279
% Backward superposition: 6543
% Self superposition: 225
% Equality factoring: 160
% Equality resolution: 19
% 
% Split clauses: 411
% Split components: 822
% SAT solver clauses: 1006
% SAT solver unit clauses: 37
% SAT solver binary clauses: 150
% 
% Sat splits: 411
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 1004
% TWLsolver calls for satisfiability: 165
% 
% Memory used [KB]: 13688
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ring
% (86134)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 117
% Passive clauses: 1004
% Generated clauses: 14519
% Final active clauses: 99
% Final passive clauses: 792
% Input formulas: 68
% Initial clauses: 80
% 
% Duplicate literals: 570
% Trivial inequalities: 94
% Fw subsumption resolutions: 785
% Fw demodulations: 1737
% Bw demodulations: 112
% Condensations: 1135
% 
% Simple tautologies: 151
% Equational tautologies: 47
% Forward subsumptions: 6660
% Fw demodulations to eq. taut.: 1
% 
% Binary resolution: 253
% Factoring: 1
% Forward superposition: 4292
% Backward superposition: 4143
% Self superposition: 178
% Equality factoring: 98
% Equality resolution: 12
% 
% Split clauses: 56
% Split components: 112
% SAT solver clauses: 1767
% SAT solver unit clauses: 129
% SAT solver binary clauses: 37
% 
% Sat splits: 56
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1767
% TWLsolver calls for satisfiability: 249
% 
% Memory used [KB]: 17270
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.ring
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ring
% SZS output start Proof for thm.ring
fof(f21533,plain,(
  $false),
  inference(subsumption_resolution,[],[f21532,f217])).
fof(f217,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.ring.ring_fupdfupds_comp|split|1',aHL_TRUTH)).
fof(f21532,plain,(
  ~p(s(bool,'T_0'))),
  inference(forward_demodulation,[],[f21528,f3966])).
fof(f3966,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f3899,f302])).
fof(f302,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f219,f220])).
fof(f220,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.ring.ring_fupdfupds_comp|split|1',aHL_BOOL_CASES)).
fof(f219,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.ring.ring_fupdfupds_comp|split|1',aHL_FALSITY)).
fof(f3899,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X1))))) )),
  inference(backward_demodulation,[],[f3760,f1535])).
fof(f1535,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'const.combin.o_3'(s(fun(X0,fun(X0,bool)),'=_0'),s(fun(X0,X0),'_i_0_0'),s(X0,X1))),s(X0,X1))))) )),
  inference(forward_demodulation,[],[f1477,f272])).
fof(f272,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))) )),
  inference(cnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))),
  inference(flattening,[],[f144])).
fof(f144,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,'const.combin.o_3'(s(fun(X1,X0),X3),s(fun(X2,X1),X4),s(X2,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5)))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6,X7,X13] : ! [X10,X9,X15] : s(X6,'const.combin.o_3'(s(fun(X7,X6),X10),s(fun(X13,X7),X9),s(X13,X15))) = s(X6,happ(s(fun(X7,X6),X10),s(X7,happ(s(fun(X13,X7),X9),s(X13,X15)))))),
  file('thm.ring.ring_fupdfupds_comp|split|1','thm.combin.o_THM')).
fof(f1477,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1))))),s(X0,X1))))) )),
  inference(unit_resulting_resolution,[],[f221,f270])).
fof(f270,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f215])).
fof(f215,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X36,X37,X38] : (p(s(bool,happ(s(fun(X36,bool),happ(s(fun(X36,fun(X36,bool)),'=_0'),s(X36,X37))),s(X36,X38)))) <=> s(X36,X37) = s(X36,X38))),
  file('thm.ring.ring_fupdfupds_comp|split|1',arity_0_HL_EQ)).
fof(f221,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f73])).
fof(f73,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X23] : ! [X31] : s(X23,X31) = s(X23,happ(s(fun(X23,X23),'_i_0_0'),s(X23,X31)))),
  file('thm.ring.ring_fupdfupds_comp|split|1','_i0_ax')).
fof(f3760,plain,(
  ( ! [X2,X0,X3,X1] : (s(X2,happ(s(fun(X0,X2),X3),s(X0,X1))) = s(X2,'const.combin.o_3'(s(fun(X0,X2),X3),s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(superposition,[],[f272,f221])).
fof(f21528,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.ring.ring'(sK1)),bool),happ(s(fun(fun(sK0,'type.ring.ring'(sK1)),fun(fun(sK0,'type.ring.ring'(sK1)),bool)),'=_0'),s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK1,sK1),'const.combin.o_2'(s(fun(sK1,sK1),sK4),s(fun(sK1,sK1),sK3))))),s(fun(sK0,'type.ring.ring'(sK1)),sK2))))),s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK1,sK1),'const.combin.o_2'(s(fun(sK1,sK1),sK4),s(fun(sK1,sK1),sK3))))),s(fun(sK0,'type.ring.ring'(sK1)),sK2))))))),
  inference(backward_demodulation,[],[f21475,f4831])).
fof(f4831,plain,(
  ~p(s(bool,happ(s(fun(fun(sK0,'type.ring.ring'(sK1)),bool),happ(s(fun(fun(sK0,'type.ring.ring'(sK1)),fun(fun(sK0,'type.ring.ring'(sK1)),bool)),'=_0'),s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK1,sK1),sK4))),s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK1,sK1),sK3))),s(fun(sK0,'type.ring.ring'(sK1)),sK2))))))),s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK1,sK1),'const.combin.o_2'(s(fun(sK1,sK1),sK4),s(fun(sK1,sK1),sK3))))),s(fun(sK0,'type.ring.ring'(sK1)),sK2))))))),
  inference(unit_resulting_resolution,[],[f216,f269])).
fof(f269,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f215])).
fof(f216,plain,(
  s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK1,sK1),'const.combin.o_2'(s(fun(sK1,sK1),sK4),s(fun(sK1,sK1),sK3))))),s(fun(sK0,'type.ring.ring'(sK1)),sK2))) != s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK1,sK1),sK4))),s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK1,sK1),sK3))),s(fun(sK0,'type.ring.ring'(sK1)),sK2)))))),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK1,sK1),'const.combin.o_2'(s(fun(sK1,sK1),sK4),s(fun(sK1,sK1),sK3))))),s(fun(sK0,'type.ring.ring'(sK1)),sK2))) != s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK1,sK1),sK4))),s(fun(sK0,'type.ring.ring'(sK1)),'const.combin.o_2'(s(fun('type.ring.ring'(sK1),'type.ring.ring'(sK1)),'const.ring.ring_R|48|_fupd_1'(s(fun(sK1,sK1),sK3))),s(fun(sK0,'type.ring.ring'(sK1)),sK2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f192])).
fof(f192,plain,(
  ? [X0,X1,X2,X3,X4] : s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_R|48|_fupd_1'(s(fun(X1,X1),'const.combin.o_2'(s(fun(X1,X1),X4),s(fun(X1,X1),X3))))),s(fun(X0,'type.ring.ring'(X1)),X2))) != s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_R|48|_fupd_1'(s(fun(X1,X1),X4))),s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_R|48|_fupd_1'(s(fun(X1,X1),X3))),s(fun(X0,'type.ring.ring'(X1)),X2)))))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1,X2,X3,X4] : s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_R|48|_fupd_1'(s(fun(X1,X1),'const.combin.o_2'(s(fun(X1,X1),X4),s(fun(X1,X1),X3))))),s(fun(X0,'type.ring.ring'(X1)),X2))) = s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_R|48|_fupd_1'(s(fun(X1,X1),X4))),s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_R|48|_fupd_1'(s(fun(X1,X1),X3))),s(fun(X0,'type.ring.ring'(X1)),X2)))))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_R|48|_fupd_1'(s(fun(X1,X1),'const.combin.o_2'(s(fun(X1,X1),X4),s(fun(X1,X1),X3))))),s(fun(X0,'type.ring.ring'(X1)),X2))) = s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_R|48|_fupd_1'(s(fun(X1,X1),X4))),s(fun(X0,'type.ring.ring'(X1)),'const.combin.o_2'(s(fun('type.ring.ring'(X1),'type.ring.ring'(X1)),'const.ring.ring_R|48|_fupd_1'(s(fun(X1,X1),X3))),s(fun(X0,'type.ring.ring'(X1)),X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_R|48|_fupd_1'(s(fun(X7,X7),'const.combin.o_2'(s(fun(X7,X7),X10),s(fun(X7,X7),X9))))),s(fun(X6,'type.ring.ring'(X7)),X8))) = s(fun(X6,'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_R|48|_fupd_1'(s(fun(X7,X7),X10))),s(fun(X6,'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_R|48|_fupd_1'(s(fun(X7,X7),X9))),s(fun(X6,'type.ring.ring'(X7)),X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_R|48|_fupd_1'(s(fun(X7,X7),'const.combin.o_2'(s(fun(X7,X7),X10),s(fun(X7,X7),X9))))),s(fun(X6,'type.ring.ring'(X7)),X8))) = s(fun(X6,'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_R|48|_fupd_1'(s(fun(X7,X7),X10))),s(fun(X6,'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_R|48|_fupd_1'(s(fun(X7,X7),X9))),s(fun(X6,'type.ring.ring'(X7)),X8)))))),
  file('thm.ring.ring_fupdfupds_comp|split|1','thm.ring.ring_fupdfupds_comp|split|1')).
fof(f21475,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s(fun(X11,'type.ring.ring'(X8)),'const.combin.o_2'(s(fun('type.ring.ring'(X8),'type.ring.ring'(X8)),'const.ring.ring_R|48|_fupd_1'(s(fun(X8,X8),'const.combin.o_2'(s(fun(X8,X8),X9),s(fun(X8,X8),X10))))),s(fun(X11,'type.ring.ring'(X8)),X12))) = s(fun(X11,'type.ring.ring'(X8)),'const.combin.o_2'(s(fun('type.ring.ring'(X8),'type.ring.ring'(X8)),'const.ring.ring_R|48|_fupd_1'(s(fun(X8,X8),X9))),s(fun(X11,'type.ring.ring'(X8)),'const.combin.o_2'(s(fun('type.ring.ring'(X8),'type.ring.ring'(X8)),'const.ring.ring_R|48|_fupd_1'(s(fun(X8,X8),X10))),s(fun(X11,'type.ring.ring'(X8)),X12)))))) )),
  inference(superposition,[],[f275,f232])).
fof(f232,plain,(
  ( ! [X2,X0,X1] : (s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|48|_fupd_1'(s(fun(X0,X0),'const.combin.o_2'(s(fun(X0,X0),X2),s(fun(X0,X0),X1))))) = s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.combin.o_2'(s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|48|_fupd_1'(s(fun(X0,X0),X2))),s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|48|_fupd_1'(s(fun(X0,X0),X1)))))) )),
  inference(cnf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0,X1,X2] : s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|48|_fupd_1'(s(fun(X0,X0),'const.combin.o_2'(s(fun(X0,X0),X2),s(fun(X0,X0),X1))))) = s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.combin.o_2'(s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|48|_fupd_1'(s(fun(X0,X0),X2))),s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|48|_fupd_1'(s(fun(X0,X0),X1)))))),
  inference(flattening,[],[f90])).
fof(f90,plain,(
  ! [X0] : ! [X1,X2] : s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|48|_fupd_1'(s(fun(X0,X0),'const.combin.o_2'(s(fun(X0,X0),X2),s(fun(X0,X0),X1))))) = s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.combin.o_2'(s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|48|_fupd_1'(s(fun(X0,X0),X2))),s(fun('type.ring.ring'(X0),'type.ring.ring'(X0)),'const.ring.ring_R|48|_fupd_1'(s(fun(X0,X0),X1)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7] : ! [X9,X10] : s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_R|48|_fupd_1'(s(fun(X7,X7),'const.combin.o_2'(s(fun(X7,X7),X10),s(fun(X7,X7),X9))))) = s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.combin.o_2'(s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_R|48|_fupd_1'(s(fun(X7,X7),X10))),s(fun('type.ring.ring'(X7),'type.ring.ring'(X7)),'const.ring.ring_R|48|_fupd_1'(s(fun(X7,X7),X9)))))),
  file('thm.ring.ring_fupdfupds_comp|split|1','thm.ring.ring_fupdfupds_comp|split|0')).
fof(f275,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))) )),
  inference(cnf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(flattening,[],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6] : s(fun(X2,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X2,X1),'const.combin.o_2'(s(fun(X3,X1),X5),s(fun(X2,X3),X6))))) = s(fun(X2,X0),'const.combin.o_2'(s(fun(X3,X0),'const.combin.o_2'(s(fun(X1,X0),X4),s(fun(X3,X1),X5))),s(fun(X2,X3),X6)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6,X7,X17,X13] : ! [X10,X9,X8] : s(fun(X17,X6),'const.combin.o_2'(s(fun(X7,X6),X10),s(fun(X17,X7),'const.combin.o_2'(s(fun(X13,X7),X9),s(fun(X17,X13),X8))))) = s(fun(X17,X6),'const.combin.o_2'(s(fun(X13,X6),'const.combin.o_2'(s(fun(X7,X6),X10),s(fun(X13,X7),X9))),s(fun(X17,X13),X8)))),
  file('thm.ring.ring_fupdfupds_comp|split|1','thm.combin.o_ASSOC')).
% SZS output end Proof for thm.ring
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 209
% Passive clauses: 2996
% Generated clauses: 26511
% Final active clauses: 69
% Final passive clauses: 625
% Input formulas: 68
% Initial clauses: 80
% 
% Duplicate literals: 717
% Trivial inequalities: 106
% Fw subsumption resolutions: 1190
% Fw demodulations: 4226
% Bw demodulations: 1956
% Condensations: 1319
% Global subsumptions: 6
% 
% Simple tautologies: 132
% Equational tautologies: 291
% Forward subsumptions: 15052
% Fw demodulations to eq. taut.: 21
% Bw demodulations to eq. taut.: 30
% 
% Unit resulting resolution: 1312
% Factoring: 1
% Forward superposition: 5588
% Backward superposition: 4296
% Self superposition: 173
% Equality factoring: 112
% Equality resolution: 9
% 
% Split clauses: 31
% Split components: 62
% SAT solver clauses: 3445
% SAT solver unit clauses: 721
% SAT solver binary clauses: 1407
% 
% Sat splits: 31
% Sat splitting refutations: 379
% 
% Memory used [KB]: 14456
% Time elapsed: 1.036 s
% ------------------------------
% ------------------------------
% Success in time 2.361 s
