Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.integer
% (113070)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 114
% Passive clauses: 2520
% Generated clauses: 17713
% Final active clauses: 111
% Final passive clauses: 2292
% Input formulas: 42
% Initial clauses: 54
% 
% Duplicate literals: 825
% Trivial inequalities: 209
% Fw subsumption resolutions: 2205
% Fw demodulations: 2917
% Bw demodulations: 1
% 
% Simple tautologies: 323
% Equational tautologies: 28
% Forward subsumptions: 7963
% 
% Binary resolution: 20
% Forward superposition: 6927
% Backward superposition: 4108
% Self superposition: 219
% Equality factoring: 86
% Equality resolution: 45
% 
% Split clauses: 462
% Split components: 924
% SAT solver clauses: 764
% SAT solver unit clauses: 56
% SAT solver binary clauses: 212
% 
% Sat splits: 462
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 764
% TWLsolver calls for satisfiability: 108
% 
% Memory used [KB]: 13176
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.integer
% (113091)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 109
% Passive clauses: 1081
% Generated clauses: 8865
% Final active clauses: 88
% Final passive clauses: 874
% Input formulas: 42
% Initial clauses: 54
% 
% Duplicate literals: 258
% Trivial inequalities: 111
% Fw subsumption resolutions: 263
% Fw demodulations: 1844
% Bw demodulations: 113
% Condensations: 374
% 
% Simple tautologies: 76
% Equational tautologies: 22
% Forward subsumptions: 3219
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 19
% Forward superposition: 3037
% Backward superposition: 1550
% Self superposition: 121
% Equality factoring: 46
% Equality resolution: 12
% 
% Split clauses: 19
% Split components: 38
% SAT solver clauses: 1615
% SAT solver unit clauses: 67
% SAT solver binary clauses: 44
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 19
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 1614
% TWLsolver calls for satisfiability: 327
% 
% Memory used [KB]: 17526
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.integer
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.integer
% SZS output start Proof for thm.integer
fof(f58486,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f139,f196,f194,f138,f195,f193,f143,f200,f198,f207,f219,f222,f220,f724,f2810,f2751,f2813,f2830,f4394,f4412,f4410,f14729,f14735,f14752,f14754,f14764,f14584,f14731,f14826,f17702,f17708,f17709,f17742,f17750,f17545,f17704,f17805,f20412,f20416,f20414,f21422,f21423,f21424,f21425,f35345,f35351,f35352,f35389,f35400,f35401,f35198,f35347,f35463,f36697,f36703,f36704,f36742,f36754,f36755,f36546,f36699,f36817,f38039,f38045,f38046,f38091,f38113,f38114,f37873,f38041,f38202,f39387,f39393,f39394,f39440,f39462,f39463,f39389,f2811,f39569,f39570,f39571,f39572,f39573,f39590,f39574,f39575,f39576,f39597,f39577,f39599,f39578,f39600,f39579,f39601,f39580,f39593,f39603,f39581,f39604,f39586,f39582,f39592,f39605,f39583,f39584,f39596,f39587,f39595,f40250,f40253,f40310,f40254,f40256,f40315,f40257,f40316,f40259,f40317,f40260,f40318,f40262,f40312,f40319,f40320,f40314,f40267,f40268,f40321,f40322,f40323,f40324,f39589,f40313,f58485,f58458])).
fof(f58458,plain,(
  $false | ($spl0 | $spl3)),
  inference(unit_resulting_resolution,[],[f193,f58289])).
fof(f58289,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK0),s('type.integer.int',X0))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',X0)))) ) | $spl3),
  inference(superposition,[],[f58081,f38235])).
fof(f38235,plain,(
  ( ! [X2,X3] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X2),s('type.integer.int',X3))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X3),s('type.integer.int',X2)))) )),
  inference(superposition,[],[f1660,f158])).
fof(f158,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_ABS_1'(s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.tint_add_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.int_REP_1'(s('type.integer.int',X0))),s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.int_REP_1'(s('type.integer.int',X1)))))))) )),
  inference(cnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1] : s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_ABS_1'(s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.tint_add_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.int_REP_1'(s('type.integer.int',X0))),s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.int_REP_1'(s('type.integer.int',X1)))))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X19,X20] : s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X19),s('type.integer.int',X20))) = s('type.integer.int','const.integer.int_ABS_1'(s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.tint_add_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.int_REP_1'(s('type.integer.int',X19))),s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.int_REP_1'(s('type.integer.int',X20)))))))),
  file('thm.integer.INT_EQ_RADD','thm.integer.int_add')).
fof(f1660,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_ABS_1'(s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.tint_add_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.int_REP_1'(s('type.integer.int',X1))),s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.int_REP_1'(s('type.integer.int',X0)))))))) )),
  inference(superposition,[],[f158,f157])).
fof(f157,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.tint_add_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X1))) = s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.tint_add_2'(s('type.pair.prod'('type.num.num','type.num.num'),X1),s('type.pair.prod'('type.num.num','type.num.num'),X0)))) )),
  inference(cnf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0,X1] : s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.tint_add_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X1))) = s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.tint_add_2'(s('type.pair.prod'('type.num.num','type.num.num'),X1),s('type.pair.prod'('type.num.num','type.num.num'),X0)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6,X7] : s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.tint_add_2'(s('type.pair.prod'('type.num.num','type.num.num'),X6),s('type.pair.prod'('type.num.num','type.num.num'),X7))) = s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.tint_add_2'(s('type.pair.prod'('type.num.num','type.num.num'),X7),s('type.pair.prod'('type.num.num','type.num.num'),X6)))),
  file('thm.integer.INT_EQ_RADD','thm.integer.TINT_ADD_SYM')).
fof(f58081,plain,(
  ( ! [X125] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X125),s('type.integer.int',sK0))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',X125)))) ) | $spl3),
  inference(superposition,[],[f769,f38235])).
fof(f769,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',sK0))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',sK1)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f196,f173])).
fof(f173,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X2))) | s('type.integer.int',X1) = s('type.integer.int',X2)) )),
  inference(cnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1,X2] : ((s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X2))) | s('type.integer.int',X1) = s('type.integer.int',X2)) & (s('type.integer.int',X1) != s('type.integer.int',X2) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X2)))))),
  inference(nnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0,X1,X2] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X2))) <=> s('type.integer.int',X1) = s('type.integer.int',X2))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6,X7,X8] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X6),s('type.integer.int',X7))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X6),s('type.integer.int',X8))) <=> s('type.integer.int',X7) = s('type.integer.int',X8))),
  file('thm.integer.INT_EQ_RADD','thm.integer.INT_EQ_LADD')).
fof(f58485,plain,(
  $false | ($spl0 | $spl3)),
  inference(trivial_inequality_removal,[],[f58480])).
fof(f58480,plain,(
  s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | ($spl0 | $spl3)),
  inference(superposition,[],[f58289,f193])).
fof(f40313,plain,(
  ( ! [X14,X15] : (s(fun(bool,bool),X14) = s(fun(bool,bool),X15) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,sK5(X14,X15,bool,bool)))) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X14,X15,bool,bool))))) ) | $spl80),
  inference(cnf_transformation,[],[f40313_D])).
fof(f40313_D,plain,(
  ( ! [X14,X15] : (s(fun(bool,bool),X14) = s(fun(bool,bool),X15) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,sK5(X14,X15,bool,bool)))) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X14,X15,bool,bool))))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f39589,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl51),
  inference(cnf_transformation,[],[f39589_D])).
fof(f39589_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f40324,plain,(
  ( ! [X70,X69] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X69) = s(fun(bool,bool),X70) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X69,X70,bool,bool)))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,sK5(X69,X70,bool,bool))))) )),
  inference(forward_demodulation,[],[f40274,f201])).
fof(f201,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(forward_demodulation,[],[f151,f149])).
fof(f149,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X9] : ! [X6] : s(X9,X6) = s(X9,'const.combin.I_1'(s(X9,X6)))),
  file('thm.integer.INT_EQ_RADD','thm.combin.I_THM')).
fof(f151,plain,(
  ( ! [X0,X1] : (s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ! [X0] : ! [X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X9] : ! [X36] : s(X9,'const.combin.I_1'(s(X9,X36))) = s(X9,happ(s(fun(X9,X9),'const.combin.I_0'),s(X9,X36)))),
  file('thm.integer.INT_EQ_RADD','ARITY_0_1_const.combin.I')).
fof(f40274,plain,(
  ( ! [X70,X69] : (s(bool,'F_0') != s(bool,f(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))))) | s(fun(bool,bool),X69) = s(fun(bool,bool),X70) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X69,X70,bool,bool)))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,sK5(X69,X70,bool,bool))))) )),
  inference(superposition,[],[f1004,f17627])).
fof(f17627,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,X0) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0')))) )),
  inference(condensation,[],[f17626])).
fof(f17626,plain,(
  ( ! [X12,X11] : (s(bool,X12) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X12)))) )),
  inference(duplicate_literal_removal,[],[f17215])).
fof(f17215,plain,(
  ( ! [X12,X11] : (s(bool,X12) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X12)))) )),
  inference(superposition,[],[f16945,f288])).
fof(f288,plain,(
  ( ! [X23,X22] : (s(bool,X22) = s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X22)))) )),
  inference(superposition,[],[f210,f149])).
fof(f210,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f148,f148])).
fof(f148,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.integer.INT_EQ_RADD',aHL_BOOL_CASES)).
fof(f16945,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(condensation,[],[f16944])).
fof(f16944,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f16491])).
fof(f16491,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f315,f148])).
fof(f315,plain,(
  ( ! [X19,X20] : (s(bool,X20) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X19))) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f210,f201])).
fof(f1004,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,f(s(X16,sK5(X17,X18,bool,X16)))) | s(fun(X16,bool),X17) = s(fun(X16,bool),X18) | s(bool,'T_0') = s(bool,g(s(X16,sK5(X17,X18,bool,X16))))) )),
  inference(superposition,[],[f162,f148])).
fof(f162,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK5(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK5(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK5(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK5(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f109])).
fof(f109,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f67])).
fof(f67,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.integer.INT_EQ_RADD',aHL_EXT)).
fof(f40323,plain,(
  ( ! [X66,X65] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X65) = s(fun(bool,bool),X66) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X65,X66,bool,bool)))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,sK5(X65,X66,bool,bool))))) )),
  inference(forward_demodulation,[],[f40272,f150])).
fof(f150,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X38] : ! [X39] : s(X38,X39) = s(X38,happ(s(fun(X38,X38),'_i_0_0'),s(X38,X39)))),
  file('thm.integer.INT_EQ_RADD','_i0_ax')).
fof(f40272,plain,(
  ( ! [X66,X65] : (s(bool,'F_0') != s(bool,f(s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))))) | s(fun(bool,bool),X65) = s(fun(bool,bool),X66) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X65,X66,bool,bool)))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,sK5(X65,X66,bool,bool))))) )),
  inference(superposition,[],[f1004,f14657])).
fof(f14657,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0')))) )),
  inference(condensation,[],[f14656])).
fof(f14656,plain,(
  ( ! [X12,X11] : (s(bool,X12) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X12)))) )),
  inference(duplicate_literal_removal,[],[f14279])).
fof(f14279,plain,(
  ( ! [X12,X11] : (s(bool,X12) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X12)))) )),
  inference(superposition,[],[f13930,f288])).
fof(f13930,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(condensation,[],[f13929])).
fof(f13929,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f13514])).
fof(f13514,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f314,f148])).
fof(f314,plain,(
  ( ! [X17,X18] : (s(bool,X18) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X17))) | s(bool,'T_0') = s(bool,X17) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f210,f150])).
fof(f40322,plain,(
  ( ! [X61,X62] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X61) = s(fun(bool,bool),X62) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X61,X62,bool,bool)))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,sK5(X61,X62,bool,bool))))) )),
  inference(forward_demodulation,[],[f40270,f149])).
fof(f40270,plain,(
  ( ! [X61,X62] : (s(bool,'F_0') != s(bool,f(s(bool,'const.combin.I_1'(s(bool,'F_0'))))) | s(fun(bool,bool),X61) = s(fun(bool,bool),X62) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X61,X62,bool,bool)))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,sK5(X61,X62,bool,bool))))) )),
  inference(superposition,[],[f1004,f2787])).
fof(f2787,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,X0) = s(bool,'const.combin.I_1'(s(bool,'F_0')))) )),
  inference(condensation,[],[f2786])).
fof(f2786,plain,(
  ( ! [X12,X11] : (s(bool,X12) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X12)))) )),
  inference(duplicate_literal_removal,[],[f2624])).
fof(f2624,plain,(
  ( ! [X12,X11] : (s(bool,X12) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X12)))) )),
  inference(superposition,[],[f2318,f288])).
fof(f2318,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(condensation,[],[f2317])).
fof(f2317,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f2146])).
fof(f2146,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f313,f148])).
fof(f313,plain,(
  ( ! [X15,X16] : (s(bool,X16) = s(bool,'const.combin.I_1'(s(bool,X15))) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f210,f149])).
fof(f40321,plain,(
  ( ! [X55] : (s(bool,'F_0') != s(bool,f(s(bool,X55))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X55)))) ) | $spl94),
  inference(cnf_transformation,[],[f40321_D])).
fof(f40321_D,plain,(
  ( ! [X55] : (s(bool,'F_0') != s(bool,f(s(bool,X55))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X55)))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f40268,plain,(
  ( ! [X57,X58,X56] : (s(bool,'F_0') != s(bool,f(s(bool,X58))) | s(fun(bool,bool),X56) = s(fun(bool,bool),X57) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X56,X57,bool,bool)))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X58))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,sK5(X56,X57,bool,bool))))) )),
  inference(superposition,[],[f1004,f1392])).
fof(f1392,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X1)))) )),
  inference(condensation,[],[f1391])).
fof(f1391,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X20))) | s(bool,X20) = s(bool,X22) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X22)))) )),
  inference(duplicate_literal_removal,[],[f1037])).
fof(f1037,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X20))) | s(bool,X20) = s(bool,X22) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X22)))) )),
  inference(superposition,[],[f288,f288])).
fof(f40267,plain,(
  ( ! [X54,X55,X53] : (s(bool,'F_0') != s(bool,f(s(bool,X55))) | s(fun(bool,bool),X53) = s(fun(bool,bool),X54) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X53,X54,bool,bool)))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,sK5(X53,X54,bool,bool)))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X55)))) )),
  inference(superposition,[],[f1004,f1392])).
fof(f40314,plain,(
  ( ! [X16] : (s(bool,'F_0') != s(bool,f(s(bool,X16))) | s(bool,'T_0') = s(bool,X16)) ) | $spl82),
  inference(cnf_transformation,[],[f40314_D])).
fof(f40314_D,plain,(
  ( ! [X16] : (s(bool,'F_0') != s(bool,f(s(bool,X16))) | s(bool,'T_0') = s(bool,X16)) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f40320,plain,(
  ( ! [X52,X50,X51] : (s(bool,'F_0') != s(bool,f(s(bool,X52))) | s(fun(bool,bool),X50) = s(fun(bool,bool),X51) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X50,X51,bool,bool)))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,sK5(X50,X51,bool,bool)))) | s(bool,'T_0') = s(bool,X52)) )),
  inference(forward_demodulation,[],[f40266,f149])).
fof(f40266,plain,(
  ( ! [X52,X50,X51] : (s(bool,'F_0') != s(bool,f(s(bool,'const.combin.I_1'(s(bool,X52))))) | s(fun(bool,bool),X50) = s(fun(bool,bool),X51) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X50,X51,bool,bool)))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,sK5(X50,X51,bool,bool)))) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f1004,f1094])).
fof(f1094,plain,(
  ( ! [X33,X32] : (s(bool,X33) = s(bool,'const.combin.I_1'(s(bool,X32))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X33))) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f288,f149])).
fof(f40319,plain,(
  ( ! [X39,X38,X40] : (s(fun(bool,bool),X38) = s(fun(bool,bool),X39) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X40))) | s(bool,X40) = s(bool,sK5(X38,X39,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X38,X39,bool,bool))))) ) | $spl92),
  inference(cnf_transformation,[],[f40319_D])).
fof(f40319_D,plain,(
  ( ! [X39,X38,X40] : (s(fun(bool,bool),X38) = s(fun(bool,bool),X39) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X40))) | s(bool,X40) = s(bool,sK5(X38,X39,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X38,X39,bool,bool))))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f40312,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl79),
  inference(cnf_transformation,[],[f40312_D])).
fof(f40312_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl79),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl79])])).
fof(f40262,plain,(
  ( ! [X39,X38,X40] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X38) = s(fun(bool,bool),X39) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X38,X39,bool,bool)))) | s(bool,X40) = s(bool,sK5(X38,X39,bool,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X40)))) )),
  inference(superposition,[],[f1004,f290])).
fof(f290,plain,(
  ( ! [X26,X27] : (s(bool,X26) = s(bool,X27) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X26)))) )),
  inference(superposition,[],[f210,f201])).
fof(f40318,plain,(
  ( ! [X33,X32] : (s(fun(bool,bool),X32) = s(fun(bool,bool),X33) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,sK5(X32,X33,bool,bool)))) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X32,X33,bool,bool))))) ) | $spl90),
  inference(cnf_transformation,[],[f40318_D])).
fof(f40318_D,plain,(
  ( ! [X33,X32] : (s(fun(bool,bool),X32) = s(fun(bool,bool),X33) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,sK5(X32,X33,bool,bool)))) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X32,X33,bool,bool))))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f40260,plain,(
  ( ! [X33,X34,X32] : (s(bool,'F_0') != s(bool,f(s(bool,X34))) | s(fun(bool,bool),X32) = s(fun(bool,bool),X33) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X32,X33,bool,bool)))) | s(bool,'T_0') = s(bool,X34) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,sK5(X32,X33,bool,bool))))) )),
  inference(superposition,[],[f1004,f290])).
fof(f40317,plain,(
  ( ! [X30,X31,X29] : (s(fun(bool,bool),X29) = s(fun(bool,bool),X30) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X31))) | s(bool,X31) = s(bool,sK5(X29,X30,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X29,X30,bool,bool))))) ) | $spl88),
  inference(cnf_transformation,[],[f40317_D])).
fof(f40317_D,plain,(
  ( ! [X30,X31,X29] : (s(fun(bool,bool),X29) = s(fun(bool,bool),X30) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X31))) | s(bool,X31) = s(bool,sK5(X29,X30,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X29,X30,bool,bool))))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f40259,plain,(
  ( ! [X30,X31,X29] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X29) = s(fun(bool,bool),X30) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X29,X30,bool,bool)))) | s(bool,X31) = s(bool,sK5(X29,X30,bool,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X31)))) )),
  inference(superposition,[],[f1004,f289])).
fof(f289,plain,(
  ( ! [X24,X25] : (s(bool,X24) = s(bool,X25) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X24)))) )),
  inference(superposition,[],[f210,f150])).
fof(f40316,plain,(
  ( ! [X24,X23] : (s(fun(bool,bool),X23) = s(fun(bool,bool),X24) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK5(X23,X24,bool,bool)))) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X23,X24,bool,bool))))) ) | $spl86),
  inference(cnf_transformation,[],[f40316_D])).
fof(f40316_D,plain,(
  ( ! [X24,X23] : (s(fun(bool,bool),X23) = s(fun(bool,bool),X24) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK5(X23,X24,bool,bool)))) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X23,X24,bool,bool))))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f40257,plain,(
  ( ! [X24,X23,X25] : (s(bool,'F_0') != s(bool,f(s(bool,X25))) | s(fun(bool,bool),X23) = s(fun(bool,bool),X24) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X23,X24,bool,bool)))) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,sK5(X23,X24,bool,bool))))) )),
  inference(superposition,[],[f1004,f289])).
fof(f40315,plain,(
  ( ! [X21,X22,X20] : (s(fun(bool,bool),X20) = s(fun(bool,bool),X21) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X22))) | s(bool,X22) = s(bool,sK5(X20,X21,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X20,X21,bool,bool))))) ) | $spl84),
  inference(cnf_transformation,[],[f40315_D])).
fof(f40315_D,plain,(
  ( ! [X21,X22,X20] : (s(fun(bool,bool),X20) = s(fun(bool,bool),X21) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X22))) | s(bool,X22) = s(bool,sK5(X20,X21,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X20,X21,bool,bool))))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f40256,plain,(
  ( ! [X21,X22,X20] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X20) = s(fun(bool,bool),X21) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X20,X21,bool,bool)))) | s(bool,X22) = s(bool,sK5(X20,X21,bool,bool)) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X22)))) )),
  inference(superposition,[],[f1004,f288])).
fof(f40254,plain,(
  ( ! [X14,X15,X16] : (s(bool,'F_0') != s(bool,f(s(bool,X16))) | s(fun(bool,bool),X14) = s(fun(bool,bool),X15) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X14,X15,bool,bool)))) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,sK5(X14,X15,bool,bool))))) )),
  inference(superposition,[],[f1004,f288])).
fof(f40310,plain,(
  ( ! [X4,X2,X3] : (s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,sK5(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X2,X3,bool,bool))))) ) | $spl76),
  inference(cnf_transformation,[],[f40310_D])).
fof(f40310_D,plain,(
  ( ! [X4,X2,X3] : (s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,sK5(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X2,X3,bool,bool))))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f40253,plain,(
  ( ! [X12,X13,X11] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X11,X12,bool,bool)))) | s(bool,X13) = s(bool,sK5(X11,X12,bool,bool)) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f1004,f210])).
fof(f40250,plain,(
  ( ! [X4,X2,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X2,X3,bool,bool)))) | s(bool,X4) = s(bool,sK5(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f1004,f210])).
fof(f39595,plain,(
  s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,g(s(bool,'F_0'))))) | $spl58),
  inference(cnf_transformation,[],[f39595_D])).
fof(f39595_D,plain,(
  s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,g(s(bool,'F_0'))))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f39587,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK5(X0,X1,bool,bool))) ) | $spl48),
  inference(cnf_transformation,[],[f39587_D])).
fof(f39587_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK5(X0,X1,bool,bool))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f39596,plain,(
  ( ! [X17] : (s(bool,X17) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X17)))) ) | $spl60),
  inference(cnf_transformation,[],[f39596_D])).
fof(f39596_D,plain,(
  ( ! [X17] : (s(bool,X17) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X17)))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f39584,plain,(
  ( ! [X50,X48,X49] : (s(bool,X48) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X49) = s(fun(bool,bool),X50) | s(bool,'T_0') = s(bool,sK5(X49,X50,bool,bool)) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X48))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,g(s(bool,'F_0')))))) )),
  inference(superposition,[],[f999,f1392])).
fof(f999,plain,(
  ( ! [X17,X18,X16] : (s(X18,f(s(bool,'F_0'))) != s(X18,g(s(bool,'F_0'))) | s(fun(bool,X18),X16) = s(fun(bool,X18),X17) | s(bool,'T_0') = s(bool,sK5(X16,X17,X18,bool))) )),
  inference(superposition,[],[f162,f148])).
fof(f39583,plain,(
  ( ! [X47,X45,X46] : (s(bool,X45) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X46) = s(fun(bool,bool),X47) | s(bool,'T_0') = s(bool,sK5(X46,X47,bool,bool)) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,g(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X45)))) )),
  inference(superposition,[],[f999,f1392])).
fof(f39605,plain,(
  ( ! [X40] : (s(bool,X40) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X40)))) ) | $spl74),
  inference(cnf_transformation,[],[f39605_D])).
fof(f39605_D,plain,(
  ( ! [X40] : (s(bool,X40) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X40)))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f39592,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl55),
  inference(cnf_transformation,[],[f39592_D])).
fof(f39592_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl55),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl55])])).
fof(f39582,plain,(
  ( ! [X39,X38,X40] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X38) = s(fun(bool,bool),X39) | s(bool,'T_0') = s(bool,sK5(X38,X39,bool,bool)) | s(bool,X40) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X40)))) )),
  inference(superposition,[],[f999,f290])).
fof(f39586,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl46),
  inference(cnf_transformation,[],[f39586_D])).
fof(f39586_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f39604,plain,(
  ( ! [X35] : (s(bool,X35) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X35)))) ) | $spl72),
  inference(cnf_transformation,[],[f39604_D])).
fof(f39604_D,plain,(
  ( ! [X35] : (s(bool,X35) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X35)))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f39581,plain,(
  ( ! [X37,X35,X36] : (s(bool,X35) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X36) = s(fun(bool,bool),X37) | s(bool,'T_0') = s(bool,sK5(X36,X37,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X35)))) )),
  inference(superposition,[],[f999,f290])).
fof(f39603,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,g(s(bool,'F_0'))))) | $spl70),
  inference(cnf_transformation,[],[f39603_D])).
fof(f39603_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,g(s(bool,'F_0'))))) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f39593,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) ) | $spl56),
  inference(cnf_transformation,[],[f39593_D])).
fof(f39593_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f39580,plain,(
  ( ! [X33,X34,X32] : (s(bool,X32) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X33) = s(fun(bool,bool),X34) | s(bool,'T_0') = s(bool,sK5(X33,X34,bool,bool)) | s(bool,'T_0') = s(bool,X32) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,g(s(bool,'F_0')))))) )),
  inference(superposition,[],[f999,f290])).
fof(f39601,plain,(
  ( ! [X31] : (s(bool,X31) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X31)))) ) | $spl68),
  inference(cnf_transformation,[],[f39601_D])).
fof(f39601_D,plain,(
  ( ! [X31] : (s(bool,X31) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X31)))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f39579,plain,(
  ( ! [X30,X31,X29] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X29) = s(fun(bool,bool),X30) | s(bool,'T_0') = s(bool,sK5(X29,X30,bool,bool)) | s(bool,X31) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X31)))) )),
  inference(superposition,[],[f999,f289])).
fof(f39600,plain,(
  ( ! [X26] : (s(bool,X26) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X26)))) ) | $spl66),
  inference(cnf_transformation,[],[f39600_D])).
fof(f39600_D,plain,(
  ( ! [X26] : (s(bool,X26) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X26)))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f39578,plain,(
  ( ! [X28,X26,X27] : (s(bool,X26) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X27) = s(fun(bool,bool),X28) | s(bool,'T_0') = s(bool,sK5(X27,X28,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X26)))) )),
  inference(superposition,[],[f999,f289])).
fof(f39599,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,g(s(bool,'F_0'))))) | $spl64),
  inference(cnf_transformation,[],[f39599_D])).
fof(f39599_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,g(s(bool,'F_0'))))) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f39577,plain,(
  ( ! [X24,X23,X25] : (s(bool,X23) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X24) = s(fun(bool,bool),X25) | s(bool,'T_0') = s(bool,sK5(X24,X25,bool,bool)) | s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,g(s(bool,'F_0')))))) )),
  inference(superposition,[],[f999,f289])).
fof(f39597,plain,(
  ( ! [X22] : (s(bool,X22) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X22)))) ) | $spl62),
  inference(cnf_transformation,[],[f39597_D])).
fof(f39597_D,plain,(
  ( ! [X22] : (s(bool,X22) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X22)))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f39576,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X20) = s(fun(bool,bool),X21) | s(bool,'T_0') = s(bool,sK5(X20,X21,bool,bool)) | s(bool,X22) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X22)))) )),
  inference(superposition,[],[f999,f288])).
fof(f39575,plain,(
  ( ! [X19,X17,X18] : (s(bool,X17) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X18) = s(fun(bool,bool),X19) | s(bool,'T_0') = s(bool,sK5(X18,X19,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X17)))) )),
  inference(superposition,[],[f999,f288])).
fof(f39574,plain,(
  ( ! [X14,X15,X16] : (s(bool,X14) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X15) = s(fun(bool,bool),X16) | s(bool,'T_0') = s(bool,sK5(X15,X16,bool,bool)) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,g(s(bool,'F_0')))))) )),
  inference(superposition,[],[f999,f288])).
fof(f39590,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) ) | $spl52),
  inference(cnf_transformation,[],[f39590_D])).
fof(f39590_D,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f39573,plain,(
  ( ! [X12,X13,X11] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,'T_0') = s(bool,sK5(X11,X12,bool,bool)) | s(bool,X13) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f999,f210])).
fof(f39572,plain,(
  ( ! [X10,X8,X9] : (s(bool,X8) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X9) = s(fun(bool,bool),X10) | s(bool,'T_0') = s(bool,sK5(X9,X10,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f999,f210])).
fof(f39571,plain,(
  ( ! [X6,X7,X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7) | s(bool,'T_0') = s(bool,sK5(X6,X7,bool,bool)) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f999,f210])).
fof(f39570,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,sK5(X2,X3,bool,bool)) | s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f999,f210])).
fof(f39569,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK5(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f999,f148])).
fof(f2811,plain,(
  ( ! [X30] : (s(bool,'T_0') = s(bool,X30)) ) | $spl12),
  inference(cnf_transformation,[],[f2811_D])).
fof(f2811_D,plain,(
  ( ! [X30] : (s(bool,'T_0') = s(bool,X30)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f39389,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK1))) | $spl45),
  inference(cnf_transformation,[],[f39389_D])).
fof(f39389_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK1))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f39463,plain,(
  ( ! [X245] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X245)) ) | ($spl3 | $spl11 | $spl27 | $spl33 | $spl37)),
  inference(subsumption_resolution,[],[f39076,f17081])).
fof(f17081,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(forward_demodulation,[],[f16876,f149])).
fof(f16876,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X0)))) )),
  inference(condensation,[],[f16875])).
fof(f16875,plain,(
  ( ! [X92,X93] : (s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X93))) | s(bool,'T_0') = s(bool,X92) | s(bool,'T_0') = s(bool,X93)) )),
  inference(duplicate_literal_removal,[],[f16621])).
fof(f16621,plain,(
  ( ! [X92,X93] : (s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X93))) | s(bool,'T_0') = s(bool,X92) | s(bool,'T_0') = s(bool,X93) | s(bool,'T_0') = s(bool,X92)) )),
  inference(superposition,[],[f1270,f315])).
fof(f1270,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f1252])).
fof(f1252,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f288])).
fof(f39076,plain,(
  ( ! [X245] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X245) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X245)))) ) | ($spl3 | $spl11 | $spl27 | $spl33 | $spl37)),
  inference(superposition,[],[f17081,f37485])).
fof(f37485,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl3 | $spl11 | $spl27 | $spl33 | $spl37)),
  inference(subsumption_resolution,[],[f37330,f222])).
fof(f37330,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X0) = s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl3 | $spl27 | $spl33 | $spl37)),
  inference(superposition,[],[f37169,f210])).
fof(f37169,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK1))) | ($spl3 | $spl27 | $spl33 | $spl37)),
  inference(forward_demodulation,[],[f36926,f34188])).
fof(f34188,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK1))) | ($spl3 | $spl27)),
  inference(forward_demodulation,[],[f34161,f201])).
fof(f34161,plain,(
  s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) = s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK1))) | ($spl3 | $spl27)),
  inference(unit_resulting_resolution,[],[f17704,f33704,f357])).
fof(f357,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f140,f210])).
fof(f140,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.integer.INT_EQ_RADD',aHL_TRUTH)).
fof(f33704,plain,(
  ~p(s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK1)))) | $spl3),
  inference(unit_resulting_resolution,[],[f196,f7816])).
fof(f7816,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(forward_demodulation,[],[f7815,f201])).
fof(f7815,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1))))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(forward_demodulation,[],[f7814,f201])).
fof(f7814,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1))))),s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X2)))))) | s(X0,X1) = s(X0,X2)) )),
  inference(subsumption_resolution,[],[f7789,f140])).
fof(f7789,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1))))),s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X2)))))) | s(X0,X1) = s(X0,X2)) )),
  inference(superposition,[],[f167,f256])).
fof(f256,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(fun(X0,X0),'const.combin.I_0'),s(fun(X0,X0),'const.combin.I_0')))) )),
  inference(unit_resulting_resolution,[],[f145,f213])).
fof(f213,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f144,f148])).
fof(f144,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f46])).
fof(f46,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.integer.INT_EQ_RADD',aHL_FALSITY)).
fof(f145,plain,(
  ( ! [X0] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(fun(X0,X0),'const.combin.I_0'),s(fun(X0,X0),'const.combin.I_0'))))) )),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0] : p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(fun(X0,X0),'const.combin.I_0'),s(fun(X0,X0),'const.combin.I_0'))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X9] : p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X9,fun(X9,bool)),'|dollar|equals_0'),s(fun(X9,X9),'const.combin.I_0'),s(fun(X9,X9),'const.combin.I_0'))))),
  file('thm.integer.INT_EQ_RADD','thm.quotient.IDENTITY_QUOTIENT')).
fof(f167,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),X2),s(fun(X0,X1),X3),s(fun(X1,X0),X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))),s(X0,happ(s(fun(X1,X0),X4),s(X1,X6)))))) | s(X1,X5) = s(X1,X6)) )),
  inference(cnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),X2),s(fun(X0,X1),X3),s(fun(X1,X0),X4)))) | ! [X5,X6] : ((s(X1,X5) != s(X1,X6) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))),s(X0,happ(s(fun(X1,X0),X4),s(X1,X6))))))) & (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))),s(X0,happ(s(fun(X1,X0),X4),s(X1,X6)))))) | s(X1,X5) = s(X1,X6))))),
  inference(nnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),X2),s(fun(X0,X1),X3),s(fun(X1,X0),X4)))) | ! [X5,X6] : (s(X1,X5) = s(X1,X6) <=> p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))),s(X0,happ(s(fun(X1,X0),X4),s(X1,X6))))))))),
  inference(ennf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),X2),s(fun(X0,X1),X3),s(fun(X1,X0),X4)))) => ! [X5,X6] : (s(X1,X5) = s(X1,X6) <=> p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))),s(X0,happ(s(fun(X1,X0),X4),s(X1,X6))))))))),
  inference(flattening,[],[f75])).
fof(f75,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),X2),s(fun(X0,X1),X3),s(fun(X1,X0),X4)))) => ! [X5,X6] : (s(X1,X5) = s(X1,X6) <=> p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))),s(X0,happ(s(fun(X1,X0),X4),s(X1,X6))))))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X9,X21] : ! [X32,X33,X34] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X9,fun(X9,bool)),X32),s(fun(X9,X21),X33),s(fun(X21,X9),X34)))) => ! [X6,X7] : (s(X21,X6) = s(X21,X7) <=> p(s(bool,happ(s(fun(X9,bool),happ(s(fun(X9,fun(X9,bool)),X32),s(X9,happ(s(fun(X21,X9),X34),s(X21,X6))))),s(X9,happ(s(fun(X21,X9),X34),s(X21,X7))))))))),
  file('thm.integer.INT_EQ_RADD','thm.quotient.EQUALS_PRS')).
fof(f36926,plain,(
  s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK1))) = s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK1))) | ($spl33 | $spl37)),
  inference(unit_resulting_resolution,[],[f35347,f182,f36699,f4096])).
fof(f4096,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f4033,f148])).
fof(f4033,plain,(
  ( ! [X47,X46] : (s(bool,'F_0') != s(bool,X46) | s(bool,X46) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f3888,f149])).
fof(f3888,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X1))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f3782])).
fof(f3782,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.combin.I_1'(s(bool,X1))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1258,f148])).
fof(f1258,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,'const.combin.I_1'(s(bool,X1))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f288])).
fof(f182,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))) )),
  inference(cnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))),
  inference(flattening,[],[f87])).
fof(f87,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X40,X38,X41] : ! [X42,X39,X43] : s(X40,happ(s(fun(X38,X40),'_dst_x1x0_2'(s(fun(X38,fun(X41,X40)),X42),s(X41,X39))),s(X38,X43))) = s(X40,happ(s(fun(X41,X40),happ(s(fun(X38,fun(X41,X40)),X42),s(X38,X43))),s(X41,X39)))),
  file('thm.integer.INT_EQ_RADD','_dst_x1x0_ax')).
fof(f39462,plain,(
  ( ! [X238] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X238)) ) | ($spl3 | $spl11 | $spl27 | $spl33 | $spl37)),
  inference(subsumption_resolution,[],[f39071,f14037])).
fof(f14037,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(forward_demodulation,[],[f13869,f149])).
fof(f13869,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X0)))) )),
  inference(condensation,[],[f13868])).
fof(f13868,plain,(
  ( ! [X87,X86] : (s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X87))) | s(bool,'T_0') = s(bool,X86) | s(bool,'T_0') = s(bool,X87)) )),
  inference(duplicate_literal_removal,[],[f13632])).
fof(f13632,plain,(
  ( ! [X87,X86] : (s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X87))) | s(bool,'T_0') = s(bool,X86) | s(bool,'T_0') = s(bool,X87) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f1270,f314])).
fof(f39071,plain,(
  ( ! [X238] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X238) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X238)))) ) | ($spl3 | $spl11 | $spl27 | $spl33 | $spl37)),
  inference(superposition,[],[f14037,f37485])).
fof(f39440,plain,(
  ( ! [X142] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X142)) ) | ($spl3 | $spl11 | $spl27 | $spl33 | $spl37)),
  inference(subsumption_resolution,[],[f39047,f2354])).
fof(f2354,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(forward_demodulation,[],[f2282,f149])).
fof(f2282,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,'const.combin.I_1'(s(bool,X0)))) )),
  inference(condensation,[],[f2281])).
fof(f2281,plain,(
  ( ! [X41,X42] : (s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,'const.combin.I_1'(s(bool,X42))) | s(bool,'T_0') = s(bool,X41) | s(bool,'T_0') = s(bool,X42)) )),
  inference(duplicate_literal_removal,[],[f2203])).
fof(f2203,plain,(
  ( ! [X41,X42] : (s(bool,'const.combin.I_1'(s(bool,'T_0'))) != s(bool,'const.combin.I_1'(s(bool,X42))) | s(bool,'T_0') = s(bool,X41) | s(bool,'T_0') = s(bool,X42) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f1270,f313])).
fof(f39047,plain,(
  ( ! [X142] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X142) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X142)))) ) | ($spl3 | $spl11 | $spl27 | $spl33 | $spl37)),
  inference(superposition,[],[f2354,f37485])).
fof(f39394,plain,(
  ( ! [X75] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X75)) ) | ($spl3 | $spl11 | $spl27 | $spl33 | $spl37)),
  inference(subsumption_resolution,[],[f39003,f17081])).
fof(f39003,plain,(
  ( ! [X75] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X75))) | s(bool,'T_0') = s(bool,X75)) ) | ($spl3 | $spl11 | $spl27 | $spl33 | $spl37)),
  inference(superposition,[],[f1544,f37485])).
fof(f1544,plain,(
  ( ! [X33] : (s(bool,'T_0') != s(bool,X33) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X33)))) )),
  inference(forward_demodulation,[],[f1500,f149])).
fof(f1500,plain,(
  ( ! [X33] : (s(bool,X33) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X33)))) )),
  inference(superposition,[],[f1270,f201])).
fof(f39393,plain,(
  ( ! [X74] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X74)) ) | ($spl3 | $spl11 | $spl27 | $spl33 | $spl37)),
  inference(subsumption_resolution,[],[f39002,f14037])).
fof(f39002,plain,(
  ( ! [X74] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X74))) | s(bool,'T_0') = s(bool,X74)) ) | ($spl3 | $spl11 | $spl27 | $spl33 | $spl37)),
  inference(superposition,[],[f1543,f37485])).
fof(f1543,plain,(
  ( ! [X32] : (s(bool,'T_0') != s(bool,X32) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X32)))) )),
  inference(forward_demodulation,[],[f1499,f149])).
fof(f1499,plain,(
  ( ! [X32] : (s(bool,X32) != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X32)))) )),
  inference(superposition,[],[f1270,f150])).
fof(f39387,plain,(
  ( ! [X72] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X72)) ) | ($spl3 | $spl11 | $spl27 | $spl33 | $spl37)),
  inference(subsumption_resolution,[],[f39000,f2354])).
fof(f39000,plain,(
  ( ! [X72] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X72))) | s(bool,'T_0') = s(bool,X72)) ) | ($spl3 | $spl11 | $spl27 | $spl33 | $spl37)),
  inference(superposition,[],[f1429,f37485])).
fof(f1429,plain,(
  ( ! [X37] : (s(bool,'T_0') != s(bool,X37) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X37)))) ) | $spl11),
  inference(subsumption_resolution,[],[f1175,f222])).
fof(f1175,plain,(
  ( ! [X37] : (s(bool,'T_0') != s(bool,X37) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X37)))) ) | $spl11),
  inference(superposition,[],[f222,f288])).
fof(f38202,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0)))) ) | $spl42),
  inference(cnf_transformation,[],[f38202_D])).
fof(f38202_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0)))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f38041,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0))) | $spl41),
  inference(cnf_transformation,[],[f38041_D])).
fof(f38041_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f37873,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0))) | s(bool,X0) = s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0)))) ) | ($spl11 | $spl27 | $spl33)),
  inference(equality_factoring,[],[f36112])).
fof(f36112,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl11 | $spl27 | $spl33)),
  inference(subsumption_resolution,[],[f35964,f222])).
fof(f35964,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X0) = s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl27 | $spl33)),
  inference(superposition,[],[f35878,f210])).
fof(f35878,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0))) | ($spl27 | $spl33)),
  inference(forward_demodulation,[],[f35669,f201])).
fof(f35669,plain,(
  s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) = s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0))) | ($spl27 | $spl33)),
  inference(unit_resulting_resolution,[],[f17704,f182,f35347,f4096])).
fof(f38114,plain,(
  ( ! [X244] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0))) | s(bool,'T_0') = s(bool,X244)) ) | ($spl11 | $spl27 | $spl33)),
  inference(subsumption_resolution,[],[f37735,f17081])).
fof(f37735,plain,(
  ( ! [X244] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0))) | s(bool,'T_0') = s(bool,X244) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X244)))) ) | ($spl11 | $spl27 | $spl33)),
  inference(superposition,[],[f17081,f36112])).
fof(f38113,plain,(
  ( ! [X237] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0))) | s(bool,'T_0') = s(bool,X237)) ) | ($spl11 | $spl27 | $spl33)),
  inference(subsumption_resolution,[],[f37730,f14037])).
fof(f37730,plain,(
  ( ! [X237] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0))) | s(bool,'T_0') = s(bool,X237) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X237)))) ) | ($spl11 | $spl27 | $spl33)),
  inference(superposition,[],[f14037,f36112])).
fof(f38091,plain,(
  ( ! [X141] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0))) | s(bool,'T_0') = s(bool,X141)) ) | ($spl11 | $spl27 | $spl33)),
  inference(subsumption_resolution,[],[f37706,f2354])).
fof(f37706,plain,(
  ( ! [X141] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0))) | s(bool,'T_0') = s(bool,X141) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X141)))) ) | ($spl11 | $spl27 | $spl33)),
  inference(superposition,[],[f2354,f36112])).
fof(f38046,plain,(
  ( ! [X75] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0))) | s(bool,'T_0') = s(bool,X75)) ) | ($spl11 | $spl27 | $spl33)),
  inference(subsumption_resolution,[],[f37663,f17081])).
fof(f37663,plain,(
  ( ! [X75] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X75))) | s(bool,'T_0') = s(bool,X75)) ) | ($spl11 | $spl27 | $spl33)),
  inference(superposition,[],[f1544,f36112])).
fof(f38045,plain,(
  ( ! [X74] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0))) | s(bool,'T_0') = s(bool,X74)) ) | ($spl11 | $spl27 | $spl33)),
  inference(subsumption_resolution,[],[f37662,f14037])).
fof(f37662,plain,(
  ( ! [X74] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X74))) | s(bool,'T_0') = s(bool,X74)) ) | ($spl11 | $spl27 | $spl33)),
  inference(superposition,[],[f1543,f36112])).
fof(f38039,plain,(
  ( ! [X72] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0))) | s(bool,'T_0') = s(bool,X72)) ) | ($spl11 | $spl27 | $spl33)),
  inference(subsumption_resolution,[],[f37660,f2354])).
fof(f37660,plain,(
  ( ! [X72] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),'_dst_x1x0_2'(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X72))) | s(bool,'T_0') = s(bool,X72)) ) | ($spl11 | $spl27 | $spl33)),
  inference(superposition,[],[f1429,f36112])).
fof(f36817,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0)))) ) | $spl38),
  inference(cnf_transformation,[],[f36817_D])).
fof(f36817_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0)))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f36699,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0))) | $spl37),
  inference(cnf_transformation,[],[f36699_D])).
fof(f36699_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0))) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f36546,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0))) | s(bool,X0) = s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0)))) ) | $spl3),
  inference(equality_factoring,[],[f34357])).
fof(f34357,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(subsumption_resolution,[],[f34328,f140])).
fof(f34328,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f33716,f210])).
fof(f33716,plain,(
  ~p(s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0)))) | $spl3),
  inference(unit_resulting_resolution,[],[f196,f7816])).
fof(f36755,plain,(
  ( ! [X243] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0))) | s(bool,'T_0') = s(bool,X243)) ) | $spl3),
  inference(subsumption_resolution,[],[f36407,f17081])).
fof(f36407,plain,(
  ( ! [X243] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0))) | s(bool,'T_0') = s(bool,X243) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X243)))) ) | $spl3),
  inference(superposition,[],[f17081,f34357])).
fof(f36754,plain,(
  ( ! [X236] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0))) | s(bool,'T_0') = s(bool,X236)) ) | $spl3),
  inference(subsumption_resolution,[],[f36402,f14037])).
fof(f36402,plain,(
  ( ! [X236] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0))) | s(bool,'T_0') = s(bool,X236) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X236)))) ) | $spl3),
  inference(superposition,[],[f14037,f34357])).
fof(f36742,plain,(
  ( ! [X140] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0))) | s(bool,'T_0') = s(bool,X140)) ) | $spl3),
  inference(subsumption_resolution,[],[f36378,f2354])).
fof(f36378,plain,(
  ( ! [X140] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0))) | s(bool,'T_0') = s(bool,X140) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X140)))) ) | $spl3),
  inference(superposition,[],[f2354,f34357])).
fof(f36704,plain,(
  ( ! [X75] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0))) | s(bool,'T_0') = s(bool,X75)) ) | $spl3),
  inference(subsumption_resolution,[],[f36336,f17081])).
fof(f36336,plain,(
  ( ! [X75] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X75))) | s(bool,'T_0') = s(bool,X75)) ) | $spl3),
  inference(superposition,[],[f1544,f34357])).
fof(f36703,plain,(
  ( ! [X74] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0))) | s(bool,'T_0') = s(bool,X74)) ) | $spl3),
  inference(subsumption_resolution,[],[f36335,f14037])).
fof(f36335,plain,(
  ( ! [X74] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X74))) | s(bool,'T_0') = s(bool,X74)) ) | $spl3),
  inference(superposition,[],[f1543,f34357])).
fof(f36697,plain,(
  ( ! [X72] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0))) | s(bool,'T_0') = s(bool,X72)) ) | ($spl3 | $spl11)),
  inference(subsumption_resolution,[],[f36333,f2354])).
fof(f36333,plain,(
  ( ! [X72] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK1))),s('type.integer.int',sK0))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X72))) | s(bool,'T_0') = s(bool,X72)) ) | ($spl3 | $spl11)),
  inference(superposition,[],[f1429,f34357])).
fof(f35463,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK1)))) ) | $spl34),
  inference(cnf_transformation,[],[f35463_D])).
fof(f35463_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK1)))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f35347,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK1))) | $spl33),
  inference(cnf_transformation,[],[f35347_D])).
fof(f35347_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK1))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f35198,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK1))) | s(bool,X0) = s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK1)))) ) | $spl3),
  inference(equality_factoring,[],[f34192])).
fof(f34192,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(subsumption_resolution,[],[f34163,f140])).
fof(f34163,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f33704,f210])).
fof(f35401,plain,(
  ( ! [X232] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X232)) ) | $spl3),
  inference(subsumption_resolution,[],[f35060,f17081])).
fof(f35060,plain,(
  ( ! [X232] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X232) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X232)))) ) | $spl3),
  inference(superposition,[],[f17081,f34192])).
fof(f35400,plain,(
  ( ! [X225] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X225)) ) | $spl3),
  inference(subsumption_resolution,[],[f35055,f14037])).
fof(f35055,plain,(
  ( ! [X225] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X225) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X225)))) ) | $spl3),
  inference(superposition,[],[f14037,f34192])).
fof(f35389,plain,(
  ( ! [X139] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X139)) ) | $spl3),
  inference(subsumption_resolution,[],[f35032,f2354])).
fof(f35032,plain,(
  ( ! [X139] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X139) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X139)))) ) | $spl3),
  inference(superposition,[],[f2354,f34192])).
fof(f35352,plain,(
  ( ! [X75] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X75)) ) | $spl3),
  inference(subsumption_resolution,[],[f34991,f17081])).
fof(f34991,plain,(
  ( ! [X75] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X75))) | s(bool,'T_0') = s(bool,X75)) ) | $spl3),
  inference(superposition,[],[f1544,f34192])).
fof(f35351,plain,(
  ( ! [X74] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X74)) ) | $spl3),
  inference(subsumption_resolution,[],[f34990,f14037])).
fof(f34990,plain,(
  ( ! [X74] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X74))) | s(bool,'T_0') = s(bool,X74)) ) | $spl3),
  inference(superposition,[],[f1543,f34192])).
fof(f35345,plain,(
  ( ! [X72] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,X72)) ) | ($spl3 | $spl11)),
  inference(subsumption_resolution,[],[f34988,f2354])).
fof(f34988,plain,(
  ( ! [X72] : (s(bool,'T_0') != s(bool,happ(s(fun('type.integer.int',bool),happ(s(fun('type.integer.int',fun('type.integer.int',bool)),'|dollar|equals_0'),s('type.integer.int',sK0))),s('type.integer.int',sK1))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X72))) | s(bool,'T_0') = s(bool,X72)) ) | ($spl3 | $spl11)),
  inference(superposition,[],[f1429,f34192])).
fof(f21425,plain,(
  $false | ($spl12 | $spl27)),
  inference(subsumption_resolution,[],[f20653,f2811])).
fof(f20653,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'T_0'))) | ($spl12 | $spl27)),
  inference(backward_demodulation,[],[f2811,f17704])).
fof(f21424,plain,(
  $false | ($spl12 | $spl23)),
  inference(subsumption_resolution,[],[f20645,f2811])).
fof(f20645,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'T_0'))) | ($spl12 | $spl23)),
  inference(backward_demodulation,[],[f2811,f14731])).
fof(f21423,plain,(
  $false | ($spl12 | $spl15)),
  inference(subsumption_resolution,[],[f20543,f2811])).
fof(f20543,plain,(
  s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'T_0'))) | ($spl12 | $spl15)),
  inference(backward_demodulation,[],[f2811,f2813])).
fof(f21422,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f20502,f140])).
fof(f20502,plain,(
  ~p(s(bool,'T_0')) | $spl12),
  inference(backward_demodulation,[],[f2811,f144])).
fof(f20414,plain,(
  p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|dollar|equals_0'),s(fun(bool,bool),'const.combin.I_0'),s(fun(bool,bool),'const.combin.I_0'))),s(bool,'F_0')))) | $spl30),
  inference(cnf_transformation,[],[f20414_D])).
fof(f20414_D,plain,(
  p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|dollar|equals_0'),s(fun(bool,bool),'const.combin.I_0'),s(fun(bool,bool),'const.combin.I_0'))),s(bool,'F_0')))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f20416,plain,(
  ( ! [X90] : (s(bool,'T_0') = s(bool,X90) | p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|dollar|equals_0'),s(fun(bool,bool),'const.combin.I_0'),s(fun(bool,bool),'const.combin.I_0'))),s(bool,'F_0'))))) )),
  inference(forward_demodulation,[],[f20415,f149])).
fof(f20415,plain,(
  ( ! [X90] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|dollar|equals_0'),s(fun(bool,bool),'const.combin.I_0'),s(fun(bool,bool),'const.combin.I_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X90)))) )),
  inference(forward_demodulation,[],[f20259,f190])).
fof(f190,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,happ(s(fun(X3,X1),X5),s(X3,X7))))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))) )),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,happ(s(fun(X3,X1),X5),s(X3,X7))))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(flattening,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,happ(s(fun(X3,X1),X5),s(X3,X7))))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X40,X41,X44,X38] : ! [X42,X39,X45,X43] : s(X40,happ(s(fun(X38,X40),'_dst_x0x1x1_3'(s(fun(X41,fun(X44,X40)),X42),s(fun(X38,X41),X39),s(fun(X38,X44),X45))),s(X38,X43))) = s(X40,happ(s(fun(X44,X40),happ(s(fun(X41,fun(X44,X40)),X42),s(X41,happ(s(fun(X38,X41),X39),s(X38,X43))))),s(X44,happ(s(fun(X38,X44),X45),s(X38,X43)))))),
  file('thm.integer.INT_EQ_RADD','_dst_x0x1x1_ax')).
fof(f20259,plain,(
  ( ! [X90] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|dollar|equals_0'),s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))))),s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0')))))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X90)))) )),
  inference(superposition,[],[f19878,f17627])).
fof(f19878,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(X0,X1))),s(X0,X1))))) )),
  inference(forward_demodulation,[],[f19877,f182])).
fof(f19877,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(X0,X1))),s(X0,X1))))) )),
  inference(forward_demodulation,[],[f19876,f201])).
fof(f19876,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(X0,X1))),s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1))))))) )),
  inference(forward_demodulation,[],[f19618,f182])).
fof(f19618,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1))))),s(X0,X1))))) )),
  inference(unit_resulting_resolution,[],[f201,f7547])).
fof(f7547,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(forward_demodulation,[],[f7546,f201])).
fof(f7546,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1))))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(forward_demodulation,[],[f7545,f201])).
fof(f7545,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1))))),s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X2))))))) )),
  inference(subsumption_resolution,[],[f7520,f140])).
fof(f7520,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'|dollar|equals_0'),s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1))))),s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X2))))))) )),
  inference(superposition,[],[f166,f256])).
fof(f166,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),X2),s(fun(X0,X1),X3),s(fun(X1,X0),X4)))) | s(X1,X5) != s(X1,X6) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))),s(X0,happ(s(fun(X1,X0),X4),s(X1,X6))))))) )),
  inference(cnf_transformation,[],[f131])).
fof(f20412,plain,(
  ( ! [X89] : (p(s(bool,happ(s(fun(bool,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'|dollar|equals_0'),s(fun(bool,bool),'const.combin.I_0'),s(fun(bool,bool),'const.combin.I_0'))),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X89)) )),
  inference(forward_demodulation,[],[f20258,f190])).
fof(f20258,plain,(
  ( ! [X89] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'|dollar|equals_0'),s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))))),s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0')))))) | s(bool,'T_0') = s(bool,X89)) )),
  inference(superposition,[],[f19878,f16945])).
fof(f17805,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0')))) ) | $spl28),
  inference(cnf_transformation,[],[f17805_D])).
fof(f17805_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0')))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f17704,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | $spl27),
  inference(cnf_transformation,[],[f17704_D])).
fof(f17704_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f17545,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,X0) = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0')))) )),
  inference(equality_factoring,[],[f16945])).
fof(f17750,plain,(
  ( ! [X205] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X205)) )),
  inference(subsumption_resolution,[],[f17414,f14037])).
fof(f17414,plain,(
  ( ! [X205] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X205) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X205)))) )),
  inference(superposition,[],[f14037,f16945])).
fof(f17742,plain,(
  ( ! [X137] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X137)) )),
  inference(subsumption_resolution,[],[f17397,f2354])).
fof(f17397,plain,(
  ( ! [X137] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X137) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X137)))) )),
  inference(superposition,[],[f2354,f16945])).
fof(f17709,plain,(
  ( ! [X75] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X75)) )),
  inference(subsumption_resolution,[],[f17359,f17081])).
fof(f17359,plain,(
  ( ! [X75] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X75))) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f1544,f16945])).
fof(f17708,plain,(
  ( ! [X74] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f17358,f14037])).
fof(f17358,plain,(
  ( ! [X74] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X74))) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f1543,f16945])).
fof(f17702,plain,(
  ( ! [X72] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X72)) ) | $spl11),
  inference(subsumption_resolution,[],[f17356,f2354])).
fof(f17356,plain,(
  ( ! [X72] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X72))) | s(bool,'T_0') = s(bool,X72)) ) | $spl11),
  inference(superposition,[],[f1429,f16945])).
fof(f14826,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0')))) ) | $spl24),
  inference(cnf_transformation,[],[f14826_D])).
fof(f14826_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0')))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f14731,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | $spl23),
  inference(cnf_transformation,[],[f14731_D])).
fof(f14731_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f14584,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,X0) = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0')))) )),
  inference(equality_factoring,[],[f13930])).
fof(f14764,plain,(
  ( ! [X129] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X129)) )),
  inference(subsumption_resolution,[],[f14447,f2354])).
fof(f14447,plain,(
  ( ! [X129] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X129) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X129)))) )),
  inference(superposition,[],[f2354,f13930])).
fof(f14754,plain,(
  ( ! [X108] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X108)) )),
  inference(forward_demodulation,[],[f14753,f201])).
fof(f14753,plain,(
  ( ! [X108] : (s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X108)) )),
  inference(subsumption_resolution,[],[f14435,f2354])).
fof(f14435,plain,(
  ( ! [X108] : (s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X108))) | s(bool,'T_0') = s(bool,X108)) )),
  inference(superposition,[],[f13026,f13930])).
fof(f13026,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X0)))) )),
  inference(condensation,[],[f13025])).
fof(f13025,plain,(
  ( ! [X12,X11] : (s(bool,X12) != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X12)))) )),
  inference(duplicate_literal_removal,[],[f12918])).
fof(f12918,plain,(
  ( ! [X12,X11] : (s(bool,X12) != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X12)))) )),
  inference(superposition,[],[f12219,f288])).
fof(f12219,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f12201])).
fof(f12201,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f290])).
fof(f14752,plain,(
  ( ! [X107] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X107)) )),
  inference(forward_demodulation,[],[f14598,f201])).
fof(f14598,plain,(
  ( ! [X107] : (s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X107)) )),
  inference(duplicate_literal_removal,[],[f14434])).
fof(f14434,plain,(
  ( ! [X107] : (s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) != s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X107) | s(bool,'T_0') = s(bool,X107)) )),
  inference(superposition,[],[f12219,f13930])).
fof(f14735,plain,(
  ( ! [X70] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f14412,f14037])).
fof(f14412,plain,(
  ( ! [X70] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,X70))) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f1543,f13930])).
fof(f14729,plain,(
  ( ! [X68] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X68)) ) | $spl11),
  inference(subsumption_resolution,[],[f14410,f2354])).
fof(f14410,plain,(
  ( ! [X68] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),'_i_0_0'),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X68))) | s(bool,'T_0') = s(bool,X68)) ) | $spl11),
  inference(superposition,[],[f1429,f13930])).
fof(f4410,plain,(
  s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_ABS_1'(s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.tint_add_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.int_REP_1'(s('type.integer.int',sK1))),s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.int_REP_1'(s('type.integer.int',sK2))))))) | $spl18),
  inference(cnf_transformation,[],[f4410_D])).
fof(f4410_D,plain,(
  s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_ABS_1'(s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.tint_add_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.int_REP_1'(s('type.integer.int',sK1))),s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.int_REP_1'(s('type.integer.int',sK2))))))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f4412,plain,(
  s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))))) | $spl21),
  inference(cnf_transformation,[],[f4412_D])).
fof(f4412_D,plain,(
  s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f4394,plain,(
  s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))))) | s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_ABS_1'(s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.tint_add_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.int_REP_1'(s('type.integer.int',sK1))),s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.int_REP_1'(s('type.integer.int',sK2))))))) | $spl0),
  inference(superposition,[],[f1670,f193])).
fof(f1670,plain,(
  ( ! [X26,X27] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X26),s('type.integer.int',X27))))) | s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_ABS_1'(s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.tint_add_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.int_REP_1'(s('type.integer.int',X26))),s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.int_REP_1'(s('type.integer.int',X27)))))))) ) | $spl0),
  inference(superposition,[],[f775,f158])).
fof(f775,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',X0))) | s('type.integer.int',sK2) = s('type.integer.int',X0)) ) | $spl0),
  inference(superposition,[],[f173,f193])).
fof(f2830,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) | $spl16),
  inference(cnf_transformation,[],[f2830_D])).
fof(f2830_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.combin.I_1'(s(bool,'F_0')))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f2813,plain,(
  s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | $spl15),
  inference(cnf_transformation,[],[f2813_D])).
fof(f2813_D,plain,(
  s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f2751,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,X0) = s(bool,'const.combin.I_1'(s(bool,'F_0')))) )),
  inference(equality_factoring,[],[f2318])).
fof(f2810,plain,(
  ( ! [X30] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X30)) ) | $spl11),
  inference(subsumption_resolution,[],[f2688,f2354])).
fof(f2688,plain,(
  ( ! [X30] : (s(bool,'T_0') != s(bool,'const.combin.I_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.I_1'(s(bool,X30))) | s(bool,'T_0') = s(bool,X30)) ) | $spl11),
  inference(superposition,[],[f1429,f2318])).
fof(f724,plain,(
  $false | $spl5),
  inference(subsumption_resolution,[],[f723,f140])).
fof(f723,plain,(
  ~p(s(bool,'T_0')) | $spl5),
  inference(backward_demodulation,[],[f712,f205])).
fof(f205,plain,(
  ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.tint_add_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK3),s('type.pair.prod'('type.num.num','type.num.num'),sK4))),s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.tint_add_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK3),s('type.pair.prod'('type.num.num','type.num.num'),sK4)))))) | $spl5),
  inference(forward_demodulation,[],[f204,f157])).
fof(f204,plain,(
  ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.tint_add_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK4),s('type.pair.prod'('type.num.num','type.num.num'),sK3))),s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.tint_add_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK3),s('type.pair.prod'('type.num.num','type.num.num'),sK4)))))) | $spl5),
  inference(forward_demodulation,[],[f203,f160])).
fof(f160,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.tint_add_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X1))) = s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s('type.pair.prod'('type.num.num','type.num.num'),X0))),s('type.pair.prod'('type.num.num','type.num.num'),X1)))) )),
  inference(cnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0,X1] : s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.tint_add_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X1))) = s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s('type.pair.prod'('type.num.num','type.num.num'),X0))),s('type.pair.prod'('type.num.num','type.num.num'),X1)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X36,X37] : s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.tint_add_2'(s('type.pair.prod'('type.num.num','type.num.num'),X36),s('type.pair.prod'('type.num.num','type.num.num'),X37))) = s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s('type.pair.prod'('type.num.num','type.num.num'),X36))),s('type.pair.prod'('type.num.num','type.num.num'),X37)))),
  file('thm.integer.INT_EQ_RADD','ARITY_0_2_const.integer.tint_add')).
fof(f203,plain,(
  ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s('type.pair.prod'('type.num.num','type.num.num'),sK4))),s('type.pair.prod'('type.num.num','type.num.num'),sK3))),s('type.pair.prod'('type.num.num','type.num.num'),'const.integer.tint_add_2'(s('type.pair.prod'('type.num.num','type.num.num'),sK3),s('type.pair.prod'('type.num.num','type.num.num'),sK4)))))) | $spl5),
  inference(forward_demodulation,[],[f202,f160])).
fof(f202,plain,(
  ~p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s('type.pair.prod'('type.num.num','type.num.num'),sK4))),s('type.pair.prod'('type.num.num','type.num.num'),sK3))),s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s('type.pair.prod'('type.num.num','type.num.num'),sK3))),s('type.pair.prod'('type.num.num','type.num.num'),sK4)))))) | $spl5),
  inference(forward_demodulation,[],[f198,f159])).
fof(f159,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X1))) = s(bool,happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'),s('type.pair.prod'('type.num.num','type.num.num'),X0))),s('type.pair.prod'('type.num.num','type.num.num'),X1)))) )),
  inference(cnf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0,X1] : s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X1))) = s(bool,happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'),s('type.pair.prod'('type.num.num','type.num.num'),X0))),s('type.pair.prod'('type.num.num','type.num.num'),X1)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X36,X37] : s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X36),s('type.pair.prod'('type.num.num','type.num.num'),X37))) = s(bool,happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'),s('type.pair.prod'('type.num.num','type.num.num'),X36))),s('type.pair.prod'('type.num.num','type.num.num'),X37)))),
  file('thm.integer.INT_EQ_RADD','ARITY_0_2_const.integer.tint_eq')).
fof(f712,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X0)))) )),
  inference(unit_resulting_resolution,[],[f692,f213])).
fof(f692,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X0))))) )),
  inference(forward_demodulation,[],[f678,f201])).
fof(f678,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),'const.combin.I_0'),s('type.pair.prod'('type.num.num','type.num.num'),X0))),s('type.pair.prod'('type.num.num','type.num.num'),X0))))) )),
  inference(unit_resulting_resolution,[],[f201,f161])).
fof(f161,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X1)))) | s('type.pair.prod'('type.num.num','type.num.num'),X0) != s('type.pair.prod'('type.num.num','type.num.num'),X1)) )),
  inference(cnf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0,X1] : (s('type.pair.prod'('type.num.num','type.num.num'),X0) != s('type.pair.prod'('type.num.num','type.num.num'),X1) | p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X1)))))),
  inference(ennf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0,X1] : (s('type.pair.prod'('type.num.num','type.num.num'),X0) = s('type.pair.prod'('type.num.num','type.num.num'),X1) => p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X0),s('type.pair.prod'('type.num.num','type.num.num'),X1)))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X13,X14] : (s('type.pair.prod'('type.num.num','type.num.num'),X13) = s('type.pair.prod'('type.num.num','type.num.num'),X14) => p(s(bool,'const.integer.tint_eq_2'(s('type.pair.prod'('type.num.num','type.num.num'),X13),s('type.pair.prod'('type.num.num','type.num.num'),X14)))))),
  file('thm.integer.INT_EQ_RADD','thm.integer.TINT_EQ_AP')).
fof(f220,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl8),
  inference(cnf_transformation,[],[f220_D])).
fof(f220_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f222,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl11),
  inference(cnf_transformation,[],[f222_D])).
fof(f222_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f219,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f148])).
fof(f207,plain,(
  $false | ($spl1 | $spl2)),
  inference(trivial_inequality_removal,[],[f206])).
fof(f206,plain,(
  s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f194,f195])).
fof(f198,plain,(
  ~p(s(bool,happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'),s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s('type.pair.prod'('type.num.num','type.num.num'),sK4))),s('type.pair.prod'('type.num.num','type.num.num'),sK3))))),s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s('type.pair.prod'('type.num.num','type.num.num'),sK3))),s('type.pair.prod'('type.num.num','type.num.num'),sK4)))))) | $spl5),
  inference(cnf_transformation,[],[f198_D])).
fof(f198_D,plain,(
  ~p(s(bool,happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'),s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s('type.pair.prod'('type.num.num','type.num.num'),sK4))),s('type.pair.prod'('type.num.num','type.num.num'),sK3))))),s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s('type.pair.prod'('type.num.num','type.num.num'),sK3))),s('type.pair.prod'('type.num.num','type.num.num'),sK4)))))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f200,plain,(
  p(s(bool,happ(s(fun(fun('type.pair.prod'('type.num.num','type.num.num'),bool),bool),'const.bool.RES_FORALL_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.quotient.respects_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'))))),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'_dst_x0x1_2'(s(fun(fun('type.pair.prod'('type.num.num','type.num.num'),bool),bool),'const.bool.RES_FORALL_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.quotient.respects_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'))))),s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'_dst_x00x11x11_3'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'_dst_x01x10_2'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),'_i_0_0'))),s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0')))))))) | $spl6),
  inference(cnf_transformation,[],[f200_D])).
fof(f200_D,plain,(
  p(s(bool,happ(s(fun(fun('type.pair.prod'('type.num.num','type.num.num'),bool),bool),'const.bool.RES_FORALL_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.quotient.respects_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'))))),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'_dst_x0x1_2'(s(fun(fun('type.pair.prod'('type.num.num','type.num.num'),bool),bool),'const.bool.RES_FORALL_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.quotient.respects_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'))))),s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'_dst_x00x11x11_3'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'_dst_x01x10_2'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),'_i_0_0'))),s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0')))))))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f143,plain,(
  p(s(bool,happ(s(fun(fun('type.pair.prod'('type.num.num','type.num.num'),bool),bool),'const.bool.RES_FORALL_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.quotient.respects_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'))))),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'_dst_x0x1_2'(s(fun(fun('type.pair.prod'('type.num.num','type.num.num'),bool),bool),'const.bool.RES_FORALL_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.quotient.respects_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'))))),s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'_dst_x00x11x11_3'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'_dst_x01x10_2'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),'_i_0_0'))),s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0')))))))) | ~p(s(bool,happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'),s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s('type.pair.prod'('type.num.num','type.num.num'),sK4))),s('type.pair.prod'('type.num.num','type.num.num'),sK3))))),s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s('type.pair.prod'('type.num.num','type.num.num'),sK3))),s('type.pair.prod'('type.num.num','type.num.num'),sK4))))))),
  inference(cnf_transformation,[],[f127])).
fof(f127,plain,(
  ~p(s(bool,happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'),s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s('type.pair.prod'('type.num.num','type.num.num'),sK4))),s('type.pair.prod'('type.num.num','type.num.num'),sK3))))),s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s('type.pair.prod'('type.num.num','type.num.num'),sK3))),s('type.pair.prod'('type.num.num','type.num.num'),sK4)))))) | p(s(bool,happ(s(fun(fun('type.pair.prod'('type.num.num','type.num.num'),bool),bool),'const.bool.RES_FORALL_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.quotient.respects_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'))))),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'_dst_x0x1_2'(s(fun(fun('type.pair.prod'('type.num.num','type.num.num'),bool),bool),'const.bool.RES_FORALL_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.quotient.respects_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'))))),s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'_dst_x00x11x11_3'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'_dst_x01x10_2'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),'_i_0_0'))),s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3,sK4])],[f107])).
fof(f107,plain,(
  ? [X0,X1] : ~p(s(bool,happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'),s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s('type.pair.prod'('type.num.num','type.num.num'),X1))),s('type.pair.prod'('type.num.num','type.num.num'),X0))))),s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s('type.pair.prod'('type.num.num','type.num.num'),X0))),s('type.pair.prod'('type.num.num','type.num.num'),X1)))))) | p(s(bool,happ(s(fun(fun('type.pair.prod'('type.num.num','type.num.num'),bool),bool),'const.bool.RES_FORALL_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.quotient.respects_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'))))),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'_dst_x0x1_2'(s(fun(fun('type.pair.prod'('type.num.num','type.num.num'),bool),bool),'const.bool.RES_FORALL_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.quotient.respects_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'))))),s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'_dst_x00x11x11_3'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'_dst_x01x10_2'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),'_i_0_0'))),s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'))))))))),
  inference(ennf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0,X1] : p(s(bool,happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'),s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s('type.pair.prod'('type.num.num','type.num.num'),X1))),s('type.pair.prod'('type.num.num','type.num.num'),X0))))),s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s('type.pair.prod'('type.num.num','type.num.num'),X0))),s('type.pair.prod'('type.num.num','type.num.num'),X1)))))) => p(s(bool,happ(s(fun(fun('type.pair.prod'('type.num.num','type.num.num'),bool),bool),'const.bool.RES_FORALL_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.quotient.respects_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'))))),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'_dst_x0x1_2'(s(fun(fun('type.pair.prod'('type.num.num','type.num.num'),bool),bool),'const.bool.RES_FORALL_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.quotient.respects_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'))))),s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'_dst_x00x11x11_3'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'_dst_x01x10_2'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),'_i_0_0'))),s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'))))))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X7,X6] : p(s(bool,happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'),s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s('type.pair.prod'('type.num.num','type.num.num'),X6))),s('type.pair.prod'('type.num.num','type.num.num'),X7))))),s('type.pair.prod'('type.num.num','type.num.num'),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),happ(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s('type.pair.prod'('type.num.num','type.num.num'),X7))),s('type.pair.prod'('type.num.num','type.num.num'),X6)))))) => p(s(bool,happ(s(fun(fun('type.pair.prod'('type.num.num','type.num.num'),bool),bool),'const.bool.RES_FORALL_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.quotient.respects_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'))))),s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'_dst_x0x1_2'(s(fun(fun('type.pair.prod'('type.num.num','type.num.num'),bool),bool),'const.bool.RES_FORALL_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),bool),'const.quotient.respects_1'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'))))),s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'_dst_x00x11x11_3'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),bool)),'const.integer.tint_eq_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'_dst_x01x10_2'(s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'),s(fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num')),'_i_0_0'))),s(fun('type.pair.prod'('type.num.num','type.num.num'),fun('type.pair.prod'('type.num.num','type.num.num'),'type.pair.prod'('type.num.num','type.num.num'))),'const.integer.tint_add_0'))))))))),
  file('thm.integer.INT_EQ_RADD','thm.integer.integer_66_prove_ltb')).
fof(f193,plain,(
  s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) | $spl0),
  inference(cnf_transformation,[],[f193_D])).
fof(f193_D,plain,(
  s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f195,plain,(
  s('type.integer.int',sK0) = s('type.integer.int',sK1) | $spl2),
  inference(cnf_transformation,[],[f195_D])).
fof(f195_D,plain,(
  s('type.integer.int',sK0) = s('type.integer.int',sK1) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f138,plain,(
  s('type.integer.int',sK0) = s('type.integer.int',sK1) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))),
  inference(cnf_transformation,[],[f126])).
fof(f126,plain,(
  (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) | s('type.integer.int',sK0) = s('type.integer.int',sK1)) & (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) | s('type.integer.int',sK0) != s('type.integer.int',sK1))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f125])).
fof(f125,plain,(
  ? [X0,X1,X2] : ((s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X2))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X2))) | s('type.integer.int',X0) = s('type.integer.int',X1)) & (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X2))) | s('type.integer.int',X0) != s('type.integer.int',X1)))),
  inference(nnf_transformation,[],[f106])).
fof(f106,plain,(
  ? [X0,X1,X2] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X2))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X2))) <~> s('type.integer.int',X0) = s('type.integer.int',X1))),
  inference(ennf_transformation,[],[f44])).
fof(f44,plain,(
  ~! [X0,X1,X2] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X2))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X2))) <=> s('type.integer.int',X0) = s('type.integer.int',X1))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X6),s('type.integer.int',X8))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X7),s('type.integer.int',X8))) <=> s('type.integer.int',X6) = s('type.integer.int',X7))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X6),s('type.integer.int',X8))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X7),s('type.integer.int',X8))) <=> s('type.integer.int',X6) = s('type.integer.int',X7))),
  file('thm.integer.INT_EQ_RADD','thm.integer.INT_EQ_RADD')).
fof(f194,plain,(
  s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) | $spl1),
  inference(cnf_transformation,[],[f194_D])).
fof(f194_D,plain,(
  s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f196,plain,(
  s('type.integer.int',sK0) != s('type.integer.int',sK1) | $spl3),
  inference(cnf_transformation,[],[f196_D])).
fof(f196_D,plain,(
  s('type.integer.int',sK0) != s('type.integer.int',sK1) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f139,plain,(
  s('type.integer.int',sK0) != s('type.integer.int',sK1) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK0),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))),
  inference(cnf_transformation,[],[f126])).
% SZS output end Proof for thm.integer
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 216
% Passive clauses: 3214
% Generated clauses: 58790
% Final active clauses: 162
% Final passive clauses: 1958
% Input formulas: 42
% Initial clauses: 54
% Discarded non-redundant clauses: 222
% 
% Duplicate literals: 2167
% Trivial inequalities: 422
% Fw subsumption resolutions: 4596
% Fw demodulations: 12991
% Bw demodulations: 845
% Condensations: 4630
% 
% Simple tautologies: 368
% Equational tautologies: 283
% Forward subsumptions: 29938
% Fw demodulations to eq. taut.: 100
% Bw demodulations to eq. taut.: 26
% 
% Unit resulting resolution: 2625
% Forward superposition: 16236
% Backward superposition: 12879
% Self superposition: 360
% Equality factoring: 107
% Equality resolution: 31
% 
% Split clauses: 78
% Split components: 172
% SAT solver clauses: 3380
% SAT solver unit clauses: 251
% SAT solver binary clauses: 1226
% 
% Sat splits: 78
% Sat splitting refutations: 8
% 
% Memory used [KB]: 43751
% Time elapsed: 3.138 s
% ------------------------------
% ------------------------------
% Success in time 4.494 s
