Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.pred_set
% (104820)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 111
% Passive clauses: 2461
% Generated clauses: 22182
% Final active clauses: 96
% Final passive clauses: 2106
% Input formulas: 21
% Initial clauses: 42
% 
% Duplicate literals: 1126
% Trivial inequalities: 303
% Fw subsumption resolutions: 3370
% Fw demodulations: 2626
% Bw demodulations: 7
% 
% Simple tautologies: 367
% Equational tautologies: 74
% Forward subsumptions: 9800
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 284
% Factoring: 3
% Forward superposition: 6424
% Backward superposition: 7286
% Self superposition: 309
% Equality factoring: 221
% Equality resolution: 20
% 
% Split clauses: 1203
% Split components: 2423
% SAT solver clauses: 2044
% SAT solver unit clauses: 14
% SAT solver binary clauses: 156
% 
% Sat splits: 1203
% Sat splitting refutations: 12
% 
% TWLsolver clauses: 2038
% TWLsolver calls for satisfiability: 152
% 
% Memory used [KB]: 9594
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (104847)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 73
% Passive clauses: 569
% Generated clauses: 9638
% Final active clauses: 65
% Final passive clauses: 484
% Input formulas: 21
% Initial clauses: 42
% 
% Duplicate literals: 815
% Trivial inequalities: 111
% Fw subsumption resolutions: 896
% Fw demodulations: 493
% Bw demodulations: 14
% Condensations: 1332
% Global subsumptions: 9
% 
% Simple tautologies: 178
% Equational tautologies: 52
% Forward subsumptions: 4178
% 
% Binary resolution: 197
% Factoring: 1
% Forward superposition: 2694
% Backward superposition: 2172
% Self superposition: 175
% Equality factoring: 133
% Equality resolution: 10
% 
% Split clauses: 97
% Split components: 195
% SAT solver clauses: 1023
% SAT solver unit clauses: 57
% SAT solver binary clauses: 47
% SAT solver learnt clauses: 18
% 
% Sat splits: 97
% Sat splitting refutations: 11
% 
% TWLsolver clauses: 987
% TWLsolver calls for satisfiability: 189
% 
% Memory used [KB]: 7164
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f8029,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f129,f132,f130,f155,f159,f157,f167,f172,f213,f217,f215,f375,f379,f377,f1048,f1049,f1050,f1051,f1062,f1052,f1061,f1059,f1063,f1064,f1058,f1056,f1148,f1149,f1150,f1151,f1162,f1152,f1161,f1159,f1163,f1164,f1158,f1156,f1256,f1265,f1257,f1258,f1259,f1269,f1262,f1260,f1268,f1263,f1266,f1364,f1367,f1384,f1382,f2544,f2917,f4116,f4119,f4117,f156,f4518,f4520,f4523,f4527,f4585,f4600,f4601,f4602,f4603,f4604,f4605,f4609,f4613,f4614,f4615,f4616,f4617,f4618,f4619,f4620,f4621,f4622,f4623,f4624,f4625,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4646,f4647,f4648,f4649,f4650,f4651,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4667,f4668,f4669,f4670,f4671,f4672,f4673,f4674,f4675,f4676,f4677,f5198,f5199,f5200,f5202,f5205,f5208,f5211,f7914,f7934,f7935,f7936,f7937,f7941,f7945,f7946,f7947,f7948,f7949,f7950,f7951,f7952,f7953,f7954,f7955,f7956,f7957,f7958,f7959,f7960,f7961,f7962,f7963,f7964,f7965,f7966,f7967,f7968,f7969,f7970,f7971,f7972,f7973,f7974,f7975,f7976,f7977,f7978,f7979,f7980,f7981,f7982,f7983,f7984,f7985,f7986,f7987,f7988,f7989,f7990,f7991,f7992,f7993,f7994,f7995,f7996,f7997,f7998,f7999,f8000,f8001,f8002,f8003,f8004,f8005,f8006,f8007,f8008,f8009,f8010,f8011,f8012,f8013,f8014,f8015,f8016,f8017,f8018,f8019,f8020,f8021,f8022,f8023,f8024,f8025,f8026,f8027,f8028])).
fof(f8028,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786,f1668,f1511,f425,f387,f382,f336,f260,f259,f212,f211,f144,f120,f1061])).
fof(f120,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f80,f77])).
fof(f80,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f54])).
fof(f54,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f25])).
fof(f25,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f18])).
fof(f18,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.pred_set.DIFF_EMPTY','thm.bool.NOT_CLAUSES|split|2')).
fof(f144,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f120,f107,f113])).
fof(f211,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK3(X0,X1,X2,bool))) )),
  inference(superposition,[],[f117,f84])).
fof(f212,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK3(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK3(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f117,f84])).
fof(f259,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X1),s(fun(bool,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X1)))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f110,f84])).
fof(f260,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X4),s(fun(bool,bool),X5)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X4)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f110,f84])).
fof(f336,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X8)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X7),s(fun(bool,bool),X8)))))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f111,f84])).
fof(f382,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X0,sK2(X1,'const.pred_set.EMPTY_0',X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) )),
  inference(superposition,[],[f109,f144])).
fof(f387,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X0,sK2('const.pred_set.EMPTY_0',X1,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) )),
  inference(superposition,[],[f109,f144])).
fof(f425,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X7)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X8)))) | p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X7),s(fun(bool,bool),X8)))))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f112,f84])).
fof(f1511,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,'const.bool.IN_2'(s(bool,X5),s(fun(bool,bool),X3)))) | s(bool,'T_0') = s(bool,X5) | s(bool,X6) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X3),s(fun(bool,bool),X4))))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f1505,f77])).
fof(f1505,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(bool,X5),s(fun(bool,bool),X3)))) | s(bool,'T_0') = s(bool,X5) | s(bool,X6) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X3),s(fun(bool,bool),X4))))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f259,f121])).
fof(f1668,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X4)))) | s(bool,'T_0') = s(bool,X3) | s(bool,X6) = s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X4),s(fun(bool,bool),X5))))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f1653,f77])).
fof(f1653,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X4)))) | s(bool,'T_0') = s(bool,X3) | s(bool,X6) = s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X4),s(fun(bool,bool),X5))))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f260,f121])).
fof(f1786,plain,(
  ( ! [X6,X8,X7,X5] : (~p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X7),s(fun(bool,bool),X5)))))) | s(bool,'T_0') = s(bool,X6) | s(bool,X8) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X5))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f1780,f77])).
fof(f1780,plain,(
  ( ! [X6,X8,X7,X5] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X7),s(fun(bool,bool),X5)))))) | s(bool,'T_0') = s(bool,X6) | s(bool,X8) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X5))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f336,f121])).
fof(f2211,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X1),s(fun(bool,bool),X2)))))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f425,f84])).
fof(f2226,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X6)))) | p(s(bool,'const.bool.IN_2'(s(bool,X7),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X5),s(fun(bool,bool),X6)))))) | s(bool,'T_0') = s(bool,X7) | s(bool,X8) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X5))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f2219,f77])).
fof(f2219,plain,(
  ( ! [X6,X8,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X6)))) | p(s(bool,'const.bool.IN_2'(s(bool,X7),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X5),s(fun(bool,bool),X6)))))) | s(bool,'T_0') = s(bool,X7) | s(bool,X8) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X5))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f425,f121])).
fof(f4196,plain,(
  ( ! [X280,X279,X281] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X279,sK2(X280,X281,X279)),s(fun(X279,bool),X280))) | s(fun(X279,bool),X280) = s(fun(X279,bool),X281) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X279,sK2(X280,X281,X279)),s(fun(X279,bool),X281)))) )),
  inference(forward_demodulation,[],[f3981,f144])).
fof(f3981,plain,(
  ( ! [X280,X282,X279,X281,X283] : (s(bool,'const.bool.IN_2'(s(X282,X283),s(fun(X282,bool),'const.pred_set.EMPTY_0'))) != s(bool,'const.bool.IN_2'(s(X279,sK2(X280,X281,X279)),s(fun(X279,bool),X280))) | s(fun(X279,bool),X280) = s(fun(X279,bool),X281) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X279,sK2(X280,X281,X279)),s(fun(X279,bool),X281)))) )),
  inference(superposition,[],[f109,f208])).
fof(f4197,plain,(
  ( ! [X285,X284,X286] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X284,sK2(X285,X286,X284)),s(fun(X284,bool),X286))) | s(fun(X284,bool),X285) = s(fun(X284,bool),X286) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X284,sK2(X285,X286,X284)),s(fun(X284,bool),X285)))) )),
  inference(forward_demodulation,[],[f3982,f144])).
fof(f3982,plain,(
  ( ! [X288,X285,X287,X284,X286] : (s(bool,'const.bool.IN_2'(s(X287,X288),s(fun(X287,bool),'const.pred_set.EMPTY_0'))) != s(bool,'const.bool.IN_2'(s(X284,sK2(X285,X286,X284)),s(fun(X284,bool),X286))) | s(fun(X284,bool),X285) = s(fun(X284,bool),X286) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X284,sK2(X285,X286,X284)),s(fun(X284,bool),X285)))) )),
  inference(superposition,[],[f109,f208])).
fof(f4202,plain,(
  ( ! [X316,X320,X319] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X319),s(fun(bool,bool),X320)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X316),s(fun(bool,bool),X320)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X316),s(fun(bool,bool),X319)))) | s(bool,'T_0') = s(bool,X316)) )),
  inference(forward_demodulation,[],[f3990,f144])).
fof(f3990,plain,(
  ( ! [X316,X318,X320,X317,X319] : (p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(X317,X318),s(fun(X317,bool),'const.pred_set.EMPTY_0'))),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X319),s(fun(bool,bool),X320)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X316),s(fun(bool,bool),X320)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X316),s(fun(bool,bool),X319)))) | s(bool,'T_0') = s(bool,X316)) )),
  inference(superposition,[],[f112,f208])).
fof(f4203,plain,(
  ( ! [X325,X321,X324] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X324),s(fun(bool,bool),X325)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X321),s(fun(bool,bool),X325)))) | s(bool,'T_0') = s(bool,X321)) )),
  inference(forward_demodulation,[],[f3991,f144])).
fof(f3991,plain,(
  ( ! [X325,X321,X323,X324,X322] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(X322,X323),s(fun(X322,bool),'const.pred_set.EMPTY_0'))),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X324),s(fun(bool,bool),X325)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X321),s(fun(bool,bool),X325)))) | s(bool,'T_0') = s(bool,X321)) )),
  inference(superposition,[],[f111,f208])).
fof(f4206,plain,(
  ( ! [X349,X348] : (s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X348))) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X349))) | s(fun(bool,bool),X348) = s(fun(bool,bool),X349) | s(bool,'T_0') = s(bool,sK2(X348,X349,bool))) )),
  inference(forward_demodulation,[],[f3995,f144])).
fof(f3995,plain,(
  ( ! [X350,X349,X351,X348] : (s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(X350,X351),s(fun(X350,bool),'const.pred_set.EMPTY_0'))),s(fun(bool,bool),X348))) != s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(X350,X351),s(fun(X350,bool),'const.pred_set.EMPTY_0'))),s(fun(bool,bool),X349))) | s(fun(bool,bool),X348) = s(fun(bool,bool),X349) | s(bool,'T_0') = s(bool,sK2(X348,X349,bool))) )),
  inference(superposition,[],[f109,f208])).
fof(f593,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,X87) != s(bool,f(s(X84,sK3(X85,X86,bool,X84)))) | s(fun(X84,bool),X85) = s(fun(X84,bool),X86) | s(bool,'T_0') = s(bool,X87) | s(bool,'T_0') = s(bool,g(s(X84,sK3(X85,X86,bool,X84))))) )),
  inference(superposition,[],[f117,f121])).
fof(f670,plain,(
  ( ! [X80,X78,X81,X79] : (s(bool,'T_0') != s(bool,f(s(X78,sK3(X79,X80,bool,X78)))) | s(fun(X78,bool),X79) = s(fun(X78,bool),X80) | s(bool,X81) = s(bool,g(s(X78,sK3(X79,X80,bool,X78)))) | s(bool,'T_0') = s(bool,X81)) )),
  inference(superposition,[],[f117,f121])).
fof(f1379,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK3(X1,X2,bool,X0)))) | s(bool,'T_0') = s(bool,f(s(X0,sK3(X1,X2,bool,X0))))) )),
  inference(trivial_inequality_removal,[],[f1372])).
fof(f1372,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK3(X1,X2,bool,X0)))) | s(bool,'T_0') = s(bool,f(s(X0,sK3(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f212,f84])).
fof(f109,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,sK2(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK2(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) )),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (s(bool,'const.bool.IN_2'(s(X0,sK2(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK2(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f71])).
fof(f71,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f41])).
fof(f41,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X7,X10] : (s(fun(X6,bool),X7) = s(fun(X6,bool),X10) <=> ! [X11] : s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X7))) = s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X10))))),
  file('thm.pred_set.DIFF_EMPTY','thm.pred_set.EXTENSION')).
fof(f108,plain,(
  ( ! [X4,X2,X0,X1] : (s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) )),
  inference(cnf_transformation,[],[f72])).
fof(f110,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(flattening,[],[f73])).
fof(f73,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(nnf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f43])).
fof(f43,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : ! [X7,X10,X11] : (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X10)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X7)))) & ~p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X10))))))),
  file('thm.pred_set.DIFF_EMPTY','thm.pred_set.IN_DIFF')).
fof(f111,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f74])).
fof(f112,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f74])).
fof(f865,plain,(
  ( ! [X47,X45,X43,X46,X44] : (~p(s(bool,'const.bool.IN_2'(s(X43,X44),s(fun(X43,bool),X46)))) | s(bool,X47) = s(bool,'const.bool.IN_2'(s(X43,X44),s(fun(X43,bool),'const.pred_set.DIFF_2'(s(fun(X43,bool),X45),s(fun(X43,bool),X46))))) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f662,f77])).
fof(f662,plain,(
  ( ! [X47,X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(X43,X44),s(fun(X43,bool),X46)))) | s(bool,X47) = s(bool,'const.bool.IN_2'(s(X43,X44),s(fun(X43,bool),'const.pred_set.DIFF_2'(s(fun(X43,bool),X45),s(fun(X43,bool),X46))))) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f111,f121])).
fof(f866,plain,(
  ( ! [X52,X50,X48,X51,X49] : (p(s(bool,'const.bool.IN_2'(s(X48,X49),s(fun(X48,bool),X50)))) | s(bool,X52) = s(bool,'const.bool.IN_2'(s(X48,X49),s(fun(X48,bool),'const.pred_set.DIFF_2'(s(fun(X48,bool),X50),s(fun(X48,bool),X51))))) | s(bool,'T_0') = s(bool,X52)) )),
  inference(subsumption_resolution,[],[f663,f77])).
fof(f663,plain,(
  ( ! [X52,X50,X48,X51,X49] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X48,X49),s(fun(X48,bool),X50)))) | s(bool,X52) = s(bool,'const.bool.IN_2'(s(X48,X49),s(fun(X48,bool),'const.pred_set.DIFF_2'(s(fun(X48,bool),X50),s(fun(X48,bool),X51))))) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f110,f121])).
fof(f4220,plain,(
  ( ! [X261,X263,X262,X264] : (p(s(bool,'const.bool.IN_2'(s(X261,X262),s(fun(X261,bool),X264)))) | ~p(s(bool,'const.bool.IN_2'(s(X261,X262),s(fun(X261,bool),X263)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X261,X262),s(fun(X261,bool),'const.pred_set.DIFF_2'(s(fun(X261,bool),X263),s(fun(X261,bool),X264)))))) )),
  inference(subsumption_resolution,[],[f4219,f120])).
fof(f4219,plain,(
  ( ! [X261,X263,X262,X264] : (p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(X261,X262),s(fun(X261,bool),X264)))) | ~p(s(bool,'const.bool.IN_2'(s(X261,X262),s(fun(X261,bool),X263)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X261,X262),s(fun(X261,bool),'const.pred_set.DIFF_2'(s(fun(X261,bool),X263),s(fun(X261,bool),X264)))))) )),
  inference(forward_demodulation,[],[f3978,f144])).
fof(f3978,plain,(
  ( ! [X261,X263,X265,X262,X264,X266] : (p(s(bool,'const.bool.IN_2'(s(X265,X266),s(fun(X265,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.bool.IN_2'(s(X261,X262),s(fun(X261,bool),X264)))) | ~p(s(bool,'const.bool.IN_2'(s(X261,X262),s(fun(X261,bool),X263)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X261,X262),s(fun(X261,bool),'const.pred_set.DIFF_2'(s(fun(X261,bool),X263),s(fun(X261,bool),X264)))))) )),
  inference(superposition,[],[f112,f208])).
fof(f637,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,'const.bool.IN_2'(s(bool,X23),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X24),s(fun(bool,bool),X25)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X25)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X24)))) | s(bool,X23) = s(bool,X26) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f112,f121])).
fof(f638,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,'const.bool.IN_2'(s(bool,X27),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X28),s(fun(bool,bool),X29)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X29)))) | s(bool,X27) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f111,f121])).
fof(f639,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'const.bool.IN_2'(s(bool,X31),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X32),s(fun(bool,bool),X33)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X32)))) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f110,f121])).
fof(f671,plain,(
  ( ! [X83,X84,X82] : (s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X82))) != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X83))) | s(fun(bool,bool),X82) = s(fun(bool,bool),X83) | s(bool,X84) = s(bool,sK2(X82,X83,bool)) | s(bool,'T_0') = s(bool,X84)) )),
  inference(superposition,[],[f109,f121])).
fof(f589,plain,(
  ( ! [X68,X66,X69,X67] : (p(s(bool,'const.bool.IN_2'(s(bool,X67),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X68),s(fun(bool,bool),X69)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X66),s(fun(bool,bool),X69)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X66),s(fun(bool,bool),X68)))) | s(bool,'T_0') = s(bool,X67) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f112,f121])).
fof(f590,plain,(
  ( ! [X70,X72,X71,X73] : (~p(s(bool,'const.bool.IN_2'(s(bool,X71),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X72),s(fun(bool,bool),X73)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X70),s(fun(bool,bool),X73)))) | s(bool,'T_0') = s(bool,X71) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f111,f121])).
fof(f591,plain,(
  ( ! [X76,X74,X77,X75] : (~p(s(bool,'const.bool.IN_2'(s(bool,X75),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X76),s(fun(bool,bool),X77)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X74),s(fun(bool,bool),X76)))) | s(bool,'T_0') = s(bool,X75) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f110,f121])).
fof(f594,plain,(
  ( ! [X90,X88,X89] : (s(bool,'const.bool.IN_2'(s(bool,X90),s(fun(bool,bool),X88))) != s(bool,'const.bool.IN_2'(s(bool,X90),s(fun(bool,bool),X89))) | s(fun(bool,bool),X88) = s(fun(bool,bool),X89) | s(bool,'T_0') = s(bool,X90) | s(bool,'T_0') = s(bool,sK2(X88,X89,bool))) )),
  inference(superposition,[],[f109,f121])).
fof(f666,plain,(
  ( ! [X61,X64,X62,X63] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X62),s(fun(bool,bool),X63)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X61),s(fun(bool,bool),X63)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X61),s(fun(bool,bool),X62)))) | s(bool,X61) = s(bool,X64) | s(bool,'T_0') = s(bool,X64)) )),
  inference(superposition,[],[f112,f121])).
fof(f667,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X66),s(fun(bool,bool),X67)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X65),s(fun(bool,bool),X67)))) | s(bool,X65) = s(bool,X68) | s(bool,'T_0') = s(bool,X68)) )),
  inference(superposition,[],[f111,f121])).
fof(f668,plain,(
  ( ! [X70,X72,X71,X69] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X70),s(fun(bool,bool),X71)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X69),s(fun(bool,bool),X70)))) | s(bool,X69) = s(bool,X72) | s(bool,'T_0') = s(bool,X72)) )),
  inference(superposition,[],[f110,f121])).
fof(f76,plain,(
  s(fun(sK0,bool),sK1) != s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  s(fun(sK0,bool),sK1) != s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f48])).
fof(f48,plain,(
  ? [X0,X1] : s(fun(X0,bool),X1) != s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))),
  inference(ennf_transformation,[],[f24])).
fof(f24,plain,(
  ~! [X0,X1] : s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))),
  inference(flattening,[],[f23])).
fof(f23,plain,(
  ~! [X0] : ! [X1] : s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : s(fun(X6,bool),X7) = s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),X7),s(fun(X6,bool),'const.pred_set.EMPTY_0')))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : s(fun(X6,bool),X7) = s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),X7),s(fun(X6,bool),'const.pred_set.EMPTY_0')))),
  file('thm.pred_set.DIFF_EMPTY','thm.pred_set.DIFF_EMPTY')).
fof(f5218,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl19),
  inference(unit_resulting_resolution,[],[f120,f1058,f158])).
fof(f5219,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl19),
  inference(unit_resulting_resolution,[],[f120,f1058,f113])).
fof(f5006,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))),
  inference(unit_resulting_resolution,[],[f76,f109])).
fof(f380,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))),
  inference(unit_resulting_resolution,[],[f76,f109])).
fof(f210,plain,(
  s(bool,f(s(sK0,sK3(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) != s(bool,g(s(sK0,sK3(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0))))),
  inference(unit_resulting_resolution,[],[f76,f117])).
fof(f209,plain,(
  s(bool,f(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) != s(bool,g(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0))))),
  inference(unit_resulting_resolution,[],[f76,f117])).
fof(f4039,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0'))) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')))) )),
  inference(equality_factoring,[],[f208])).
fof(f4000,plain,(
  ( ! [X6,X8,X7] : (s(bool,X8) != s(bool,'const.bool.IN_2'(s(X6,sK2('const.pred_set.EMPTY_0',X7,X6)),s(fun(X6,bool),X7))) | s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),X7) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f109,f208])).
fof(f3999,plain,(
  ( ! [X4,X5,X3] : (s(bool,X5) != s(bool,'const.bool.IN_2'(s(X3,sK2(X4,'const.pred_set.EMPTY_0',X3)),s(fun(X3,bool),X4))) | s(fun(X3,bool),'const.pred_set.EMPTY_0') = s(fun(X3,bool),X4) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f109,f208])).
fof(f3331,plain,(
  ( ! [X6] : (s(bool,sK2('const.pred_set.EMPTY_0',X6,bool)) != s(bool,'const.bool.IN_2'(s(bool,sK2('const.pred_set.EMPTY_0',X6,bool)),s(fun(bool,bool),X6))) | s(fun(bool,bool),'const.pred_set.EMPTY_0') = s(fun(bool,bool),X6) | s(bool,'T_0') = s(bool,sK2('const.pred_set.EMPTY_0',X6,bool))) )),
  inference(superposition,[],[f109,f190])).
fof(f3330,plain,(
  ( ! [X5] : (s(bool,sK2(X5,'const.pred_set.EMPTY_0',bool)) != s(bool,'const.bool.IN_2'(s(bool,sK2(X5,'const.pred_set.EMPTY_0',bool)),s(fun(bool,bool),X5))) | s(fun(bool,bool),'const.pred_set.EMPTY_0') = s(fun(bool,bool),X5) | s(bool,'T_0') = s(bool,sK2(X5,'const.pred_set.EMPTY_0',bool))) )),
  inference(superposition,[],[f109,f190])).
fof(f1097,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl23),
  inference(trivial_inequality_removal,[],[f1090])).
fof(f1090,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl23),
  inference(superposition,[],[f1061,f121])).
fof(f1088,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl23),
  inference(unit_resulting_resolution,[],[f77,f1061,f116])).
fof(f672,plain,(
  ( ! [X88,X87,X85,X86] : (s(X87,f(s(bool,'T_0'))) != s(X87,g(s(bool,'T_0'))) | s(fun(bool,X87),X85) = s(fun(bool,X87),X86) | s(bool,X88) = s(bool,sK3(X85,X86,X87,bool)) | s(bool,'T_0') = s(bool,X88)) )),
  inference(superposition,[],[f117,f121])).
fof(f665,plain,(
  ( ! [X59,X57,X60,X58] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X57,sK2(X58,X59,X57)),s(fun(X57,bool),X59))) | s(fun(X57,bool),X58) = s(fun(X57,bool),X59) | s(bool,X60) = s(bool,'const.bool.IN_2'(s(X57,sK2(X58,X59,X57)),s(fun(X57,bool),X58))) | s(bool,'T_0') = s(bool,X60)) )),
  inference(superposition,[],[f109,f121])).
fof(f664,plain,(
  ( ! [X54,X56,X55,X53] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X53,sK2(X54,X55,X53)),s(fun(X53,bool),X54))) | s(fun(X53,bool),X54) = s(fun(X53,bool),X55) | s(bool,X56) = s(bool,'const.bool.IN_2'(s(X53,sK2(X54,X55,X53)),s(fun(X53,bool),X55))) | s(bool,'T_0') = s(bool,X56)) )),
  inference(superposition,[],[f109,f121])).
fof(f864,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,X9)) | s(bool,X8) = s(bool,X9) | s(bool,X8) = s(bool,X10) | s(bool,'T_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f644,f77])).
fof(f644,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X9)) | s(bool,X8) = s(bool,X9) | s(bool,X8) = s(bool,X10) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f116,f121])).
fof(f595,plain,(
  ( ! [X94,X92,X93,X91] : (s(X93,f(s(bool,X94))) != s(X93,g(s(bool,X94))) | s(fun(bool,X93),X91) = s(fun(bool,X93),X92) | s(bool,'T_0') = s(bool,X94) | s(bool,'T_0') = s(bool,sK3(X91,X92,X93,bool))) )),
  inference(superposition,[],[f117,f121])).
fof(f588,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,X65) != s(bool,'const.bool.IN_2'(s(X62,sK2(X63,X64,X62)),s(fun(X62,bool),X64))) | s(fun(X62,bool),X63) = s(fun(X62,bool),X64) | s(bool,'T_0') = s(bool,X65) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X62,sK2(X63,X64,X62)),s(fun(X62,bool),X63)))) )),
  inference(superposition,[],[f109,f121])).
fof(f587,plain,(
  ( ! [X61,X59,X60,X58] : (s(bool,X61) != s(bool,'const.bool.IN_2'(s(X58,sK2(X59,X60,X58)),s(fun(X58,bool),X59))) | s(fun(X58,bool),X59) = s(fun(X58,bool),X60) | s(bool,'T_0') = s(bool,X61) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X58,sK2(X59,X60,X58)),s(fun(X58,bool),X60)))) )),
  inference(superposition,[],[f109,f121])).
fof(f532,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f121,f121])).
fof(f381,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))),
  inference(unit_resulting_resolution,[],[f76,f109])).
fof(f208,plain,(
  ( ! [X6,X4,X5] : (s(bool,X6) = s(bool,'const.bool.IN_2'(s(X4,X5),s(fun(X4,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f207,f120])).
fof(f207,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,'F_0')) | s(bool,X6) = s(bool,'const.bool.IN_2'(s(X4,X5),s(fun(X4,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f158,f144])).
fof(f190,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f144,f84])).
fof(f158,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f150,f120])).
fof(f150,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f113,f84])).
fof(f145,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f107,f107,f113])).
fof(f123,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f120,f84])).
fof(f121,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f84,f84])).
fof(f117,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK3(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK3(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK3(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK3(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f52])).
fof(f52,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.pred_set.DIFF_EMPTY',aHL_EXT)).
fof(f113,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X8,X9] : ((p(s(bool,X8)) => p(s(bool,X9))) => ((p(s(bool,X9)) => p(s(bool,X8))) => s(bool,X8) = s(bool,X9)))),
  file('thm.pred_set.DIFF_EMPTY','thm.bool.IMP_ANTISYM_AX')).
fof(f116,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f51])).
fof(f107,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(cnf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0,X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ! [X0] : ! [X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6] : ! [X11] : ~p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),'const.pred_set.EMPTY_0'))))),
  file('thm.pred_set.DIFF_EMPTY','thm.pred_set.NOT_IN_EMPTY')).
fof(f84,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.DIFF_EMPTY',aHL_BOOL_CASES)).
fof(f77,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.DIFF_EMPTY',aHL_TRUTH)).
fof(f8027,plain,(
  $false | ($spl17 | $spl19 | $spl23)),
  inference(global_subsumption,[],[f7386,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786,f1668,f1511,f425,f387,f382,f336,f260,f259,f212,f211,f144,f120])).
fof(f7386,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl17),
  inference(trivial_inequality_removal,[],[f7377])).
fof(f7377,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl17),
  inference(superposition,[],[f1055,f121])).
fof(f1055,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl17),
  inference(cnf_transformation,[],[f1055_D])).
fof(f1055_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f8026,plain,(
  $false | ($spl17 | $spl19 | $spl23)),
  inference(global_subsumption,[],[f7385,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786,f1668,f1511,f425,f387,f382,f336,f260,f259,f212,f211,f144,f120])).
fof(f7385,plain,(
  ( ! [X3] : (s(bool,X3) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | s(bool,'T_0') = s(bool,X3)) ) | $spl17),
  inference(trivial_inequality_removal,[],[f7380])).
fof(f7380,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X3) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | s(bool,'T_0') = s(bool,X3)) ) | $spl17),
  inference(superposition,[],[f1055,f121])).
fof(f8025,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786,f1668,f1511,f425,f387,f382,f336,f260,f259,f212,f211,f144,f120])).
fof(f8024,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f120,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786,f1668,f1511,f425,f387,f382,f336,f260,f259,f212,f211,f144])).
fof(f8023,plain,(
  $false | ($spl6 | $spl19 | $spl23)),
  inference(global_subsumption,[],[f157,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786,f1668,f1511,f425,f387,f382,f336,f260,f259,f212,f211])).
fof(f8022,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786,f1668,f1511,f425,f387,f382,f336,f260,f259,f212,f211])).
fof(f8021,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786,f1668,f1511,f425,f387,f382,f336,f260,f259,f212])).
fof(f8020,plain,(
  $false | ($spl10 | $spl19 | $spl23)),
  inference(global_subsumption,[],[f216,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786,f1668,f1511,f425,f387,f382,f336,f260,f259])).
fof(f216,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) | $spl10),
  inference(cnf_transformation,[],[f216_D])).
fof(f216_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f8019,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786,f1668,f1511,f425,f387,f382,f336,f260,f259])).
fof(f8018,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786,f1668,f1511,f425,f387,f382,f336,f260])).
fof(f8017,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786,f1668,f1511,f425,f387,f382,f336])).
fof(f8016,plain,(
  $false | ($spl14 | $spl19 | $spl23)),
  inference(global_subsumption,[],[f378,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786,f1668,f1511,f425,f387,f382])).
fof(f378,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) | $spl14),
  inference(cnf_transformation,[],[f378_D])).
fof(f378_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f8015,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786,f1668,f1511,f425,f387,f382])).
fof(f8014,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786,f1668,f1511,f425,f387])).
fof(f8013,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786,f1668,f1511,f425])).
fof(f8012,plain,(
  $false | ($spl19 | $spl23 | $spl36)),
  inference(global_subsumption,[],[f1262,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786,f1668,f1511])).
fof(f8011,plain,(
  $false | ($spl19 | $spl23 | $spl45)),
  inference(global_subsumption,[],[f1268,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786,f1668,f1511])).
fof(f8010,plain,(
  $false | ($spl19 | $spl23 | $spl45)),
  inference(global_subsumption,[],[f1346,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786,f1668,f1511])).
fof(f1346,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl45),
  inference(trivial_inequality_removal,[],[f1341])).
fof(f1341,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl45),
  inference(superposition,[],[f1268,f121])).
fof(f8009,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786,f1668,f1511])).
fof(f8008,plain,(
  $false | ($spl19 | $spl23 | $spl51)),
  inference(global_subsumption,[],[f1602,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786,f1668])).
fof(f1602,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl51),
  inference(superposition,[],[f1384,f121])).
fof(f8007,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786,f1668])).
fof(f8006,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786])).
fof(f8005,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593,f4206,f4203,f4202,f4197,f4196,f2226,f2211])).
fof(f8004,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593,f4206,f4203,f4202,f4197,f4196,f2226])).
fof(f8003,plain,(
  $false | ($spl19 | $spl23 | $spl45)),
  inference(global_subsumption,[],[f4103,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593,f4206,f4203,f4202,f4197,f4196])).
fof(f4103,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | $spl45),
  inference(forward_demodulation,[],[f3833,f144])).
fof(f3833,plain,(
  ( ! [X0,X1] : (s(bool,f(s(bool,'F_0'))) = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | $spl45),
  inference(unit_resulting_resolution,[],[f1268,f208])).
fof(f8002,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593,f4206,f4203,f4202,f4197,f4196])).
fof(f8001,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593,f4206,f4203,f4202,f4197])).
fof(f8000,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593,f4206,f4203,f4202])).
fof(f7999,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593,f4206,f4203])).
fof(f7998,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593,f4206])).
fof(f7997,plain,(
  $false | ($spl8 | $spl19 | $spl23 | $spl45)),
  inference(global_subsumption,[],[f6578,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593])).
fof(f6578,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,X1)) ) | ($spl8 | $spl45)),
  inference(condensation,[],[f6577])).
fof(f6577,plain,(
  ( ! [X39,X37,X38] : (s(bool,X38) != s(bool,f(s(bool,'F_0'))) | s(bool,X38) = s(bool,X39) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,X37)) ) | ($spl8 | $spl45)),
  inference(duplicate_literal_removal,[],[f6495])).
fof(f6495,plain,(
  ( ! [X39,X37,X38] : (s(bool,X38) != s(bool,f(s(bool,'F_0'))) | s(bool,X38) = s(bool,X39) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) ) | ($spl8 | $spl45)),
  inference(superposition,[],[f5442,f1346])).
fof(f5442,plain,(
  ( ! [X6,X4,X5] : (s(bool,X4) != s(bool,X5) | s(bool,X4) = s(bool,X6) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) ) | $spl8),
  inference(superposition,[],[f5063,f208])).
fof(f5063,plain,(
  ( ! [X12,X13] : (s(bool,X12) != s(bool,'const.bool.IN_2'(s(bool,X12),s(fun(bool,bool),'const.pred_set.EMPTY_0'))) | s(bool,X12) = s(bool,X13) | s(bool,'T_0') = s(bool,X13)) ) | $spl8),
  inference(superposition,[],[f5004,f121])).
fof(f5004,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X4,X5),s(fun(X4,bool),'const.pred_set.EMPTY_0')))) ) | $spl8),
  inference(subsumption_resolution,[],[f4994,f4295])).
fof(f4295,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) | $spl8),
  inference(subsumption_resolution,[],[f4279,f76])).
fof(f4279,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | $spl8),
  inference(superposition,[],[f117,f215])).
fof(f4994,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X4,X5),s(fun(X4,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0))))) ) | $spl8),
  inference(superposition,[],[f4295,f208])).
fof(f7996,plain,(
  $false | ($spl6 | $spl19 | $spl23 | $spl45)),
  inference(global_subsumption,[],[f7070,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593])).
fof(f7070,plain,(
  ( ! [X25] : (s(bool,'F_0') = s(bool,f(s(bool,X25))) | s(bool,'T_0') = s(bool,X25)) ) | ($spl6 | $spl45)),
  inference(subsumption_resolution,[],[f7019,f123])).
fof(f7019,plain,(
  ( ! [X25] : (p(s(bool,X25)) | s(bool,'F_0') = s(bool,f(s(bool,X25))) | s(bool,'T_0') = s(bool,X25)) ) | ($spl6 | $spl45)),
  inference(superposition,[],[f157,f4227])).
fof(f4227,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl45),
  inference(superposition,[],[f4103,f84])).
fof(f7995,plain,(
  $false | ($spl6 | $spl19 | $spl23 | $spl45)),
  inference(global_subsumption,[],[f7157,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593])).
fof(f7157,plain,(
  ( ! [X68,X69,X67] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X68),s(fun(bool,bool),X69)))))) | p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X67))),s(fun(bool,bool),X68)))) | s(bool,'T_0') = s(bool,X67)) ) | ($spl6 | $spl45)),
  inference(superposition,[],[f110,f7070])).
fof(f7994,plain,(
  $false | ($spl6 | $spl19 | $spl23 | $spl45)),
  inference(global_subsumption,[],[f7158,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593])).
fof(f7158,plain,(
  ( ! [X70,X72,X71] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X71),s(fun(bool,bool),X72)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X70))),s(fun(bool,bool),X72)))) | s(bool,'T_0') = s(bool,X70)) ) | ($spl6 | $spl45)),
  inference(superposition,[],[f111,f7070])).
fof(f7993,plain,(
  $false | ($spl6 | $spl19 | $spl23 | $spl45)),
  inference(global_subsumption,[],[f7159,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593])).
fof(f7159,plain,(
  ( ! [X74,X75,X73] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X74),s(fun(bool,bool),X75)))))) | p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X73))),s(fun(bool,bool),X75)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X73))),s(fun(bool,bool),X74)))) | s(bool,'T_0') = s(bool,X73)) ) | ($spl6 | $spl45)),
  inference(superposition,[],[f112,f7070])).
fof(f7992,plain,(
  $false | ($spl3 | $spl8 | $spl17 | $spl19 | $spl23 | $spl45)),
  inference(global_subsumption,[],[f7339,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593])).
fof(f7339,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | ($spl3 | $spl8 | $spl17 | $spl45)),
  inference(unit_resulting_resolution,[],[f132,f4103,f1055,f6704])).
fof(f6704,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X2) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) ) | ($spl8 | $spl45)),
  inference(superposition,[],[f6578,f84])).
fof(f7991,plain,(
  $false | ($spl3 | $spl8 | $spl19 | $spl23 | $spl27 | $spl45)),
  inference(global_subsumption,[],[f7509,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593])).
fof(f7509,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | ($spl3 | $spl8 | $spl27 | $spl45)),
  inference(unit_resulting_resolution,[],[f132,f4103,f1155,f6704])).
fof(f1155,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl27),
  inference(cnf_transformation,[],[f1155_D])).
fof(f1155_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f7990,plain,(
  $false | ($spl19 | $spl23 | $spl45)),
  inference(global_subsumption,[],[f7038,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593])).
fof(f7038,plain,(
  ( ! [X66,X67,X65] : (~p(s(bool,'const.bool.IN_2'(s(bool,X65),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X66),s(fun(bool,bool),X67)))))) | p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X65))),s(fun(bool,bool),X66)))) | s(bool,'T_0') = s(bool,X65)) ) | $spl45),
  inference(superposition,[],[f110,f4227])).
fof(f7989,plain,(
  $false | ($spl19 | $spl23 | $spl45)),
  inference(global_subsumption,[],[f7039,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593])).
fof(f7039,plain,(
  ( ! [X70,X68,X69] : (~p(s(bool,'const.bool.IN_2'(s(bool,X68),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X69),s(fun(bool,bool),X70)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X68))),s(fun(bool,bool),X70)))) | s(bool,'T_0') = s(bool,X68)) ) | $spl45),
  inference(superposition,[],[f111,f4227])).
fof(f7988,plain,(
  $false | ($spl19 | $spl23 | $spl45)),
  inference(global_subsumption,[],[f7040,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593])).
fof(f7040,plain,(
  ( ! [X72,X71,X73] : (p(s(bool,'const.bool.IN_2'(s(bool,X71),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X72),s(fun(bool,bool),X73)))))) | p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X71))),s(fun(bool,bool),X73)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,f(s(bool,X71))),s(fun(bool,bool),X72)))) | s(bool,'T_0') = s(bool,X71)) ) | $spl45),
  inference(superposition,[],[f112,f4227])).
fof(f7987,plain,(
  $false | ($spl19 | $spl23 | $spl45)),
  inference(global_subsumption,[],[f7069,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593])).
fof(f7069,plain,(
  ( ! [X10,X11] : (p(s(bool,X11)) | s(bool,X11) = s(bool,f(s(bool,X10))) | s(bool,'T_0') = s(bool,X10)) ) | $spl45),
  inference(subsumption_resolution,[],[f7011,f123])).
fof(f7011,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,X11) = s(bool,f(s(bool,X10))) | s(bool,'T_0') = s(bool,X10)) ) | $spl45),
  inference(superposition,[],[f113,f4227])).
fof(f7986,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670,f593])).
fof(f7985,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f593,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379,f670])).
fof(f7984,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f670,f593,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109,f1379])).
fof(f7983,plain,(
  $false | ($spl19 | $spl23 | $spl36)),
  inference(global_subsumption,[],[f6315,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f1379,f670,f593,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109])).
fof(f6315,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X51),s(fun(bool,bool),X52)))))) | p(s(bool,'const.bool.IN_2'(s(bool,g(s(bool,X50))),s(fun(bool,bool),X52)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,g(s(bool,X50))),s(fun(bool,bool),X51)))) | s(bool,'T_0') = s(bool,X50)) ) | $spl36),
  inference(superposition,[],[f112,f1270])).
fof(f1270,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,g(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl36),
  inference(superposition,[],[f1262,f84])).
fof(f7982,plain,(
  $false | ($spl19 | $spl23 | $spl36)),
  inference(global_subsumption,[],[f6314,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f1379,f670,f593,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109])).
fof(f6314,plain,(
  ( ! [X47,X48,X49] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X48),s(fun(bool,bool),X49)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,g(s(bool,X47))),s(fun(bool,bool),X49)))) | s(bool,'T_0') = s(bool,X47)) ) | $spl36),
  inference(superposition,[],[f111,f1270])).
fof(f7981,plain,(
  $false | ($spl19 | $spl23 | $spl36)),
  inference(global_subsumption,[],[f6313,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f1379,f670,f593,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109])).
fof(f6313,plain,(
  ( ! [X45,X46,X44] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.DIFF_2'(s(fun(bool,bool),X45),s(fun(bool,bool),X46)))))) | p(s(bool,'const.bool.IN_2'(s(bool,g(s(bool,X44))),s(fun(bool,bool),X45)))) | s(bool,'T_0') = s(bool,X44)) ) | $spl36),
  inference(superposition,[],[f110,f1270])).
fof(f7980,plain,(
  $false | ($spl19 | $spl23 | $spl36)),
  inference(global_subsumption,[],[f6349,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f1379,f670,f593,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109])).
fof(f6349,plain,(
  ( ! [X6,X5] : (~p(s(bool,X6)) | s(bool,X6) = s(bool,g(s(bool,X5))) | s(bool,'T_0') = s(bool,X5)) ) | $spl36),
  inference(subsumption_resolution,[],[f6295,f77])).
fof(f6295,plain,(
  ( ! [X6,X5] : (~p(s(bool,'T_0')) | ~p(s(bool,X6)) | s(bool,X6) = s(bool,g(s(bool,X5))) | s(bool,'T_0') = s(bool,X5)) ) | $spl36),
  inference(superposition,[],[f116,f1270])).
fof(f7979,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f1379,f670,f593,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108,f109])).
fof(f7978,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f1379,f670,f593,f109,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110,f108])).
fof(f7977,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f1379,f670,f593,f109,f108,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111,f110])).
fof(f7976,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f1379,f670,f593,f109,f108,f110,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112,f111])).
fof(f7975,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f1379,f670,f593,f109,f108,f110,f111,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865,f112])).
fof(f7974,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f1379,f670,f593,f109,f108,f110,f111,f112,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866,f865])).
fof(f7973,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f1379,f670,f593,f109,f108,f110,f111,f112,f865,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220,f866])).
fof(f7972,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f1379,f670,f593,f109,f108,f110,f111,f112,f865,f866,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637,f4220])).
fof(f7971,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f1379,f670,f593,f109,f108,f110,f111,f112,f865,f866,f4220,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638,f637])).
fof(f7970,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f1379,f670,f593,f109,f108,f110,f111,f112,f865,f866,f4220,f637,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639,f638])).
fof(f7969,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f1379,f670,f593,f109,f108,f110,f111,f112,f865,f866,f4220,f638,f637,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671,f639])).
fof(f7968,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f1379,f670,f593,f109,f108,f110,f111,f112,f865,f866,f4220,f639,f638,f637,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589,f671])).
fof(f7967,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f1379,f670,f593,f109,f108,f110,f111,f112,f865,f866,f4220,f671,f639,f638,f637,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590,f589])).
fof(f7966,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f1379,f670,f593,f109,f108,f110,f111,f112,f865,f866,f4220,f671,f639,f638,f637,f589,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591,f590])).
fof(f7965,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f1379,f670,f593,f109,f108,f110,f111,f112,f865,f866,f4220,f671,f639,f638,f637,f589,f590,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594,f591])).
fof(f7964,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f1379,f670,f593,f109,f108,f110,f111,f112,f865,f866,f4220,f671,f639,f638,f637,f589,f590,f591,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666,f594])).
fof(f7963,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f1379,f670,f593,f109,f108,f110,f111,f112,f865,f866,f4220,f671,f639,f638,f637,f589,f590,f591,f594,f380,f1058,f5006,f5219,f5218,f76,f668,f667,f666])).
fof(f7962,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f1379,f670,f593,f109,f108,f110,f111,f112,f865,f866,f4220,f671,f639,f638,f637,f589,f590,f591,f594,f666,f380,f1058,f5006,f5219,f5218,f76,f668,f667])).
fof(f7961,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f1379,f670,f593,f109,f108,f110,f111,f112,f865,f866,f4220,f671,f639,f638,f637,f589,f590,f591,f594,f666,f667,f380,f1058,f5006,f5219,f5218,f76,f668])).
fof(f7960,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f1379,f670,f593,f109,f108,f110,f111,f112,f865,f866,f4220,f671,f639,f638,f637,f589,f590,f591,f594,f666,f667,f668,f380,f1058,f5006,f5219,f5218,f76])).
fof(f7959,plain,(
  $false | ($spl17 | $spl19 | $spl23)),
  inference(global_subsumption,[],[f1055,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f1379,f670,f593,f109,f108,f110,f111,f112,f865,f866,f4220,f671,f639,f638,f637,f589,f590,f591,f594,f666,f667,f668,f76,f380,f1058,f5006,f5219,f5218])).
fof(f7958,plain,(
  $false | ($spl19 | $spl23 | $spl27)),
  inference(global_subsumption,[],[f1155,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f1379,f670,f593,f109,f108,f110,f111,f112,f865,f866,f4220,f671,f639,f638,f637,f589,f590,f591,f594,f666,f667,f668,f76,f380,f1058,f5006,f5219,f5218])).
fof(f7957,plain,(
  $false | ($spl19 | $spl23 | $spl30)),
  inference(global_subsumption,[],[f1159,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f1379,f670,f593,f109,f108,f110,f111,f112,f865,f866,f4220,f671,f639,f638,f637,f589,f590,f591,f594,f666,f667,f668,f76,f380,f1058,f5006,f5219,f5218])).
fof(f7956,plain,(
  $false | ($spl8 | $spl19 | $spl23)),
  inference(global_subsumption,[],[f215,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f1379,f670,f593,f109,f108,f110,f111,f112,f865,f866,f4220,f671,f639,f638,f637,f589,f590,f591,f594,f666,f667,f668,f76,f380,f1058,f5006,f5219,f5218])).
fof(f7955,plain,(
  $false | ($spl12 | $spl19 | $spl23)),
  inference(global_subsumption,[],[f377,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f1379,f670,f593,f109,f108,f110,f111,f112,f865,f866,f4220,f671,f639,f638,f637,f589,f590,f591,f594,f666,f667,f668,f76,f380,f1058,f5006,f5219,f5218])).
fof(f7954,plain,(
  $false | ($spl19 | $spl23 | $spl24)),
  inference(global_subsumption,[],[f1062,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f1379,f670,f593,f109,f108,f110,f111,f112,f865,f866,f4220,f671,f639,f638,f637,f589,f590,f591,f594,f666,f667,f668,f76,f380,f1058,f5006,f5219,f5218])).
fof(f7953,plain,(
  $false | ($spl19 | $spl23 | $spl34)),
  inference(global_subsumption,[],[f1162,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f1379,f670,f593,f109,f108,f110,f111,f112,f865,f866,f4220,f671,f639,f638,f637,f589,f590,f591,f594,f666,f667,f668,f76,f380,f1058,f5006,f5219,f5218])).
fof(f7952,plain,(
  $false | ($spl8 | $spl19 | $spl23)),
  inference(global_subsumption,[],[f4295,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f1379,f670,f593,f109,f108,f110,f111,f112,f865,f866,f4220,f671,f639,f638,f637,f589,f590,f591,f594,f666,f667,f668,f76,f380,f1058,f5006,f5219,f5218])).
fof(f7951,plain,(
  $false | ($spl3 | $spl10 | $spl19 | $spl23)),
  inference(global_subsumption,[],[f4326,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f1379,f670,f593,f109,f108,f110,f111,f112,f865,f866,f4220,f671,f639,f638,f637,f589,f590,f591,f594,f666,f667,f668,f76,f380,f1058,f5006,f5219,f5218])).
fof(f4326,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl3 | $spl10)),
  inference(subsumption_resolution,[],[f4301,f132])).
fof(f4301,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X0) = s(bool,f(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl10),
  inference(superposition,[],[f216,f121])).
fof(f7950,plain,(
  $false | ($spl12 | $spl19 | $spl23)),
  inference(global_subsumption,[],[f4367,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f1379,f670,f593,f109,f108,f110,f111,f112,f865,f866,f4220,f671,f639,f638,f637,f589,f590,f591,f594,f666,f667,f668,f76,f380,f1058,f5006,f5219,f5218])).
fof(f4367,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) | $spl12),
  inference(subsumption_resolution,[],[f4351,f76])).
fof(f4351,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | $spl12),
  inference(superposition,[],[f117,f377])).
fof(f7949,plain,(
  $false | ($spl3 | $spl14 | $spl19 | $spl23)),
  inference(global_subsumption,[],[f4398,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f1379,f670,f593,f109,f108,f110,f111,f112,f865,f866,f4220,f671,f639,f638,f637,f589,f590,f591,f594,f666,f667,f668,f76,f380,f1058,f5006,f5219,f5218])).
fof(f4398,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl3 | $spl14)),
  inference(subsumption_resolution,[],[f4373,f132])).
fof(f4373,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X0) = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(superposition,[],[f378,f121])).
fof(f7948,plain,(
  $false | ($spl19 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f1379,f670,f593,f109,f108,f110,f111,f112,f865,f866,f4220,f671,f639,f638,f637,f589,f590,f591,f594,f666,f667,f668,f76,f380,f1058,f5006,f5219,f5218])).
fof(f7947,plain,(
  $false | ($spl19 | $spl23 | $spl29)),
  inference(global_subsumption,[],[f5303,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f1379,f670,f593,f109,f108,f110,f111,f112,f865,f866,f4220,f671,f639,f638,f637,f589,f590,f591,f594,f666,f667,f668,f76,f380,f1058,f5006,f5219,f5218])).
fof(f5303,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | $spl29),
  inference(unit_resulting_resolution,[],[f120,f1158,f158])).
fof(f7946,plain,(
  $false | ($spl3 | $spl8 | $spl10 | $spl19 | $spl23)),
  inference(global_subsumption,[],[f5978,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f1379,f670,f593,f109,f108,f110,f111,f112,f865,f866,f4220,f671,f639,f638,f637,f589,f590,f591,f594,f666,f667,f668,f76,f380,f1058,f5006,f5219,f5218])).
fof(f5978,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) != s(bool,f(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) | s(bool,X0) = s(bool,X1)) ) | ($spl3 | $spl8 | $spl10)),
  inference(condensation,[],[f5977])).
fof(f5977,plain,(
  ( ! [X30,X31,X32] : (s(bool,X31) != s(bool,f(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) | s(bool,X31) = s(bool,X32) | s(bool,'T_0') = s(bool,X32) | s(bool,'T_0') = s(bool,X30)) ) | ($spl3 | $spl8 | $spl10)),
  inference(duplicate_literal_removal,[],[f5909])).
fof(f5909,plain,(
  ( ! [X30,X31,X32] : (s(bool,X31) != s(bool,f(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) | s(bool,X31) = s(bool,X32) | s(bool,'T_0') = s(bool,X32) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) ) | ($spl3 | $spl8 | $spl10)),
  inference(superposition,[],[f5442,f4326])).
fof(f7945,plain,(
  $false | ($spl3 | $spl8 | $spl14 | $spl19 | $spl23)),
  inference(global_subsumption,[],[f6195,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f1061,f144,f120,f211,f212,f259,f260,f336,f382,f387,f425,f1511,f1668,f1786,f2211,f2226,f4196,f4197,f4202,f4203,f4206,f1379,f670,f593,f109,f108,f110,f111,f112,f865,f866,f4220,f671,f639,f638,f637,f589,f590,f591,f594,f666,f667,f668,f76,f380,f1058,f5006,f5219,f5218])).
fof(f6195,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) | s(bool,X0) = s(bool,X1)) ) | ($spl3 | $spl8 | $spl14)),
  inference(condensation,[],[f6194])).
fof(f6194,plain,(
  ( ! [X35,X36,X34] : (s(bool,X35) != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36) | s(bool,'T_0') = s(bool,X34)) ) | ($spl3 | $spl8 | $spl14)),
  inference(duplicate_literal_removal,[],[f6121])).
fof(f6121,plain,(
  ( ! [X35,X36,X34] : (s(bool,X35) != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36) | s(bool,'T_0') = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) ) | ($spl3 | $spl8 | $spl14)),
  inference(superposition,[],[f5442,f4398])).
fof(f7941,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f7940,f77])).
fof(f7940,plain,(
  ~p(s(bool,'T_0')) | ($spl19 | $spl20)),
  inference(forward_demodulation,[],[f7939,f7842])).
fof(f7842,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2)) ) | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f7841,f77])).
fof(f7841,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,X2)) ) | ($spl19 | $spl20)),
  inference(forward_demodulation,[],[f7840,f5218])).
fof(f7840,plain,(
  ( ! [X2] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl20),
  inference(subsumption_resolution,[],[f7839,f120])).
fof(f7839,plain,(
  ( ! [X2] : (p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl20),
  inference(forward_demodulation,[],[f7838,f144])).
fof(f7838,plain,(
  ( ! [X2] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl20),
  inference(subsumption_resolution,[],[f7711,f123])).
fof(f7711,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl20),
  inference(superposition,[],[f112,f1059])).
fof(f7939,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))) ) | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f7891,f77])).
fof(f7891,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))) ) | ($spl19 | $spl20)),
  inference(backward_demodulation,[],[f7842,f111])).
fof(f7937,plain,(
  $false | ($spl19 | $spl20 | $spl45)),
  inference(subsumption_resolution,[],[f7857,f7842])).
fof(f7857,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl19 | $spl20 | $spl45)),
  inference(backward_demodulation,[],[f7842,f1268])).
fof(f7936,plain,(
  $false | ($spl12 | $spl19 | $spl20)),
  inference(subsumption_resolution,[],[f7928,f7842])).
fof(f7928,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) | ($spl12 | $spl19 | $spl20)),
  inference(backward_demodulation,[],[f7915,f4367])).
fof(f7915,plain,(
  ( ! [X0,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) ) | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f7853,f7842])).
fof(f7853,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK2(X1,'const.pred_set.EMPTY_0',X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) ) | ($spl19 | $spl20)),
  inference(backward_demodulation,[],[f7842,f382])).
fof(f7935,plain,(
  $false | ($spl8 | $spl19 | $spl20)),
  inference(subsumption_resolution,[],[f7926,f7842])).
fof(f7926,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) | ($spl8 | $spl19 | $spl20)),
  inference(backward_demodulation,[],[f7915,f4295])).
fof(f7934,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f7916,f7915])).
fof(f7916,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),sK1) | ($spl19 | $spl20)),
  inference(backward_demodulation,[],[f7915,f76])).
fof(f7914,plain,(
  $false | ($spl19 | $spl20)),
  inference(subsumption_resolution,[],[f7843,f77])).
fof(f7843,plain,(
  ~p(s(bool,'T_0')) | ($spl19 | $spl20)),
  inference(backward_demodulation,[],[f7842,f120])).
fof(f5211,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f5210,f120])).
fof(f5210,plain,(
  p(s(bool,'F_0')) | $spl26),
  inference(forward_demodulation,[],[f5209,f144])).
fof(f5209,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.bool.IN_2'(s(X4,X5),s(fun(X4,bool),'const.pred_set.EMPTY_0'))))) ) | $spl26),
  inference(subsumption_resolution,[],[f5197,f2926])).
fof(f2926,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | $spl26),
  inference(subsumption_resolution,[],[f2871,f76])).
fof(f2871,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | $spl26),
  inference(superposition,[],[f109,f1156])).
fof(f5197,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.bool.IN_2'(s(X4,X5),s(fun(X4,bool),'const.pred_set.EMPTY_0')))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1)))) ) | $spl26),
  inference(superposition,[],[f2915,f208])).
fof(f2915,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1)))) | $spl26),
  inference(subsumption_resolution,[],[f2874,f77])).
fof(f2874,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1)))) | $spl26),
  inference(superposition,[],[f110,f1156])).
fof(f5208,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f5207,f2926])).
fof(f5207,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | $spl26),
  inference(condensation,[],[f5206])).
fof(f5206,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X2)) ) | $spl26),
  inference(subsumption_resolution,[],[f5195,f123])).
fof(f5195,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X2)) ) | $spl26),
  inference(superposition,[],[f2915,f121])).
fof(f5205,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f5204,f2926])).
fof(f5204,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | $spl26),
  inference(condensation,[],[f5203])).
fof(f5203,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1)))) ) | $spl26),
  inference(subsumption_resolution,[],[f5194,f123])).
fof(f5194,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1)))) ) | $spl26),
  inference(superposition,[],[f2915,f121])).
fof(f5202,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f5201,f2926])).
fof(f5201,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | $spl26),
  inference(subsumption_resolution,[],[f5192,f120])).
fof(f5192,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | $spl26),
  inference(superposition,[],[f2915,f84])).
fof(f5200,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f5187,f2926])).
fof(f5187,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | $spl26),
  inference(unit_resulting_resolution,[],[f77,f2915,f116])).
fof(f5199,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f5189,f2926])).
fof(f5189,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | $spl26),
  inference(unit_resulting_resolution,[],[f77,f2915,f116])).
fof(f5198,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f5191,f2926])).
fof(f5191,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | $spl26),
  inference(unit_resulting_resolution,[],[f2915,f123])).
fof(f4677,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f380,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638,f639,f671,f4220,f866,f865,f112,f111,f110,f108,f109,f593,f670,f1379,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786,f1668,f1511,f425,f387,f382,f336,f260,f259,f212,f211,f120,f144,f1061])).
fof(f2542,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl16),
  inference(subsumption_resolution,[],[f2503,f77])).
fof(f2503,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl16),
  inference(superposition,[],[f110,f1056])).
fof(f2553,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl16),
  inference(subsumption_resolution,[],[f2500,f76])).
fof(f2500,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | $spl16),
  inference(superposition,[],[f109,f1056])).
fof(f4676,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f190,f208,f213,f375,f381,f380,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638,f639,f671,f4220,f866,f865,f112,f111,f110,f108,f109,f593,f670,f1379,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786,f1668,f1511,f425,f387,f382,f336,f260,f259,f212,f211,f120,f144])).
fof(f4675,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f144,f190,f208,f213,f375,f381,f380,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638,f639,f671,f4220,f866,f865,f112,f111,f110,f108,f109,f593,f670,f1379,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786,f1668,f1511,f425,f387,f382,f336,f260,f259,f212,f211,f120])).
fof(f4674,plain,(
  $false | ($spl6 | $spl16 | $spl23)),
  inference(global_subsumption,[],[f157,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f120,f144,f190,f208,f213,f375,f381,f380,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638,f639,f671,f4220,f866,f865,f112,f111,f110,f108,f109,f593,f670,f1379,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786,f1668,f1511,f425,f387,f382,f336,f260,f259,f212,f211])).
fof(f4673,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f120,f144,f190,f208,f213,f375,f381,f380,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638,f639,f671,f4220,f866,f865,f112,f111,f110,f108,f109,f593,f670,f1379,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786,f1668,f1511,f425,f387,f382,f336,f260,f259,f212,f211])).
fof(f4672,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f120,f144,f190,f208,f211,f213,f375,f381,f380,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638,f639,f671,f4220,f866,f865,f112,f111,f110,f108,f109,f593,f670,f1379,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786,f1668,f1511,f425,f387,f382,f336,f260,f259,f212])).
fof(f4671,plain,(
  $false | ($spl10 | $spl16 | $spl23)),
  inference(global_subsumption,[],[f216,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f120,f144,f190,f208,f211,f212,f213,f375,f381,f380,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638,f639,f671,f4220,f866,f865,f112,f111,f110,f108,f109,f593,f670,f1379,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786,f1668,f1511,f425,f387,f382,f336,f260,f259])).
fof(f4670,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f120,f144,f190,f208,f211,f212,f213,f375,f381,f380,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638,f639,f671,f4220,f866,f865,f112,f111,f110,f108,f109,f593,f670,f1379,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786,f1668,f1511,f425,f387,f382,f336,f260,f259])).
fof(f4669,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f120,f144,f190,f208,f211,f212,f213,f259,f375,f381,f380,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638,f639,f671,f4220,f866,f865,f112,f111,f110,f108,f109,f593,f670,f1379,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786,f1668,f1511,f425,f387,f382,f336,f260])).
fof(f4668,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f120,f144,f190,f208,f211,f212,f213,f259,f260,f375,f381,f380,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638,f639,f671,f4220,f866,f865,f112,f111,f110,f108,f109,f593,f670,f1379,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786,f1668,f1511,f425,f387,f382,f336])).
fof(f4667,plain,(
  $false | ($spl14 | $spl16 | $spl23)),
  inference(global_subsumption,[],[f378,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638,f639,f671,f4220,f866,f865,f112,f111,f110,f108,f109,f593,f670,f1379,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786,f1668,f1511,f425,f387,f382])).
fof(f4666,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638,f639,f671,f4220,f866,f865,f112,f111,f110,f108,f109,f593,f670,f1379,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786,f1668,f1511,f425,f387,f382])).
fof(f4665,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638,f639,f671,f4220,f866,f865,f112,f111,f110,f108,f109,f593,f670,f1379,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786,f1668,f1511,f425,f387])).
fof(f4664,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638,f639,f671,f4220,f866,f865,f112,f111,f110,f108,f109,f593,f670,f1379,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786,f1668,f1511,f425])).
fof(f4663,plain,(
  $false | ($spl16 | $spl23 | $spl36)),
  inference(global_subsumption,[],[f1262,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638,f639,f671,f4220,f866,f865,f112,f111,f110,f108,f109,f593,f670,f1379,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786,f1668,f1511])).
fof(f4662,plain,(
  $false | ($spl16 | $spl23 | $spl45)),
  inference(global_subsumption,[],[f1268,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638,f639,f671,f4220,f866,f865,f112,f111,f110,f108,f109,f593,f670,f1379,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786,f1668,f1511])).
fof(f4661,plain,(
  $false | ($spl16 | $spl23 | $spl45)),
  inference(global_subsumption,[],[f1346,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638,f639,f671,f4220,f866,f865,f112,f111,f110,f108,f109,f593,f670,f1379,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786,f1668,f1511])).
fof(f4660,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f3330,f3331,f3999,f4000,f4039,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638,f639,f671,f4220,f866,f865,f112,f111,f110,f108,f109,f593,f670,f1379,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786,f1668,f1511])).
fof(f4659,plain,(
  $false | ($spl16 | $spl23 | $spl51)),
  inference(global_subsumption,[],[f1602,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1511,f3330,f3331,f3999,f4000,f4039,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638,f639,f671,f4220,f866,f865,f112,f111,f110,f108,f109,f593,f670,f1379,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786,f1668])).
fof(f4658,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1511,f3330,f3331,f3999,f4000,f4039,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638,f639,f671,f4220,f866,f865,f112,f111,f110,f108,f109,f593,f670,f1379,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786,f1668])).
fof(f4657,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1511,f1668,f3330,f3331,f3999,f4000,f4039,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638,f639,f671,f4220,f866,f865,f112,f111,f110,f108,f109,f593,f670,f1379,f4206,f4203,f4202,f4197,f4196,f2226,f2211,f1786])).
fof(f4656,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1511,f1668,f1786,f3330,f3331,f3999,f4000,f4039,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638,f639,f671,f4220,f866,f865,f112,f111,f110,f108,f109,f593,f670,f1379,f4206,f4203,f4202,f4197,f4196,f2226,f2211])).
fof(f4655,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1511,f1668,f1786,f2211,f3330,f3331,f3999,f4000,f4039,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638,f639,f671,f4220,f866,f865,f112,f111,f110,f108,f109,f593,f670,f1379,f4206,f4203,f4202,f4197,f4196,f2226])).
fof(f4654,plain,(
  $false | ($spl16 | $spl23 | $spl45)),
  inference(global_subsumption,[],[f4103,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1511,f1668,f1786,f2211,f2226,f3330,f3331,f3999,f4000,f4039,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638,f639,f671,f4220,f866,f865,f112,f111,f110,f108,f109,f593,f670,f1379,f4206,f4203,f4202,f4197,f4196])).
fof(f4653,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1511,f1668,f1786,f2211,f2226,f3330,f3331,f3999,f4000,f4039,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638,f639,f671,f4220,f866,f865,f112,f111,f110,f108,f109,f593,f670,f1379,f4206,f4203,f4202,f4197,f4196])).
fof(f4652,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1511,f1668,f1786,f2211,f2226,f3330,f3331,f3999,f4000,f4039,f4196,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638,f639,f671,f4220,f866,f865,f112,f111,f110,f108,f109,f593,f670,f1379,f4206,f4203,f4202,f4197])).
fof(f4651,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1511,f1668,f1786,f2211,f2226,f3330,f3331,f3999,f4000,f4039,f4197,f4196,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638,f639,f671,f4220,f866,f865,f112,f111,f110,f108,f109,f593,f670,f1379,f4206,f4203,f4202])).
fof(f4650,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1511,f1668,f1786,f2211,f2226,f3330,f3331,f3999,f4000,f4039,f4202,f4197,f4196,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638,f639,f671,f4220,f866,f865,f112,f111,f110,f108,f109,f593,f670,f1379,f4206,f4203])).
fof(f4649,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1511,f1668,f1786,f2211,f2226,f3330,f3331,f3999,f4000,f4039,f4203,f4202,f4197,f4196,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638,f639,f671,f4220,f866,f865,f112,f111,f110,f108,f109,f593,f670,f1379,f4206])).
fof(f4648,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1511,f1668,f1786,f2211,f2226,f3330,f3331,f3999,f4000,f4039,f4206,f4203,f4202,f4197,f4196,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638,f639,f671,f4220,f866,f865,f112,f111,f110,f108,f109,f593,f670,f1379])).
fof(f4647,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f595,f864,f664,f665,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1379,f1511,f1668,f1786,f2211,f2226,f3330,f3331,f3999,f4000,f4039,f4206,f4203,f4202,f4197,f4196,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638,f639,f671,f4220,f866,f865,f112,f111,f110,f108,f109,f593,f670])).
fof(f4646,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f595,f864,f664,f665,f670,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1379,f1511,f1668,f1786,f2211,f2226,f3330,f3331,f3999,f4000,f4039,f4206,f4203,f4202,f4197,f4196,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638,f639,f671,f4220,f866,f865,f112,f111,f110,f108,f109,f593])).
fof(f4645,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f593,f595,f864,f664,f665,f670,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1379,f1511,f1668,f1786,f2211,f2226,f3330,f3331,f3999,f4000,f4039,f4206,f4203,f4202,f4197,f4196,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638,f639,f671,f4220,f866,f865,f112,f111,f110,f108,f109])).
fof(f4644,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f109,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f593,f595,f864,f664,f665,f670,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1379,f1511,f1668,f1786,f2211,f2226,f3330,f3331,f3999,f4000,f4039,f4206,f4203,f4202,f4197,f4196,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638,f639,f671,f4220,f866,f865,f112,f111,f110,f108])).
fof(f4643,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f109,f108,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f593,f595,f864,f664,f665,f670,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1379,f1511,f1668,f1786,f2211,f2226,f3330,f3331,f3999,f4000,f4039,f4206,f4203,f4202,f4197,f4196,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638,f639,f671,f4220,f866,f865,f112,f111,f110])).
fof(f4642,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f109,f108,f110,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f593,f595,f864,f664,f665,f670,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1379,f1511,f1668,f1786,f2211,f2226,f3330,f3331,f3999,f4000,f4039,f4206,f4203,f4202,f4197,f4196,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638,f639,f671,f4220,f866,f865,f112,f111])).
fof(f4641,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f109,f108,f111,f110,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f593,f595,f864,f664,f665,f670,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1379,f1511,f1668,f1786,f2211,f2226,f3330,f3331,f3999,f4000,f4039,f4206,f4203,f4202,f4197,f4196,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638,f639,f671,f4220,f866,f865,f112])).
fof(f4640,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f109,f108,f112,f111,f110,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f593,f595,f864,f664,f665,f670,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1379,f1511,f1668,f1786,f2211,f2226,f3330,f3331,f3999,f4000,f4039,f4206,f4203,f4202,f4197,f4196,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638,f639,f671,f4220,f866,f865])).
fof(f4639,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f109,f108,f112,f111,f110,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f593,f595,f864,f865,f664,f665,f670,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1379,f1511,f1668,f1786,f2211,f2226,f3330,f3331,f3999,f4000,f4039,f4206,f4203,f4202,f4197,f4196,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638,f639,f671,f4220,f866])).
fof(f4638,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f109,f108,f112,f111,f110,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f593,f595,f864,f865,f866,f664,f665,f670,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1379,f1511,f1668,f1786,f2211,f2226,f3330,f3331,f3999,f4000,f4039,f4206,f4203,f4202,f4197,f4196,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638,f639,f671,f4220])).
fof(f4637,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f109,f108,f112,f111,f110,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f593,f595,f864,f865,f866,f664,f665,f670,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1379,f1511,f1668,f1786,f2211,f2226,f3330,f3331,f3999,f4000,f4039,f4206,f4203,f4202,f4197,f4196,f4220,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638,f639,f671])).
fof(f4636,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f109,f108,f112,f111,f110,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f593,f595,f864,f865,f866,f664,f665,f670,f671,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1379,f1511,f1668,f1786,f2211,f2226,f3330,f3331,f3999,f4000,f4039,f4206,f4203,f4202,f4197,f4196,f4220,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638,f639])).
fof(f4635,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f109,f108,f112,f111,f110,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f593,f595,f639,f864,f865,f866,f664,f665,f670,f671,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1379,f1511,f1668,f1786,f2211,f2226,f3330,f3331,f3999,f4000,f4039,f4206,f4203,f4202,f4197,f4196,f4220,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637,f638])).
fof(f4634,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f109,f108,f112,f111,f110,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f593,f595,f638,f639,f864,f865,f866,f664,f665,f670,f671,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1379,f1511,f1668,f1786,f2211,f2226,f3330,f3331,f3999,f4000,f4039,f4206,f4203,f4202,f4197,f4196,f4220,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589,f637])).
fof(f4633,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f109,f108,f112,f111,f110,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f593,f595,f637,f638,f639,f864,f865,f866,f664,f665,f670,f671,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1379,f1511,f1668,f1786,f2211,f2226,f3330,f3331,f3999,f4000,f4039,f4206,f4203,f4202,f4197,f4196,f4220,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590,f589])).
fof(f4632,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f109,f108,f112,f111,f110,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f589,f593,f595,f637,f638,f639,f864,f865,f866,f664,f665,f670,f671,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1379,f1511,f1668,f1786,f2211,f2226,f3330,f3331,f3999,f4000,f4039,f4206,f4203,f4202,f4197,f4196,f4220,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591,f590])).
fof(f4631,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f109,f108,f112,f111,f110,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f589,f590,f593,f595,f637,f638,f639,f864,f865,f866,f664,f665,f670,f671,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1379,f1511,f1668,f1786,f2211,f2226,f3330,f3331,f3999,f4000,f4039,f4206,f4203,f4202,f4197,f4196,f4220,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594,f591])).
fof(f4630,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f109,f108,f112,f111,f110,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f589,f590,f591,f593,f595,f637,f638,f639,f864,f865,f866,f664,f665,f670,f671,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1379,f1511,f1668,f1786,f2211,f2226,f3330,f3331,f3999,f4000,f4039,f4206,f4203,f4202,f4197,f4196,f4220,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666,f594])).
fof(f4629,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f109,f108,f112,f111,f110,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f589,f590,f591,f593,f594,f595,f637,f638,f639,f864,f865,f866,f664,f665,f670,f671,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1379,f1511,f1668,f1786,f2211,f2226,f3330,f3331,f3999,f4000,f4039,f4206,f4203,f4202,f4197,f4196,f4220,f209,f210,f2553,f76,f1056,f2542,f668,f667,f666])).
fof(f4628,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f109,f108,f112,f111,f110,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f589,f590,f591,f593,f594,f595,f637,f638,f639,f864,f865,f866,f664,f665,f666,f670,f671,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1379,f1511,f1668,f1786,f2211,f2226,f3330,f3331,f3999,f4000,f4039,f4206,f4203,f4202,f4197,f4196,f4220,f209,f210,f2553,f76,f1056,f2542,f668,f667])).
fof(f4627,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f109,f108,f112,f111,f110,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f589,f590,f591,f593,f594,f595,f637,f638,f639,f864,f865,f866,f664,f665,f666,f667,f670,f671,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1379,f1511,f1668,f1786,f2211,f2226,f3330,f3331,f3999,f4000,f4039,f4206,f4203,f4202,f4197,f4196,f4220,f209,f210,f2553,f76,f1056,f2542,f668])).
fof(f4626,plain,(
  $false | ($spl16 | $spl23 | $spl26)),
  inference(global_subsumption,[],[f2915,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f109,f108,f112,f111,f110,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f589,f590,f591,f593,f594,f595,f637,f638,f639,f864,f865,f866,f664,f665,f666,f667,f668,f670,f671,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1379,f1511,f1668,f1786,f2211,f2226,f3330,f3331,f3999,f4000,f4039,f4206,f4203,f4202,f4197,f4196,f4220,f209,f210,f2553,f76,f1056,f2542])).
fof(f4625,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f109,f108,f112,f111,f110,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f589,f590,f591,f593,f594,f595,f637,f638,f639,f864,f865,f866,f664,f665,f666,f667,f668,f670,f671,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1379,f1511,f1668,f1786,f2211,f2226,f3330,f3331,f3999,f4000,f4039,f4206,f4203,f4202,f4197,f4196,f4220,f209,f210,f2553,f76,f1056,f2542])).
fof(f4624,plain,(
  $false | ($spl16 | $spl23 | $spl30)),
  inference(global_subsumption,[],[f1159,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f109,f108,f112,f111,f110,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f589,f590,f591,f593,f594,f595,f637,f638,f639,f864,f865,f866,f664,f665,f666,f667,f668,f670,f671,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1379,f1511,f1668,f1786,f2211,f2226,f2542,f3330,f3331,f3999,f4000,f4039,f4206,f4203,f4202,f4197,f4196,f4220,f209,f210,f2553,f76,f1056])).
fof(f4623,plain,(
  $false | ($spl16 | $spl23 | $spl26)),
  inference(global_subsumption,[],[f1156,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f109,f108,f112,f111,f110,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f589,f590,f591,f593,f594,f595,f637,f638,f639,f864,f865,f866,f664,f665,f666,f667,f668,f670,f671,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1379,f1511,f1668,f1786,f2211,f2226,f2542,f3330,f3331,f3999,f4000,f4039,f4206,f4203,f4202,f4197,f4196,f4220,f209,f210,f2553,f76,f1056])).
fof(f4622,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f109,f108,f112,f111,f110,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f589,f590,f591,f593,f594,f595,f637,f638,f639,f864,f865,f866,f664,f665,f666,f667,f668,f670,f671,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1379,f1511,f1668,f1786,f2211,f2226,f2542,f3330,f3331,f3999,f4000,f4039,f4206,f4203,f4202,f4197,f4196,f4220,f209,f210,f2553,f76,f1056])).
fof(f4621,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f109,f108,f112,f111,f110,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f589,f590,f591,f593,f594,f595,f637,f638,f639,f864,f865,f866,f664,f665,f666,f667,f668,f670,f671,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1056,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1379,f1511,f1668,f1786,f2211,f2226,f2542,f3330,f3331,f3999,f4000,f4039,f4206,f4203,f4202,f4197,f4196,f4220,f209,f210,f2553,f76])).
fof(f4620,plain,(
  $false | ($spl8 | $spl16 | $spl23)),
  inference(global_subsumption,[],[f215,f76,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f109,f108,f112,f111,f110,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f589,f590,f591,f593,f594,f595,f637,f638,f639,f864,f865,f866,f664,f665,f666,f667,f668,f670,f671,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1056,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1379,f1511,f1668,f1786,f2211,f2226,f2542,f3330,f3331,f3999,f4000,f4039,f4206,f4203,f4202,f4197,f4196,f4220,f209,f210,f2553])).
fof(f4619,plain,(
  $false | ($spl12 | $spl16 | $spl23)),
  inference(global_subsumption,[],[f377,f76,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f109,f108,f112,f111,f110,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f589,f590,f591,f593,f594,f595,f637,f638,f639,f864,f865,f866,f664,f665,f666,f667,f668,f670,f671,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1056,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1379,f1511,f1668,f1786,f2211,f2226,f2542,f3330,f3331,f3999,f4000,f4039,f4206,f4203,f4202,f4197,f4196,f4220,f209,f210,f2553])).
fof(f4618,plain,(
  $false | ($spl16 | $spl23)),
  inference(global_subsumption,[],[f76,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f109,f108,f112,f111,f110,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f589,f590,f591,f593,f594,f595,f637,f638,f639,f864,f865,f866,f664,f665,f666,f667,f668,f670,f671,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1056,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1379,f1511,f1668,f1786,f2211,f2226,f2542,f3330,f3331,f3999,f4000,f4039,f4206,f4203,f4202,f4197,f4196,f4220,f209,f210,f2553])).
fof(f4617,plain,(
  $false | ($spl16 | $spl23 | $spl26)),
  inference(global_subsumption,[],[f2926,f76,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f109,f108,f112,f111,f110,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f589,f590,f591,f593,f594,f595,f637,f638,f639,f864,f865,f866,f664,f665,f666,f667,f668,f670,f671,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1056,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1379,f1511,f1668,f1786,f2211,f2226,f2542,f2553,f3330,f3331,f3999,f4000,f4039,f4206,f4203,f4202,f4197,f4196,f4220,f209,f210])).
fof(f4616,plain,(
  $false | ($spl8 | $spl16 | $spl23)),
  inference(global_subsumption,[],[f4295,f76,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f109,f108,f112,f111,f110,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f589,f590,f591,f593,f594,f595,f637,f638,f639,f864,f865,f866,f664,f665,f666,f667,f668,f670,f671,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1056,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1379,f1511,f1668,f1786,f2211,f2226,f2542,f2553,f3330,f3331,f3999,f4000,f4039,f4206,f4203,f4202,f4197,f4196,f4220,f209,f210])).
fof(f4615,plain,(
  $false | ($spl3 | $spl10 | $spl16 | $spl23)),
  inference(global_subsumption,[],[f4326,f76,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f109,f108,f112,f111,f110,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f589,f590,f591,f593,f594,f595,f637,f638,f639,f864,f865,f866,f664,f665,f666,f667,f668,f670,f671,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1056,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1379,f1511,f1668,f1786,f2211,f2226,f2542,f2553,f3330,f3331,f3999,f4000,f4039,f4206,f4203,f4202,f4197,f4196,f4220,f209,f210])).
fof(f4614,plain,(
  $false | ($spl12 | $spl16 | $spl23)),
  inference(global_subsumption,[],[f4367,f76,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f109,f108,f112,f111,f110,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f589,f590,f591,f593,f594,f595,f637,f638,f639,f864,f865,f866,f664,f665,f666,f667,f668,f670,f671,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1056,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1379,f1511,f1668,f1786,f2211,f2226,f2542,f2553,f3330,f3331,f3999,f4000,f4039,f4206,f4203,f4202,f4197,f4196,f4220,f209,f210])).
fof(f4613,plain,(
  $false | ($spl3 | $spl14 | $spl16 | $spl23)),
  inference(global_subsumption,[],[f4398,f76,f77,f84,f107,f116,f113,f117,f121,f123,f129,f145,f155,f158,f159,f109,f108,f112,f111,f110,f120,f144,f190,f208,f211,f212,f213,f259,f260,f336,f375,f381,f380,f382,f387,f425,f532,f587,f588,f589,f590,f591,f593,f594,f595,f637,f638,f639,f864,f865,f866,f664,f665,f666,f667,f668,f670,f671,f672,f1048,f1049,f1050,f1051,f1052,f1061,f1056,f1088,f1097,f1148,f1149,f1150,f1151,f1152,f1256,f1257,f1258,f1259,f1260,f1364,f1367,f1379,f1511,f1668,f1786,f2211,f2226,f2542,f2553,f3330,f3331,f3999,f4000,f4039,f4206,f4203,f4202,f4197,f4196,f4220,f209,f210])).
fof(f4609,plain,(
  $false | ($spl16 | $spl20)),
  inference(subsumption_resolution,[],[f4608,f77])).
fof(f4608,plain,(
  ~p(s(bool,'T_0')) | ($spl16 | $spl20)),
  inference(forward_demodulation,[],[f4607,f4510])).
fof(f4510,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2)) ) | ($spl16 | $spl20)),
  inference(duplicate_literal_removal,[],[f4419])).
fof(f4419,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl16 | $spl20)),
  inference(superposition,[],[f1059,f1056])).
fof(f4607,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))) ) | ($spl16 | $spl20)),
  inference(subsumption_resolution,[],[f4560,f77])).
fof(f4560,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))) ) | ($spl16 | $spl20)),
  inference(backward_demodulation,[],[f4510,f111])).
fof(f4605,plain,(
  $false | ($spl16 | $spl20 | $spl45)),
  inference(subsumption_resolution,[],[f4541,f4510])).
fof(f4541,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl16 | $spl20 | $spl45)),
  inference(backward_demodulation,[],[f4510,f1268])).
fof(f4604,plain,(
  $false | ($spl12 | $spl16 | $spl20)),
  inference(subsumption_resolution,[],[f4598,f4510])).
fof(f4598,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) | ($spl12 | $spl16 | $spl20)),
  inference(backward_demodulation,[],[f4586,f4367])).
fof(f4586,plain,(
  ( ! [X0,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) ) | ($spl16 | $spl20)),
  inference(subsumption_resolution,[],[f4537,f4510])).
fof(f4537,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK2(X1,'const.pred_set.EMPTY_0',X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) ) | ($spl16 | $spl20)),
  inference(backward_demodulation,[],[f4510,f382])).
fof(f4603,plain,(
  $false | ($spl8 | $spl16 | $spl20)),
  inference(subsumption_resolution,[],[f4596,f4510])).
fof(f4596,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) | ($spl8 | $spl16 | $spl20)),
  inference(backward_demodulation,[],[f4586,f4295])).
fof(f4602,plain,(
  $false | ($spl16 | $spl20 | $spl26)),
  inference(subsumption_resolution,[],[f4595,f4510])).
fof(f4595,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ($spl16 | $spl20 | $spl26)),
  inference(backward_demodulation,[],[f4586,f2926])).
fof(f4601,plain,(
  $false | ($spl16 | $spl20)),
  inference(subsumption_resolution,[],[f4594,f4510])).
fof(f4594,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ($spl16 | $spl20)),
  inference(backward_demodulation,[],[f4586,f2553])).
fof(f4600,plain,(
  $false | ($spl16 | $spl20)),
  inference(subsumption_resolution,[],[f4591,f4586])).
fof(f4591,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),sK1) | ($spl16 | $spl20)),
  inference(backward_demodulation,[],[f4586,f76])).
fof(f4585,plain,(
  $false | ($spl16 | $spl20)),
  inference(subsumption_resolution,[],[f4528,f77])).
fof(f4528,plain,(
  ~p(s(bool,'T_0')) | ($spl16 | $spl20)),
  inference(backward_demodulation,[],[f4510,f120])).
fof(f4527,plain,(
  $false | ($spl3 | $spl16 | $spl20)),
  inference(subsumption_resolution,[],[f4526,f132])).
fof(f4526,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl16 | $spl20)),
  inference(backward_demodulation,[],[f4525,f144])).
fof(f4525,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | ($spl16 | $spl20)),
  inference(condensation,[],[f4524])).
fof(f4524,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X10,X11),s(fun(X10,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X9)) ) | ($spl16 | $spl20)),
  inference(forward_demodulation,[],[f4511,f1056])).
fof(f4511,plain,(
  ( ! [X10,X11,X9] : (s(bool,'const.bool.IN_2'(s(X10,X11),s(fun(X10,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | s(bool,'T_0') = s(bool,X9)) ) | $spl20),
  inference(duplicate_literal_removal,[],[f4414])).
fof(f4414,plain,(
  ( ! [X10,X11,X9] : (s(bool,'const.bool.IN_2'(s(X10,X11),s(fun(X10,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | $spl20),
  inference(superposition,[],[f1059,f208])).
fof(f4523,plain,(
  $false | ($spl3 | $spl16 | $spl20)),
  inference(subsumption_resolution,[],[f4522,f132])).
fof(f4522,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl16 | $spl20)),
  inference(condensation,[],[f4521])).
fof(f4521,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X0)) ) | ($spl16 | $spl20)),
  inference(forward_demodulation,[],[f4516,f1056])).
fof(f4516,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl20),
  inference(duplicate_literal_removal,[],[f4409])).
fof(f4409,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl20),
  inference(superposition,[],[f1059,f84])).
fof(f4520,plain,(
  $false | ($spl3 | $spl16 | $spl20)),
  inference(subsumption_resolution,[],[f4519,f132])).
fof(f4519,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl16 | $spl20)),
  inference(forward_demodulation,[],[f4407,f1056])).
fof(f4407,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | ($spl3 | $spl20)),
  inference(unit_resulting_resolution,[],[f132,f1059])).
fof(f4518,plain,(
  $false | ($spl16 | $spl20 | $spl45)),
  inference(subsumption_resolution,[],[f4517,f1268])).
fof(f4517,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl16 | $spl20 | $spl45)),
  inference(forward_demodulation,[],[f4408,f1056])).
fof(f4408,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | ($spl20 | $spl45)),
  inference(unit_resulting_resolution,[],[f1268,f1059])).
fof(f156,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl4),
  inference(cnf_transformation,[],[f156_D])).
fof(f156_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f4117,plain,(
  ( ! [X54,X53] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X53,X54),s(fun(X53,bool),'const.pred_set.EMPTY_0')))) ) | $spl52),
  inference(cnf_transformation,[],[f4117_D])).
fof(f4117_D,plain,(
  ( ! [X54,X53] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X53,X54),s(fun(X53,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f4119,plain,(
  ( ! [X59,X60,X58] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X59,X60),s(fun(X59,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X58)) ) | $spl15),
  inference(forward_demodulation,[],[f4060,f398])).
fof(f398,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) | $spl15),
  inference(unit_resulting_resolution,[],[f120,f379,f158])).
fof(f4060,plain,(
  ( ! [X59,X60,X58] : (s(bool,'const.bool.IN_2'(s(X59,X60),s(fun(X59,bool),'const.pred_set.EMPTY_0'))) != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) | s(bool,'T_0') = s(bool,X58)) ) | $spl15),
  inference(duplicate_literal_removal,[],[f3926])).
fof(f3926,plain,(
  ( ! [X59,X60,X58] : (s(bool,'const.bool.IN_2'(s(X59,X60),s(fun(X59,bool),'const.pred_set.EMPTY_0'))) != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) | s(bool,'T_0') = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) ) | $spl15),
  inference(superposition,[],[f985,f208])).
fof(f985,plain,(
  ( ! [X23] : (s(bool,X23) != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) | s(bool,'T_0') = s(bool,X23)) ) | $spl15),
  inference(subsumption_resolution,[],[f933,f76])).
fof(f933,plain,(
  ( ! [X23] : (s(bool,X23) != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X23)) ) | $spl15),
  inference(superposition,[],[f117,f460])).
fof(f460,plain,(
  ( ! [X3] : (s(bool,X3) = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) | s(bool,'T_0') = s(bool,X3)) ) | $spl15),
  inference(subsumption_resolution,[],[f449,f120])).
fof(f449,plain,(
  ( ! [X3] : (p(s(bool,'F_0')) | s(bool,X3) = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) | s(bool,'T_0') = s(bool,X3)) ) | $spl15),
  inference(superposition,[],[f158,f413])).
fof(f413,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) | $spl15),
  inference(unit_resulting_resolution,[],[f404,f84])).
fof(f404,plain,(
  s(bool,'T_0') != s(bool,g(s(sK0,sK3(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) | $spl15),
  inference(backward_demodulation,[],[f398,f210])).
fof(f4116,plain,(
  ( ! [X54,X52,X53] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X53,X54),s(fun(X53,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X52)) ) | $spl11),
  inference(forward_demodulation,[],[f4062,f222])).
fof(f222,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) | $spl11),
  inference(unit_resulting_resolution,[],[f120,f217,f158])).
fof(f4062,plain,(
  ( ! [X54,X52,X53] : (s(bool,'const.bool.IN_2'(s(X53,X54),s(fun(X53,bool),'const.pred_set.EMPTY_0'))) != s(bool,f(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X52)) ) | $spl11),
  inference(duplicate_literal_removal,[],[f3924])).
fof(f3924,plain,(
  ( ! [X54,X52,X53] : (s(bool,'const.bool.IN_2'(s(X53,X54),s(fun(X53,bool),'const.pred_set.EMPTY_0'))) != s(bool,f(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) ) | $spl11),
  inference(superposition,[],[f332,f208])).
fof(f332,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,f(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X13)) ) | $spl11),
  inference(subsumption_resolution,[],[f312,f76])).
fof(f312,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,f(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X13)) ) | $spl11),
  inference(superposition,[],[f117,f256])).
fof(f256,plain,(
  ( ! [X3] : (s(bool,X3) = s(bool,g(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X3)) ) | $spl11),
  inference(subsumption_resolution,[],[f249,f120])).
fof(f249,plain,(
  ( ! [X3] : (p(s(bool,'F_0')) | s(bool,X3) = s(bool,g(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X3)) ) | $spl11),
  inference(superposition,[],[f158,f236])).
fof(f236,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) | $spl11),
  inference(unit_resulting_resolution,[],[f228,f84])).
fof(f228,plain,(
  s(bool,'T_0') != s(bool,g(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) | $spl11),
  inference(backward_demodulation,[],[f222,f209])).
fof(f2917,plain,(
  $false | ($spl15 | $spl26 | $spl33)),
  inference(subsumption_resolution,[],[f2916,f120])).
fof(f2916,plain,(
  p(s(bool,'F_0')) | ($spl15 | $spl26 | $spl33)),
  inference(forward_demodulation,[],[f2915,f1201])).
fof(f1201,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | ($spl15 | $spl33)),
  inference(forward_demodulation,[],[f1184,f413])).
fof(f1184,plain,(
  s(bool,g(s(sK0,sK3(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | ($spl15 | $spl33)),
  inference(unit_resulting_resolution,[],[f1161,f460])).
fof(f2544,plain,(
  $false | ($spl15 | $spl16 | $spl23)),
  inference(subsumption_resolution,[],[f2543,f120])).
fof(f2543,plain,(
  p(s(bool,'F_0')) | ($spl15 | $spl16 | $spl23)),
  inference(forward_demodulation,[],[f2542,f1099])).
fof(f1099,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl15 | $spl23)),
  inference(forward_demodulation,[],[f1082,f413])).
fof(f1082,plain,(
  s(bool,g(s(sK0,sK3(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl15 | $spl23)),
  inference(unit_resulting_resolution,[],[f1061,f460])).
fof(f1382,plain,(
  ( ! [X4,X2,X3] : (s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,sK3(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X2,X3,bool,bool))))) ) | $spl48),
  inference(cnf_transformation,[],[f1382_D])).
fof(f1382_D,plain,(
  ( ! [X4,X2,X3] : (s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,sK3(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X2,X3,bool,bool))))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f1384,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl51),
  inference(cnf_transformation,[],[f1384_D])).
fof(f1384_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f1367,plain,(
  ( ! [X12,X13,X11] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X11,X12,bool,bool)))) | s(bool,X13) = s(bool,sK3(X11,X12,bool,bool)) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f212,f121])).
fof(f1364,plain,(
  ( ! [X4,X2,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X2,X3,bool,bool)))) | s(bool,X4) = s(bool,sK3(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f212,f121])).
fof(f1266,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) ) | $spl42),
  inference(cnf_transformation,[],[f1266_D])).
fof(f1266_D,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f1263,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK3(X0,X1,bool,bool))) ) | $spl38),
  inference(cnf_transformation,[],[f1263_D])).
fof(f1263_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK3(X0,X1,bool,bool))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f1268,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl45),
  inference(cnf_transformation,[],[f1268_D])).
fof(f1268_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f1260,plain,(
  ( ! [X12,X13,X11] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,'T_0') = s(bool,sK3(X11,X12,bool,bool)) | s(bool,X13) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f211,f121])).
fof(f1262,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl36),
  inference(cnf_transformation,[],[f1262_D])).
fof(f1262_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f1269,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) ) | $spl46),
  inference(cnf_transformation,[],[f1269_D])).
fof(f1269_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f1259,plain,(
  ( ! [X10,X8,X9] : (s(bool,X8) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X9) = s(fun(bool,bool),X10) | s(bool,'T_0') = s(bool,sK3(X9,X10,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f211,f121])).
fof(f1258,plain,(
  ( ! [X6,X7,X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7) | s(bool,'T_0') = s(bool,sK3(X6,X7,bool,bool)) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f211,f121])).
fof(f1257,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,sK3(X2,X3,bool,bool)) | s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f211,f121])).
fof(f1265,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl41),
  inference(cnf_transformation,[],[f1265_D])).
fof(f1265_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f1256,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK3(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f211,f84])).
fof(f1156,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl26),
  inference(cnf_transformation,[],[f1156_D])).
fof(f1156_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f1158,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | $spl29),
  inference(cnf_transformation,[],[f1158_D])).
fof(f1158_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f1164,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl15),
  inference(forward_demodulation,[],[f1154,f413])).
fof(f1154,plain,(
  s(bool,g(s(sK0,sK3(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl15),
  inference(superposition,[],[f381,f460])).
fof(f1163,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl11),
  inference(forward_demodulation,[],[f1153,f236])).
fof(f1153,plain,(
  s(bool,g(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl11),
  inference(superposition,[],[f381,f256])).
fof(f1159,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl30),
  inference(cnf_transformation,[],[f1159_D])).
fof(f1159_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f1161,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | $spl33),
  inference(cnf_transformation,[],[f1161_D])).
fof(f1161_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f1152,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f381,f121])).
fof(f1162,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X1)) ) | $spl34),
  inference(cnf_transformation,[],[f1162_D])).
fof(f1162_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X1)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1151,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f381,f121])).
fof(f1150,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))) )),
  inference(superposition,[],[f381,f121])).
fof(f1149,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f381,f121])).
fof(f1148,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))),
  inference(superposition,[],[f381,f84])).
fof(f1056,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl16),
  inference(cnf_transformation,[],[f1056_D])).
fof(f1056_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1058,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl19),
  inference(cnf_transformation,[],[f1058_D])).
fof(f1058_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f1064,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl15),
  inference(forward_demodulation,[],[f1054,f413])).
fof(f1054,plain,(
  s(bool,g(s(sK0,sK3(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl15),
  inference(superposition,[],[f380,f460])).
fof(f1063,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl11),
  inference(forward_demodulation,[],[f1053,f236])).
fof(f1053,plain,(
  s(bool,g(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl11),
  inference(superposition,[],[f380,f256])).
fof(f1059,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl20),
  inference(cnf_transformation,[],[f1059_D])).
fof(f1059_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f1061,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl23),
  inference(cnf_transformation,[],[f1061_D])).
fof(f1061_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f1052,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f380,f121])).
fof(f1062,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X1)) ) | $spl24),
  inference(cnf_transformation,[],[f1062_D])).
fof(f1062_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X1)) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1051,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f380,f121])).
fof(f1050,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))) )),
  inference(superposition,[],[f380,f121])).
fof(f1049,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f380,f121])).
fof(f1048,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))),
  inference(superposition,[],[f380,f84])).
fof(f377,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) | $spl12),
  inference(cnf_transformation,[],[f377_D])).
fof(f377_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f379,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) | $spl15),
  inference(cnf_transformation,[],[f379_D])).
fof(f379_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f375,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0))))),
  inference(superposition,[],[f210,f84])).
fof(f215,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) | $spl8),
  inference(cnf_transformation,[],[f215_D])).
fof(f215_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f217,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) | $spl11),
  inference(cnf_transformation,[],[f217_D])).
fof(f217_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f213,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.DIFF_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK1,bool,sK0))))),
  inference(superposition,[],[f209,f84])).
fof(f172,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f171,f77])).
fof(f171,plain,(
  ~p(s(bool,'T_0')) | $spl4),
  inference(forward_demodulation,[],[f170,f156])).
fof(f170,plain,(
  ( ! [X2,X0,X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) ) | $spl4),
  inference(subsumption_resolution,[],[f163,f77])).
fof(f163,plain,(
  ( ! [X2,X0,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) ) | $spl4),
  inference(backward_demodulation,[],[f156,f111])).
fof(f167,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f161,f77])).
fof(f161,plain,(
  ~p(s(bool,'T_0')) | $spl4),
  inference(backward_demodulation,[],[f156,f120])).
fof(f157,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1)) ) | $spl6),
  inference(cnf_transformation,[],[f157_D])).
fof(f157_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1)) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f159,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f151,f123])).
fof(f151,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f113,f84])).
fof(f155,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f149,f123])).
fof(f149,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f113,f84])).
fof(f130,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl0),
  inference(cnf_transformation,[],[f130_D])).
fof(f130_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f132,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl3),
  inference(cnf_transformation,[],[f132_D])).
fof(f132_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f129,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f84])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 122
% Passive clauses: 322
% Generated clauses: 8110
% Final active clauses: 21
% Final passive clauses: 13
% Input formulas: 21
% Initial clauses: 42
% 
% Duplicate literals: 536
% Trivial inequalities: 58
% Fw subsumption resolutions: 337
% Fw demodulations: 1313
% Bw demodulations: 188
% Condensations: 479
% Global subsumptions: 149
% 
% Simple tautologies: 107
% Equational tautologies: 167
% Forward subsumptions: 4432
% Fw demodulations to eq. taut.: 4
% Bw demodulations to eq. taut.: 2
% 
% Unit resulting resolution: 887
% Factoring: 1
% Forward superposition: 2040
% Backward superposition: 1680
% Self superposition: 65
% Equality factoring: 44
% Equality resolution: 13
% 
% Split clauses: 28
% Split components: 61
% SAT solver clauses: 707
% SAT solver unit clauses: 34
% SAT solver binary clauses: 250
% 
% Sat splits: 28
% Sat splitting refutations: 178
% 
% Memory used [KB]: 4093
% Time elapsed: 0.182 s
% ------------------------------
% ------------------------------
% Success in time 1.498 s
