Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.res_quan
% (71632)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 768
% Passive clauses: 10560
% Generated clauses: 22625
% Final active clauses: 742
% Final passive clauses: 9230
% Input formulas: 48
% Initial clauses: 192
% 
% Duplicate literals: 350
% Trivial inequalities: 83
% Fw subsumption resolutions: 1987
% Fw demodulations: 897
% Bw demodulations: 25
% 
% Simple tautologies: 235
% Equational tautologies: 78
% Forward subsumptions: 8157
% 
% Binary resolution: 12328
% Factoring: 1463
% Forward superposition: 2596
% Backward superposition: 2491
% Self superposition: 31
% Equality factoring: 28
% Equality resolution: 19
% 
% Split clauses: 319
% Split components: 640
% SAT solver clauses: 774
% SAT solver unit clauses: 17
% SAT solver binary clauses: 150
% 
% Sat splits: 319
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 774
% TWLsolver calls for satisfiability: 445
% 
% Memory used [KB]: 10362
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.res_quan
% (71642)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 425
% Passive clauses: 2894
% Generated clauses: 11172
% Final active clauses: 420
% Final passive clauses: 2469
% Input formulas: 48
% Initial clauses: 192
% 
% Duplicate literals: 190
% Trivial inequalities: 27
% Fw subsumption resolutions: 833
% Fw demodulations: 235
% Bw demodulations: 4
% Condensations: 343
% Global subsumptions: 18
% 
% Simple tautologies: 255
% Equational tautologies: 21
% Forward subsumptions: 3105
% 
% Binary resolution: 4006
% Factoring: 323
% Forward superposition: 1081
% Backward superposition: 1001
% Self superposition: 27
% Equality factoring: 26
% Equality resolution: 7
% 
% Split clauses: 34
% Split components: 69
% SAT solver clauses: 3435
% SAT solver unit clauses: 204
% SAT solver binary clauses: 85
% SAT solver learnt clauses: 36
% SAT solver learnt literals: 10
% 
% Sat splits: 34
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 3363
% TWLsolver calls for satisfiability: 1821
% 
% Memory used [KB]: 12409
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.res_quan
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.res_quan
% SZS output start Proof for thm.res_quan
fof(f15460,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f267,f470,f266,f472,f468,f265,f469,f471,f467,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f4476,f4414,f4478,f4418,f4480,f4421,f4482,f4424,f4488,f4431,f4490,f4434,f4492,f4437,f4494,f4440,f4496,f4443,f4498,f4446,f4618,f4619,f4620,f4621,f4623,f4625,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4693,f4695,f4697,f5180,f5182,f5222,f5224,f5225,f5227,f5229,f5233,f5235,f5237,f5239,f5241,f5243,f5247,f5249,f5251,f5253,f5255,f5257,f5259,f5261,f5273,f5275,f5287,f5289,f5296,f5303,f5319,f5356,f5412,f5414,f5502,f5504,f5506,f5508,f5511,f5513,f5519,f5521,f5527,f5529,f5576,f5578,f5580,f5582,f5585,f5588,f5591,f5594,f5597,f5600,f5669,f5680,f5694,f5705,f5707,f5720,f5730,f5751,f5753,f5755,f5757,f5759,f5761,f6476,f6477,f5782,f5784,f5786,f5788,f5825,f5827,f5829,f5831,f5833,f5835,f5845,f5855,f5868,f5879,f5881,f5896,f5908,f5910,f5925,f5935,f5956,f5958,f5960,f5970,f5972,f5978,f5980,f5991,f5993,f5999,f6001,f6012,f6014,f6020,f6022,f6035,f6045,f6058,f6068,f6077,f6086,f6098,f6113,f6115,f6129,f6131,f6133,f6139,f6141,f6147,f6149,f6155,f6157,f6163,f6165,f6177,f6187,f6200,f6212,f6214,f6236,f6238,f6240,f6242,f6244,f6246,f6261,f6265,f6280,f6284,f6307,f6309,f6330,f6332,f6334,f6341,f6347,f6353,f6355,f6361,f6363,f6376,f6386,f6399,f6409,f6415,f6417,f6419,f6421,f6423,f6425,f6433,f6435,f6437,f6439,f6441,f6451,f6456,f6478,f6479,f6480,f6481,f6482,f6483,f6484,f6485,f6486,f6487,f6488,f6489,f6490,f6491,f6492,f6493,f6494,f6495,f6496,f6497,f6498,f6499,f6500,f6501,f6502,f6503,f6504,f6505,f6506,f6507,f6508,f6509,f6510,f6511,f6512,f6513,f6514,f6515,f6516,f6517,f6518,f6519,f6520,f6521,f6522,f6523,f6524,f6525,f6526,f6527,f6528,f6529,f6530,f6531,f6532,f6533,f6534,f6535,f6536,f6537,f6538,f6539,f6540,f6541,f6542,f6543,f6544,f6545,f6546,f6547,f6548,f6549,f6550,f6551,f6552,f6553,f6554,f6555,f6556,f6557,f6558,f6559,f6560,f6561,f6562,f6563,f6564,f6565,f6566,f6567,f6568,f6569,f6570,f6571,f6572,f6573,f6574,f6575,f6576,f6577,f6578,f6579,f6580,f6581,f6582,f6583,f6584,f6585,f6586,f6587,f6588,f6589,f6590,f6591,f6592,f6593,f6594,f6595,f6596,f6597,f6598,f6599,f6601,f6602,f6604,f6605,f6608,f6611,f6613,f6614,f6617,f6620,f6621,f6623,f6624,f6625,f6626,f6627,f6631,f6632,f6633,f6635,f6636,f6637,f6638,f6639,f6640,f6641,f6643,f6644,f6645,f6646,f6647,f6648,f6649,f6650,f6651,f6652,f6653,f6654,f6655,f6656,f6657,f6658,f6659,f6660,f6661,f6662,f6663,f6664,f6665,f6666,f6667,f6668,f6669,f6670,f6671,f6672,f6673,f6674,f6675,f6676,f6677,f6678,f6679,f6680,f6681,f6682,f6683,f6684,f6685,f6686,f6687,f6688,f6689,f6690,f6691,f6692,f6693,f6694,f6695,f6696,f6697,f6698,f6699,f6700,f6701,f6702,f6703,f6704,f6705,f6706,f6707,f6708,f6709,f6710,f6711,f6712,f6713,f6714,f6715,f6716,f6717,f6718,f6719,f6720,f6721,f6722,f6723,f6724,f6725,f6726,f6727,f6728,f6729,f6730,f6731,f6732,f6733,f6734,f6735,f6736,f6737,f6738,f6739,f6740,f6741,f6742,f6743,f6744,f6745,f6746,f6747,f6748,f6749,f6750,f6751,f6752,f6753,f6754,f6755,f6756,f6757,f6758,f6759,f6760,f6761,f6762,f6763,f6764,f6765,f6766,f6767,f6768,f6769,f6770,f6785,f9654,f9656,f9655,f12191,f12193,f12526,f12527,f12528,f12529,f12530,f13266,f13267,f13273,f13271,f13268,f13269,f13276,f13270,f13278,f13275,f13291,f13292,f13298,f13296,f13293,f13294,f13301,f13295,f13303,f13300,f13476,f13479,f13477,f13505,f13508,f13506,f13767,f13770,f13768,f13799,f13802,f13800,f13852,f13855,f13853,f13944,f13947,f4415,f13945,f13995,f13997,f5183,f13996,f14141,f14142,f14149,f14146,f14143,f14144,f14152,f14145,f14154,f14147,f14151,f14236,f14240,f14241,f14243,f14244,f14245,f14399,f14400,f14402,f14404,f14405,f14406,f14407,f14408,f14409,f14410,f14411,f14412,f14413,f14414,f14416,f14419,f14420,f14421,f14422,f14427,f14477,f14479,f14481,f15211,f15213,f15215,f15216,f15218,f15220,f15221,f15222,f15376,f15377,f15379,f15381,f15382,f15383,f15384,f15385,f15386,f15387,f15388,f15389,f15390,f15391,f15393,f15396,f15397,f15398,f15399,f15400,f15454,f15456,f15458,f15459])).
fof(f15459,plain,(
  $false | $spl32),
  inference(subsumption_resolution,[],[f8928,f13441])).
fof(f13441,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))))))) )),
  inference(unit_resulting_resolution,[],[f268,f1538])).
fof(f1538,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X4)) | p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X3,bool),X5),s(fun(X3,bool),'_dst_x0_1'(s(bool,X4))))))) )),
  inference(superposition,[],[f339,f381])).
fof(f381,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X21,X22] : ! [X23,X24] : s(X22,X23) = s(X22,happ(s(fun(X21,X22),'_dst_x0_1'(s(X22,X23))),s(X21,X24)))),
  file('thm.res_quan.RES_FORALL_NULL','_dst_x0_ax')).
fof(f339,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK18(X2,X1,X0))))) | p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f206])).
fof(f206,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((p(s(bool,'const.bool.IN_2'(s(X0,sK18(X2,X1,X0)),s(fun(X0,bool),X1)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK18(X2,X1,X0)))))) | p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK18])],[f205])).
fof(f205,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f204])).
fof(f204,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(nnf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(ennf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(flattening,[],[f95])).
fof(f95,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X6] : ! [X14,X17] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X6,bool),X14),s(fun(X6,bool),X17)))) <=> ! [X12] : (p(s(bool,'const.bool.IN_2'(s(X6,X12),s(fun(X6,bool),X14)))) => p(s(bool,happ(s(fun(X6,bool),X17),s(X6,X12))))))),
  file('thm.res_quan.RES_FORALL_NULL','thm.bool.RES_FORALL_THM')).
fof(f268,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f8])).
fof(f8,axiom,(
  p(s(bool,'T_0'))),
  file('thm.res_quan.RES_FORALL_NULL','thm.bool.TRUTH')).
fof(f8928,plain,(
  ~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))))) | $spl32),
  inference(unit_resulting_resolution,[],[f268,f5224,f372])).
fof(f372,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f227])).
fof(f227,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f226])).
fof(f226,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f225])).
fof(f225,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f15458,plain,(
  $false | ($spl49 | $spl52)),
  inference(global_subsumption,[],[f267,f266,f265,f268,f275,f287,f286,f332,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1691,f1700,f1717,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2010,f2040,f2044,f2048,f2084,f2088,f2092,f2125,f2130,f2170,f2174,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f3420,f3421,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f13301,f13298,f13462,f13474,f13476,f13479,f13491,f13503,f13505,f13508,f13746,f13765,f13767,f13770,f13762,f13781,f13797,f13799,f13802,f13795,f13837,f13850,f13852,f13855,f13929,f13942,f13944,f13947,f13941,f13995,f13997,f14141,f14142,f14143,f14144,f14145,f3503,f2972,f1714,f1715,f336,f335,f4473,f4472,f4467,f4466,f4465,f4464,f4453,f4452,f4451,f4450,f1827,f1800,f1771,f1741,f1610,f1593,f1578,f1562,f1496,f1480,f351,f350,f349,f348,f346,f344,f343,f341,f329,f328,f3502,f3501,f2971,f2970,f13441,f1538,f13453,f13998,f13994,f13988,f4474,f4463,f1539,f1515,f1514,f339,f338,f337,f2002,f3505,f3504,f2974,f2973,f2968,f3499,f14178,f14176,f14170,f14119,f14116,f14059,f14056,f13948,f13803,f13771,f13699,f13698,f13697,f13696,f13695,f13694,f13693,f13692,f13691,f13689,f13688,f13687,f13686,f13685,f13683,f13584,f13583,f13582,f13581,f13540,f13534,f2176,f2173,f2142,f2132,f2129,f2126,f2090,f2087,f2055,f2046,f2043,f2037,f1929,f1928,f1712,f1710,f1709,f1623,f1622,f1536,f1511,f1493,f1477,f1220,f592,f473,f291,f290,f15203,f15202,f15201,f15197,f15196,f15223,f15375,f15378,f15380,f15240,f15242,f15243,f15245,f15246,f15248,f15249,f15251,f15392,f15453,f15455,f15457])).
fof(f15457,plain,(
  ( ! [X4,X2,X0] : (s(fun(X0,bool),X2) = s(fun(X0,bool),X4)) ) | ($spl49 | $spl52)),
  inference(subsumption_resolution,[],[f15362,f15223])).
fof(f15362,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK17(X4,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X2) = s(fun(X0,bool),X4)) ) | ($spl49 | $spl52)),
  inference(backward_demodulation,[],[f15223,f1714])).
fof(f15455,plain,(
  ( ! [X4,X2,X0] : (s(fun(X0,bool),X2) = s(fun(X0,bool),X4)) ) | ($spl49 | $spl52)),
  inference(subsumption_resolution,[],[f15361,f15223])).
fof(f15361,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK17(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X4,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X2) = s(fun(X0,bool),X4)) ) | ($spl49 | $spl52)),
  inference(backward_demodulation,[],[f15223,f1715])).
fof(f15453,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl49 | $spl52)),
  inference(subsumption_resolution,[],[f15360,f15223])).
fof(f15360,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK17(X2,X1,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl49 | $spl52)),
  inference(backward_demodulation,[],[f15223,f336])).
fof(f15392,plain,(
  ( ! [X35] : (s(bool,'F_0') = s(bool,X35)) ) | ($spl49 | $spl52)),
  inference(subsumption_resolution,[],[f15266,f15223])).
fof(f15266,plain,(
  ( ! [X35,X36,X34] : (s(bool,'T_0') != s(bool,happ(s(fun(X34,bool),'_dst_x0_1'(s(bool,X35))),s(X34,X36))) | s(bool,'F_0') = s(bool,X35)) ) | ($spl49 | $spl52)),
  inference(backward_demodulation,[],[f15223,f13685])).
fof(f15251,plain,(
  ( ! [X10,X8,X11,X9] : (sP5('T_0','const.bool.IN_2'(s(X8,X9),s(fun(X8,bool),'const.pred_set.EMPTY_0')),happ(s(fun(X10,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X10,X11)))) ) | ($spl49 | $spl52)),
  inference(backward_demodulation,[],[f15223,f2176])).
fof(f15249,plain,(
  ( ! [X2,X0,X3,X1] : (~sP2(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X3)),'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')),'F_0')) ) | ($spl49 | $spl52)),
  inference(backward_demodulation,[],[f15223,f2142])).
fof(f15248,plain,(
  ( ! [X10,X8,X11,X9] : (~sP5('F_0','const.bool.IN_2'(s(X8,X9),s(fun(X8,bool),'const.pred_set.EMPTY_0')),happ(s(fun(X10,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X10,X11)))) ) | ($spl49 | $spl52)),
  inference(backward_demodulation,[],[f15223,f2132])).
fof(f15246,plain,(
  ( ! [X0,X1] : (~sP2(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)),'F_0','F_0')) ) | ($spl49 | $spl52)),
  inference(backward_demodulation,[],[f15223,f2126])).
fof(f15245,plain,(
  ( ! [X10,X8,X11,X9] : (sP5('T_0',happ(s(fun(X10,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X10,X11)),'const.bool.IN_2'(s(X8,X9),s(fun(X8,bool),'const.pred_set.EMPTY_0')))) ) | ($spl49 | $spl52)),
  inference(backward_demodulation,[],[f15223,f2090])).
fof(f15243,plain,(
  ( ! [X2,X0,X3,X1] : (~sP2('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),'F_0')) ) | ($spl49 | $spl52)),
  inference(backward_demodulation,[],[f15223,f2055])).
fof(f15242,plain,(
  ( ! [X10,X8,X11,X9] : (~sP5('F_0',happ(s(fun(X10,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X10,X11)),'const.bool.IN_2'(s(X8,X9),s(fun(X8,bool),'const.pred_set.EMPTY_0')))) ) | ($spl49 | $spl52)),
  inference(backward_demodulation,[],[f15223,f2046])).
fof(f15240,plain,(
  ( ! [X0,X1] : (~sP2('F_0',happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)),'F_0')) ) | ($spl49 | $spl52)),
  inference(backward_demodulation,[],[f15223,f2037])).
fof(f15380,plain,(
  ( ! [X6,X7] : (s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),X7)) ) | ($spl49 | $spl52)),
  inference(subsumption_resolution,[],[f15239,f15223])).
fof(f15239,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X6,sK17('const.pred_set.EMPTY_0',X7,X6)),s(fun(X6,bool),X7))) | s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),X7)) ) | ($spl49 | $spl52)),
  inference(backward_demodulation,[],[f15223,f1929])).
fof(f15378,plain,(
  ( ! [X4,X5] : (s(fun(X4,bool),'const.pred_set.EMPTY_0') = s(fun(X4,bool),X5)) ) | ($spl49 | $spl52)),
  inference(subsumption_resolution,[],[f15238,f15223])).
fof(f15238,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X4,sK17(X5,'const.pred_set.EMPTY_0',X4)),s(fun(X4,bool),X5))) | s(fun(X4,bool),'const.pred_set.EMPTY_0') = s(fun(X4,bool),X5)) ) | ($spl49 | $spl52)),
  inference(backward_demodulation,[],[f15223,f1928])).
fof(f15375,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl49 | $spl52)),
  inference(subsumption_resolution,[],[f15224,f15223])).
fof(f15224,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl49 | $spl52)),
  inference(backward_demodulation,[],[f15223,f290])).
fof(f15223,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl49 | $spl52)),
  inference(subsumption_resolution,[],[f15210,f13301])).
fof(f15210,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl49),
  inference(trivial_inequality_removal,[],[f15204])).
fof(f15204,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,f(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl49),
  inference(superposition,[],[f13298,f747])).
fof(f15196,plain,(
  ~sP5('T_0',f(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12))),'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f268,f13298,f404])).
fof(f15197,plain,(
  sP5('F_0',f(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12))),'T_0') | $spl49),
  inference(unit_resulting_resolution,[],[f473,f13298,f402])).
fof(f15201,plain,(
  ~sP5('T_0','T_0',f(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12)))) | $spl49),
  inference(unit_resulting_resolution,[],[f268,f13298,f404])).
fof(f15202,plain,(
  sP5('F_0','T_0',f(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12)))) | $spl49),
  inference(unit_resulting_resolution,[],[f473,f13298,f402])).
fof(f15203,plain,(
  ~p(s(bool,f(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12))))) | $spl49),
  inference(unit_resulting_resolution,[],[f268,f13298,f355])).
fof(f290,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f67])).
fof(f67,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X9] : (s(bool,'F_0') = s(bool,X9) <=> ~p(s(bool,X9)))),
  file('thm.res_quan.RES_FORALL_NULL','thm.bool.EQ_CLAUSES|split|2')).
fof(f291,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f171])).
fof(f473,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f271,f268])).
fof(f271,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f167])).
fof(f167,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f52])).
fof(f52,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f26])).
fof(f26,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.res_quan.RES_FORALL_NULL','thm.bool.NOT_CLAUSES|split|2')).
fof(f592,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f332,f291])).
fof(f1220,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1))),s(X0,X2))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f381,f275])).
fof(f1477,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | s(bool,'T_0') = s(bool,sK15(X0,bool))) )),
  inference(superposition,[],[f328,f275])).
fof(f1493,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | s(bool,'T_0') = s(bool,sK16(X0,bool))) )),
  inference(superposition,[],[f329,f275])).
fof(f1511,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(bool,bool),X1),s(fun(bool,bool),X0)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))) | s(bool,'T_0') = s(bool,sK18(X0,X1,bool))) )),
  inference(superposition,[],[f338,f275])).
fof(f1536,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,'const.bool.RES_FORALL_2'(s(fun(bool,bool),X1),s(fun(bool,bool),X0)))) | s(bool,'T_0') = s(bool,sK18(X0,X1,bool))) )),
  inference(superposition,[],[f339,f275])).
fof(f1622,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK26(X0,X1,X2,bool))) )),
  inference(superposition,[],[f382,f275])).
fof(f1623,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK26(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK26(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f382,f275])).
fof(f1709,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X0,sK17(X1,X2,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,sK17(X1,X2,X0)),s(fun(X0,bool),X2)))) )),
  inference(superposition,[],[f336,f275])).
fof(f1710,plain,(
  ( ! [X0,X1] : (s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X0))) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK17(X0,X1,bool))) )),
  inference(superposition,[],[f336,f275])).
fof(f1712,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X0,sK17(X1,X2,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,sK17(X1,X2,X0)),s(fun(X0,bool),X1)))) )),
  inference(superposition,[],[f336,f275])).
fof(f1928,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X4,sK17(X5,'const.pred_set.EMPTY_0',X4)),s(fun(X4,bool),X5))) | s(fun(X4,bool),'const.pred_set.EMPTY_0') = s(fun(X4,bool),X5)) )),
  inference(superposition,[],[f336,f592])).
fof(f1929,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X6,sK17('const.pred_set.EMPTY_0',X7,X6)),s(fun(X6,bool),X7))) | s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),X7)) )),
  inference(superposition,[],[f336,f592])).
fof(f2037,plain,(
  ( ! [X0,X1] : (~sP2('F_0',happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X0,X1)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f473,f485,f481,f1206,f389])).
fof(f2043,plain,(
  ( ! [X2,X0,X1] : (~sP5('F_0',happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X1,X2)),X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1206,f275])).
fof(f2046,plain,(
  ( ! [X10,X8,X11,X9] : (~sP5('F_0',happ(s(fun(X10,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X10,X11)),'const.bool.IN_2'(s(X8,X9),s(fun(X8,bool),'const.pred_set.EMPTY_0')))) )),
  inference(superposition,[],[f1206,f592])).
fof(f2055,plain,(
  ( ! [X2,X0,X3,X1] : (~sP2('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X2,X3)),'F_0')) )),
  inference(forward_demodulation,[],[f2039,f592])).
fof(f2039,plain,(
  ( ! [X2,X0,X3,X1] : (~sP2('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),s(X2,X3)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f473,f601,f599,f1206,f389])).
fof(f2087,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0',happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X1,X2)),X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1207,f275])).
fof(f2090,plain,(
  ( ! [X10,X8,X11,X9] : (sP5('T_0',happ(s(fun(X10,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X10,X11)),'const.bool.IN_2'(s(X8,X9),s(fun(X8,bool),'const.pred_set.EMPTY_0')))) )),
  inference(superposition,[],[f1207,f592])).
fof(f2126,plain,(
  ( ! [X0,X1] : (~sP2(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X0,X1)),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f473,f486,f482,f1208,f388])).
fof(f2129,plain,(
  ( ! [X2,X0,X1] : (~sP5('F_0',X0,happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X1,X2))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1208,f275])).
fof(f2132,plain,(
  ( ! [X10,X8,X11,X9] : (~sP5('F_0','const.bool.IN_2'(s(X8,X9),s(fun(X8,bool),'const.pred_set.EMPTY_0')),happ(s(fun(X10,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X10,X11)))) )),
  inference(superposition,[],[f1208,f592])).
fof(f2142,plain,(
  ( ! [X2,X0,X3,X1] : (~sP2(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X0,X3)),'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')),'F_0')) )),
  inference(forward_demodulation,[],[f2123,f592])).
fof(f2123,plain,(
  ( ! [X2,X0,X3,X1] : (~sP2(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0'))))),s(X0,X3)),'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f473,f602,f482,f1208,f389])).
fof(f2173,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0',X0,happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X1,X2))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1209,f275])).
fof(f2176,plain,(
  ( ! [X10,X8,X11,X9] : (sP5('T_0','const.bool.IN_2'(s(X8,X9),s(fun(X8,bool),'const.pred_set.EMPTY_0')),happ(s(fun(X10,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X10,X11)))) )),
  inference(superposition,[],[f1209,f592])).
fof(f13534,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s(bool,'T_0') = s(bool,happ(s(fun(X44,bool),'_dst_x0_1'(s(bool,X42))),s(X44,X45))) | s(bool,'F_0') = s(bool,happ(s(fun(X41,bool),'_dst_x0_1'(s(bool,X42))),s(X41,X43)))) )),
  inference(superposition,[],[f1220,f381])).
fof(f13540,plain,(
  ( ! [X14,X15,X16] : (s(bool,X14) = s(bool,happ(s(fun(bool,bool),'_dst_x0_1'(s(bool,X15))),s(bool,X14))) | s(bool,'F_0') = s(bool,X15) | s(bool,X14) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f1220,f747])).
fof(f13581,plain,(
  ( ! [X134,X132,X130,X133,X131] : (~sP5('F_0',happ(s(fun(X133,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X133,X134)),happ(s(fun(X130,bool),'_dst_x0_1'(s(bool,X131))),s(X130,X132))) | s(bool,'F_0') = s(bool,X131)) )),
  inference(superposition,[],[f1206,f1220])).
fof(f13582,plain,(
  ( ! [X138,X136,X135,X139,X137] : (sP5('T_0',happ(s(fun(X138,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X138,X139)),happ(s(fun(X135,bool),'_dst_x0_1'(s(bool,X136))),s(X135,X137))) | s(bool,'F_0') = s(bool,X136)) )),
  inference(superposition,[],[f1207,f1220])).
fof(f13583,plain,(
  ( ! [X144,X142,X140,X143,X141] : (~sP5('F_0',happ(s(fun(X140,bool),'_dst_x0_1'(s(bool,X141))),s(X140,X142)),happ(s(fun(X143,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X143,X144))) | s(bool,'F_0') = s(bool,X141)) )),
  inference(superposition,[],[f1208,f1220])).
fof(f13584,plain,(
  ( ! [X146,X149,X147,X145,X148] : (sP5('T_0',happ(s(fun(X145,bool),'_dst_x0_1'(s(bool,X146))),s(X145,X147)),happ(s(fun(X148,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X148,X149))) | s(bool,'F_0') = s(bool,X146)) )),
  inference(superposition,[],[f1209,f1220])).
fof(f13683,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X22,bool),X24),s(fun(X22,bool),'_dst_x0_1'(s(bool,X23)))))) | s(bool,'F_0') = s(bool,X23)) )),
  inference(subsumption_resolution,[],[f13557,f268])).
fof(f13557,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X22,bool),X24),s(fun(X22,bool),'_dst_x0_1'(s(bool,X23)))))) | s(bool,'F_0') = s(bool,X23)) )),
  inference(superposition,[],[f339,f1220])).
fof(f13685,plain,(
  ( ! [X35,X36,X34] : (s(bool,'F_0') != s(bool,happ(s(fun(X34,bool),'_dst_x0_1'(s(bool,X35))),s(X34,X36))) | s(bool,'F_0') = s(bool,X35)) )),
  inference(subsumption_resolution,[],[f13561,f268])).
fof(f13561,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,happ(s(fun(X34,bool),'_dst_x0_1'(s(bool,X35))),s(X34,X36))) | s(bool,'F_0') = s(bool,X35)) )),
  inference(superposition,[],[f290,f1220])).
fof(f13686,plain,(
  ( ! [X45,X43,X41,X44,X42,X40] : (~p(s(bool,happ(s(fun(X43,bool),X44),s(X43,X45)))) | p(s(bool,happ(s(fun(X43,bool),X44),s(X43,sK19(X44,happ(s(fun(X40,bool),'_dst_x0_1'(s(bool,X41))),s(X40,X42)),X43))))) | s(bool,'F_0') = s(bool,X41)) )),
  inference(subsumption_resolution,[],[f13563,f268])).
fof(f13563,plain,(
  ( ! [X45,X43,X41,X44,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X43,bool),X44),s(X43,X45)))) | p(s(bool,happ(s(fun(X43,bool),X44),s(X43,sK19(X44,happ(s(fun(X40,bool),'_dst_x0_1'(s(bool,X41))),s(X40,X42)),X43))))) | s(bool,'F_0') = s(bool,X41)) )),
  inference(superposition,[],[f341,f1220])).
fof(f13687,plain,(
  ( ! [X47,X50,X48,X46,X51,X49] : (~p(s(bool,happ(s(fun(X49,bool),X50),s(X49,X51)))) | p(s(bool,happ(s(fun(X49,bool),X50),s(X49,sK20(X50,happ(s(fun(X46,bool),'_dst_x0_1'(s(bool,X47))),s(X46,X48)),X49))))) | s(bool,'F_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f13564,f268])).
fof(f13564,plain,(
  ( ! [X47,X50,X48,X46,X51,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X49,bool),X50),s(X49,X51)))) | p(s(bool,happ(s(fun(X49,bool),X50),s(X49,sK20(X50,happ(s(fun(X46,bool),'_dst_x0_1'(s(bool,X47))),s(X46,X48)),X49))))) | s(bool,'F_0') = s(bool,X47)) )),
  inference(superposition,[],[f343,f1220])).
fof(f13688,plain,(
  ( ! [X57,X54,X52,X56,X55,X53] : (p(s(bool,happ(s(fun(X55,bool),X56),s(X55,sK21(happ(s(fun(X52,bool),'_dst_x0_1'(s(bool,X53))),s(X52,X54)),X56,X55))))) | ~p(s(bool,happ(s(fun(X55,bool),X56),s(X55,X57)))) | s(bool,'F_0') = s(bool,X53)) )),
  inference(subsumption_resolution,[],[f13565,f268])).
fof(f13565,plain,(
  ( ! [X57,X54,X52,X56,X55,X53] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X55,bool),X56),s(X55,sK21(happ(s(fun(X52,bool),'_dst_x0_1'(s(bool,X53))),s(X52,X54)),X56,X55))))) | ~p(s(bool,happ(s(fun(X55,bool),X56),s(X55,X57)))) | s(bool,'F_0') = s(bool,X53)) )),
  inference(superposition,[],[f344,f1220])).
fof(f13689,plain,(
  ( ! [X61,X59,X62,X60,X58,X63] : (p(s(bool,happ(s(fun(X61,bool),X62),s(X61,sK22(happ(s(fun(X58,bool),'_dst_x0_1'(s(bool,X59))),s(X58,X60)),X62,X61))))) | ~p(s(bool,happ(s(fun(X61,bool),X62),s(X61,X63)))) | s(bool,'F_0') = s(bool,X59)) )),
  inference(subsumption_resolution,[],[f13566,f268])).
fof(f13566,plain,(
  ( ! [X61,X59,X62,X60,X58,X63] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X61,bool),X62),s(X61,sK22(happ(s(fun(X58,bool),'_dst_x0_1'(s(bool,X59))),s(X58,X60)),X62,X61))))) | ~p(s(bool,happ(s(fun(X61,bool),X62),s(X61,X63)))) | s(bool,'F_0') = s(bool,X59)) )),
  inference(superposition,[],[f346,f1220])).
fof(f13691,plain,(
  ( ! [X75,X73] : (s(bool,X73) = s(bool,X75) | ~p(s(bool,X75)) | s(bool,'F_0') = s(bool,X73)) )),
  inference(forward_demodulation,[],[f13690,f381])).
fof(f13690,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,X75)) | s(bool,X75) = s(bool,happ(s(fun(X72,bool),'_dst_x0_1'(s(bool,X73))),s(X72,X74))) | s(bool,'F_0') = s(bool,X73)) )),
  inference(subsumption_resolution,[],[f13569,f268])).
fof(f13569,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,'T_0')) | ~p(s(bool,X75)) | s(bool,X75) = s(bool,happ(s(fun(X72,bool),'_dst_x0_1'(s(bool,X73))),s(X72,X74))) | s(bool,'F_0') = s(bool,X73)) )),
  inference(superposition,[],[f355,f1220])).
fof(f13692,plain,(
  ( ! [X78,X76,X79,X77] : (sP0(X79,happ(s(fun(X76,bool),'_dst_x0_1'(s(bool,X77))),s(X76,X78))) | s(bool,'F_0') = s(bool,X77)) )),
  inference(subsumption_resolution,[],[f13570,f268])).
fof(f13570,plain,(
  ( ! [X78,X76,X79,X77] : (~p(s(bool,'T_0')) | sP0(X79,happ(s(fun(X76,bool),'_dst_x0_1'(s(bool,X77))),s(X76,X78))) | s(bool,'F_0') = s(bool,X77)) )),
  inference(superposition,[],[f357,f1220])).
fof(f13693,plain,(
  ( ! [X80,X83,X81,X82] : (sP0(happ(s(fun(X80,bool),'_dst_x0_1'(s(bool,X81))),s(X80,X82)),X83) | s(bool,'F_0') = s(bool,X81)) )),
  inference(subsumption_resolution,[],[f13571,f268])).
fof(f13571,plain,(
  ( ! [X80,X83,X81,X82] : (~p(s(bool,'T_0')) | sP0(happ(s(fun(X80,bool),'_dst_x0_1'(s(bool,X81))),s(X80,X82)),X83) | s(bool,'F_0') = s(bool,X81)) )),
  inference(superposition,[],[f358,f1220])).
fof(f13694,plain,(
  ( ! [X87,X85,X86,X84] : (p(s(bool,X87)) | sP1(happ(s(fun(X84,bool),'_dst_x0_1'(s(bool,X85))),s(X84,X86)),X87) | s(bool,'F_0') = s(bool,X85)) )),
  inference(subsumption_resolution,[],[f13572,f268])).
fof(f13572,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,'T_0')) | p(s(bool,X87)) | sP1(happ(s(fun(X84,bool),'_dst_x0_1'(s(bool,X85))),s(X84,X86)),X87) | s(bool,'F_0') = s(bool,X85)) )),
  inference(superposition,[],[f373,f1220])).
fof(f13695,plain,(
  ( ! [X90,X88,X91,X89] : (p(s(bool,X91)) | sP1(X91,happ(s(fun(X88,bool),'_dst_x0_1'(s(bool,X89))),s(X88,X90))) | s(bool,'F_0') = s(bool,X89)) )),
  inference(subsumption_resolution,[],[f13573,f268])).
fof(f13573,plain,(
  ( ! [X90,X88,X91,X89] : (~p(s(bool,'T_0')) | p(s(bool,X91)) | sP1(X91,happ(s(fun(X88,bool),'_dst_x0_1'(s(bool,X89))),s(X88,X90))) | s(bool,'F_0') = s(bool,X89)) )),
  inference(superposition,[],[f376,f1220])).
fof(f13696,plain,(
  ( ! [X94,X92,X95,X93,X96] : (~p(s(bool,X95)) | p(s(bool,X96)) | ~sP4(X95,happ(s(fun(X92,bool),'_dst_x0_1'(s(bool,X93))),s(X92,X94)),X96) | s(bool,'F_0') = s(bool,X93)) )),
  inference(subsumption_resolution,[],[f13574,f268])).
fof(f13574,plain,(
  ( ! [X94,X92,X95,X93,X96] : (~p(s(bool,'T_0')) | ~p(s(bool,X95)) | p(s(bool,X96)) | ~sP4(X95,happ(s(fun(X92,bool),'_dst_x0_1'(s(bool,X93))),s(X92,X94)),X96) | s(bool,'F_0') = s(bool,X93)) )),
  inference(superposition,[],[f390,f1220])).
fof(f13697,plain,(
  ( ! [X101,X99,X97,X100,X98] : (~p(s(bool,X100)) | p(s(bool,X101)) | ~sP7(X101,X100,happ(s(fun(X97,bool),'_dst_x0_1'(s(bool,X98))),s(X97,X99))) | s(bool,'F_0') = s(bool,X98)) )),
  inference(subsumption_resolution,[],[f13575,f268])).
fof(f13575,plain,(
  ( ! [X101,X99,X97,X100,X98] : (~p(s(bool,'T_0')) | ~p(s(bool,X100)) | p(s(bool,X101)) | ~sP7(X101,X100,happ(s(fun(X97,bool),'_dst_x0_1'(s(bool,X98))),s(X97,X99))) | s(bool,'F_0') = s(bool,X98)) )),
  inference(superposition,[],[f422,f1220])).
fof(f13698,plain,(
  ( ! [X103,X105,X102,X106,X104] : (p(s(bool,X105)) | p(s(bool,X106)) | ~sP9(happ(s(fun(X102,bool),'_dst_x0_1'(s(bool,X103))),s(X102,X104)),X106,X105) | s(bool,'F_0') = s(bool,X103)) )),
  inference(subsumption_resolution,[],[f13576,f268])).
fof(f13576,plain,(
  ( ! [X103,X105,X102,X106,X104] : (~p(s(bool,'T_0')) | p(s(bool,X105)) | p(s(bool,X106)) | ~sP9(happ(s(fun(X102,bool),'_dst_x0_1'(s(bool,X103))),s(X102,X104)),X106,X105) | s(bool,'F_0') = s(bool,X103)) )),
  inference(superposition,[],[f437,f1220])).
fof(f13699,plain,(
  ( ! [X111,X109,X107,X110,X108] : (~p(s(bool,X110)) | p(s(bool,X111)) | ~sP11(happ(s(fun(X107,bool),'_dst_x0_1'(s(bool,X108))),s(X107,X109)),X111,X110) | s(bool,'F_0') = s(bool,X108)) )),
  inference(subsumption_resolution,[],[f13577,f268])).
fof(f13577,plain,(
  ( ! [X111,X109,X107,X110,X108] : (~p(s(bool,'T_0')) | ~p(s(bool,X110)) | p(s(bool,X111)) | ~sP11(happ(s(fun(X107,bool),'_dst_x0_1'(s(bool,X108))),s(X107,X109)),X111,X110) | s(bool,'F_0') = s(bool,X108)) )),
  inference(superposition,[],[f454,f1220])).
fof(f13771,plain,(
  ( ! [X47,X52,X50,X48,X46,X51,X49] : (~p(s(bool,X49)) | p(s(bool,X50)) | ~sP4(happ(s(fun(X46,bool),'_dst_x0_1'(s(bool,X47))),s(X46,X48)),happ(s(fun(X51,bool),'_dst_x0_1'(s(bool,X49))),s(X51,X52)),X50) | s(bool,'F_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f13759,f268])).
fof(f13759,plain,(
  ( ! [X47,X52,X50,X48,X46,X51,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X49)) | p(s(bool,X50)) | ~sP4(happ(s(fun(X46,bool),'_dst_x0_1'(s(bool,X47))),s(X46,X48)),happ(s(fun(X51,bool),'_dst_x0_1'(s(bool,X49))),s(X51,X52)),X50) | s(bool,'F_0') = s(bool,X47)) )),
  inference(superposition,[],[f1233,f1220])).
fof(f13803,plain,(
  ( ! [X47,X52,X50,X48,X46,X51,X49] : (~p(s(bool,X49)) | p(s(bool,X50)) | ~sP7(X50,happ(s(fun(X46,bool),'_dst_x0_1'(s(bool,X47))),s(X46,X48)),happ(s(fun(X51,bool),'_dst_x0_1'(s(bool,X49))),s(X51,X52))) | s(bool,'F_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f13792,f268])).
fof(f13792,plain,(
  ( ! [X47,X52,X50,X48,X46,X51,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X49)) | p(s(bool,X50)) | ~sP7(X50,happ(s(fun(X46,bool),'_dst_x0_1'(s(bool,X47))),s(X46,X48)),happ(s(fun(X51,bool),'_dst_x0_1'(s(bool,X49))),s(X51,X52))) | s(bool,'F_0') = s(bool,X47)) )),
  inference(superposition,[],[f1234,f1220])).
fof(f13948,plain,(
  ( ! [X47,X52,X50,X48,X46,X51,X49] : (~p(s(bool,X49)) | p(s(bool,X50)) | ~sP11(happ(s(fun(X51,bool),'_dst_x0_1'(s(bool,X49))),s(X51,X52)),X50,happ(s(fun(X46,bool),'_dst_x0_1'(s(bool,X47))),s(X46,X48))) | s(bool,'F_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f13938,f268])).
fof(f13938,plain,(
  ( ! [X47,X52,X50,X48,X46,X51,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X49)) | p(s(bool,X50)) | ~sP11(happ(s(fun(X51,bool),'_dst_x0_1'(s(bool,X49))),s(X51,X52)),X50,happ(s(fun(X46,bool),'_dst_x0_1'(s(bool,X47))),s(X46,X48))) | s(bool,'F_0') = s(bool,X47)) )),
  inference(superposition,[],[f1236,f1220])).
fof(f14056,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X3)))) | ~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK15(happ(s(fun(X0,fun(bool,bool)),'_dst_x0_1'(s(fun(bool,bool),X1))),s(X0,X2)),bool))) )),
  inference(forward_demodulation,[],[f14042,f381])).
fof(f14042,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(X0,fun(bool,bool)),'_dst_x0_1'(s(fun(bool,bool),X1))),s(X0,X2))),s(bool,X3)))) | s(bool,'T_0') = s(bool,sK15(happ(s(fun(X0,fun(bool,bool)),'_dst_x0_1'(s(fun(bool,bool),X1))),s(X0,X2)),bool))) )),
  inference(superposition,[],[f1477,f381])).
fof(f14059,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,happ(s(fun(bool,bool),X2),s(bool,X3)))) | s(bool,'T_0') = s(bool,sK15(X2,bool)) | s(bool,X4) = s(bool,happ(s(fun(bool,bool),X2),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f14050,f268])).
fof(f14050,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),X2),s(bool,X3)))) | s(bool,'T_0') = s(bool,sK15(X2,bool)) | s(bool,X4) = s(bool,happ(s(fun(bool,bool),X2),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f1477,f747])).
fof(f14116,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X3)))) | ~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK16(happ(s(fun(X0,fun(bool,bool)),'_dst_x0_1'(s(fun(bool,bool),X1))),s(X0,X2)),bool))) )),
  inference(forward_demodulation,[],[f14102,f381])).
fof(f14102,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(X0,fun(bool,bool)),'_dst_x0_1'(s(fun(bool,bool),X1))),s(X0,X2))),s(bool,X3)))) | s(bool,'T_0') = s(bool,sK16(happ(s(fun(X0,fun(bool,bool)),'_dst_x0_1'(s(fun(bool,bool),X1))),s(X0,X2)),bool))) )),
  inference(superposition,[],[f1493,f381])).
fof(f14119,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,happ(s(fun(bool,bool),X2),s(bool,X3)))) | s(bool,'T_0') = s(bool,sK16(X2,bool)) | s(bool,X4) = s(bool,happ(s(fun(bool,bool),X2),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f14110,f268])).
fof(f14110,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),X2),s(bool,X3)))) | s(bool,'T_0') = s(bool,sK16(X2,bool)) | s(bool,X4) = s(bool,happ(s(fun(bool,bool),X2),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f1493,f747])).
fof(f14170,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(bool,bool),X3),s(fun(bool,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X3)))) | s(bool,'T_0') = s(bool,sK18(happ(s(fun(X0,fun(bool,bool)),'_dst_x0_1'(s(fun(bool,bool),X1))),s(X0,X2)),X3,bool))) )),
  inference(superposition,[],[f1511,f381])).
fof(f14176,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))) | p(s(bool,'const.bool.RES_FORALL_2'(s(fun(bool,bool),X1),s(fun(bool,bool),X3)))) | s(bool,'T_0') = s(bool,sK18(X3,happ(s(fun(X0,fun(bool,bool)),'_dst_x0_1'(s(fun(bool,bool),X1))),s(X0,X2)),bool))) )),
  inference(forward_demodulation,[],[f14169,f381])).
fof(f14169,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(bool,bool),X1),s(fun(bool,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),happ(s(fun(X0,fun(bool,bool)),'_dst_x0_1'(s(fun(bool,bool),X1))),s(X0,X2)))))) | s(bool,'T_0') = s(bool,sK18(X3,happ(s(fun(X0,fun(bool,bool)),'_dst_x0_1'(s(fun(bool,bool),X1))),s(X0,X2)),bool))) )),
  inference(superposition,[],[f1511,f381])).
fof(f14178,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))) | s(bool,'T_0') = s(bool,sK18(X1,X0,bool)) | s(bool,'T_0') = s(bool,'const.bool.RES_FORALL_2'(s(fun(bool,bool),X0),s(fun(bool,bool),X1)))) )),
  inference(condensation,[],[f14177])).
fof(f14177,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X8)) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X6)))) | s(bool,'T_0') = s(bool,sK18(X7,X6,bool)) | s(bool,'T_0') = s(bool,'const.bool.RES_FORALL_2'(s(fun(bool,bool),X6),s(fun(bool,bool),X7)))) )),
  inference(subsumption_resolution,[],[f14173,f286])).
fof(f14173,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X8)) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X6)))) | s(bool,'T_0') = s(bool,sK18(X7,X6,bool)) | s(bool,'T_0') = s(bool,'const.bool.RES_FORALL_2'(s(fun(bool,bool),X6),s(fun(bool,bool),X7))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f1511,f747])).
fof(f3499,plain,(
  ( ! [X505,X502,X504,X503] : (s(bool,'T_0') != s(bool,f(s(X502,sK26(X503,X504,bool,X502)))) | s(fun(X502,bool),X503) = s(fun(X502,bool),X504) | s(bool,X505) = s(bool,g(s(X502,sK26(X503,X504,bool,X502)))) | s(bool,'T_0') = s(bool,X505)) )),
  inference(superposition,[],[f382,f747])).
fof(f2968,plain,(
  ( ! [X509,X507,X508,X506] : (s(bool,X509) != s(bool,f(s(X506,sK26(X507,X508,bool,X506)))) | s(fun(X506,bool),X507) = s(fun(X506,bool),X508) | s(bool,'T_0') = s(bool,X509) | s(bool,'T_0') = s(bool,g(s(X506,sK26(X507,X508,bool,X506))))) )),
  inference(superposition,[],[f382,f747])).
fof(f2973,plain,(
  ( ! [X524,X525,X523] : (~p(s(bool,happ(s(fun(bool,bool),X523),s(bool,X525)))) | p(s(bool,'const.bool.RES_FORALL_2'(s(fun(bool,bool),X524),s(fun(bool,bool),X523)))) | s(bool,'T_0') = s(bool,X525) | s(bool,'T_0') = s(bool,sK18(X523,X524,bool))) )),
  inference(superposition,[],[f339,f747])).
fof(f2974,plain,(
  ( ! [X528,X526,X527] : (p(s(bool,'const.bool.IN_2'(s(bool,X528),s(fun(bool,bool),X527)))) | p(s(bool,'const.bool.RES_FORALL_2'(s(fun(bool,bool),X527),s(fun(bool,bool),X526)))) | s(bool,'T_0') = s(bool,X528) | s(bool,'T_0') = s(bool,sK18(X526,X527,bool))) )),
  inference(superposition,[],[f338,f747])).
fof(f3504,plain,(
  ( ! [X520,X521,X519] : (~p(s(bool,happ(s(fun(bool,bool),X519),s(bool,'T_0')))) | p(s(bool,'const.bool.RES_FORALL_2'(s(fun(bool,bool),X520),s(fun(bool,bool),X519)))) | s(bool,X521) = s(bool,sK18(X519,X520,bool)) | s(bool,'T_0') = s(bool,X521)) )),
  inference(superposition,[],[f339,f747])).
fof(f3505,plain,(
  ( ! [X524,X522,X523] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X523)))) | p(s(bool,'const.bool.RES_FORALL_2'(s(fun(bool,bool),X523),s(fun(bool,bool),X522)))) | s(bool,X524) = s(bool,sK18(X522,X523,bool)) | s(bool,'T_0') = s(bool,X524)) )),
  inference(superposition,[],[f338,f747])).
fof(f2002,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X2,bool),'const.pred_set.EMPTY_0'),s(fun(X2,bool),X3))))) )),
  inference(subsumption_resolution,[],[f1927,f473])).
fof(f1927,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X2,bool),'const.pred_set.EMPTY_0'),s(fun(X2,bool),X3))))) )),
  inference(superposition,[],[f338,f592])).
fof(f337,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f206])).
fof(f338,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK18(X2,X1,X0)),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f206])).
fof(f1514,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X4)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK18(X4,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X0)),s(fun(X0,bool),X2))))) )),
  inference(forward_demodulation,[],[f1512,f381])).
fof(f1512,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK18(X4,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X0)),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))),s(fun(X0,bool),X4))))) )),
  inference(superposition,[],[f338,f381])).
fof(f1515,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,sK18(X1,X2,X0)),s(fun(X0,bool),X2)))) )),
  inference(subsumption_resolution,[],[f1513,f473])).
fof(f1513,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,sK18(X1,X2,X0)),s(fun(X0,bool),X2)))) )),
  inference(superposition,[],[f338,f275])).
fof(f1539,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X4),s(fun(X0,bool),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK18(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X4,X0)))))) )),
  inference(forward_demodulation,[],[f1535,f381])).
fof(f1535,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK18(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X4,X0))))) | p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X4),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))))))) )),
  inference(superposition,[],[f339,f381])).
fof(f4463,plain,(
  ( ! [X222,X225,X223,X221,X224] : (~p(s(bool,'const.bool.IN_2'(s(X221,X224),s(fun(X221,bool),X222)))) | p(s(bool,happ(s(fun(X221,bool),X223),s(X221,X224)))) | s(bool,X225) = s(bool,'const.bool.RES_FORALL_2'(s(fun(X221,bool),X222),s(fun(X221,bool),X223))) | s(bool,'T_0') = s(bool,X225)) )),
  inference(subsumption_resolution,[],[f3407,f268])).
fof(f3407,plain,(
  ( ! [X222,X225,X223,X221,X224] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(X221,X224),s(fun(X221,bool),X222)))) | p(s(bool,happ(s(fun(X221,bool),X223),s(X221,X224)))) | s(bool,X225) = s(bool,'const.bool.RES_FORALL_2'(s(fun(X221,bool),X222),s(fun(X221,bool),X223))) | s(bool,'T_0') = s(bool,X225)) )),
  inference(superposition,[],[f337,f747])).
fof(f4474,plain,(
  ( ! [X271,X273,X270,X272] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X270,bool),X272),s(fun(X270,bool),X271)))) | s(bool,X273) = s(bool,happ(s(fun(X270,bool),X271),s(X270,sK18(X271,X272,X270)))) | s(bool,'T_0') = s(bool,X273)) )),
  inference(subsumption_resolution,[],[f3418,f268])).
fof(f3418,plain,(
  ( ! [X271,X273,X270,X272] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X270,bool),X272),s(fun(X270,bool),X271)))) | s(bool,X273) = s(bool,happ(s(fun(X270,bool),X271),s(X270,sK18(X271,X272,X270)))) | s(bool,'T_0') = s(bool,X273)) )),
  inference(superposition,[],[f339,f747])).
fof(f13988,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X4),s(fun(X0,bool),X2)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,sK18(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X4,X0)),s(fun(X0,bool),X4)))) )),
  inference(superposition,[],[f1515,f381])).
fof(f13994,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,sK18(X4,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X0)),s(fun(X0,bool),X2))) | p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X4))))) )),
  inference(forward_demodulation,[],[f13987,f381])).
fof(f13987,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X4)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,sK18(X4,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X0)),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)))))) )),
  inference(superposition,[],[f1515,f381])).
fof(f13998,plain,(
  ( ! [X17,X18,X16] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X16,sK18(X18,X17,X16)),s(fun(X16,bool),X17))) | s(bool,'T_0') = s(bool,'const.bool.RES_FORALL_2'(s(fun(X16,bool),X17),s(fun(X16,bool),X18)))) )),
  inference(subsumption_resolution,[],[f13993,f473])).
fof(f13993,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X16,sK18(X18,X17,X16)),s(fun(X16,bool),X17))) | s(bool,'T_0') = s(bool,'const.bool.RES_FORALL_2'(s(fun(X16,bool),X17),s(fun(X16,bool),X18)))) )),
  inference(superposition,[],[f1515,f275])).
fof(f13453,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X1,bool),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X0)))))) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f13442,f268])).
fof(f13442,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X1,bool),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X0)))))) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1538,f747])).
fof(f2970,plain,(
  ( ! [X516,X514,X515] : (~p(s(bool,happ(s(fun(bool,bool),X514),s(bool,X515)))) | p(s(bool,happ(s(fun(bool,bool),X514),s(bool,X516)))) | s(bool,'T_0') = s(bool,X515) | s(bool,'T_0') = s(bool,sK15(X514,bool))) )),
  inference(superposition,[],[f328,f747])).
fof(f2971,plain,(
  ( ! [X518,X517,X519] : (~p(s(bool,happ(s(fun(bool,bool),X517),s(bool,X518)))) | p(s(bool,happ(s(fun(bool,bool),X517),s(bool,X519)))) | s(bool,'T_0') = s(bool,X518) | s(bool,'T_0') = s(bool,sK16(X517,bool))) )),
  inference(superposition,[],[f329,f747])).
fof(f3501,plain,(
  ( ! [X511,X512,X510] : (~p(s(bool,happ(s(fun(bool,bool),X510),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X510),s(bool,X511)))) | s(bool,X512) = s(bool,sK15(X510,bool)) | s(bool,'T_0') = s(bool,X512)) )),
  inference(superposition,[],[f328,f747])).
fof(f3502,plain,(
  ( ! [X514,X513,X515] : (~p(s(bool,happ(s(fun(bool,bool),X513),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X513),s(bool,X514)))) | s(bool,X515) = s(bool,sK16(X513,bool)) | s(bool,'T_0') = s(bool,X515)) )),
  inference(superposition,[],[f329,f747])).
fof(f328,plain,(
  ( ! [X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK15(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK15(X1,X0)))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK16(X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK15,sK16])],[f196])).
fof(f196,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f195])).
fof(f195,plain,(
  ! [X0,X1] : ((! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0,X1] : (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(ennf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0,X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(flattening,[],[f86])).
fof(f86,plain,(
  ! [X0] : ! [X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X6] : ! [X14] : (~! [X12] : p(s(bool,happ(s(fun(X6,bool),X14),s(X6,X12)))) <=> ? [X12] : ~p(s(bool,happ(s(fun(X6,bool),X14),s(X6,X12)))))),
  file('thm.res_quan.RES_FORALL_NULL','thm.bool.NOT_FORALL_THM')).
fof(f329,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK16(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f197])).
fof(f341,plain,(
  ( ! [X6,X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK19(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f210])).
fof(f210,plain,(
  ! [X0,X1,X2] : ((! [X6] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK19(X2,X1,X0))))))) & (~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK20(X2,X1,X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK19,sK20])],[f209])).
fof(f209,plain,(
  ! [X0,X1,X2] : ((! [X6] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | (p(s(bool,X1)) & ? [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & (~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f208])).
fof(f208,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f207])).
fof(f207,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f97])).
fof(f97,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X6] : ! [X14,X15] : (? [X12] : (p(s(bool,X14)) & p(s(bool,happ(s(fun(X6,bool),X15),s(X6,X12))))) <=> (p(s(bool,X14)) & ? [X12] : p(s(bool,happ(s(fun(X6,bool),X15),s(X6,X12))))))),
  file('thm.res_quan.RES_FORALL_NULL','thm.bool.RIGHT_EXISTS_AND_THM')).
fof(f343,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK20(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f210])).
fof(f344,plain,(
  ( ! [X6,X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK21(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6))))) )),
  inference(cnf_transformation,[],[f214])).
fof(f214,plain,(
  ! [X0,X1,X2] : ((! [X6] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | ~p(s(bool,X2))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK21(X2,X1,X0))))) & p(s(bool,X2)))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2)) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK22(X2,X1,X0))))) & p(s(bool,X2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK21,sK22])],[f213])).
fof(f213,plain,(
  ! [X0,X1,X2] : ((! [X6] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | ~p(s(bool,X2))) | (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & p(s(bool,X2)))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2)) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(rectify,[],[f212])).
fof(f212,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,X2))) | (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2)))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2)) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(flattening,[],[f211])).
fof(f211,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,X2))) | (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2)))) & ((! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(nnf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2))))),
  inference(flattening,[],[f99])).
fof(f99,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X6] : ! [X14,X15] : (? [X12] : (p(s(bool,happ(s(fun(X6,bool),X14),s(X6,X12)))) & p(s(bool,X15))) <=> (? [X12] : p(s(bool,happ(s(fun(X6,bool),X14),s(X6,X12)))) & p(s(bool,X15))))),
  file('thm.res_quan.RES_FORALL_NULL','thm.bool.LEFT_EXISTS_AND_THM')).
fof(f346,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK22(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f214])).
fof(f348,plain,(
  ( ! [X2,X0,X8,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK23(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK24(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f218])).
fof(f218,plain,(
  ! [X0,X1,X2] : ((! [X8] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK23(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK24(X2,X1,X0)))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK25(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK25(X2,X1,X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK23,sK24,sK25])],[f217])).
fof(f217,plain,(
  ! [X0,X1,X2] : ((! [X8] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) | ? [X7] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X7)))) | ? [X6] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f216])).
fof(f216,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f215])).
fof(f215,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f101])).
fof(f101,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X6] : ! [X14,X15] : (? [X12] : (p(s(bool,happ(s(fun(X6,bool),X14),s(X6,X12)))) | p(s(bool,happ(s(fun(X6,bool),X15),s(X6,X12))))) <=> (? [X12] : p(s(bool,happ(s(fun(X6,bool),X14),s(X6,X12)))) | ? [X12] : p(s(bool,happ(s(fun(X6,bool),X15),s(X6,X12))))))),
  file('thm.res_quan.RES_FORALL_NULL','thm.bool.EXISTS_OR_THM')).
fof(f349,plain,(
  ( ! [X2,X0,X8,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK23(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK24(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f218])).
fof(f350,plain,(
  ( ! [X2,X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK25(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK25(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f218])).
fof(f351,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK25(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK25(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f218])).
fof(f1480,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK15(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X0)))))) )),
  inference(forward_demodulation,[],[f1476,f381])).
fof(f1476,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK15(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X0))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))),s(X0,X4))))) )),
  inference(superposition,[],[f328,f381])).
fof(f1496,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK16(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X0)))))) )),
  inference(forward_demodulation,[],[f1492,f381])).
fof(f1492,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK16(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X0))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))),s(X0,X4))))) )),
  inference(superposition,[],[f329,f381])).
fof(f1562,plain,(
  ( ! [X14,X19,X17,X15,X18,X16] : (~p(s(bool,X15)) | ~p(s(bool,happ(s(fun(X17,bool),X18),s(X17,X19)))) | p(s(bool,happ(s(fun(X17,bool),X18),s(X17,sK19(X18,happ(s(fun(X14,bool),'_dst_x0_1'(s(bool,X15))),s(X14,X16)),X17)))))) )),
  inference(superposition,[],[f341,f381])).
fof(f1578,plain,(
  ( ! [X14,X19,X17,X15,X18,X16] : (~p(s(bool,X15)) | ~p(s(bool,happ(s(fun(X17,bool),X18),s(X17,X19)))) | p(s(bool,happ(s(fun(X17,bool),X18),s(X17,sK20(X18,happ(s(fun(X14,bool),'_dst_x0_1'(s(bool,X15))),s(X14,X16)),X17)))))) )),
  inference(superposition,[],[f343,f381])).
fof(f1593,plain,(
  ( ! [X14,X19,X17,X15,X18,X16] : (~p(s(bool,X15)) | p(s(bool,happ(s(fun(X17,bool),X18),s(X17,sK21(happ(s(fun(X14,bool),'_dst_x0_1'(s(bool,X15))),s(X14,X16)),X18,X17))))) | ~p(s(bool,happ(s(fun(X17,bool),X18),s(X17,X19))))) )),
  inference(superposition,[],[f344,f381])).
fof(f1610,plain,(
  ( ! [X14,X19,X17,X15,X18,X16] : (~p(s(bool,X15)) | p(s(bool,happ(s(fun(X17,bool),X18),s(X17,sK22(happ(s(fun(X14,bool),'_dst_x0_1'(s(bool,X15))),s(X14,X16)),X18,X17))))) | ~p(s(bool,happ(s(fun(X17,bool),X18),s(X17,X19))))) )),
  inference(superposition,[],[f346,f381])).
fof(f1741,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK23(X5,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X5),s(X0,sK24(X5,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X0)))))) )),
  inference(forward_demodulation,[],[f1733,f381])).
fof(f1733,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))),s(X0,sK23(X5,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X0))))) | p(s(bool,happ(s(fun(X0,bool),X5),s(X0,sK24(X5,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X0)))))) )),
  inference(superposition,[],[f348,f381])).
fof(f1771,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK24(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X5),s(X0,sK23(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X0)))))) )),
  inference(forward_demodulation,[],[f1763,f381])).
fof(f1763,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X5),s(X0,sK23(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X0))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))),s(X0,sK24(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X0)))))) )),
  inference(superposition,[],[f349,f381])).
fof(f1800,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK25(X5,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X5),s(X0,sK25(X5,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X0)))))) )),
  inference(forward_demodulation,[],[f1792,f381])).
fof(f1792,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))),s(X0,sK25(X5,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X0))))) | p(s(bool,happ(s(fun(X0,bool),X5),s(X0,sK25(X5,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X0)))))) )),
  inference(superposition,[],[f350,f381])).
fof(f1827,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK25(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X5),s(X0,sK25(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X0)))))) )),
  inference(forward_demodulation,[],[f1819,f381])).
fof(f1819,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X5),s(X0,sK25(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X0))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))),s(X0,sK25(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X0)))))) )),
  inference(superposition,[],[f351,f381])).
fof(f4450,plain,(
  ( ! [X14,X12,X10,X13,X11] : (~p(s(bool,happ(s(fun(X11,bool),X12),s(X11,X13)))) | p(s(bool,happ(s(fun(X11,bool),X12),s(X11,sK19(X12,X10,X11))))) | s(bool,X10) = s(bool,X14) | s(bool,'T_0') = s(bool,X14)) )),
  inference(subsumption_resolution,[],[f3324,f268])).
fof(f3324,plain,(
  ( ! [X14,X12,X10,X13,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X11,bool),X12),s(X11,X13)))) | p(s(bool,happ(s(fun(X11,bool),X12),s(X11,sK19(X12,X10,X11))))) | s(bool,X10) = s(bool,X14) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f341,f747])).
fof(f4451,plain,(
  ( ! [X19,X17,X15,X18,X16] : (~p(s(bool,happ(s(fun(X16,bool),X17),s(X16,X18)))) | p(s(bool,happ(s(fun(X16,bool),X17),s(X16,sK20(X17,X15,X16))))) | s(bool,X15) = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) )),
  inference(subsumption_resolution,[],[f3325,f268])).
fof(f3325,plain,(
  ( ! [X19,X17,X15,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X16,bool),X17),s(X16,X18)))) | p(s(bool,happ(s(fun(X16,bool),X17),s(X16,sK20(X17,X15,X16))))) | s(bool,X15) = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f343,f747])).
fof(f4452,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,happ(s(fun(X21,bool),X22),s(X21,sK21(X20,X22,X21))))) | ~p(s(bool,happ(s(fun(X21,bool),X22),s(X21,X23)))) | s(bool,X20) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f3326,f268])).
fof(f3326,plain,(
  ( ! [X24,X23,X21,X22,X20] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X21,bool),X22),s(X21,sK21(X20,X22,X21))))) | ~p(s(bool,happ(s(fun(X21,bool),X22),s(X21,X23)))) | s(bool,X20) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f344,f747])).
fof(f4453,plain,(
  ( ! [X28,X26,X29,X27,X25] : (p(s(bool,happ(s(fun(X26,bool),X27),s(X26,sK22(X25,X27,X26))))) | ~p(s(bool,happ(s(fun(X26,bool),X27),s(X26,X28)))) | s(bool,X25) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(subsumption_resolution,[],[f3327,f268])).
fof(f3327,plain,(
  ( ! [X28,X26,X29,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X26,bool),X27),s(X26,sK22(X25,X27,X26))))) | ~p(s(bool,happ(s(fun(X26,bool),X27),s(X26,X28)))) | s(bool,X25) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(superposition,[],[f346,f747])).
fof(f4464,plain,(
  ( ! [X229,X227,X230,X228,X226] : (p(s(bool,happ(s(fun(X226,bool),X227),s(X226,sK23(X229,X227,X226))))) | p(s(bool,happ(s(fun(X226,bool),X229),s(X226,sK24(X229,X227,X226))))) | s(bool,X230) = s(bool,happ(s(fun(X226,bool),X227),s(X226,X228))) | s(bool,'T_0') = s(bool,X230)) )),
  inference(subsumption_resolution,[],[f3408,f268])).
fof(f3408,plain,(
  ( ! [X229,X227,X230,X228,X226] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X226,bool),X227),s(X226,sK23(X229,X227,X226))))) | p(s(bool,happ(s(fun(X226,bool),X229),s(X226,sK24(X229,X227,X226))))) | s(bool,X230) = s(bool,happ(s(fun(X226,bool),X227),s(X226,X228))) | s(bool,'T_0') = s(bool,X230)) )),
  inference(superposition,[],[f348,f747])).
fof(f4465,plain,(
  ( ! [X231,X235,X233,X234,X232] : (p(s(bool,happ(s(fun(X231,bool),X234),s(X231,sK23(X232,X234,X231))))) | p(s(bool,happ(s(fun(X231,bool),X232),s(X231,sK24(X232,X234,X231))))) | s(bool,X235) = s(bool,happ(s(fun(X231,bool),X232),s(X231,X233))) | s(bool,'T_0') = s(bool,X235)) )),
  inference(subsumption_resolution,[],[f3409,f268])).
fof(f3409,plain,(
  ( ! [X231,X235,X233,X234,X232] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X231,bool),X234),s(X231,sK23(X232,X234,X231))))) | p(s(bool,happ(s(fun(X231,bool),X232),s(X231,sK24(X232,X234,X231))))) | s(bool,X235) = s(bool,happ(s(fun(X231,bool),X232),s(X231,X233))) | s(bool,'T_0') = s(bool,X235)) )),
  inference(superposition,[],[f349,f747])).
fof(f4466,plain,(
  ( ! [X239,X237,X240,X238,X236] : (p(s(bool,happ(s(fun(X236,bool),X237),s(X236,sK25(X239,X237,X236))))) | p(s(bool,happ(s(fun(X236,bool),X239),s(X236,sK25(X239,X237,X236))))) | s(bool,X240) = s(bool,happ(s(fun(X236,bool),X237),s(X236,X238))) | s(bool,'T_0') = s(bool,X240)) )),
  inference(subsumption_resolution,[],[f3410,f268])).
fof(f3410,plain,(
  ( ! [X239,X237,X240,X238,X236] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X236,bool),X237),s(X236,sK25(X239,X237,X236))))) | p(s(bool,happ(s(fun(X236,bool),X239),s(X236,sK25(X239,X237,X236))))) | s(bool,X240) = s(bool,happ(s(fun(X236,bool),X237),s(X236,X238))) | s(bool,'T_0') = s(bool,X240)) )),
  inference(superposition,[],[f350,f747])).
fof(f4467,plain,(
  ( ! [X243,X241,X244,X242,X245] : (p(s(bool,happ(s(fun(X241,bool),X244),s(X241,sK25(X242,X244,X241))))) | p(s(bool,happ(s(fun(X241,bool),X242),s(X241,sK25(X242,X244,X241))))) | s(bool,X245) = s(bool,happ(s(fun(X241,bool),X242),s(X241,X243))) | s(bool,'T_0') = s(bool,X245)) )),
  inference(subsumption_resolution,[],[f3411,f268])).
fof(f3411,plain,(
  ( ! [X243,X241,X244,X242,X245] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X241,bool),X244),s(X241,sK25(X242,X244,X241))))) | p(s(bool,happ(s(fun(X241,bool),X242),s(X241,sK25(X242,X244,X241))))) | s(bool,X245) = s(bool,happ(s(fun(X241,bool),X242),s(X241,X243))) | s(bool,'T_0') = s(bool,X245)) )),
  inference(superposition,[],[f351,f747])).
fof(f4472,plain,(
  ( ! [X263,X265,X262,X264] : (p(s(bool,happ(s(fun(X262,bool),X263),s(X262,X264)))) | s(bool,X265) = s(bool,happ(s(fun(X262,bool),X263),s(X262,sK15(X263,X262)))) | s(bool,'T_0') = s(bool,X265)) )),
  inference(subsumption_resolution,[],[f3416,f268])).
fof(f3416,plain,(
  ( ! [X263,X265,X262,X264] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X262,bool),X263),s(X262,X264)))) | s(bool,X265) = s(bool,happ(s(fun(X262,bool),X263),s(X262,sK15(X263,X262)))) | s(bool,'T_0') = s(bool,X265)) )),
  inference(superposition,[],[f328,f747])).
fof(f4473,plain,(
  ( ! [X269,X267,X268,X266] : (p(s(bool,happ(s(fun(X266,bool),X267),s(X266,X268)))) | s(bool,X269) = s(bool,happ(s(fun(X266,bool),X267),s(X266,sK16(X267,X266)))) | s(bool,'T_0') = s(bool,X269)) )),
  inference(subsumption_resolution,[],[f3417,f268])).
fof(f3417,plain,(
  ( ! [X269,X267,X268,X266] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X266,bool),X267),s(X266,X268)))) | s(bool,X269) = s(bool,happ(s(fun(X266,bool),X267),s(X266,sK16(X267,X266)))) | s(bool,'T_0') = s(bool,X269)) )),
  inference(superposition,[],[f329,f747])).
fof(f335,plain,(
  ( ! [X4,X2,X0,X1] : (s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) )),
  inference(cnf_transformation,[],[f203])).
fof(f203,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (s(bool,'const.bool.IN_2'(s(X0,sK17(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK17(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK17])],[f202])).
fof(f202,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f201])).
fof(f201,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f93])).
fof(f93,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X6] : ! [X16,X9] : (s(fun(X6,bool),X9) = s(fun(X6,bool),X16) <=> ! [X12] : s(bool,'const.bool.IN_2'(s(X6,X12),s(fun(X6,bool),X9))) = s(bool,'const.bool.IN_2'(s(X6,X12),s(fun(X6,bool),X16))))),
  file('thm.res_quan.RES_FORALL_NULL','thm.pred_set.EXTENSION')).
fof(f336,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,sK17(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK17(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) )),
  inference(cnf_transformation,[],[f203])).
fof(f1715,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X0,bool),X2) = s(fun(X0,bool),X4) | s(bool,'const.bool.IN_2'(s(X0,sK17(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X4,X0)),s(fun(X0,bool),X2))) != s(bool,'const.bool.IN_2'(s(X0,sK17(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X4,X0)),s(fun(X0,bool),X4)))) )),
  inference(forward_demodulation,[],[f1711,f381])).
fof(f1711,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'const.bool.IN_2'(s(X0,sK17(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X4,X0)),s(fun(X0,bool),X2))) != s(bool,'const.bool.IN_2'(s(X0,sK17(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X4,X0)),s(fun(X0,bool),X4))) | s(fun(X0,bool),X4) = s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)))) )),
  inference(superposition,[],[f336,f381])).
fof(f1714,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X0,bool),X2) = s(fun(X0,bool),X4) | s(bool,'const.bool.IN_2'(s(X0,sK17(X4,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X0)),s(fun(X0,bool),X2))) != s(bool,'const.bool.IN_2'(s(X0,sK17(X4,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X0)),s(fun(X0,bool),X4)))) )),
  inference(forward_demodulation,[],[f1708,f381])).
fof(f1708,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'const.bool.IN_2'(s(X0,sK17(X4,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X0)),s(fun(X0,bool),X2))) != s(bool,'const.bool.IN_2'(s(X0,sK17(X4,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X0)),s(fun(X0,bool),X4))) | s(fun(X0,bool),X4) = s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)))) )),
  inference(superposition,[],[f336,f381])).
fof(f2972,plain,(
  ( ! [X520,X522,X521] : (s(bool,'const.bool.IN_2'(s(bool,X522),s(fun(bool,bool),X520))) != s(bool,'const.bool.IN_2'(s(bool,X522),s(fun(bool,bool),X521))) | s(fun(bool,bool),X520) = s(fun(bool,bool),X521) | s(bool,'T_0') = s(bool,X522) | s(bool,'T_0') = s(bool,sK17(X520,X521,bool))) )),
  inference(superposition,[],[f336,f747])).
fof(f3503,plain,(
  ( ! [X516,X518,X517] : (s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X516))) != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X517))) | s(fun(bool,bool),X516) = s(fun(bool,bool),X517) | s(bool,X518) = s(bool,sK17(X516,X517,bool)) | s(bool,'T_0') = s(bool,X518)) )),
  inference(superposition,[],[f336,f747])).
fof(f13941,plain,(
  ( ! [X68,X66,X64,X69,X67,X65,X63] : (~p(s(bool,X64)) | ~p(s(bool,X66)) | p(s(bool,X67)) | ~sP11(happ(s(fun(X68,bool),'_dst_x0_1'(s(bool,X66))),s(X68,X69)),X67,happ(s(fun(X63,bool),'_dst_x0_1'(s(bool,X64))),s(X63,X65)))) )),
  inference(superposition,[],[f1236,f381])).
fof(f13942,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11(happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2,X0) | s(bool,X0) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f13930,f268])).
fof(f13930,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11(happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2,X0) | s(bool,X0) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1236,f747])).
fof(f13929,plain,(
  ( ! [X0,X1] : (~sP11(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f268,f473,f268,f1236])).
fof(f13850,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | ~sP9(happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X0))),s(X3,X4)),X2,X1) | s(bool,X0) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f13838,f268])).
fof(f13838,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9(happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X0))),s(X3,X4)),X2,X1) | s(bool,X0) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1235,f747])).
fof(f13837,plain,(
  ( ! [X0,X1] : (~sP9(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f268,f473,f473,f1235])).
fof(f13795,plain,(
  ( ! [X68,X66,X64,X69,X67,X65,X63] : (~p(s(bool,X64)) | ~p(s(bool,X66)) | p(s(bool,X67)) | ~sP7(X67,happ(s(fun(X63,bool),'_dst_x0_1'(s(bool,X64))),s(X63,X65)),happ(s(fun(X68,bool),'_dst_x0_1'(s(bool,X66))),s(X68,X69)))) )),
  inference(superposition,[],[f1234,f381])).
fof(f13797,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X0,happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4))) | s(bool,X0) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f13784,f268])).
fof(f13784,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X0,happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4))) | s(bool,X0) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1234,f747])).
fof(f13781,plain,(
  ( ! [X0,X1] : (~sP7('F_0','T_0',happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f268,f473,f268,f1234])).
fof(f13762,plain,(
  ( ! [X68,X66,X64,X69,X67,X65,X63] : (~p(s(bool,X64)) | ~p(s(bool,X66)) | p(s(bool,X67)) | ~sP4(happ(s(fun(X63,bool),'_dst_x0_1'(s(bool,X64))),s(X63,X65)),happ(s(fun(X68,bool),'_dst_x0_1'(s(bool,X66))),s(X68,X69)),X67)) )),
  inference(superposition,[],[f1233,f381])).
fof(f13765,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X0,happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2) | s(bool,X0) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f13751,f268])).
fof(f13751,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X0,happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2) | s(bool,X0) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1233,f747])).
fof(f13746,plain,(
  ( ! [X0,X1] : (~sP4('T_0',happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f268,f473,f268,f1233])).
fof(f13503,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,X1)) | sP1(X1,happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,X0))),s(X2,X3))) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f13492,f268])).
fof(f13492,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,X1)) | sP1(X1,happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,X0))),s(X2,X3))) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f1232,f747])).
fof(f13491,plain,(
  ( ! [X0,X1] : (sP1('F_0',happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f268,f473,f1232])).
fof(f13474,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,X1)) | sP1(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,X0))),s(X2,X3)),X1) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f13463,f268])).
fof(f13463,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,X1)) | sP1(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,X0))),s(X2,X3)),X1) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f1231,f747])).
fof(f13462,plain,(
  ( ! [X0,X1] : (sP1(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f268,f473,f1231])).
fof(f3506,plain,(
  ( ! [X528,X526,X525,X527] : (s(X527,f(s(bool,'T_0'))) != s(X527,g(s(bool,'T_0'))) | s(fun(bool,X527),X525) = s(fun(bool,X527),X526) | s(bool,X528) = s(bool,sK26(X525,X526,X527,bool)) | s(bool,'T_0') = s(bool,X528)) )),
  inference(superposition,[],[f382,f747])).
fof(f3421,plain,(
  ( ! [X282,X285,X283,X284] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X282,sK17(X283,X284,X282)),s(fun(X282,bool),X284))) | s(fun(X282,bool),X283) = s(fun(X282,bool),X284) | s(bool,X285) = s(bool,'const.bool.IN_2'(s(X282,sK17(X283,X284,X282)),s(fun(X282,bool),X283))) | s(bool,'T_0') = s(bool,X285)) )),
  inference(superposition,[],[f336,f747])).
fof(f3420,plain,(
  ( ! [X280,X279,X281,X278] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X278,sK17(X279,X280,X278)),s(fun(X278,bool),X279))) | s(fun(X278,bool),X279) = s(fun(X278,bool),X280) | s(bool,X281) = s(bool,'const.bool.IN_2'(s(X278,sK17(X279,X280,X278)),s(fun(X278,bool),X280))) | s(bool,'T_0') = s(bool,X281)) )),
  inference(superposition,[],[f336,f747])).
fof(f3345,plain,(
  ( ! [X88,X87,X89,X86] : (sP5('T_0',X86,happ(s(fun(X87,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X87,X88))) | s(bool,X86) = s(bool,X89) | s(bool,'T_0') = s(bool,X89)) )),
  inference(superposition,[],[f1209,f747])).
fof(f3344,plain,(
  ( ! [X85,X83,X84,X82] : (~sP5('F_0',X82,happ(s(fun(X83,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X83,X84))) | s(bool,X82) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) )),
  inference(superposition,[],[f1208,f747])).
fof(f3343,plain,(
  ( ! [X80,X78,X81,X79] : (sP5('T_0',happ(s(fun(X79,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X79,X80)),X78) | s(bool,X78) = s(bool,X81) | s(bool,'T_0') = s(bool,X81)) )),
  inference(superposition,[],[f1207,f747])).
fof(f3342,plain,(
  ( ! [X76,X74,X77,X75] : (~sP5('F_0',happ(s(fun(X75,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X75,X76)),X74) | s(bool,X74) = s(bool,X77) | s(bool,'T_0') = s(bool,X77)) )),
  inference(superposition,[],[f1206,f747])).
fof(f4462,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X64)) | p(s(bool,X65)) | ~sP11(X63,X65,X64) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f3338,f268])).
fof(f3338,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'T_0')) | ~p(s(bool,X64)) | p(s(bool,X65)) | ~sP11(X63,X65,X64) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f454,f747])).
fof(f4461,plain,(
  ( ! [X61,X59,X62,X60] : (p(s(bool,X60)) | p(s(bool,X61)) | ~sP9(X59,X61,X60) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f3337,f268])).
fof(f3337,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | p(s(bool,X60)) | p(s(bool,X61)) | ~sP9(X59,X61,X60) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f437,f747])).
fof(f4460,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X56)) | p(s(bool,X57)) | ~sP7(X57,X56,X55) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f3336,f268])).
fof(f3336,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X56)) | p(s(bool,X57)) | ~sP7(X57,X56,X55) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f422,f747])).
fof(f4459,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X52)) | p(s(bool,X53)) | ~sP4(X52,X51,X53) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f3335,f268])).
fof(f3335,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X52)) | p(s(bool,X53)) | ~sP4(X52,X51,X53) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f390,f747])).
fof(f4458,plain,(
  ( ! [X50,X48,X49] : (p(s(bool,X49)) | sP1(X49,X48) | s(bool,X48) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f3334,f268])).
fof(f3334,plain,(
  ( ! [X50,X48,X49] : (~p(s(bool,'T_0')) | p(s(bool,X49)) | sP1(X49,X48) | s(bool,X48) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f376,f747])).
fof(f4457,plain,(
  ( ! [X47,X45,X46] : (p(s(bool,X46)) | sP1(X45,X46) | s(bool,X45) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f3333,f268])).
fof(f3333,plain,(
  ( ! [X47,X45,X46] : (~p(s(bool,'T_0')) | p(s(bool,X46)) | sP1(X45,X46) | s(bool,X45) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f373,f747])).
fof(f4456,plain,(
  ( ! [X43,X44,X42] : (sP0(X42,X43) | s(bool,X42) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(subsumption_resolution,[],[f3332,f268])).
fof(f3332,plain,(
  ( ! [X43,X44,X42] : (~p(s(bool,'T_0')) | sP0(X42,X43) | s(bool,X42) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f358,f747])).
fof(f4455,plain,(
  ( ! [X39,X41,X40] : (sP0(X40,X39) | s(bool,X39) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) )),
  inference(subsumption_resolution,[],[f3331,f268])).
fof(f3331,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'T_0')) | sP0(X40,X39) | s(bool,X39) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f357,f747])).
fof(f4454,plain,(
  ( ! [X37,X38,X36] : (~p(s(bool,X37)) | s(bool,X36) = s(bool,X37) | s(bool,X36) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f3330,f268])).
fof(f3330,plain,(
  ( ! [X37,X38,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X37)) | s(bool,X36) = s(bool,X37) | s(bool,X36) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f355,f747])).
fof(f4449,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f3322,f268])).
fof(f3322,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f290,f747])).
fof(f3250,plain,(
  ( ! [X275,X274] : (s(bool,X274) != s(bool,'const.bool.IN_2'(s(bool,X274),s(fun(bool,bool),'const.pred_set.EMPTY_0'))) | s(bool,X274) = s(bool,X275) | s(bool,'T_0') = s(bool,X275)) )),
  inference(superposition,[],[f591,f747])).
fof(f4448,plain,(
  ( ! [X231,X233,X230,X234,X232] : (sP0(happ(s(fun(X231,bool),'_dst_x0_1'(s(bool,X230))),s(X231,X232)),X233) | s(bool,X230) = s(bool,X234) | s(bool,'T_0') = s(bool,X234)) )),
  inference(subsumption_resolution,[],[f3238,f268])).
fof(f3238,plain,(
  ( ! [X231,X233,X230,X234,X232] : (sP0(happ(s(fun(X231,bool),'_dst_x0_1'(s(bool,X230))),s(X231,X232)),X233) | ~p(s(bool,'T_0')) | s(bool,X230) = s(bool,X234) | s(bool,'T_0') = s(bool,X234)) )),
  inference(superposition,[],[f1230,f747])).
fof(f4447,plain,(
  ( ! [X229,X227,X225,X228,X226] : (sP0(X226,happ(s(fun(X227,bool),'_dst_x0_1'(s(bool,X225))),s(X227,X228))) | s(bool,X225) = s(bool,X229) | s(bool,'T_0') = s(bool,X229)) )),
  inference(subsumption_resolution,[],[f3237,f268])).
fof(f3237,plain,(
  ( ! [X229,X227,X225,X228,X226] : (sP0(X226,happ(s(fun(X227,bool),'_dst_x0_1'(s(bool,X225))),s(X227,X228))) | ~p(s(bool,'T_0')) | s(bool,X225) = s(bool,X229) | s(bool,'T_0') = s(bool,X229)) )),
  inference(superposition,[],[f1229,f747])).
fof(f3236,plain,(
  ( ! [X222,X223,X221,X224] : (sP5('T_0','T_0',happ(s(fun(X222,bool),'_dst_x0_1'(s(bool,X221))),s(X222,X223))) | s(bool,X221) = s(bool,X224) | s(bool,'T_0') = s(bool,X224)) )),
  inference(superposition,[],[f1209,f747])).
fof(f3235,plain,(
  ( ! [X218,X220,X219,X217] : (~sP5('F_0','T_0',happ(s(fun(X218,bool),'_dst_x0_1'(s(bool,X217))),s(X218,X219))) | s(bool,X217) = s(bool,X220) | s(bool,'T_0') = s(bool,X220)) )),
  inference(superposition,[],[f1208,f747])).
fof(f3234,plain,(
  ( ! [X216,X215,X213,X214] : (sP5('T_0',happ(s(fun(X214,bool),'_dst_x0_1'(s(bool,X213))),s(X214,X215)),'T_0') | s(bool,X213) = s(bool,X216) | s(bool,'T_0') = s(bool,X216)) )),
  inference(superposition,[],[f1207,f747])).
fof(f3233,plain,(
  ( ! [X212,X210,X211,X209] : (~sP5('F_0',happ(s(fun(X210,bool),'_dst_x0_1'(s(bool,X209))),s(X210,X211)),'T_0') | s(bool,X209) = s(bool,X212) | s(bool,'T_0') = s(bool,X212)) )),
  inference(superposition,[],[f1206,f747])).
fof(f3166,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f286,f747])).
fof(f2975,plain,(
  ( ! [X532,X530,X529,X531] : (s(X531,f(s(bool,X532))) != s(X531,g(s(bool,X532))) | s(fun(bool,X531),X529) = s(fun(bool,X531),X530) | s(bool,'T_0') = s(bool,X532) | s(bool,'T_0') = s(bool,sK26(X529,X530,X531,bool))) )),
  inference(superposition,[],[f382,f747])).
fof(f2889,plain,(
  ( ! [X285,X287,X284,X286] : (s(bool,X287) != s(bool,'const.bool.IN_2'(s(X284,sK17(X285,X286,X284)),s(fun(X284,bool),X286))) | s(fun(X284,bool),X285) = s(fun(X284,bool),X286) | s(bool,'T_0') = s(bool,X287) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X284,sK17(X285,X286,X284)),s(fun(X284,bool),X285)))) )),
  inference(superposition,[],[f336,f747])).
fof(f2888,plain,(
  ( ! [X280,X282,X281,X283] : (s(bool,X283) != s(bool,'const.bool.IN_2'(s(X280,sK17(X281,X282,X280)),s(fun(X280,bool),X281))) | s(fun(X280,bool),X281) = s(fun(X280,bool),X282) | s(bool,'T_0') = s(bool,X283) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X280,sK17(X281,X282,X280)),s(fun(X280,bool),X282)))) )),
  inference(superposition,[],[f336,f747])).
fof(f2813,plain,(
  ( ! [X88,X87,X89,X86] : (sP5('T_0',X86,happ(s(fun(X88,bool),'_dst_x0_1'(s(bool,X87))),s(X88,X89))) | s(bool,'T_0') = s(bool,X87) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f1209,f747])).
fof(f2812,plain,(
  ( ! [X85,X83,X84,X82] : (~sP5('F_0',X82,happ(s(fun(X84,bool),'_dst_x0_1'(s(bool,X83))),s(X84,X85))) | s(bool,'T_0') = s(bool,X83) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f1208,f747])).
fof(f2811,plain,(
  ( ! [X80,X78,X81,X79] : (sP5('T_0',happ(s(fun(X80,bool),'_dst_x0_1'(s(bool,X79))),s(X80,X81)),X78) | s(bool,'T_0') = s(bool,X79) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f1207,f747])).
fof(f2810,plain,(
  ( ! [X76,X74,X77,X75] : (~sP5('F_0',happ(s(fun(X76,bool),'_dst_x0_1'(s(bool,X75))),s(X76,X77)),X74) | s(bool,'T_0') = s(bool,X75) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f1206,f747])).
fof(f2771,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'T_0') = s(bool,X23) | s(bool,X25) = s(bool,happ(s(fun(X22,bool),'_dst_x0_1'(s(bool,X23))),s(X22,X24))) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f747,f381])).
fof(f2757,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f747,f747])).
fof(f2742,plain,(
  ( ! [X30,X31,X29,X32] : (s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32) | s(bool,'T_0') = s(bool,happ(s(fun(X29,bool),'_dst_x0_1'(s(bool,X30))),s(X29,X31)))) )),
  inference(superposition,[],[f747,f381])).
fof(f2290,plain,(
  ( ! [X2,X0,X1] : (sP0(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)),X2)) )),
  inference(unit_resulting_resolution,[],[f268,f1230])).
fof(f2252,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X1,X2)))) )),
  inference(unit_resulting_resolution,[],[f268,f1229])).
fof(f2228,plain,(
  ( ! [X6,X7,X5] : (s(bool,X7) = s(bool,'const.bool.IN_2'(s(X5,X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f2223,f473])).
fof(f2223,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,'F_0')) | s(bool,X7) = s(bool,'const.bool.IN_2'(s(X5,X6),s(fun(X5,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f810,f592])).
fof(f2204,plain,(
  ( ! [X0,X1] : (~sP2('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f473,f601,f599,f1916,f388])).
fof(f2174,plain,(
  ( ! [X4,X5,X3] : (sP5('T_0','F_0',happ(s(fun(X4,bool),'_dst_x0_1'(s(bool,X3))),s(X4,X5))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1209,f275])).
fof(f2170,plain,(
  ( ! [X2,X0,X1] : (sP4('T_0',happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1))),s(X0,X2)),X1)) )),
  inference(unit_resulting_resolution,[],[f1209,f385])).
fof(f2130,plain,(
  ( ! [X4,X5,X3] : (~sP5('F_0','F_0',happ(s(fun(X4,bool),'_dst_x0_1'(s(bool,X3))),s(X4,X5))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1208,f275])).
fof(f2125,plain,(
  ( ! [X0,X1] : (~sP3('T_0',happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f473,f490,f480,f1208,f389])).
fof(f2092,plain,(
  ( ! [X14,X17,X15,X18,X16] : (sP5('T_0',happ(s(fun(X17,bool),'_dst_x0_1'(s(bool,X15))),s(X17,X18)),happ(s(fun(X14,bool),'_dst_x0_1'(s(bool,X15))),s(X14,X16)))) )),
  inference(superposition,[],[f1207,f381])).
fof(f2088,plain,(
  ( ! [X4,X5,X3] : (sP5('T_0',happ(s(fun(X4,bool),'_dst_x0_1'(s(bool,X3))),s(X4,X5)),'F_0') | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1207,f275])).
fof(f2084,plain,(
  ( ! [X2,X0,X1] : (sP4('T_0',X0,happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,X0))),s(X1,X2)))) )),
  inference(unit_resulting_resolution,[],[f1207,f385])).
fof(f2048,plain,(
  ( ! [X14,X17,X15,X18,X16] : (~sP5('F_0',happ(s(fun(X17,bool),'_dst_x0_1'(s(bool,X15))),s(X17,X18)),happ(s(fun(X14,bool),'_dst_x0_1'(s(bool,X15))),s(X14,X16)))) )),
  inference(superposition,[],[f1206,f381])).
fof(f2044,plain,(
  ( ! [X4,X5,X3] : (~sP5('F_0',happ(s(fun(X4,bool),'_dst_x0_1'(s(bool,X3))),s(X4,X5)),'F_0') | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1206,f275])).
fof(f2040,plain,(
  ( ! [X0,X1] : (~sP3(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f268,f492,f482,f1206,f387])).
fof(f2010,plain,(
  ( ! [X0,X1] : (~sP2('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f473,f485,f481,f1914,f389])).
fof(f1921,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f592,f275])).
fof(f1914,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f473,f592,f405])).
fof(f1915,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f268,f592,f403])).
fof(f1916,plain,(
  ( ! [X0,X1] : (~sP5('F_0','F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f473,f592,f405])).
fof(f1874,plain,(
  ( ! [X0,X1] : (~sP4('T_0','T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f268,f488,f484,f1852,f387])).
fof(f1852,plain,(
  ( ! [X0,X1] : (~sP5('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f268,f591,f404])).
fof(f1856,plain,(
  ( ! [X0,X1] : (~sP5('T_0','T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f268,f591,f404])).
fof(f1717,plain,(
  ( ! [X0,X1] : (sP11('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f268,f473,f607,f446])).
fof(f1700,plain,(
  ( ! [X2,X0,X1] : (sP11('T_0',X0,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f268,f606,f457])).
fof(f1691,plain,(
  ( ! [X0,X1] : (sP9('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f473,f473,f605,f429])).
fof(f1453,plain,(
  ( ! [X0,X1] : (sP5('F_0','T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f601,f599,f490,f473,f389])).
fof(f1381,plain,(
  sP5('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f486,f481,f488,f473,f388])).
fof(f1385,plain,(
  ( ! [X0,X1] : (sP5('T_0','F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f484,f599,f488,f473,f388])).
fof(f1311,plain,(
  ( ! [X0,X1] : (sP5('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f602,f482,f492,f268,f387])).
fof(f1315,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f486,f600,f492,f268,f387])).
fof(f1317,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.EMPTY_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f602,f600,f492,f268,f387])).
fof(f1332,plain,(
  sP5('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f484,f480,f490,f268,f387])).
fof(f1236,plain,(
  ( ! [X61,X59,X62,X60,X58] : (~p(s(bool,X61)) | ~p(s(bool,X59)) | p(s(bool,X62)) | ~sP11(happ(s(fun(X58,bool),'_dst_x0_1'(s(bool,X59))),s(X58,X60)),X62,X61)) )),
  inference(superposition,[],[f454,f381])).
fof(f1235,plain,(
  ( ! [X57,X54,X56,X55,X53] : (~p(s(bool,X54)) | p(s(bool,X56)) | p(s(bool,X57)) | ~sP9(happ(s(fun(X53,bool),'_dst_x0_1'(s(bool,X54))),s(X53,X55)),X57,X56)) )),
  inference(superposition,[],[f437,f381])).
fof(f1234,plain,(
  ( ! [X52,X50,X48,X51,X49] : (~p(s(bool,X51)) | ~p(s(bool,X49)) | p(s(bool,X52)) | ~sP7(X52,X51,happ(s(fun(X48,bool),'_dst_x0_1'(s(bool,X49))),s(X48,X50)))) )),
  inference(superposition,[],[f422,f381])).
fof(f1233,plain,(
  ( ! [X47,X45,X43,X46,X44] : (~p(s(bool,X46)) | ~p(s(bool,X44)) | p(s(bool,X47)) | ~sP4(X46,happ(s(fun(X43,bool),'_dst_x0_1'(s(bool,X44))),s(X43,X45)),X47)) )),
  inference(superposition,[],[f390,f381])).
fof(f1232,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X40)) | p(s(bool,X42)) | sP1(X42,happ(s(fun(X39,bool),'_dst_x0_1'(s(bool,X40))),s(X39,X41)))) )),
  inference(superposition,[],[f376,f381])).
fof(f1231,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X36)) | p(s(bool,X38)) | sP1(happ(s(fun(X35,bool),'_dst_x0_1'(s(bool,X36))),s(X35,X37)),X38)) )),
  inference(superposition,[],[f373,f381])).
fof(f1230,plain,(
  ( ! [X33,X31,X34,X32] : (sP0(happ(s(fun(X31,bool),'_dst_x0_1'(s(bool,X32))),s(X31,X33)),X34) | ~p(s(bool,X32))) )),
  inference(superposition,[],[f358,f381])).
fof(f1229,plain,(
  ( ! [X30,X28,X29,X27] : (sP0(X30,happ(s(fun(X27,bool),'_dst_x0_1'(s(bool,X28))),s(X27,X29))) | ~p(s(bool,X28))) )),
  inference(superposition,[],[f357,f381])).
fof(f1206,plain,(
  ( ! [X2,X0,X1] : (~sP5('F_0',happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1))),s(X0,X2)),X1)) )),
  inference(unit_resulting_resolution,[],[f473,f381,f405])).
fof(f1207,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0',happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1))),s(X0,X2)),X1)) )),
  inference(unit_resulting_resolution,[],[f268,f381,f403])).
fof(f1208,plain,(
  ( ! [X2,X0,X1] : (~sP5('F_0',X0,happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,X0))),s(X1,X2)))) )),
  inference(unit_resulting_resolution,[],[f473,f381,f405])).
fof(f1209,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0',X0,happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,X0))),s(X1,X2)))) )),
  inference(unit_resulting_resolution,[],[f268,f381,f403])).
fof(f1178,plain,(
  ( ! [X2,X0,X1] : (sP7('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),X2)) )),
  inference(unit_resulting_resolution,[],[f473,f604,f459])).
fof(f1170,plain,(
  ( ! [X2,X0,X1] : (sP7('F_0',X0,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f473,f603,f459])).
fof(f1158,plain,(
  ~sP11('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1149,f442])).
fof(f1149,plain,(
  ~sP10('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f268,f268,f473,f447])).
fof(f1139,plain,(
  sP11('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f505,f268,f473,f446])).
fof(f1079,plain,(
  ~sP9('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1070,f425])).
fof(f1070,plain,(
  ~sP8('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f268,f473,f473,f430])).
fof(f1059,plain,(
  sP9('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f501,f473,f473,f429])).
fof(f996,plain,(
  ~sP7('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f987,f406])).
fof(f987,plain,(
  ~sP6('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f268,f473,f268,f413])).
fof(f980,plain,(
  sP7('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f494,f268,f268,f409])).
fof(f976,plain,(
  ~sP5('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f963,f383])).
fof(f972,plain,(
  ~sP5('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f946,f384])).
fof(f963,plain,(
  ~sP2('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f473,f473,f473,f398])).
fof(f946,plain,(
  ~sP3('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f268,f473,f268,f394])).
fof(f918,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f473,f268,f268,f390])).
fof(f875,plain,(
  ~sP5('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f268,f515,f404])).
fof(f878,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f268,f515,f404])).
fof(f852,plain,(
  sP5('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f473,f515,f402])).
fof(f854,plain,(
  sP5('F_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f473,f515,f402])).
fof(f810,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f802,f473])).
fof(f802,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f352,f275])).
fof(f747,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f275,f275])).
fof(f696,plain,(
  ( ! [X0] : (sP11('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f268,f268,f453])).
fof(f686,plain,(
  ( ! [X0] : (sP11('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f268,f473,f452])).
fof(f679,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f268,f473,f441])).
fof(f674,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f473,f473,f440])).
fof(f670,plain,(
  ( ! [X0] : (sP9('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f268,f268,f436])).
fof(f664,plain,(
  ( ! [X0] : (sP9('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f268,f268,f435])).
fof(f644,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f268,f473,f424])).
fof(f641,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f268,f473,f423])).
fof(f637,plain,(
  ( ! [X0] : (sP7('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f473,f473,f418])).
fof(f631,plain,(
  ( ! [X0] : (sP7('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f473,f473,f417])).
fof(f628,plain,(
  ( ! [X0] : (~sP7('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f268,f473,f408])).
fof(f626,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f268,f473,f407])).
fof(f591,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f332,f286])).
fof(f593,plain,(
  ( ! [X0,X1] : (~sP1('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f473,f332,f371])).
fof(f595,plain,(
  ( ! [X0,X1] : (~sP1('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f473,f332,f371])).
fof(f596,plain,(
  ( ! [X2,X0,X3,X1] : (~sP1('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f332,f332,f371])).
fof(f597,plain,(
  ( ! [X0,X1] : (sP1('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f268,f332,f373])).
fof(f598,plain,(
  ( ! [X0,X1] : (sP1('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f268,f332,f376])).
fof(f599,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')),X3)) )),
  inference(unit_resulting_resolution,[],[f332,f392])).
fof(f600,plain,(
  ( ! [X2,X0,X3,X1] : (sP4('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),X2,X3)) )),
  inference(unit_resulting_resolution,[],[f332,f393])).
fof(f601,plain,(
  ( ! [X2,X0,X3,X1] : (sP3(X0,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')),X3)) )),
  inference(unit_resulting_resolution,[],[f332,f396])).
fof(f602,plain,(
  ( ! [X2,X0,X3,X1] : (sP3('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),X2,X3)) )),
  inference(unit_resulting_resolution,[],[f332,f397])).
fof(f603,plain,(
  ( ! [X2,X0,X3,X1] : (sP6(X0,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')),X3)) )),
  inference(unit_resulting_resolution,[],[f332,f415])).
fof(f604,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),X2,X3)) )),
  inference(unit_resulting_resolution,[],[f332,f416])).
fof(f605,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),X2,X3)) )),
  inference(unit_resulting_resolution,[],[f332,f433])).
fof(f606,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,X1,'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f332,f448])).
fof(f607,plain,(
  ( ! [X2,X0,X3,X1] : (sP10('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),X2,X3)) )),
  inference(unit_resulting_resolution,[],[f332,f450])).
fof(f608,plain,(
  ( ! [X0,X1] : (~sP0('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f473,f332,f474])).
fof(f612,plain,(
  ( ! [X0,X1] : (~sP0('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f473,f332,f474])).
fof(f613,plain,(
  ( ! [X2,X0,X3,X1] : (~sP0('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f332,f332,f474])).
fof(f584,plain,(
  ~sP0('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f473,f473,f474])).
fof(f573,plain,(
  sP1('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f473,f268,f376])).
fof(f564,plain,(
  sP1('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f473,f268,f373])).
fof(f561,plain,(
  ~sP1('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f268,f268,f372])).
fof(f560,plain,(
  ~sP1('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f473,f473,f371])).
fof(f515,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f473,f286])).
fof(f505,plain,(
  ( ! [X0,X1] : (sP10('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f473,f450])).
fof(f504,plain,(
  ( ! [X0,X1] : (sP10(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f268,f449])).
fof(f502,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f473,f448])).
fof(f501,plain,(
  ( ! [X0,X1] : (sP8('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f473,f433])).
fof(f500,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f268,f432])).
fof(f498,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f268,f431])).
fof(f496,plain,(
  ( ! [X0,X1] : (sP6('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f473,f416])).
fof(f495,plain,(
  ( ! [X0,X1] : (sP6(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f473,f415])).
fof(f494,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f268,f414])).
fof(f492,plain,(
  ( ! [X0,X1] : (sP2('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f268,f401])).
fof(f490,plain,(
  ( ! [X0,X1] : (sP2(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f268,f400])).
fof(f488,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f268,f399])).
fof(f486,plain,(
  ( ! [X0,X1] : (sP3('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f473,f397])).
fof(f485,plain,(
  ( ! [X0,X1] : (sP3(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f473,f396])).
fof(f484,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f268,f395])).
fof(f482,plain,(
  ( ! [X0,X1] : (sP4('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f473,f393])).
fof(f481,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f473,f392])).
fof(f480,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f268,f391])).
fof(f478,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f268,f358])).
fof(f476,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f268,f357])).
fof(f451,plain,(
  ( ! [X2,X0,X1] : (sP11(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f264])).
fof(f264,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f263])).
fof(f263,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP11(X2,X0,X1))),
  inference(definition_folding,[],[f145,f162,f161])).
fof(f161,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f162,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f145,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X19,X20,X7] : ((p(s(bool,X7)) <=> (p(s(bool,X20)) => p(s(bool,X19)))) <=> ((p(s(bool,X7)) | p(s(bool,X20))) & (p(s(bool,X7)) | ~p(s(bool,X19))) & (~p(s(bool,X20)) | p(s(bool,X19)) | ~p(s(bool,X7)))))),
  file('thm.res_quan.RES_FORALL_NULL','thm.sat.sat_27_prove_ltb')).
fof(f452,plain,(
  ( ! [X2,X0,X1] : (sP11(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f264])).
fof(f453,plain,(
  ( ! [X2,X0,X1] : (sP11(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f264])).
fof(f454,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f264])).
fof(f447,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f262])).
fof(f262,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f261])).
fof(f261,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f260])).
fof(f260,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f161])).
fof(f448,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f262])).
fof(f449,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f262])).
fof(f450,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f262])).
fof(f440,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f259])).
fof(f259,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP10(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f258])).
fof(f258,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f257])).
fof(f257,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1)) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f162])).
fof(f441,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f259])).
fof(f442,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f259])).
fof(f457,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f443])).
fof(f443,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f259])).
fof(f444,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f259])).
fof(f445,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f259])).
fof(f446,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f259])).
fof(f434,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f256])).
fof(f256,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f255])).
fof(f255,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f119,f159,f158])).
fof(f158,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f159,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f119,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X19,X20,X7] : ((p(s(bool,X7)) <=> (p(s(bool,X20)) | p(s(bool,X19)))) <=> ((p(s(bool,X7)) | ~p(s(bool,X20))) & (p(s(bool,X7)) | ~p(s(bool,X19))) & (p(s(bool,X20)) | p(s(bool,X19)) | ~p(s(bool,X7)))))),
  file('thm.res_quan.RES_FORALL_NULL','thm.sat.sat_25_prove_ltb')).
fof(f435,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f256])).
fof(f436,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f256])).
fof(f437,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | ~sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f256])).
fof(f430,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f254])).
fof(f254,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f253])).
fof(f253,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f252])).
fof(f252,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f158])).
fof(f431,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f254])).
fof(f432,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f254])).
fof(f433,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f254])).
fof(f423,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f251])).
fof(f251,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP8(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f250])).
fof(f250,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f249])).
fof(f249,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1)) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f159])).
fof(f424,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f251])).
fof(f425,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f251])).
fof(f458,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f426])).
fof(f426,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f251])).
fof(f427,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f251])).
fof(f428,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f251])).
fof(f429,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f251])).
fof(f417,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f248])).
fof(f248,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f247])).
fof(f247,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f117,f156,f155])).
fof(f155,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f156,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f117,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f116])).
fof(f116,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X19,X20,X7] : ((p(s(bool,X7)) <=> (p(s(bool,X20)) & p(s(bool,X19)))) <=> ((p(s(bool,X7)) | ~p(s(bool,X20)) | ~p(s(bool,X19))) & (p(s(bool,X20)) | ~p(s(bool,X7))) & (p(s(bool,X19)) | ~p(s(bool,X7)))))),
  file('thm.res_quan.RES_FORALL_NULL','thm.sat.sat_23_prove_ltb')).
fof(f418,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f248])).
fof(f419,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f248])).
fof(f422,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f248])).
fof(f413,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f246])).
fof(f246,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f245])).
fof(f245,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f155])).
fof(f414,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f246])).
fof(f415,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f246])).
fof(f416,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f246])).
fof(f406,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP6(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f243])).
fof(f243,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f242])).
fof(f242,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f156])).
fof(f407,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f244])).
fof(f408,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f244])).
fof(f409,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f244])).
fof(f459,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f412])).
fof(f412,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f244])).
fof(f402,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP5(X2,X1,X0)) & (~sP5(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP5(X2,X1,X0))),
  inference(definition_folding,[],[f115,f153,f152,f151,f150])).
fof(f150,plain,(
  ! [X0,X1,X2] : (sP2(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f151,plain,(
  ! [X1,X0,X2] : (sP3(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f152,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f153,plain,(
  ! [X2,X1,X0] : (sP5(X2,X1,X0) <=> (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f115,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f114])).
fof(f114,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X19,X20,X7] : ((p(s(bool,X7)) <=> s(bool,X19) = s(bool,X20)) <=> ((p(s(bool,X7)) | p(s(bool,X20)) | p(s(bool,X19))) & (p(s(bool,X7)) | ~p(s(bool,X19)) | ~p(s(bool,X20))) & (p(s(bool,X20)) | ~p(s(bool,X19)) | ~p(s(bool,X7))) & (p(s(bool,X19)) | ~p(s(bool,X20)) | ~p(s(bool,X7)))))),
  file('thm.res_quan.RES_FORALL_NULL','thm.sat.sat_2|49|_prove_ltb')).
fof(f403,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f241])).
fof(f404,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f241])).
fof(f405,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f241])).
fof(f398,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(flattening,[],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(nnf_transformation,[],[f150])).
fof(f399,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f240])).
fof(f400,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f240])).
fof(f401,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f240])).
fof(f394,plain,(
  ( ! [X2,X0,X1] : (~sP3(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f238])).
fof(f238,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f237])).
fof(f237,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(flattening,[],[f236])).
fof(f236,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(nnf_transformation,[],[f151])).
fof(f395,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f238])).
fof(f396,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f238])).
fof(f397,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f238])).
fof(f390,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f235])).
fof(f235,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f234])).
fof(f234,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f233])).
fof(f233,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f152])).
fof(f391,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f235])).
fof(f392,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f235])).
fof(f393,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f235])).
fof(f383,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f232])).
fof(f232,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP2(X2,X1,X0) & sP3(X1,X2,X0) & sP4(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP2(X2,X1,X0) | ~sP3(X1,X2,X0) | ~sP4(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f231])).
fof(f231,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP5(X2,X1,X0)))),
  inference(flattening,[],[f230])).
fof(f230,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP5(X2,X1,X0)))),
  inference(nnf_transformation,[],[f153])).
fof(f384,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP3(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f232])).
fof(f385,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP4(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f232])).
fof(f386,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f232])).
fof(f387,plain,(
  ( ! [X2,X0,X1] : (~sP2(X2,X1,X0) | ~p(s(bool,X2)) | ~sP4(X0,X2,X1) | ~sP3(X1,X2,X0) | sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f232])).
fof(f388,plain,(
  ( ! [X2,X0,X1] : (~sP2(X2,X1,X0) | p(s(bool,X1)) | ~sP4(X0,X2,X1) | ~sP3(X1,X2,X0) | sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f232])).
fof(f389,plain,(
  ( ! [X2,X0,X1] : (~sP2(X2,X1,X0) | p(s(bool,X0)) | ~sP4(X0,X2,X1) | ~sP3(X1,X2,X0) | sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f232])).
fof(f382,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK26(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK26(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f229])).
fof(f229,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK26(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK26(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK26])],[f144])).
fof(f144,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f112])).
fof(f112,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.res_quan.RES_FORALL_NULL',aHL_EXT)).
fof(f371,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f227])).
fof(f373,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f227])).
fof(f376,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f227])).
fof(f474,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f356,f473])).
fof(f356,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f220])).
fof(f220,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f219])).
fof(f219,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f357,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f220])).
fof(f358,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f220])).
fof(f352,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f138])).
fof(f138,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X10,X11] : ((p(s(bool,X10)) => p(s(bool,X11))) => ((p(s(bool,X11)) => p(s(bool,X10))) => s(bool,X10) = s(bool,X11)))),
  file('thm.res_quan.RES_FORALL_NULL','thm.bool.IMP_ANTISYM_AX')).
fof(f355,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f139])).
fof(f332,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0,X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(flattening,[],[f89])).
fof(f89,plain,(
  ! [X0] : ! [X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X6] : ! [X12] : ~p(s(bool,'const.bool.IN_2'(s(X6,X12),s(fun(X6,bool),'const.pred_set.EMPTY_0'))))),
  file('thm.res_quan.RES_FORALL_NULL','thm.pred_set.NOT_IN_EMPTY')).
fof(f286,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X9] : (s(bool,'T_0') = s(bool,X9) <=> p(s(bool,X9)))),
  file('thm.res_quan.RES_FORALL_NULL','thm.bool.EQ_CLAUSES|split|0')).
fof(f287,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f169])).
fof(f275,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.res_quan.RES_FORALL_NULL',aHL_BOOL_CASES)).
fof(f15456,plain,(
  $false | ($spl49 | $spl52)),
  inference(global_subsumption,[],[f267,f266,f265,f268,f275,f287,f286,f332,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1691,f1700,f1717,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2010,f2040,f2044,f2048,f2084,f2088,f2092,f2125,f2130,f2170,f2174,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f3420,f3421,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f13301,f13298,f13462,f13474,f13476,f13479,f13491,f13503,f13505,f13508,f13746,f13765,f13767,f13770,f13762,f13781,f13797,f13799,f13802,f13795,f13837,f13850,f13852,f13855,f13929,f13942,f13944,f13947,f13941,f13995,f13997,f14141,f14142,f14143,f14144,f14145,f3503,f2972,f1714,f1715,f336,f335,f4473,f4472,f4467,f4466,f4465,f4464,f4453,f4452,f4451,f4450,f1827,f1800,f1771,f1741,f1610,f1593,f1578,f1562,f1496,f1480,f351,f350,f349,f348,f346,f344,f343,f341,f329,f328,f3502,f3501,f2971,f2970,f13441,f1538,f13453,f13998,f13994,f13988,f4474,f4463,f1539,f1515,f1514,f339,f338,f337,f2002,f3505,f3504,f2974,f2973,f2968,f3499,f14178,f14176,f14170,f14119,f14116,f14059,f14056,f13948,f13803,f13771,f13699,f13698,f13697,f13696,f13695,f13694,f13693,f13692,f13691,f13689,f13688,f13687,f13686,f13685,f13683,f13584,f13583,f13582,f13581,f13540,f13534,f2176,f2173,f2142,f2132,f2129,f2126,f2090,f2087,f2055,f2046,f2043,f2037,f1929,f1928,f1712,f1710,f1709,f1623,f1622,f1536,f1511,f1493,f1477,f1220,f592,f473,f291,f290,f15203,f15202,f15201,f15197,f15196,f15223,f15375,f15378,f15380,f15240,f15242,f15243,f15245,f15246,f15248,f15249,f15251,f15392,f15453,f15455])).
fof(f15454,plain,(
  $false | ($spl49 | $spl52)),
  inference(global_subsumption,[],[f267,f266,f265,f268,f275,f287,f286,f332,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1691,f1700,f1717,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2010,f2040,f2044,f2048,f2084,f2088,f2092,f2125,f2130,f2170,f2174,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f3420,f3421,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f13301,f13298,f13462,f13474,f13476,f13479,f13491,f13503,f13505,f13508,f13746,f13765,f13767,f13770,f13762,f13781,f13797,f13799,f13802,f13795,f13837,f13850,f13852,f13855,f13929,f13942,f13944,f13947,f13941,f13995,f13997,f14141,f14142,f14143,f14144,f14145,f3503,f2972,f1714,f1715,f336,f335,f4473,f4472,f4467,f4466,f4465,f4464,f4453,f4452,f4451,f4450,f1827,f1800,f1771,f1741,f1610,f1593,f1578,f1562,f1496,f1480,f351,f350,f349,f348,f346,f344,f343,f341,f329,f328,f3502,f3501,f2971,f2970,f13441,f1538,f13453,f13998,f13994,f13988,f4474,f4463,f1539,f1515,f1514,f339,f338,f337,f2002,f3505,f3504,f2974,f2973,f2968,f3499,f14178,f14176,f14170,f14119,f14116,f14059,f14056,f13948,f13803,f13771,f13699,f13698,f13697,f13696,f13695,f13694,f13693,f13692,f13691,f13689,f13688,f13687,f13686,f13685,f13683,f13584,f13583,f13582,f13581,f13540,f13534,f2176,f2173,f2142,f2132,f2129,f2126,f2090,f2087,f2055,f2046,f2043,f2037,f1929,f1928,f1712,f1710,f1709,f1623,f1622,f1536,f1511,f1493,f1477,f1220,f592,f473,f291,f290,f15203,f15202,f15201,f15197,f15196,f15223,f15375,f15378,f15380,f15240,f15242,f15243,f15245,f15246,f15248,f15249,f15251,f15392,f15453])).
fof(f15400,plain,(
  $false | ($spl32 | $spl49 | $spl52)),
  inference(subsumption_resolution,[],[f15298,f268])).
fof(f15298,plain,(
  ~p(s(bool,'T_0')) | ($spl32 | $spl49 | $spl52)),
  inference(backward_demodulation,[],[f15223,f8928])).
fof(f15399,plain,(
  $false | ($spl5 | $spl49 | $spl52)),
  inference(subsumption_resolution,[],[f15295,f15223])).
fof(f15295,plain,(
  s(bool,'T_0') != s(bool,g(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12)))) | ($spl5 | $spl49 | $spl52)),
  inference(backward_demodulation,[],[f15223,f9875])).
fof(f9875,plain,(
  s(bool,f(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12)))) != s(bool,g(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12)))) | $spl5),
  inference(unit_resulting_resolution,[],[f472,f382])).
fof(f15398,plain,(
  $false | ($spl49 | $spl52)),
  inference(subsumption_resolution,[],[f15294,f268])).
fof(f15294,plain,(
  ~p(s(bool,'T_0')) | ($spl49 | $spl52)),
  inference(backward_demodulation,[],[f15223,f15203])).
fof(f15397,plain,(
  $false | ($spl49 | $spl52 | $spl81)),
  inference(subsumption_resolution,[],[f15289,f15223])).
fof(f15289,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl49 | $spl52 | $spl81)),
  inference(backward_demodulation,[],[f15223,f14154])).
fof(f15396,plain,(
  $false | ($spl49 | $spl52 | $spl75)),
  inference(subsumption_resolution,[],[f15287,f15223])).
fof(f15287,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl49 | $spl52 | $spl75)),
  inference(backward_demodulation,[],[f15223,f14149])).
fof(f15393,plain,(
  $false | ($spl49 | $spl52)),
  inference(global_subsumption,[],[f267,f266,f265,f268,f275,f287,f286,f332,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1691,f1700,f1717,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2010,f2040,f2044,f2048,f2084,f2088,f2092,f2125,f2130,f2170,f2174,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f3420,f3421,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f13301,f13298,f13462,f13474,f13476,f13479,f13491,f13503,f13505,f13508,f13746,f13765,f13767,f13770,f13762,f13781,f13797,f13799,f13802,f13795,f13837,f13850,f13852,f13855,f13929,f13942,f13944,f13947,f13941,f13995,f13997,f14141,f14142,f14143,f14144,f14145,f3503,f2972,f1714,f1715,f336,f335,f4473,f4472,f4467,f4466,f4465,f4464,f4453,f4452,f4451,f4450,f1827,f1800,f1771,f1741,f1610,f1593,f1578,f1562,f1496,f1480,f351,f350,f349,f348,f346,f344,f343,f341,f329,f328,f3502,f3501,f2971,f2970,f13441,f1538,f13453,f13998,f13994,f13988,f4474,f4463,f1539,f1515,f1514,f339,f338,f337,f2002,f3505,f3504,f2974,f2973,f2968,f3499,f14178,f14176,f14170,f14119,f14116,f14059,f14056,f13948,f13803,f13771,f13699,f13698,f13697,f13696,f13695,f13694,f13693,f13692,f13691,f13689,f13688,f13687,f13686,f13685,f13683,f13584,f13583,f13582,f13581,f13540,f13534,f2176,f2173,f2142,f2132,f2129,f2126,f2090,f2087,f2055,f2046,f2043,f2037,f1929,f1928,f1712,f1710,f1709,f1623,f1622,f1536,f1511,f1493,f1477,f1220,f592,f473,f291,f290,f15203,f15202,f15201,f15197,f15196,f15223,f15375,f15378,f15380,f15240,f15242,f15243,f15245,f15246,f15248,f15249,f15251,f15392])).
fof(f15391,plain,(
  $false | ($spl2 | $spl49 | $spl52)),
  inference(global_subsumption,[],[f15256,f267,f266,f265,f268,f275,f287,f286,f332,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1691,f1700,f1717,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2010,f2040,f2044,f2048,f2084,f2088,f2092,f2125,f2130,f2170,f2174,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f3420,f3421,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f13301,f13298,f13462,f13474,f13476,f13479,f13491,f13503,f13505,f13508,f13746,f13765,f13767,f13770,f13762,f13781,f13797,f13799,f13802,f13795,f13837,f13850,f13852,f13855,f13929,f13942,f13944,f13947,f13941,f13995,f13997,f14141,f14142,f14143,f14144,f14145,f3503,f2972,f1714,f1715,f336,f335,f4473,f4472,f4467,f4466,f4465,f4464,f4453,f4452,f4451,f4450,f1827,f1800,f1771,f1741,f1610,f1593,f1578,f1562,f1496,f1480,f351,f350,f349,f348,f346,f344,f343,f341,f329,f328,f3502,f3501,f2971,f2970,f13441,f1538,f13453,f13998,f13994,f13988,f4474,f4463,f1539,f1515,f1514,f339,f338,f337,f2002,f3505,f3504,f2974,f2973,f2968,f3499,f14178,f14176,f14170,f14119,f14116,f14059,f14056,f13948,f13803,f13771,f13699,f13698,f13697,f13696,f13695,f13694,f13693,f13692,f13691,f13689,f13688,f13687,f13686,f13685,f13683,f13584,f13583,f13582,f13581,f13540,f13534,f2176,f2173,f2142,f2132,f2129,f2126,f2090,f2087,f2055,f2046,f2043,f2037,f1929,f1928,f1712,f1710,f1709,f1623,f1622,f1536,f1511,f1493,f1477,f1220,f592,f473,f291,f290,f15203,f15202,f15201,f15197,f15196,f15223,f15375,f15378,f15380,f15240,f15242,f15243,f15245,f15246,f15248,f15249,f15251])).
fof(f15256,plain,(
  ( ! [X127,X125,X126,X124] : (sP5(sK14,'const.bool.IN_2'(s(X124,X125),s(fun(X124,bool),'const.pred_set.EMPTY_0')),happ(s(fun(X126,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X126,X127)))) ) | ($spl2 | $spl49 | $spl52)),
  inference(backward_demodulation,[],[f15223,f13371])).
fof(f13371,plain,(
  ( ! [X127,X125,X126,X124] : (sP5(sK14,'const.bool.IN_2'(s(X124,X125),s(fun(X124,bool),'const.pred_set.EMPTY_0')),happ(s(fun(X126,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X126,X127)))) ) | $spl2),
  inference(superposition,[],[f12590,f592])).
fof(f12590,plain,(
  ( ! [X2,X0,X1] : (sP5(sK14,X0,happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,X0))),s(X1,X2)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f381,f469,f403])).
fof(f15390,plain,(
  $false | ($spl2 | $spl49 | $spl52)),
  inference(global_subsumption,[],[f15255,f267,f266,f265,f268,f275,f287,f286,f332,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1691,f1700,f1717,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2010,f2040,f2044,f2048,f2084,f2088,f2092,f2125,f2130,f2170,f2174,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f3420,f3421,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f13301,f13298,f13462,f13474,f13476,f13479,f13491,f13503,f13505,f13508,f13746,f13765,f13767,f13770,f13762,f13781,f13797,f13799,f13802,f13795,f13837,f13850,f13852,f13855,f13929,f13942,f13944,f13947,f13941,f13995,f13997,f14141,f14142,f14143,f14144,f14145,f3503,f2972,f1714,f1715,f336,f335,f4473,f4472,f4467,f4466,f4465,f4464,f4453,f4452,f4451,f4450,f1827,f1800,f1771,f1741,f1610,f1593,f1578,f1562,f1496,f1480,f351,f350,f349,f348,f346,f344,f343,f341,f329,f328,f3502,f3501,f2971,f2970,f13441,f1538,f13453,f13998,f13994,f13988,f4474,f4463,f1539,f1515,f1514,f339,f338,f337,f2002,f3505,f3504,f2974,f2973,f2968,f3499,f14178,f14176,f14170,f14119,f14116,f14059,f14056,f13948,f13803,f13771,f13699,f13698,f13697,f13696,f13695,f13694,f13693,f13692,f13691,f13689,f13688,f13687,f13686,f13685,f13683,f13584,f13583,f13582,f13581,f13540,f13534,f2176,f2173,f2142,f2132,f2129,f2126,f2090,f2087,f2055,f2046,f2043,f2037,f1929,f1928,f1712,f1710,f1709,f1623,f1622,f1536,f1511,f1493,f1477,f1220,f592,f473,f291,f290,f15203,f15202,f15201,f15197,f15196,f15223,f15375,f15378,f15380,f15240,f15242,f15243,f15245,f15246,f15248,f15249,f15251])).
fof(f15255,plain,(
  ( ! [X123,X121,X122,X120] : (sP5(sK14,happ(s(fun(X122,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X122,X123)),'const.bool.IN_2'(s(X120,X121),s(fun(X120,bool),'const.pred_set.EMPTY_0')))) ) | ($spl2 | $spl49 | $spl52)),
  inference(backward_demodulation,[],[f15223,f13370])).
fof(f13370,plain,(
  ( ! [X123,X121,X122,X120] : (sP5(sK14,happ(s(fun(X122,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X122,X123)),'const.bool.IN_2'(s(X120,X121),s(fun(X120,bool),'const.pred_set.EMPTY_0')))) ) | $spl2),
  inference(superposition,[],[f12589,f592])).
fof(f12589,plain,(
  ( ! [X2,X0,X1] : (sP5(sK14,happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1))),s(X0,X2)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f381,f469,f403])).
fof(f15389,plain,(
  $false | ($spl49 | $spl52)),
  inference(global_subsumption,[],[f267,f266,f265,f268,f275,f287,f286,f332,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1691,f1700,f1717,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2010,f2040,f2044,f2048,f2084,f2088,f2092,f2125,f2130,f2170,f2174,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f3420,f3421,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f13301,f13298,f13462,f13474,f13476,f13479,f13491,f13503,f13505,f13508,f13746,f13765,f13767,f13770,f13762,f13781,f13797,f13799,f13802,f13795,f13837,f13850,f13852,f13855,f13929,f13942,f13944,f13947,f13941,f13995,f13997,f14141,f14142,f14143,f14144,f14145,f3503,f2972,f1714,f1715,f336,f335,f4473,f4472,f4467,f4466,f4465,f4464,f4453,f4452,f4451,f4450,f1827,f1800,f1771,f1741,f1610,f1593,f1578,f1562,f1496,f1480,f351,f350,f349,f348,f346,f344,f343,f341,f329,f328,f3502,f3501,f2971,f2970,f13441,f1538,f13453,f13998,f13994,f13988,f4474,f4463,f1539,f1515,f1514,f339,f338,f337,f2002,f3505,f3504,f2974,f2973,f2968,f3499,f14178,f14176,f14170,f14119,f14116,f14059,f14056,f13948,f13803,f13771,f13699,f13698,f13697,f13696,f13695,f13694,f13693,f13692,f13691,f13689,f13688,f13687,f13686,f13685,f13683,f13584,f13583,f13582,f13581,f13540,f13534,f2176,f2173,f2142,f2132,f2129,f2126,f2090,f2087,f2055,f2046,f2043,f2037,f1929,f1928,f1712,f1710,f1709,f1623,f1622,f1536,f1511,f1493,f1477,f1220,f592,f473,f291,f290,f15203,f15202,f15201,f15197,f15196,f15223,f15375,f15378,f15380,f15240,f15242,f15243,f15245,f15246,f15248,f15249,f15251])).
fof(f15388,plain,(
  $false | ($spl49 | $spl52)),
  inference(global_subsumption,[],[f267,f266,f265,f268,f275,f287,f286,f332,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1691,f1700,f1717,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2010,f2040,f2044,f2048,f2084,f2088,f2092,f2125,f2130,f2170,f2174,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f3420,f3421,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f13301,f13298,f13462,f13474,f13476,f13479,f13491,f13503,f13505,f13508,f13746,f13765,f13767,f13770,f13762,f13781,f13797,f13799,f13802,f13795,f13837,f13850,f13852,f13855,f13929,f13942,f13944,f13947,f13941,f13995,f13997,f14141,f14142,f14143,f14144,f14145,f3503,f2972,f1714,f1715,f336,f335,f4473,f4472,f4467,f4466,f4465,f4464,f4453,f4452,f4451,f4450,f1827,f1800,f1771,f1741,f1610,f1593,f1578,f1562,f1496,f1480,f351,f350,f349,f348,f346,f344,f343,f341,f329,f328,f3502,f3501,f2971,f2970,f13441,f1538,f13453,f13998,f13994,f13988,f4474,f4463,f1539,f1515,f1514,f339,f338,f337,f2002,f3505,f3504,f2974,f2973,f2968,f3499,f14178,f14176,f14170,f14119,f14116,f14059,f14056,f13948,f13803,f13771,f13699,f13698,f13697,f13696,f13695,f13694,f13693,f13692,f13691,f13689,f13688,f13687,f13686,f13685,f13683,f13584,f13583,f13582,f13581,f13540,f13534,f2176,f2173,f2142,f2132,f2129,f2126,f2090,f2087,f2055,f2046,f2043,f2037,f1929,f1928,f1712,f1710,f1709,f1623,f1622,f1536,f1511,f1493,f1477,f1220,f592,f473,f291,f290,f15203,f15202,f15201,f15197,f15196,f15223,f15375,f15378,f15380,f15240,f15242,f15243,f15245,f15246,f15248,f15249])).
fof(f15387,plain,(
  $false | ($spl49 | $spl52)),
  inference(global_subsumption,[],[f267,f266,f265,f268,f275,f287,f286,f332,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1691,f1700,f1717,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2010,f2040,f2044,f2048,f2084,f2088,f2092,f2125,f2130,f2170,f2174,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f3420,f3421,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f13301,f13298,f13462,f13474,f13476,f13479,f13491,f13503,f13505,f13508,f13746,f13765,f13767,f13770,f13762,f13781,f13797,f13799,f13802,f13795,f13837,f13850,f13852,f13855,f13929,f13942,f13944,f13947,f13941,f13995,f13997,f14141,f14142,f14143,f14144,f14145,f3503,f2972,f1714,f1715,f336,f335,f4473,f4472,f4467,f4466,f4465,f4464,f4453,f4452,f4451,f4450,f1827,f1800,f1771,f1741,f1610,f1593,f1578,f1562,f1496,f1480,f351,f350,f349,f348,f346,f344,f343,f341,f329,f328,f3502,f3501,f2971,f2970,f13441,f1538,f13453,f13998,f13994,f13988,f4474,f4463,f1539,f1515,f1514,f339,f338,f337,f2002,f3505,f3504,f2974,f2973,f2968,f3499,f14178,f14176,f14170,f14119,f14116,f14059,f14056,f13948,f13803,f13771,f13699,f13698,f13697,f13696,f13695,f13694,f13693,f13692,f13691,f13689,f13688,f13687,f13686,f13685,f13683,f13584,f13583,f13582,f13581,f13540,f13534,f2176,f2173,f2142,f2132,f2129,f2126,f2090,f2087,f2055,f2046,f2043,f2037,f1929,f1928,f1712,f1710,f1709,f1623,f1622,f1536,f1511,f1493,f1477,f1220,f592,f473,f291,f290,f15203,f15202,f15201,f15197,f15196,f15223,f15375,f15378,f15380,f15240,f15242,f15243,f15245,f15246,f15248])).
fof(f15386,plain,(
  $false | ($spl49 | $spl52)),
  inference(global_subsumption,[],[f267,f266,f265,f268,f275,f287,f286,f332,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1691,f1700,f1717,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2010,f2040,f2044,f2048,f2084,f2088,f2092,f2125,f2130,f2170,f2174,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f3420,f3421,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f13301,f13298,f13462,f13474,f13476,f13479,f13491,f13503,f13505,f13508,f13746,f13765,f13767,f13770,f13762,f13781,f13797,f13799,f13802,f13795,f13837,f13850,f13852,f13855,f13929,f13942,f13944,f13947,f13941,f13995,f13997,f14141,f14142,f14143,f14144,f14145,f3503,f2972,f1714,f1715,f336,f335,f4473,f4472,f4467,f4466,f4465,f4464,f4453,f4452,f4451,f4450,f1827,f1800,f1771,f1741,f1610,f1593,f1578,f1562,f1496,f1480,f351,f350,f349,f348,f346,f344,f343,f341,f329,f328,f3502,f3501,f2971,f2970,f13441,f1538,f13453,f13998,f13994,f13988,f4474,f4463,f1539,f1515,f1514,f339,f338,f337,f2002,f3505,f3504,f2974,f2973,f2968,f3499,f14178,f14176,f14170,f14119,f14116,f14059,f14056,f13948,f13803,f13771,f13699,f13698,f13697,f13696,f13695,f13694,f13693,f13692,f13691,f13689,f13688,f13687,f13686,f13685,f13683,f13584,f13583,f13582,f13581,f13540,f13534,f2176,f2173,f2142,f2132,f2129,f2126,f2090,f2087,f2055,f2046,f2043,f2037,f1929,f1928,f1712,f1710,f1709,f1623,f1622,f1536,f1511,f1493,f1477,f1220,f592,f473,f291,f290,f15203,f15202,f15201,f15197,f15196,f15223,f15375,f15378,f15380,f15240,f15242,f15243,f15245,f15246])).
fof(f15385,plain,(
  $false | ($spl49 | $spl52)),
  inference(global_subsumption,[],[f267,f266,f265,f268,f275,f287,f286,f332,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1691,f1700,f1717,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2010,f2040,f2044,f2048,f2084,f2088,f2092,f2125,f2130,f2170,f2174,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f3420,f3421,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f13301,f13298,f13462,f13474,f13476,f13479,f13491,f13503,f13505,f13508,f13746,f13765,f13767,f13770,f13762,f13781,f13797,f13799,f13802,f13795,f13837,f13850,f13852,f13855,f13929,f13942,f13944,f13947,f13941,f13995,f13997,f14141,f14142,f14143,f14144,f14145,f3503,f2972,f1714,f1715,f336,f335,f4473,f4472,f4467,f4466,f4465,f4464,f4453,f4452,f4451,f4450,f1827,f1800,f1771,f1741,f1610,f1593,f1578,f1562,f1496,f1480,f351,f350,f349,f348,f346,f344,f343,f341,f329,f328,f3502,f3501,f2971,f2970,f13441,f1538,f13453,f13998,f13994,f13988,f4474,f4463,f1539,f1515,f1514,f339,f338,f337,f2002,f3505,f3504,f2974,f2973,f2968,f3499,f14178,f14176,f14170,f14119,f14116,f14059,f14056,f13948,f13803,f13771,f13699,f13698,f13697,f13696,f13695,f13694,f13693,f13692,f13691,f13689,f13688,f13687,f13686,f13685,f13683,f13584,f13583,f13582,f13581,f13540,f13534,f2176,f2173,f2142,f2132,f2129,f2126,f2090,f2087,f2055,f2046,f2043,f2037,f1929,f1928,f1712,f1710,f1709,f1623,f1622,f1536,f1511,f1493,f1477,f1220,f592,f473,f291,f290,f15203,f15202,f15201,f15197,f15196,f15223,f15375,f15378,f15380,f15240,f15242,f15243,f15245])).
fof(f15384,plain,(
  $false | ($spl49 | $spl52)),
  inference(global_subsumption,[],[f267,f266,f265,f268,f275,f287,f286,f332,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1691,f1700,f1717,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2010,f2040,f2044,f2048,f2084,f2088,f2092,f2125,f2130,f2170,f2174,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f3420,f3421,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f13301,f13298,f13462,f13474,f13476,f13479,f13491,f13503,f13505,f13508,f13746,f13765,f13767,f13770,f13762,f13781,f13797,f13799,f13802,f13795,f13837,f13850,f13852,f13855,f13929,f13942,f13944,f13947,f13941,f13995,f13997,f14141,f14142,f14143,f14144,f14145,f3503,f2972,f1714,f1715,f336,f335,f4473,f4472,f4467,f4466,f4465,f4464,f4453,f4452,f4451,f4450,f1827,f1800,f1771,f1741,f1610,f1593,f1578,f1562,f1496,f1480,f351,f350,f349,f348,f346,f344,f343,f341,f329,f328,f3502,f3501,f2971,f2970,f13441,f1538,f13453,f13998,f13994,f13988,f4474,f4463,f1539,f1515,f1514,f339,f338,f337,f2002,f3505,f3504,f2974,f2973,f2968,f3499,f14178,f14176,f14170,f14119,f14116,f14059,f14056,f13948,f13803,f13771,f13699,f13698,f13697,f13696,f13695,f13694,f13693,f13692,f13691,f13689,f13688,f13687,f13686,f13685,f13683,f13584,f13583,f13582,f13581,f13540,f13534,f2176,f2173,f2142,f2132,f2129,f2126,f2090,f2087,f2055,f2046,f2043,f2037,f1929,f1928,f1712,f1710,f1709,f1623,f1622,f1536,f1511,f1493,f1477,f1220,f592,f473,f291,f290,f15203,f15202,f15201,f15197,f15196,f15223,f15375,f15378,f15380,f15240,f15242,f15243])).
fof(f15383,plain,(
  $false | ($spl49 | $spl52)),
  inference(global_subsumption,[],[f267,f266,f265,f268,f275,f287,f286,f332,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1691,f1700,f1717,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2010,f2040,f2044,f2048,f2084,f2088,f2092,f2125,f2130,f2170,f2174,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f3420,f3421,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f13301,f13298,f13462,f13474,f13476,f13479,f13491,f13503,f13505,f13508,f13746,f13765,f13767,f13770,f13762,f13781,f13797,f13799,f13802,f13795,f13837,f13850,f13852,f13855,f13929,f13942,f13944,f13947,f13941,f13995,f13997,f14141,f14142,f14143,f14144,f14145,f3503,f2972,f1714,f1715,f336,f335,f4473,f4472,f4467,f4466,f4465,f4464,f4453,f4452,f4451,f4450,f1827,f1800,f1771,f1741,f1610,f1593,f1578,f1562,f1496,f1480,f351,f350,f349,f348,f346,f344,f343,f341,f329,f328,f3502,f3501,f2971,f2970,f13441,f1538,f13453,f13998,f13994,f13988,f4474,f4463,f1539,f1515,f1514,f339,f338,f337,f2002,f3505,f3504,f2974,f2973,f2968,f3499,f14178,f14176,f14170,f14119,f14116,f14059,f14056,f13948,f13803,f13771,f13699,f13698,f13697,f13696,f13695,f13694,f13693,f13692,f13691,f13689,f13688,f13687,f13686,f13685,f13683,f13584,f13583,f13582,f13581,f13540,f13534,f2176,f2173,f2142,f2132,f2129,f2126,f2090,f2087,f2055,f2046,f2043,f2037,f1929,f1928,f1712,f1710,f1709,f1623,f1622,f1536,f1511,f1493,f1477,f1220,f592,f473,f291,f290,f15203,f15202,f15201,f15197,f15196,f15223,f15375,f15378,f15380,f15240,f15242])).
fof(f15382,plain,(
  $false | ($spl49 | $spl52)),
  inference(global_subsumption,[],[f267,f266,f265,f268,f275,f287,f286,f332,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1691,f1700,f1717,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2010,f2040,f2044,f2048,f2084,f2088,f2092,f2125,f2130,f2170,f2174,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f3420,f3421,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f13301,f13298,f13462,f13474,f13476,f13479,f13491,f13503,f13505,f13508,f13746,f13765,f13767,f13770,f13762,f13781,f13797,f13799,f13802,f13795,f13837,f13850,f13852,f13855,f13929,f13942,f13944,f13947,f13941,f13995,f13997,f14141,f14142,f14143,f14144,f14145,f3503,f2972,f1714,f1715,f336,f335,f4473,f4472,f4467,f4466,f4465,f4464,f4453,f4452,f4451,f4450,f1827,f1800,f1771,f1741,f1610,f1593,f1578,f1562,f1496,f1480,f351,f350,f349,f348,f346,f344,f343,f341,f329,f328,f3502,f3501,f2971,f2970,f13441,f1538,f13453,f13998,f13994,f13988,f4474,f4463,f1539,f1515,f1514,f339,f338,f337,f2002,f3505,f3504,f2974,f2973,f2968,f3499,f14178,f14176,f14170,f14119,f14116,f14059,f14056,f13948,f13803,f13771,f13699,f13698,f13697,f13696,f13695,f13694,f13693,f13692,f13691,f13689,f13688,f13687,f13686,f13685,f13683,f13584,f13583,f13582,f13581,f13540,f13534,f2176,f2173,f2142,f2132,f2129,f2126,f2090,f2087,f2055,f2046,f2043,f2037,f1929,f1928,f1712,f1710,f1709,f1623,f1622,f1536,f1511,f1493,f1477,f1220,f592,f473,f291,f290,f15203,f15202,f15201,f15197,f15196,f15223,f15375,f15378,f15380,f15240])).
fof(f15381,plain,(
  $false | ($spl49 | $spl52)),
  inference(global_subsumption,[],[f267,f266,f265,f268,f275,f287,f286,f332,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1691,f1700,f1717,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2010,f2040,f2044,f2048,f2084,f2088,f2092,f2125,f2130,f2170,f2174,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f3420,f3421,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f13301,f13298,f13462,f13474,f13476,f13479,f13491,f13503,f13505,f13508,f13746,f13765,f13767,f13770,f13762,f13781,f13797,f13799,f13802,f13795,f13837,f13850,f13852,f13855,f13929,f13942,f13944,f13947,f13941,f13995,f13997,f14141,f14142,f14143,f14144,f14145,f3503,f2972,f1714,f1715,f336,f335,f4473,f4472,f4467,f4466,f4465,f4464,f4453,f4452,f4451,f4450,f1827,f1800,f1771,f1741,f1610,f1593,f1578,f1562,f1496,f1480,f351,f350,f349,f348,f346,f344,f343,f341,f329,f328,f3502,f3501,f2971,f2970,f13441,f1538,f13453,f13998,f13994,f13988,f4474,f4463,f1539,f1515,f1514,f339,f338,f337,f2002,f3505,f3504,f2974,f2973,f2968,f3499,f14178,f14176,f14170,f14119,f14116,f14059,f14056,f13948,f13803,f13771,f13699,f13698,f13697,f13696,f13695,f13694,f13693,f13692,f13691,f13689,f13688,f13687,f13686,f13685,f13683,f13584,f13583,f13582,f13581,f13540,f13534,f2176,f2173,f2142,f2132,f2129,f2126,f2090,f2087,f2055,f2046,f2043,f2037,f1929,f1928,f1712,f1710,f1709,f1623,f1622,f1536,f1511,f1493,f1477,f1220,f592,f473,f291,f290,f15203,f15202,f15201,f15197,f15196,f15223,f15375,f15378,f15380])).
fof(f15379,plain,(
  $false | ($spl49 | $spl52)),
  inference(global_subsumption,[],[f267,f266,f265,f268,f275,f287,f286,f332,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1691,f1700,f1717,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2010,f2040,f2044,f2048,f2084,f2088,f2092,f2125,f2130,f2170,f2174,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f3420,f3421,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f13301,f13298,f13462,f13474,f13476,f13479,f13491,f13503,f13505,f13508,f13746,f13765,f13767,f13770,f13762,f13781,f13797,f13799,f13802,f13795,f13837,f13850,f13852,f13855,f13929,f13942,f13944,f13947,f13941,f13995,f13997,f14141,f14142,f14143,f14144,f14145,f3503,f2972,f1714,f1715,f336,f335,f4473,f4472,f4467,f4466,f4465,f4464,f4453,f4452,f4451,f4450,f1827,f1800,f1771,f1741,f1610,f1593,f1578,f1562,f1496,f1480,f351,f350,f349,f348,f346,f344,f343,f341,f329,f328,f3502,f3501,f2971,f2970,f13441,f1538,f13453,f13998,f13994,f13988,f4474,f4463,f1539,f1515,f1514,f339,f338,f337,f2002,f3505,f3504,f2974,f2973,f2968,f3499,f14178,f14176,f14170,f14119,f14116,f14059,f14056,f13948,f13803,f13771,f13699,f13698,f13697,f13696,f13695,f13694,f13693,f13692,f13691,f13689,f13688,f13687,f13686,f13685,f13683,f13584,f13583,f13582,f13581,f13540,f13534,f2176,f2173,f2142,f2132,f2129,f2126,f2090,f2087,f2055,f2046,f2043,f2037,f1929,f1928,f1712,f1710,f1709,f1623,f1622,f1536,f1511,f1493,f1477,f1220,f592,f473,f291,f290,f15203,f15202,f15201,f15197,f15196,f15223,f15375,f15378])).
fof(f15377,plain,(
  $false | ($spl49 | $spl52)),
  inference(subsumption_resolution,[],[f15226,f268])).
fof(f15226,plain,(
  ~p(s(bool,'T_0')) | ($spl49 | $spl52)),
  inference(backward_demodulation,[],[f15223,f473])).
fof(f15376,plain,(
  $false | ($spl49 | $spl52)),
  inference(global_subsumption,[],[f267,f266,f265,f268,f275,f287,f286,f332,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1691,f1700,f1717,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2010,f2040,f2044,f2048,f2084,f2088,f2092,f2125,f2130,f2170,f2174,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f3420,f3421,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f13301,f13298,f13462,f13474,f13476,f13479,f13491,f13503,f13505,f13508,f13746,f13765,f13767,f13770,f13762,f13781,f13797,f13799,f13802,f13795,f13837,f13850,f13852,f13855,f13929,f13942,f13944,f13947,f13941,f13995,f13997,f14141,f14142,f14143,f14144,f14145,f3503,f2972,f1714,f1715,f336,f335,f4473,f4472,f4467,f4466,f4465,f4464,f4453,f4452,f4451,f4450,f1827,f1800,f1771,f1741,f1610,f1593,f1578,f1562,f1496,f1480,f351,f350,f349,f348,f346,f344,f343,f341,f329,f328,f3502,f3501,f2971,f2970,f13441,f1538,f13453,f13998,f13994,f13988,f4474,f4463,f1539,f1515,f1514,f339,f338,f337,f2002,f3505,f3504,f2974,f2973,f2968,f3499,f14178,f14176,f14170,f14119,f14116,f14059,f14056,f13948,f13803,f13771,f13699,f13698,f13697,f13696,f13695,f13694,f13693,f13692,f13691,f13689,f13688,f13687,f13686,f13685,f13683,f13584,f13583,f13582,f13581,f13540,f13534,f2176,f2173,f2142,f2132,f2129,f2126,f2090,f2087,f2055,f2046,f2043,f2037,f1929,f1928,f1712,f1710,f1709,f1623,f1622,f1536,f1511,f1493,f1477,f1220,f592,f473,f291,f290,f15203,f15202,f15201,f15197,f15196,f15223,f15375])).
fof(f15222,plain,(
  $false | ($spl49 | $spl55)),
  inference(subsumption_resolution,[],[f15179,f13303])).
fof(f15179,plain,(
  s(bool,'F_0') = s(bool,f(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12)))) | $spl49),
  inference(unit_resulting_resolution,[],[f13298,f275])).
fof(f15221,plain,(
  $false | ($spl49 | $spl55)),
  inference(subsumption_resolution,[],[f15181,f13303])).
fof(f15181,plain,(
  s(bool,'F_0') = s(bool,f(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12)))) | $spl49),
  inference(unit_resulting_resolution,[],[f515,f13298,f747])).
fof(f15220,plain,(
  $false | ($spl5 | $spl45 | $spl49 | $spl55)),
  inference(subsumption_resolution,[],[f15219,f13303])).
fof(f15219,plain,(
  s(bool,'F_0') = s(bool,f(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12)))) | ($spl5 | $spl45 | $spl49)),
  inference(forward_demodulation,[],[f15183,f15065])).
fof(f15065,plain,(
  s(bool,'F_0') = s(bool,g(s(sK12,sK26(sK13,'const.pred_set.EMPTY_0',bool,sK12)))) | ($spl5 | $spl45)),
  inference(unit_resulting_resolution,[],[f473,f14956,f810])).
fof(f14956,plain,(
  s(bool,'T_0') != s(bool,g(s(sK12,sK26(sK13,'const.pred_set.EMPTY_0',bool,sK12)))) | ($spl5 | $spl45)),
  inference(backward_demodulation,[],[f14945,f9873])).
fof(f9873,plain,(
  s(bool,f(s(sK12,sK26(sK13,'const.pred_set.EMPTY_0',bool,sK12)))) != s(bool,g(s(sK12,sK26(sK13,'const.pred_set.EMPTY_0',bool,sK12)))) | $spl5),
  inference(unit_resulting_resolution,[],[f472,f382])).
fof(f14945,plain,(
  s(bool,'T_0') = s(bool,f(s(sK12,sK26(sK13,'const.pred_set.EMPTY_0',bool,sK12)))) | $spl45),
  inference(unit_resulting_resolution,[],[f473,f13278,f810])).
fof(f15183,plain,(
  s(bool,f(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12)))) = s(bool,g(s(sK12,sK26(sK13,'const.pred_set.EMPTY_0',bool,sK12)))) | ($spl5 | $spl45 | $spl49)),
  inference(unit_resulting_resolution,[],[f14956,f13298,f747])).
fof(f15218,plain,(
  $false | ($spl49 | $spl55)),
  inference(subsumption_resolution,[],[f15217,f13303])).
fof(f15217,plain,(
  s(bool,'F_0') = s(bool,f(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12)))) | $spl49),
  inference(forward_demodulation,[],[f15184,f592])).
fof(f15184,plain,(
  ( ! [X0,X1] : (s(bool,f(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12)))) = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | $spl49),
  inference(unit_resulting_resolution,[],[f591,f13298,f747])).
fof(f15216,plain,(
  $false | ($spl49 | $spl55)),
  inference(subsumption_resolution,[],[f15186,f13303])).
fof(f15186,plain,(
  s(bool,'F_0') = s(bool,f(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12)))) | $spl49),
  inference(unit_resulting_resolution,[],[f515,f13298,f747])).
fof(f15215,plain,(
  $false | ($spl5 | $spl45 | $spl49 | $spl55)),
  inference(subsumption_resolution,[],[f15214,f13303])).
fof(f15214,plain,(
  s(bool,'F_0') = s(bool,f(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12)))) | ($spl5 | $spl45 | $spl49)),
  inference(forward_demodulation,[],[f15188,f15065])).
fof(f15188,plain,(
  s(bool,f(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12)))) = s(bool,g(s(sK12,sK26(sK13,'const.pred_set.EMPTY_0',bool,sK12)))) | ($spl5 | $spl45 | $spl49)),
  inference(unit_resulting_resolution,[],[f14956,f13298,f747])).
fof(f15213,plain,(
  $false | ($spl49 | $spl55)),
  inference(subsumption_resolution,[],[f15212,f13303])).
fof(f15212,plain,(
  s(bool,'F_0') = s(bool,f(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12)))) | $spl49),
  inference(forward_demodulation,[],[f15189,f592])).
fof(f15189,plain,(
  ( ! [X0,X1] : (s(bool,f(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12)))) = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | $spl49),
  inference(unit_resulting_resolution,[],[f591,f13298,f747])).
fof(f15211,plain,(
  $false | ($spl49 | $spl55)),
  inference(subsumption_resolution,[],[f15191,f13303])).
fof(f15191,plain,(
  s(bool,'F_0') = s(bool,f(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12)))) | $spl49),
  inference(unit_resulting_resolution,[],[f473,f13298,f810])).
fof(f14481,plain,(
  $false | ($spl39 | $spl42)),
  inference(global_subsumption,[],[f267,f266,f265,f268,f275,f287,f286,f332,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1691,f1700,f1717,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2010,f2040,f2044,f2048,f2084,f2088,f2092,f2125,f2130,f2170,f2174,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f3420,f3421,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f3503,f2972,f1715,f1714,f336,f335,f3502,f3501,f2971,f2970,f4473,f4472,f4467,f4466,f4465,f4464,f4453,f4452,f4451,f4450,f1827,f1800,f1771,f1741,f1610,f1593,f1578,f1562,f1496,f1480,f351,f350,f349,f348,f346,f344,f343,f341,f329,f328,f2973,f2974,f3504,f3505,f2002,f1538,f4474,f4463,f1539,f1515,f1514,f339,f338,f337,f2968,f3499,f2176,f2173,f2142,f2132,f2129,f2126,f2090,f2087,f2055,f2046,f2043,f2037,f1929,f1928,f1712,f1710,f1709,f1623,f1622,f1536,f1511,f1493,f1477,f1220,f592,f473,f291,f290,f13276,f13273,f13441,f13453,f13462,f13474,f13476,f13479,f13491,f13503,f13505,f13508,f13534,f13540,f13683,f13685,f13686,f13687,f13688,f13689,f13691,f13692,f13693,f13694,f13695,f13696,f13697,f13698,f13699,f13581,f13582,f13583,f13584,f13746,f13765,f13767,f13770,f13771,f13762,f13781,f13797,f13799,f13802,f13803,f13795,f13837,f13850,f13852,f13855,f13929,f13942,f13944,f13947,f13948,f13941,f13994,f13988,f13995,f13997,f13998,f14056,f14059,f14116,f14119,f14141,f14142,f14143,f14144,f14145,f14176,f14170,f14178,f14228,f14227,f14226,f14222,f14221,f14246,f14398,f14401,f14403,f14263,f14265,f14266,f14268,f14269,f14271,f14272,f14274,f14415,f14476,f14478,f14480])).
fof(f14480,plain,(
  ( ! [X4,X2,X0] : (s(fun(X0,bool),X2) = s(fun(X0,bool),X4)) ) | ($spl39 | $spl42)),
  inference(subsumption_resolution,[],[f14385,f14246])).
fof(f14385,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK17(X4,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X2) = s(fun(X0,bool),X4)) ) | ($spl39 | $spl42)),
  inference(backward_demodulation,[],[f14246,f1714])).
fof(f14478,plain,(
  ( ! [X4,X2,X0] : (s(fun(X0,bool),X2) = s(fun(X0,bool),X4)) ) | ($spl39 | $spl42)),
  inference(subsumption_resolution,[],[f14384,f14246])).
fof(f14384,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK17(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X4,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X2) = s(fun(X0,bool),X4)) ) | ($spl39 | $spl42)),
  inference(backward_demodulation,[],[f14246,f1715])).
fof(f14476,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl39 | $spl42)),
  inference(subsumption_resolution,[],[f14383,f14246])).
fof(f14383,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK17(X2,X1,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl39 | $spl42)),
  inference(backward_demodulation,[],[f14246,f336])).
fof(f14415,plain,(
  ( ! [X35] : (s(bool,'F_0') = s(bool,X35)) ) | ($spl39 | $spl42)),
  inference(subsumption_resolution,[],[f14290,f14246])).
fof(f14290,plain,(
  ( ! [X35,X36,X34] : (s(bool,'T_0') != s(bool,happ(s(fun(X34,bool),'_dst_x0_1'(s(bool,X35))),s(X34,X36))) | s(bool,'F_0') = s(bool,X35)) ) | ($spl39 | $spl42)),
  inference(backward_demodulation,[],[f14246,f13685])).
fof(f14274,plain,(
  ( ! [X10,X8,X11,X9] : (sP5('T_0','const.bool.IN_2'(s(X8,X9),s(fun(X8,bool),'const.pred_set.EMPTY_0')),happ(s(fun(X10,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X10,X11)))) ) | ($spl39 | $spl42)),
  inference(backward_demodulation,[],[f14246,f2176])).
fof(f14272,plain,(
  ( ! [X2,X0,X3,X1] : (~sP2(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X3)),'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')),'F_0')) ) | ($spl39 | $spl42)),
  inference(backward_demodulation,[],[f14246,f2142])).
fof(f14271,plain,(
  ( ! [X10,X8,X11,X9] : (~sP5('F_0','const.bool.IN_2'(s(X8,X9),s(fun(X8,bool),'const.pred_set.EMPTY_0')),happ(s(fun(X10,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X10,X11)))) ) | ($spl39 | $spl42)),
  inference(backward_demodulation,[],[f14246,f2132])).
fof(f14269,plain,(
  ( ! [X0,X1] : (~sP2(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)),'F_0','F_0')) ) | ($spl39 | $spl42)),
  inference(backward_demodulation,[],[f14246,f2126])).
fof(f14268,plain,(
  ( ! [X10,X8,X11,X9] : (sP5('T_0',happ(s(fun(X10,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X10,X11)),'const.bool.IN_2'(s(X8,X9),s(fun(X8,bool),'const.pred_set.EMPTY_0')))) ) | ($spl39 | $spl42)),
  inference(backward_demodulation,[],[f14246,f2090])).
fof(f14266,plain,(
  ( ! [X2,X0,X3,X1] : (~sP2('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),'F_0')) ) | ($spl39 | $spl42)),
  inference(backward_demodulation,[],[f14246,f2055])).
fof(f14265,plain,(
  ( ! [X10,X8,X11,X9] : (~sP5('F_0',happ(s(fun(X10,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X10,X11)),'const.bool.IN_2'(s(X8,X9),s(fun(X8,bool),'const.pred_set.EMPTY_0')))) ) | ($spl39 | $spl42)),
  inference(backward_demodulation,[],[f14246,f2046])).
fof(f14263,plain,(
  ( ! [X0,X1] : (~sP2('F_0',happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)),'F_0')) ) | ($spl39 | $spl42)),
  inference(backward_demodulation,[],[f14246,f2037])).
fof(f14403,plain,(
  ( ! [X6,X7] : (s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),X7)) ) | ($spl39 | $spl42)),
  inference(subsumption_resolution,[],[f14262,f14246])).
fof(f14262,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X6,sK17('const.pred_set.EMPTY_0',X7,X6)),s(fun(X6,bool),X7))) | s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),X7)) ) | ($spl39 | $spl42)),
  inference(backward_demodulation,[],[f14246,f1929])).
fof(f14401,plain,(
  ( ! [X4,X5] : (s(fun(X4,bool),'const.pred_set.EMPTY_0') = s(fun(X4,bool),X5)) ) | ($spl39 | $spl42)),
  inference(subsumption_resolution,[],[f14261,f14246])).
fof(f14261,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X4,sK17(X5,'const.pred_set.EMPTY_0',X4)),s(fun(X4,bool),X5))) | s(fun(X4,bool),'const.pred_set.EMPTY_0') = s(fun(X4,bool),X5)) ) | ($spl39 | $spl42)),
  inference(backward_demodulation,[],[f14246,f1928])).
fof(f14398,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl39 | $spl42)),
  inference(subsumption_resolution,[],[f14247,f14246])).
fof(f14247,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl39 | $spl42)),
  inference(backward_demodulation,[],[f14246,f290])).
fof(f14246,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl39 | $spl42)),
  inference(subsumption_resolution,[],[f14235,f13276])).
fof(f14235,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(sK12,sK26(sK13,'const.pred_set.EMPTY_0',bool,sK12)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl39),
  inference(trivial_inequality_removal,[],[f14229])).
fof(f14229,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,f(s(sK12,sK26(sK13,'const.pred_set.EMPTY_0',bool,sK12)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl39),
  inference(superposition,[],[f13273,f747])).
fof(f14221,plain,(
  ~sP5('T_0',f(s(sK12,sK26(sK13,'const.pred_set.EMPTY_0',bool,sK12))),'T_0') | $spl39),
  inference(unit_resulting_resolution,[],[f268,f13273,f404])).
fof(f14222,plain,(
  sP5('F_0',f(s(sK12,sK26(sK13,'const.pred_set.EMPTY_0',bool,sK12))),'T_0') | $spl39),
  inference(unit_resulting_resolution,[],[f473,f13273,f402])).
fof(f14226,plain,(
  ~sP5('T_0','T_0',f(s(sK12,sK26(sK13,'const.pred_set.EMPTY_0',bool,sK12)))) | $spl39),
  inference(unit_resulting_resolution,[],[f268,f13273,f404])).
fof(f14227,plain,(
  sP5('F_0','T_0',f(s(sK12,sK26(sK13,'const.pred_set.EMPTY_0',bool,sK12)))) | $spl39),
  inference(unit_resulting_resolution,[],[f473,f13273,f402])).
fof(f14228,plain,(
  ~p(s(bool,f(s(sK12,sK26(sK13,'const.pred_set.EMPTY_0',bool,sK12))))) | $spl39),
  inference(unit_resulting_resolution,[],[f268,f13273,f355])).
fof(f14479,plain,(
  $false | ($spl39 | $spl42)),
  inference(global_subsumption,[],[f267,f266,f265,f268,f275,f287,f286,f332,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1691,f1700,f1717,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2010,f2040,f2044,f2048,f2084,f2088,f2092,f2125,f2130,f2170,f2174,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f3420,f3421,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f3503,f2972,f1715,f1714,f336,f335,f3502,f3501,f2971,f2970,f4473,f4472,f4467,f4466,f4465,f4464,f4453,f4452,f4451,f4450,f1827,f1800,f1771,f1741,f1610,f1593,f1578,f1562,f1496,f1480,f351,f350,f349,f348,f346,f344,f343,f341,f329,f328,f2973,f2974,f3504,f3505,f2002,f1538,f4474,f4463,f1539,f1515,f1514,f339,f338,f337,f2968,f3499,f2176,f2173,f2142,f2132,f2129,f2126,f2090,f2087,f2055,f2046,f2043,f2037,f1929,f1928,f1712,f1710,f1709,f1623,f1622,f1536,f1511,f1493,f1477,f1220,f592,f473,f291,f290,f13276,f13273,f13441,f13453,f13462,f13474,f13476,f13479,f13491,f13503,f13505,f13508,f13534,f13540,f13683,f13685,f13686,f13687,f13688,f13689,f13691,f13692,f13693,f13694,f13695,f13696,f13697,f13698,f13699,f13581,f13582,f13583,f13584,f13746,f13765,f13767,f13770,f13771,f13762,f13781,f13797,f13799,f13802,f13803,f13795,f13837,f13850,f13852,f13855,f13929,f13942,f13944,f13947,f13948,f13941,f13994,f13988,f13995,f13997,f13998,f14056,f14059,f14116,f14119,f14141,f14142,f14143,f14144,f14145,f14176,f14170,f14178,f14228,f14227,f14226,f14222,f14221,f14246,f14398,f14401,f14403,f14263,f14265,f14266,f14268,f14269,f14271,f14272,f14274,f14415,f14476,f14478])).
fof(f14477,plain,(
  $false | ($spl39 | $spl42)),
  inference(global_subsumption,[],[f267,f266,f265,f268,f275,f287,f286,f332,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1691,f1700,f1717,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2010,f2040,f2044,f2048,f2084,f2088,f2092,f2125,f2130,f2170,f2174,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f3420,f3421,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f3503,f2972,f1715,f1714,f336,f335,f3502,f3501,f2971,f2970,f4473,f4472,f4467,f4466,f4465,f4464,f4453,f4452,f4451,f4450,f1827,f1800,f1771,f1741,f1610,f1593,f1578,f1562,f1496,f1480,f351,f350,f349,f348,f346,f344,f343,f341,f329,f328,f2973,f2974,f3504,f3505,f2002,f1538,f4474,f4463,f1539,f1515,f1514,f339,f338,f337,f2968,f3499,f2176,f2173,f2142,f2132,f2129,f2126,f2090,f2087,f2055,f2046,f2043,f2037,f1929,f1928,f1712,f1710,f1709,f1623,f1622,f1536,f1511,f1493,f1477,f1220,f592,f473,f291,f290,f13276,f13273,f13441,f13453,f13462,f13474,f13476,f13479,f13491,f13503,f13505,f13508,f13534,f13540,f13683,f13685,f13686,f13687,f13688,f13689,f13691,f13692,f13693,f13694,f13695,f13696,f13697,f13698,f13699,f13581,f13582,f13583,f13584,f13746,f13765,f13767,f13770,f13771,f13762,f13781,f13797,f13799,f13802,f13803,f13795,f13837,f13850,f13852,f13855,f13929,f13942,f13944,f13947,f13948,f13941,f13994,f13988,f13995,f13997,f13998,f14056,f14059,f14116,f14119,f14141,f14142,f14143,f14144,f14145,f14176,f14170,f14178,f14228,f14227,f14226,f14222,f14221,f14246,f14398,f14401,f14403,f14263,f14265,f14266,f14268,f14269,f14271,f14272,f14274,f14415,f14476])).
fof(f14427,plain,(
  $false | ($spl32 | $spl39 | $spl42)),
  inference(subsumption_resolution,[],[f14339,f268])).
fof(f14339,plain,(
  ~p(s(bool,'T_0')) | ($spl32 | $spl39 | $spl42)),
  inference(backward_demodulation,[],[f14246,f8928])).
fof(f14422,plain,(
  $false | ($spl39 | $spl42)),
  inference(subsumption_resolution,[],[f14320,f268])).
fof(f14320,plain,(
  ~p(s(bool,'T_0')) | ($spl39 | $spl42)),
  inference(backward_demodulation,[],[f14246,f14228])).
fof(f14421,plain,(
  $false | ($spl5 | $spl39 | $spl42)),
  inference(subsumption_resolution,[],[f14319,f14246])).
fof(f14319,plain,(
  s(bool,'T_0') != s(bool,g(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12)))) | ($spl5 | $spl39 | $spl42)),
  inference(backward_demodulation,[],[f14246,f9875])).
fof(f14420,plain,(
  $false | ($spl39 | $spl42 | $spl81)),
  inference(subsumption_resolution,[],[f14313,f14246])).
fof(f14313,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl39 | $spl42 | $spl81)),
  inference(backward_demodulation,[],[f14246,f14154])).
fof(f14419,plain,(
  $false | ($spl39 | $spl42 | $spl75)),
  inference(subsumption_resolution,[],[f14311,f14246])).
fof(f14311,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl39 | $spl42 | $spl75)),
  inference(backward_demodulation,[],[f14246,f14149])).
fof(f14416,plain,(
  $false | ($spl39 | $spl42)),
  inference(global_subsumption,[],[f267,f266,f265,f268,f275,f287,f286,f332,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1691,f1700,f1717,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2010,f2040,f2044,f2048,f2084,f2088,f2092,f2125,f2130,f2170,f2174,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f3420,f3421,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f3503,f2972,f1715,f1714,f336,f335,f3502,f3501,f2971,f2970,f4473,f4472,f4467,f4466,f4465,f4464,f4453,f4452,f4451,f4450,f1827,f1800,f1771,f1741,f1610,f1593,f1578,f1562,f1496,f1480,f351,f350,f349,f348,f346,f344,f343,f341,f329,f328,f2973,f2974,f3504,f3505,f2002,f1538,f4474,f4463,f1539,f1515,f1514,f339,f338,f337,f2968,f3499,f2176,f2173,f2142,f2132,f2129,f2126,f2090,f2087,f2055,f2046,f2043,f2037,f1929,f1928,f1712,f1710,f1709,f1623,f1622,f1536,f1511,f1493,f1477,f1220,f592,f473,f291,f290,f13276,f13273,f13441,f13453,f13462,f13474,f13476,f13479,f13491,f13503,f13505,f13508,f13534,f13540,f13683,f13685,f13686,f13687,f13688,f13689,f13691,f13692,f13693,f13694,f13695,f13696,f13697,f13698,f13699,f13581,f13582,f13583,f13584,f13746,f13765,f13767,f13770,f13771,f13762,f13781,f13797,f13799,f13802,f13803,f13795,f13837,f13850,f13852,f13855,f13929,f13942,f13944,f13947,f13948,f13941,f13994,f13988,f13995,f13997,f13998,f14056,f14059,f14116,f14119,f14141,f14142,f14143,f14144,f14145,f14176,f14170,f14178,f14228,f14227,f14226,f14222,f14221,f14246,f14398,f14401,f14403,f14263,f14265,f14266,f14268,f14269,f14271,f14272,f14274,f14415])).
fof(f14414,plain,(
  $false | ($spl2 | $spl39 | $spl42)),
  inference(global_subsumption,[],[f14280,f267,f266,f265,f268,f275,f287,f286,f332,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1691,f1700,f1717,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2010,f2040,f2044,f2048,f2084,f2088,f2092,f2125,f2130,f2170,f2174,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f3420,f3421,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f3503,f2972,f1715,f1714,f336,f335,f3502,f3501,f2971,f2970,f4473,f4472,f4467,f4466,f4465,f4464,f4453,f4452,f4451,f4450,f1827,f1800,f1771,f1741,f1610,f1593,f1578,f1562,f1496,f1480,f351,f350,f349,f348,f346,f344,f343,f341,f329,f328,f2973,f2974,f3504,f3505,f2002,f1538,f4474,f4463,f1539,f1515,f1514,f339,f338,f337,f2968,f3499,f2176,f2173,f2142,f2132,f2129,f2126,f2090,f2087,f2055,f2046,f2043,f2037,f1929,f1928,f1712,f1710,f1709,f1623,f1622,f1536,f1511,f1493,f1477,f1220,f592,f473,f291,f290,f13276,f13273,f13441,f13453,f13462,f13474,f13476,f13479,f13491,f13503,f13505,f13508,f13534,f13540,f13683,f13685,f13686,f13687,f13688,f13689,f13691,f13692,f13693,f13694,f13695,f13696,f13697,f13698,f13699,f13581,f13582,f13583,f13584,f13746,f13765,f13767,f13770,f13771,f13762,f13781,f13797,f13799,f13802,f13803,f13795,f13837,f13850,f13852,f13855,f13929,f13942,f13944,f13947,f13948,f13941,f13994,f13988,f13995,f13997,f13998,f14056,f14059,f14116,f14119,f14141,f14142,f14143,f14144,f14145,f14176,f14170,f14178,f14228,f14227,f14226,f14222,f14221,f14246,f14398,f14401,f14403,f14263,f14265,f14266,f14268,f14269,f14271,f14272,f14274])).
fof(f14280,plain,(
  ( ! [X127,X125,X126,X124] : (sP5(sK14,'const.bool.IN_2'(s(X124,X125),s(fun(X124,bool),'const.pred_set.EMPTY_0')),happ(s(fun(X126,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X126,X127)))) ) | ($spl2 | $spl39 | $spl42)),
  inference(backward_demodulation,[],[f14246,f13371])).
fof(f14413,plain,(
  $false | ($spl2 | $spl39 | $spl42)),
  inference(global_subsumption,[],[f14279,f267,f266,f265,f268,f275,f287,f286,f332,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1691,f1700,f1717,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2010,f2040,f2044,f2048,f2084,f2088,f2092,f2125,f2130,f2170,f2174,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f3420,f3421,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f3503,f2972,f1715,f1714,f336,f335,f3502,f3501,f2971,f2970,f4473,f4472,f4467,f4466,f4465,f4464,f4453,f4452,f4451,f4450,f1827,f1800,f1771,f1741,f1610,f1593,f1578,f1562,f1496,f1480,f351,f350,f349,f348,f346,f344,f343,f341,f329,f328,f2973,f2974,f3504,f3505,f2002,f1538,f4474,f4463,f1539,f1515,f1514,f339,f338,f337,f2968,f3499,f2176,f2173,f2142,f2132,f2129,f2126,f2090,f2087,f2055,f2046,f2043,f2037,f1929,f1928,f1712,f1710,f1709,f1623,f1622,f1536,f1511,f1493,f1477,f1220,f592,f473,f291,f290,f13276,f13273,f13441,f13453,f13462,f13474,f13476,f13479,f13491,f13503,f13505,f13508,f13534,f13540,f13683,f13685,f13686,f13687,f13688,f13689,f13691,f13692,f13693,f13694,f13695,f13696,f13697,f13698,f13699,f13581,f13582,f13583,f13584,f13746,f13765,f13767,f13770,f13771,f13762,f13781,f13797,f13799,f13802,f13803,f13795,f13837,f13850,f13852,f13855,f13929,f13942,f13944,f13947,f13948,f13941,f13994,f13988,f13995,f13997,f13998,f14056,f14059,f14116,f14119,f14141,f14142,f14143,f14144,f14145,f14176,f14170,f14178,f14228,f14227,f14226,f14222,f14221,f14246,f14398,f14401,f14403,f14263,f14265,f14266,f14268,f14269,f14271,f14272,f14274])).
fof(f14279,plain,(
  ( ! [X123,X121,X122,X120] : (sP5(sK14,happ(s(fun(X122,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X122,X123)),'const.bool.IN_2'(s(X120,X121),s(fun(X120,bool),'const.pred_set.EMPTY_0')))) ) | ($spl2 | $spl39 | $spl42)),
  inference(backward_demodulation,[],[f14246,f13370])).
fof(f14412,plain,(
  $false | ($spl39 | $spl42)),
  inference(global_subsumption,[],[f267,f266,f265,f268,f275,f287,f286,f332,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1691,f1700,f1717,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2010,f2040,f2044,f2048,f2084,f2088,f2092,f2125,f2130,f2170,f2174,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f3420,f3421,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f3503,f2972,f1715,f1714,f336,f335,f3502,f3501,f2971,f2970,f4473,f4472,f4467,f4466,f4465,f4464,f4453,f4452,f4451,f4450,f1827,f1800,f1771,f1741,f1610,f1593,f1578,f1562,f1496,f1480,f351,f350,f349,f348,f346,f344,f343,f341,f329,f328,f2973,f2974,f3504,f3505,f2002,f1538,f4474,f4463,f1539,f1515,f1514,f339,f338,f337,f2968,f3499,f2176,f2173,f2142,f2132,f2129,f2126,f2090,f2087,f2055,f2046,f2043,f2037,f1929,f1928,f1712,f1710,f1709,f1623,f1622,f1536,f1511,f1493,f1477,f1220,f592,f473,f291,f290,f13276,f13273,f13441,f13453,f13462,f13474,f13476,f13479,f13491,f13503,f13505,f13508,f13534,f13540,f13683,f13685,f13686,f13687,f13688,f13689,f13691,f13692,f13693,f13694,f13695,f13696,f13697,f13698,f13699,f13581,f13582,f13583,f13584,f13746,f13765,f13767,f13770,f13771,f13762,f13781,f13797,f13799,f13802,f13803,f13795,f13837,f13850,f13852,f13855,f13929,f13942,f13944,f13947,f13948,f13941,f13994,f13988,f13995,f13997,f13998,f14056,f14059,f14116,f14119,f14141,f14142,f14143,f14144,f14145,f14176,f14170,f14178,f14228,f14227,f14226,f14222,f14221,f14246,f14398,f14401,f14403,f14263,f14265,f14266,f14268,f14269,f14271,f14272,f14274])).
fof(f14411,plain,(
  $false | ($spl39 | $spl42)),
  inference(global_subsumption,[],[f267,f266,f265,f268,f275,f287,f286,f332,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1691,f1700,f1717,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2010,f2040,f2044,f2048,f2084,f2088,f2092,f2125,f2130,f2170,f2174,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f3420,f3421,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f3503,f2972,f1715,f1714,f336,f335,f3502,f3501,f2971,f2970,f4473,f4472,f4467,f4466,f4465,f4464,f4453,f4452,f4451,f4450,f1827,f1800,f1771,f1741,f1610,f1593,f1578,f1562,f1496,f1480,f351,f350,f349,f348,f346,f344,f343,f341,f329,f328,f2973,f2974,f3504,f3505,f2002,f1538,f4474,f4463,f1539,f1515,f1514,f339,f338,f337,f2968,f3499,f2176,f2173,f2142,f2132,f2129,f2126,f2090,f2087,f2055,f2046,f2043,f2037,f1929,f1928,f1712,f1710,f1709,f1623,f1622,f1536,f1511,f1493,f1477,f1220,f592,f473,f291,f290,f13276,f13273,f13441,f13453,f13462,f13474,f13476,f13479,f13491,f13503,f13505,f13508,f13534,f13540,f13683,f13685,f13686,f13687,f13688,f13689,f13691,f13692,f13693,f13694,f13695,f13696,f13697,f13698,f13699,f13581,f13582,f13583,f13584,f13746,f13765,f13767,f13770,f13771,f13762,f13781,f13797,f13799,f13802,f13803,f13795,f13837,f13850,f13852,f13855,f13929,f13942,f13944,f13947,f13948,f13941,f13994,f13988,f13995,f13997,f13998,f14056,f14059,f14116,f14119,f14141,f14142,f14143,f14144,f14145,f14176,f14170,f14178,f14228,f14227,f14226,f14222,f14221,f14246,f14398,f14401,f14403,f14263,f14265,f14266,f14268,f14269,f14271,f14272])).
fof(f14410,plain,(
  $false | ($spl39 | $spl42)),
  inference(global_subsumption,[],[f267,f266,f265,f268,f275,f287,f286,f332,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1691,f1700,f1717,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2010,f2040,f2044,f2048,f2084,f2088,f2092,f2125,f2130,f2170,f2174,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f3420,f3421,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f3503,f2972,f1715,f1714,f336,f335,f3502,f3501,f2971,f2970,f4473,f4472,f4467,f4466,f4465,f4464,f4453,f4452,f4451,f4450,f1827,f1800,f1771,f1741,f1610,f1593,f1578,f1562,f1496,f1480,f351,f350,f349,f348,f346,f344,f343,f341,f329,f328,f2973,f2974,f3504,f3505,f2002,f1538,f4474,f4463,f1539,f1515,f1514,f339,f338,f337,f2968,f3499,f2176,f2173,f2142,f2132,f2129,f2126,f2090,f2087,f2055,f2046,f2043,f2037,f1929,f1928,f1712,f1710,f1709,f1623,f1622,f1536,f1511,f1493,f1477,f1220,f592,f473,f291,f290,f13276,f13273,f13441,f13453,f13462,f13474,f13476,f13479,f13491,f13503,f13505,f13508,f13534,f13540,f13683,f13685,f13686,f13687,f13688,f13689,f13691,f13692,f13693,f13694,f13695,f13696,f13697,f13698,f13699,f13581,f13582,f13583,f13584,f13746,f13765,f13767,f13770,f13771,f13762,f13781,f13797,f13799,f13802,f13803,f13795,f13837,f13850,f13852,f13855,f13929,f13942,f13944,f13947,f13948,f13941,f13994,f13988,f13995,f13997,f13998,f14056,f14059,f14116,f14119,f14141,f14142,f14143,f14144,f14145,f14176,f14170,f14178,f14228,f14227,f14226,f14222,f14221,f14246,f14398,f14401,f14403,f14263,f14265,f14266,f14268,f14269,f14271])).
fof(f14409,plain,(
  $false | ($spl39 | $spl42)),
  inference(global_subsumption,[],[f267,f266,f265,f268,f275,f287,f286,f332,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1691,f1700,f1717,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2010,f2040,f2044,f2048,f2084,f2088,f2092,f2125,f2130,f2170,f2174,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f3420,f3421,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f3503,f2972,f1715,f1714,f336,f335,f3502,f3501,f2971,f2970,f4473,f4472,f4467,f4466,f4465,f4464,f4453,f4452,f4451,f4450,f1827,f1800,f1771,f1741,f1610,f1593,f1578,f1562,f1496,f1480,f351,f350,f349,f348,f346,f344,f343,f341,f329,f328,f2973,f2974,f3504,f3505,f2002,f1538,f4474,f4463,f1539,f1515,f1514,f339,f338,f337,f2968,f3499,f2176,f2173,f2142,f2132,f2129,f2126,f2090,f2087,f2055,f2046,f2043,f2037,f1929,f1928,f1712,f1710,f1709,f1623,f1622,f1536,f1511,f1493,f1477,f1220,f592,f473,f291,f290,f13276,f13273,f13441,f13453,f13462,f13474,f13476,f13479,f13491,f13503,f13505,f13508,f13534,f13540,f13683,f13685,f13686,f13687,f13688,f13689,f13691,f13692,f13693,f13694,f13695,f13696,f13697,f13698,f13699,f13581,f13582,f13583,f13584,f13746,f13765,f13767,f13770,f13771,f13762,f13781,f13797,f13799,f13802,f13803,f13795,f13837,f13850,f13852,f13855,f13929,f13942,f13944,f13947,f13948,f13941,f13994,f13988,f13995,f13997,f13998,f14056,f14059,f14116,f14119,f14141,f14142,f14143,f14144,f14145,f14176,f14170,f14178,f14228,f14227,f14226,f14222,f14221,f14246,f14398,f14401,f14403,f14263,f14265,f14266,f14268,f14269])).
fof(f14408,plain,(
  $false | ($spl39 | $spl42)),
  inference(global_subsumption,[],[f267,f266,f265,f268,f275,f287,f286,f332,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1691,f1700,f1717,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2010,f2040,f2044,f2048,f2084,f2088,f2092,f2125,f2130,f2170,f2174,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f3420,f3421,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f3503,f2972,f1715,f1714,f336,f335,f3502,f3501,f2971,f2970,f4473,f4472,f4467,f4466,f4465,f4464,f4453,f4452,f4451,f4450,f1827,f1800,f1771,f1741,f1610,f1593,f1578,f1562,f1496,f1480,f351,f350,f349,f348,f346,f344,f343,f341,f329,f328,f2973,f2974,f3504,f3505,f2002,f1538,f4474,f4463,f1539,f1515,f1514,f339,f338,f337,f2968,f3499,f2176,f2173,f2142,f2132,f2129,f2126,f2090,f2087,f2055,f2046,f2043,f2037,f1929,f1928,f1712,f1710,f1709,f1623,f1622,f1536,f1511,f1493,f1477,f1220,f592,f473,f291,f290,f13276,f13273,f13441,f13453,f13462,f13474,f13476,f13479,f13491,f13503,f13505,f13508,f13534,f13540,f13683,f13685,f13686,f13687,f13688,f13689,f13691,f13692,f13693,f13694,f13695,f13696,f13697,f13698,f13699,f13581,f13582,f13583,f13584,f13746,f13765,f13767,f13770,f13771,f13762,f13781,f13797,f13799,f13802,f13803,f13795,f13837,f13850,f13852,f13855,f13929,f13942,f13944,f13947,f13948,f13941,f13994,f13988,f13995,f13997,f13998,f14056,f14059,f14116,f14119,f14141,f14142,f14143,f14144,f14145,f14176,f14170,f14178,f14228,f14227,f14226,f14222,f14221,f14246,f14398,f14401,f14403,f14263,f14265,f14266,f14268])).
fof(f14407,plain,(
  $false | ($spl39 | $spl42)),
  inference(global_subsumption,[],[f267,f266,f265,f268,f275,f287,f286,f332,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1691,f1700,f1717,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2010,f2040,f2044,f2048,f2084,f2088,f2092,f2125,f2130,f2170,f2174,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f3420,f3421,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f3503,f2972,f1715,f1714,f336,f335,f3502,f3501,f2971,f2970,f4473,f4472,f4467,f4466,f4465,f4464,f4453,f4452,f4451,f4450,f1827,f1800,f1771,f1741,f1610,f1593,f1578,f1562,f1496,f1480,f351,f350,f349,f348,f346,f344,f343,f341,f329,f328,f2973,f2974,f3504,f3505,f2002,f1538,f4474,f4463,f1539,f1515,f1514,f339,f338,f337,f2968,f3499,f2176,f2173,f2142,f2132,f2129,f2126,f2090,f2087,f2055,f2046,f2043,f2037,f1929,f1928,f1712,f1710,f1709,f1623,f1622,f1536,f1511,f1493,f1477,f1220,f592,f473,f291,f290,f13276,f13273,f13441,f13453,f13462,f13474,f13476,f13479,f13491,f13503,f13505,f13508,f13534,f13540,f13683,f13685,f13686,f13687,f13688,f13689,f13691,f13692,f13693,f13694,f13695,f13696,f13697,f13698,f13699,f13581,f13582,f13583,f13584,f13746,f13765,f13767,f13770,f13771,f13762,f13781,f13797,f13799,f13802,f13803,f13795,f13837,f13850,f13852,f13855,f13929,f13942,f13944,f13947,f13948,f13941,f13994,f13988,f13995,f13997,f13998,f14056,f14059,f14116,f14119,f14141,f14142,f14143,f14144,f14145,f14176,f14170,f14178,f14228,f14227,f14226,f14222,f14221,f14246,f14398,f14401,f14403,f14263,f14265,f14266])).
fof(f14406,plain,(
  $false | ($spl39 | $spl42)),
  inference(global_subsumption,[],[f267,f266,f265,f268,f275,f287,f286,f332,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1691,f1700,f1717,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2010,f2040,f2044,f2048,f2084,f2088,f2092,f2125,f2130,f2170,f2174,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f3420,f3421,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f3503,f2972,f1715,f1714,f336,f335,f3502,f3501,f2971,f2970,f4473,f4472,f4467,f4466,f4465,f4464,f4453,f4452,f4451,f4450,f1827,f1800,f1771,f1741,f1610,f1593,f1578,f1562,f1496,f1480,f351,f350,f349,f348,f346,f344,f343,f341,f329,f328,f2973,f2974,f3504,f3505,f2002,f1538,f4474,f4463,f1539,f1515,f1514,f339,f338,f337,f2968,f3499,f2176,f2173,f2142,f2132,f2129,f2126,f2090,f2087,f2055,f2046,f2043,f2037,f1929,f1928,f1712,f1710,f1709,f1623,f1622,f1536,f1511,f1493,f1477,f1220,f592,f473,f291,f290,f13276,f13273,f13441,f13453,f13462,f13474,f13476,f13479,f13491,f13503,f13505,f13508,f13534,f13540,f13683,f13685,f13686,f13687,f13688,f13689,f13691,f13692,f13693,f13694,f13695,f13696,f13697,f13698,f13699,f13581,f13582,f13583,f13584,f13746,f13765,f13767,f13770,f13771,f13762,f13781,f13797,f13799,f13802,f13803,f13795,f13837,f13850,f13852,f13855,f13929,f13942,f13944,f13947,f13948,f13941,f13994,f13988,f13995,f13997,f13998,f14056,f14059,f14116,f14119,f14141,f14142,f14143,f14144,f14145,f14176,f14170,f14178,f14228,f14227,f14226,f14222,f14221,f14246,f14398,f14401,f14403,f14263,f14265])).
fof(f14405,plain,(
  $false | ($spl39 | $spl42)),
  inference(global_subsumption,[],[f267,f266,f265,f268,f275,f287,f286,f332,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1691,f1700,f1717,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2010,f2040,f2044,f2048,f2084,f2088,f2092,f2125,f2130,f2170,f2174,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f3420,f3421,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f3503,f2972,f1715,f1714,f336,f335,f3502,f3501,f2971,f2970,f4473,f4472,f4467,f4466,f4465,f4464,f4453,f4452,f4451,f4450,f1827,f1800,f1771,f1741,f1610,f1593,f1578,f1562,f1496,f1480,f351,f350,f349,f348,f346,f344,f343,f341,f329,f328,f2973,f2974,f3504,f3505,f2002,f1538,f4474,f4463,f1539,f1515,f1514,f339,f338,f337,f2968,f3499,f2176,f2173,f2142,f2132,f2129,f2126,f2090,f2087,f2055,f2046,f2043,f2037,f1929,f1928,f1712,f1710,f1709,f1623,f1622,f1536,f1511,f1493,f1477,f1220,f592,f473,f291,f290,f13276,f13273,f13441,f13453,f13462,f13474,f13476,f13479,f13491,f13503,f13505,f13508,f13534,f13540,f13683,f13685,f13686,f13687,f13688,f13689,f13691,f13692,f13693,f13694,f13695,f13696,f13697,f13698,f13699,f13581,f13582,f13583,f13584,f13746,f13765,f13767,f13770,f13771,f13762,f13781,f13797,f13799,f13802,f13803,f13795,f13837,f13850,f13852,f13855,f13929,f13942,f13944,f13947,f13948,f13941,f13994,f13988,f13995,f13997,f13998,f14056,f14059,f14116,f14119,f14141,f14142,f14143,f14144,f14145,f14176,f14170,f14178,f14228,f14227,f14226,f14222,f14221,f14246,f14398,f14401,f14403,f14263])).
fof(f14404,plain,(
  $false | ($spl39 | $spl42)),
  inference(global_subsumption,[],[f267,f266,f265,f268,f275,f287,f286,f332,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1691,f1700,f1717,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2010,f2040,f2044,f2048,f2084,f2088,f2092,f2125,f2130,f2170,f2174,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f3420,f3421,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f3503,f2972,f1715,f1714,f336,f335,f3502,f3501,f2971,f2970,f4473,f4472,f4467,f4466,f4465,f4464,f4453,f4452,f4451,f4450,f1827,f1800,f1771,f1741,f1610,f1593,f1578,f1562,f1496,f1480,f351,f350,f349,f348,f346,f344,f343,f341,f329,f328,f2973,f2974,f3504,f3505,f2002,f1538,f4474,f4463,f1539,f1515,f1514,f339,f338,f337,f2968,f3499,f2176,f2173,f2142,f2132,f2129,f2126,f2090,f2087,f2055,f2046,f2043,f2037,f1929,f1928,f1712,f1710,f1709,f1623,f1622,f1536,f1511,f1493,f1477,f1220,f592,f473,f291,f290,f13276,f13273,f13441,f13453,f13462,f13474,f13476,f13479,f13491,f13503,f13505,f13508,f13534,f13540,f13683,f13685,f13686,f13687,f13688,f13689,f13691,f13692,f13693,f13694,f13695,f13696,f13697,f13698,f13699,f13581,f13582,f13583,f13584,f13746,f13765,f13767,f13770,f13771,f13762,f13781,f13797,f13799,f13802,f13803,f13795,f13837,f13850,f13852,f13855,f13929,f13942,f13944,f13947,f13948,f13941,f13994,f13988,f13995,f13997,f13998,f14056,f14059,f14116,f14119,f14141,f14142,f14143,f14144,f14145,f14176,f14170,f14178,f14228,f14227,f14226,f14222,f14221,f14246,f14398,f14401,f14403])).
fof(f14402,plain,(
  $false | ($spl39 | $spl42)),
  inference(global_subsumption,[],[f267,f266,f265,f268,f275,f287,f286,f332,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1691,f1700,f1717,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2010,f2040,f2044,f2048,f2084,f2088,f2092,f2125,f2130,f2170,f2174,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f3420,f3421,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f3503,f2972,f1715,f1714,f336,f335,f3502,f3501,f2971,f2970,f4473,f4472,f4467,f4466,f4465,f4464,f4453,f4452,f4451,f4450,f1827,f1800,f1771,f1741,f1610,f1593,f1578,f1562,f1496,f1480,f351,f350,f349,f348,f346,f344,f343,f341,f329,f328,f2973,f2974,f3504,f3505,f2002,f1538,f4474,f4463,f1539,f1515,f1514,f339,f338,f337,f2968,f3499,f2176,f2173,f2142,f2132,f2129,f2126,f2090,f2087,f2055,f2046,f2043,f2037,f1929,f1928,f1712,f1710,f1709,f1623,f1622,f1536,f1511,f1493,f1477,f1220,f592,f473,f291,f290,f13276,f13273,f13441,f13453,f13462,f13474,f13476,f13479,f13491,f13503,f13505,f13508,f13534,f13540,f13683,f13685,f13686,f13687,f13688,f13689,f13691,f13692,f13693,f13694,f13695,f13696,f13697,f13698,f13699,f13581,f13582,f13583,f13584,f13746,f13765,f13767,f13770,f13771,f13762,f13781,f13797,f13799,f13802,f13803,f13795,f13837,f13850,f13852,f13855,f13929,f13942,f13944,f13947,f13948,f13941,f13994,f13988,f13995,f13997,f13998,f14056,f14059,f14116,f14119,f14141,f14142,f14143,f14144,f14145,f14176,f14170,f14178,f14228,f14227,f14226,f14222,f14221,f14246,f14398,f14401])).
fof(f14400,plain,(
  $false | ($spl39 | $spl42)),
  inference(subsumption_resolution,[],[f14249,f268])).
fof(f14249,plain,(
  ~p(s(bool,'T_0')) | ($spl39 | $spl42)),
  inference(backward_demodulation,[],[f14246,f473])).
fof(f14399,plain,(
  $false | ($spl39 | $spl42)),
  inference(global_subsumption,[],[f267,f266,f265,f268,f275,f287,f286,f332,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1691,f1700,f1717,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2010,f2040,f2044,f2048,f2084,f2088,f2092,f2125,f2130,f2170,f2174,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f3420,f3421,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f3503,f2972,f1715,f1714,f336,f335,f3502,f3501,f2971,f2970,f4473,f4472,f4467,f4466,f4465,f4464,f4453,f4452,f4451,f4450,f1827,f1800,f1771,f1741,f1610,f1593,f1578,f1562,f1496,f1480,f351,f350,f349,f348,f346,f344,f343,f341,f329,f328,f2973,f2974,f3504,f3505,f2002,f1538,f4474,f4463,f1539,f1515,f1514,f339,f338,f337,f2968,f3499,f2176,f2173,f2142,f2132,f2129,f2126,f2090,f2087,f2055,f2046,f2043,f2037,f1929,f1928,f1712,f1710,f1709,f1623,f1622,f1536,f1511,f1493,f1477,f1220,f592,f473,f291,f290,f13276,f13273,f13441,f13453,f13462,f13474,f13476,f13479,f13491,f13503,f13505,f13508,f13534,f13540,f13683,f13685,f13686,f13687,f13688,f13689,f13691,f13692,f13693,f13694,f13695,f13696,f13697,f13698,f13699,f13581,f13582,f13583,f13584,f13746,f13765,f13767,f13770,f13771,f13762,f13781,f13797,f13799,f13802,f13803,f13795,f13837,f13850,f13852,f13855,f13929,f13942,f13944,f13947,f13948,f13941,f13994,f13988,f13995,f13997,f13998,f14056,f14059,f14116,f14119,f14141,f14142,f14143,f14144,f14145,f14176,f14170,f14178,f14228,f14227,f14226,f14222,f14221,f14246,f14398])).
fof(f14245,plain,(
  $false | ($spl39 | $spl45)),
  inference(subsumption_resolution,[],[f14203,f13278])).
fof(f14203,plain,(
  s(bool,'F_0') = s(bool,f(s(sK12,sK26(sK13,'const.pred_set.EMPTY_0',bool,sK12)))) | $spl39),
  inference(unit_resulting_resolution,[],[f13273,f275])).
fof(f14244,plain,(
  $false | ($spl39 | $spl45)),
  inference(subsumption_resolution,[],[f14205,f13278])).
fof(f14205,plain,(
  s(bool,'F_0') = s(bool,f(s(sK12,sK26(sK13,'const.pred_set.EMPTY_0',bool,sK12)))) | $spl39),
  inference(unit_resulting_resolution,[],[f515,f13273,f747])).
fof(f14243,plain,(
  $false | ($spl39 | $spl45)),
  inference(subsumption_resolution,[],[f14242,f13278])).
fof(f14242,plain,(
  s(bool,'F_0') = s(bool,f(s(sK12,sK26(sK13,'const.pred_set.EMPTY_0',bool,sK12)))) | $spl39),
  inference(forward_demodulation,[],[f14206,f592])).
fof(f14206,plain,(
  ( ! [X0,X1] : (s(bool,f(s(sK12,sK26(sK13,'const.pred_set.EMPTY_0',bool,sK12)))) = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | $spl39),
  inference(unit_resulting_resolution,[],[f591,f13273,f747])).
fof(f14241,plain,(
  $false | ($spl39 | $spl45)),
  inference(subsumption_resolution,[],[f14210,f13278])).
fof(f14210,plain,(
  s(bool,'F_0') = s(bool,f(s(sK12,sK26(sK13,'const.pred_set.EMPTY_0',bool,sK12)))) | $spl39),
  inference(unit_resulting_resolution,[],[f515,f13273,f747])).
fof(f14240,plain,(
  $false | ($spl39 | $spl45)),
  inference(subsumption_resolution,[],[f14239,f13278])).
fof(f14239,plain,(
  s(bool,'F_0') = s(bool,f(s(sK12,sK26(sK13,'const.pred_set.EMPTY_0',bool,sK12)))) | $spl39),
  inference(forward_demodulation,[],[f14211,f592])).
fof(f14211,plain,(
  ( ! [X0,X1] : (s(bool,f(s(sK12,sK26(sK13,'const.pred_set.EMPTY_0',bool,sK12)))) = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | $spl39),
  inference(unit_resulting_resolution,[],[f591,f13273,f747])).
fof(f14236,plain,(
  $false | ($spl39 | $spl45)),
  inference(subsumption_resolution,[],[f14215,f13278])).
fof(f14215,plain,(
  s(bool,'F_0') = s(bool,f(s(sK12,sK26(sK13,'const.pred_set.EMPTY_0',bool,sK12)))) | $spl39),
  inference(unit_resulting_resolution,[],[f473,f13273,f810])).
fof(f14151,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl76),
  inference(cnf_transformation,[],[f14151_D])).
fof(f14151_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f14147,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK26(X0,X1,bool,bool))) ) | $spl72),
  inference(cnf_transformation,[],[f14147_D])).
fof(f14147_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK26(X0,X1,bool,bool))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f14154,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl81),
  inference(cnf_transformation,[],[f14154_D])).
fof(f14154_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl81),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl81])])).
fof(f14145,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK26(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1622,f275])).
fof(f14152,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl78),
  inference(cnf_transformation,[],[f14152_D])).
fof(f14152_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f14144,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK26(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1622,f747])).
fof(f14143,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK26(X7,X8,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1622,f747])).
fof(f14146,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl70),
  inference(cnf_transformation,[],[f14146_D])).
fof(f14146_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f14149,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl75),
  inference(cnf_transformation,[],[f14149_D])).
fof(f14149_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl75),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl75])])).
fof(f14142,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK26(X3,X4,bool,bool)) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1622,f747])).
fof(f14141,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK26(X0,X1,bool,bool)) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1622,f747])).
fof(f13996,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X8,sK18(X10,X9,X8)),s(fun(X8,bool),X9))) | s(bool,'T_0') = s(bool,'const.bool.RES_FORALL_2'(s(fun(X8,bool),X9),s(fun(X8,bool),X10)))) ) | $spl68),
  inference(cnf_transformation,[],[f13996_D])).
fof(f13996_D,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X8,sK18(X10,X9,X8)),s(fun(X8,bool),X9))) | s(bool,'T_0') = s(bool,'const.bool.RES_FORALL_2'(s(fun(X8,bool),X9),s(fun(X8,bool),X10)))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f5183,plain,(
  ( ! [X15] : (p(s(bool,X15))) ) | $spl30),
  inference(cnf_transformation,[],[f5183_D])).
fof(f5183_D,plain,(
  ( ! [X15] : (p(s(bool,X15))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f13997,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X15)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X12,sK18(X14,X13,X12)),s(fun(X12,bool),X13))) | s(bool,'T_0') = s(bool,'const.bool.RES_FORALL_2'(s(fun(X12,bool),X13),s(fun(X12,bool),X14)))) )),
  inference(subsumption_resolution,[],[f13992,f286])).
fof(f13992,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X15)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X12,sK18(X14,X13,X12)),s(fun(X12,bool),X13))) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,'const.bool.RES_FORALL_2'(s(fun(X12,bool),X13),s(fun(X12,bool),X14)))) )),
  inference(superposition,[],[f1515,f747])).
fof(f13995,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X11)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X8,sK18(X10,X9,X8)),s(fun(X8,bool),X9))) | s(bool,'T_0') = s(bool,'const.bool.RES_FORALL_2'(s(fun(X8,bool),X9),s(fun(X8,bool),X10)))) )),
  inference(subsumption_resolution,[],[f13991,f286])).
fof(f13991,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X11)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X8,sK18(X10,X9,X8)),s(fun(X8,bool),X9))) | s(bool,'T_0') = s(bool,'const.bool.RES_FORALL_2'(s(fun(X8,bool),X9),s(fun(X8,bool),X10))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f1515,f747])).
fof(f13945,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,X1)) | ~sP11(happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2,'T_0') | p(s(bool,X2))) ) | $spl66),
  inference(cnf_transformation,[],[f13945_D])).
fof(f13945_D,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,X1)) | ~sP11(happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2,'T_0') | p(s(bool,X2))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f4415,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(cnf_transformation,[],[f4415_D])).
fof(f4415_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f13947,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11(happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2,'T_0')) )),
  inference(condensation,[],[f13946])).
fof(f13946,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X36)) | p(s(bool,X37)) | ~sP11(happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,X36))),s(X38,X39)),X37,'T_0') | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f13936,f810])).
fof(f13936,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | ~p(s(bool,X36)) | p(s(bool,X37)) | ~sP11(happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,X36))),s(X38,X39)),X37,'T_0') | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1236,f747])).
fof(f13944,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11(happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2,'T_0')) )),
  inference(condensation,[],[f13943])).
fof(f13943,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | ~sP11(happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,X30))),s(X32,X33)),X31,'T_0') | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f13935,f810])).
fof(f13935,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | p(s(bool,X31)) | ~sP11(happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,X30))),s(X32,X33)),X31,'T_0') | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f1236,f747])).
fof(f13853,plain,(
  ( ! [X4,X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | ~sP9(happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X3,X4)),X2,X1)) ) | $spl64),
  inference(cnf_transformation,[],[f13853_D])).
fof(f13853_D,plain,(
  ( ! [X4,X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | ~sP9(happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X3,X4)),X2,X1)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f13855,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9(happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X3,X4)),X2,X1)) )),
  inference(condensation,[],[f13854])).
fof(f13854,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (p(s(bool,X36)) | p(s(bool,X37)) | ~sP9(happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X38,X39)),X37,X36) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f13844,f810])).
fof(f13844,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | p(s(bool,X36)) | p(s(bool,X37)) | ~sP9(happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X38,X39)),X37,X36) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1235,f747])).
fof(f13852,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9(happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X3,X4)),X2,X1)) )),
  inference(condensation,[],[f13851])).
fof(f13851,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (p(s(bool,X30)) | p(s(bool,X31)) | ~sP9(happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X32,X33)),X31,X30) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f13843,f810])).
fof(f13843,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | p(s(bool,X30)) | p(s(bool,X31)) | ~sP9(happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X32,X33)),X31,X30) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f1235,f747])).
fof(f13800,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,'T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)))) ) | $spl62),
  inference(cnf_transformation,[],[f13800_D])).
fof(f13800_D,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,'T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f13802,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,'T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)))) )),
  inference(condensation,[],[f13801])).
fof(f13801,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X36)) | p(s(bool,X37)) | ~sP7(X37,'T_0',happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,X36))),s(X38,X39))) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f13790,f810])).
fof(f13790,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | ~p(s(bool,X36)) | p(s(bool,X37)) | ~sP7(X37,'T_0',happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,X36))),s(X38,X39))) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1234,f747])).
fof(f13799,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,'T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)))) )),
  inference(condensation,[],[f13798])).
fof(f13798,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | ~sP7(X31,'T_0',happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,X30))),s(X32,X33))) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f13789,f810])).
fof(f13789,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | p(s(bool,X31)) | ~sP7(X31,'T_0',happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,X30))),s(X32,X33))) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f1234,f747])).
fof(f13768,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,X1)) | ~sP4('T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2) | p(s(bool,X2))) ) | $spl60),
  inference(cnf_transformation,[],[f13768_D])).
fof(f13768_D,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,X1)) | ~sP4('T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2) | p(s(bool,X2))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f13770,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4('T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2)) )),
  inference(condensation,[],[f13769])).
fof(f13769,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X36)) | p(s(bool,X37)) | ~sP4('T_0',happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,X36))),s(X38,X39)),X37) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f13757,f810])).
fof(f13757,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | ~p(s(bool,X36)) | p(s(bool,X37)) | ~sP4('T_0',happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,X36))),s(X38,X39)),X37) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1233,f747])).
fof(f13767,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4('T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2)) )),
  inference(condensation,[],[f13766])).
fof(f13766,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | ~sP4('T_0',happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,X30))),s(X32,X33)),X31) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f13756,f810])).
fof(f13756,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | p(s(bool,X31)) | ~sP4('T_0',happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,X30))),s(X32,X33)),X31) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f1233,f747])).
fof(f13506,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | sP1(X1,happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)))) ) | $spl58),
  inference(cnf_transformation,[],[f13506_D])).
fof(f13506_D,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | sP1(X1,happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f13508,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)))) )),
  inference(condensation,[],[f13507])).
fof(f13507,plain,(
  ( ! [X30,X33,X31,X29,X32] : (p(s(bool,X30)) | sP1(X30,happ(s(fun(X31,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X31,X32))) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f13498,f810])).
fof(f13498,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | p(s(bool,X30)) | sP1(X30,happ(s(fun(X31,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X31,X32))) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f1232,f747])).
fof(f13505,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)))) )),
  inference(condensation,[],[f13504])).
fof(f13504,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,X25)) | sP1(X25,happ(s(fun(X26,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X26,X27))) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f13497,f810])).
fof(f13497,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | p(s(bool,X25)) | sP1(X25,happ(s(fun(X26,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X26,X27))) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f1232,f747])).
fof(f13477,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | sP1(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1)) ) | $spl56),
  inference(cnf_transformation,[],[f13477_D])).
fof(f13477_D,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | sP1(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f13479,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1)) )),
  inference(condensation,[],[f13478])).
fof(f13478,plain,(
  ( ! [X30,X33,X31,X29,X32] : (p(s(bool,X30)) | sP1(happ(s(fun(X31,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X31,X32)),X30) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f13469,f810])).
fof(f13469,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | p(s(bool,X30)) | sP1(happ(s(fun(X31,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X31,X32)),X30) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f1231,f747])).
fof(f13476,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1)) )),
  inference(condensation,[],[f13475])).
fof(f13475,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,X25)) | sP1(happ(s(fun(X26,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X26,X27)),X25) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f13468,f810])).
fof(f13468,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | p(s(bool,X25)) | sP1(happ(s(fun(X26,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X26,X27)),X25) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f1231,f747])).
fof(f13300,plain,(
  s(bool,'T_0') = s(bool,g(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12)))) | $spl50),
  inference(cnf_transformation,[],[f13300_D])).
fof(f13300_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12)))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f13303,plain,(
  s(bool,'F_0') != s(bool,f(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12)))) | $spl55),
  inference(cnf_transformation,[],[f13303_D])).
fof(f13303_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12)))) <=> ~$spl55),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl55])])).
fof(f13295,plain,(
  s(bool,'F_0') != s(bool,f(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12)))) | s(bool,'T_0') = s(bool,g(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12)))) | $spl5),
  inference(superposition,[],[f9875,f275])).
fof(f13301,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl52),
  inference(cnf_transformation,[],[f13301_D])).
fof(f13301_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12)))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f13294,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,f(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12)))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,g(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12))))) ) | $spl5),
  inference(superposition,[],[f9875,f747])).
fof(f13293,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12)))) | s(bool,'T_0') = s(bool,g(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl5),
  inference(superposition,[],[f9875,f747])).
fof(f13296,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl46),
  inference(cnf_transformation,[],[f13296_D])).
fof(f13296_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12)))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f13298,plain,(
  s(bool,'T_0') != s(bool,f(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12)))) | $spl49),
  inference(cnf_transformation,[],[f13298_D])).
fof(f13298_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12)))) <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f13292,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,f(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12)))) | s(bool,X1) = s(bool,g(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12)))) | s(bool,'T_0') = s(bool,X1)) ) | $spl5),
  inference(superposition,[],[f9875,f747])).
fof(f13291,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12)))) | s(bool,X0) = s(bool,g(s(sK12,sK26('const.pred_set.EMPTY_0',sK13,bool,sK12)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl5),
  inference(superposition,[],[f9875,f747])).
fof(f13275,plain,(
  s(bool,'T_0') = s(bool,g(s(sK12,sK26(sK13,'const.pred_set.EMPTY_0',bool,sK12)))) | $spl40),
  inference(cnf_transformation,[],[f13275_D])).
fof(f13275_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK12,sK26(sK13,'const.pred_set.EMPTY_0',bool,sK12)))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f13278,plain,(
  s(bool,'F_0') != s(bool,f(s(sK12,sK26(sK13,'const.pred_set.EMPTY_0',bool,sK12)))) | $spl45),
  inference(cnf_transformation,[],[f13278_D])).
fof(f13278_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK12,sK26(sK13,'const.pred_set.EMPTY_0',bool,sK12)))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f13270,plain,(
  s(bool,'F_0') != s(bool,f(s(sK12,sK26(sK13,'const.pred_set.EMPTY_0',bool,sK12)))) | s(bool,'T_0') = s(bool,g(s(sK12,sK26(sK13,'const.pred_set.EMPTY_0',bool,sK12)))) | $spl5),
  inference(superposition,[],[f9873,f275])).
fof(f13276,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(sK12,sK26(sK13,'const.pred_set.EMPTY_0',bool,sK12)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl42),
  inference(cnf_transformation,[],[f13276_D])).
fof(f13276_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(sK12,sK26(sK13,'const.pred_set.EMPTY_0',bool,sK12)))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f13269,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,f(s(sK12,sK26(sK13,'const.pred_set.EMPTY_0',bool,sK12)))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,g(s(sK12,sK26(sK13,'const.pred_set.EMPTY_0',bool,sK12))))) ) | $spl5),
  inference(superposition,[],[f9873,f747])).
fof(f13268,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(sK12,sK26(sK13,'const.pred_set.EMPTY_0',bool,sK12)))) | s(bool,'T_0') = s(bool,g(s(sK12,sK26(sK13,'const.pred_set.EMPTY_0',bool,sK12)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl5),
  inference(superposition,[],[f9873,f747])).
fof(f13271,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(sK12,sK26(sK13,'const.pred_set.EMPTY_0',bool,sK12)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl36),
  inference(cnf_transformation,[],[f13271_D])).
fof(f13271_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(sK12,sK26(sK13,'const.pred_set.EMPTY_0',bool,sK12)))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f13273,plain,(
  s(bool,'T_0') != s(bool,f(s(sK12,sK26(sK13,'const.pred_set.EMPTY_0',bool,sK12)))) | $spl39),
  inference(cnf_transformation,[],[f13273_D])).
fof(f13273_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK12,sK26(sK13,'const.pred_set.EMPTY_0',bool,sK12)))) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f13267,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,f(s(sK12,sK26(sK13,'const.pred_set.EMPTY_0',bool,sK12)))) | s(bool,X1) = s(bool,g(s(sK12,sK26(sK13,'const.pred_set.EMPTY_0',bool,sK12)))) | s(bool,'T_0') = s(bool,X1)) ) | $spl5),
  inference(superposition,[],[f9873,f747])).
fof(f13266,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK12,sK26(sK13,'const.pred_set.EMPTY_0',bool,sK12)))) | s(bool,X0) = s(bool,g(s(sK12,sK26(sK13,'const.pred_set.EMPTY_0',bool,sK12)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl5),
  inference(superposition,[],[f9873,f747])).
fof(f12530,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f12185,f488])).
fof(f12185,plain,(
  ~sP2('F_0','const.bool.IN_2'(s(sK12,sK17('const.pred_set.EMPTY_0',sK13,sK12)),s(fun(sK12,bool),sK13)),'T_0') | ($spl0 | $spl3 | $spl5)),
  inference(unit_resulting_resolution,[],[f11948,f485,f481,f12154,f388])).
fof(f12154,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f12153,f473])).
fof(f12153,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13))))) ) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f12152,f381])).
fof(f12152,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)))) | p(s(bool,happ(s(fun(sK12,bool),'_dst_x0_1'(s(bool,'F_0'))),s(sK12,X0))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f12035,f268])).
fof(f12035,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK13)))) | p(s(bool,happ(s(fun(sK12,bool),'_dst_x0_1'(s(bool,'F_0'))),s(sK12,X0))))) ) | ($spl0 | $spl3)),
  inference(superposition,[],[f337,f9886])).
fof(f9886,plain,(
  s(bool,'T_0') = s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'F_0'))))) | ($spl0 | $spl3)),
  inference(unit_resulting_resolution,[],[f268,f9851,f355])).
fof(f9851,plain,(
  p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'F_0')))))) | ($spl0 | $spl3)),
  inference(backward_demodulation,[],[f9850,f467])).
fof(f9850,plain,(
  s(bool,'F_0') = s(bool,sK14) | $spl3),
  inference(forward_demodulation,[],[f9842,f592])).
fof(f9842,plain,(
  ( ! [X0,X1] : (s(bool,sK14) = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f591,f470,f810])).
fof(f11948,plain,(
  ~sP5('T_0','const.bool.IN_2'(s(sK12,sK17('const.pred_set.EMPTY_0',sK13,sK12)),s(fun(sK12,bool),sK13)),'F_0') | $spl5),
  inference(unit_resulting_resolution,[],[f268,f9877,f404])).
fof(f9877,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK17('const.pred_set.EMPTY_0',sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl5),
  inference(forward_demodulation,[],[f9876,f592])).
fof(f9876,plain,(
  s(bool,'const.bool.IN_2'(s(sK12,sK17('const.pred_set.EMPTY_0',sK13,sK12)),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) != s(bool,'const.bool.IN_2'(s(sK12,sK17('const.pred_set.EMPTY_0',sK13,sK12)),s(fun(sK12,bool),sK13))) | $spl5),
  inference(unit_resulting_resolution,[],[f472,f336])).
fof(f12529,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f12186,f488])).
fof(f12186,plain,(
  ~sP2('F_0','const.bool.IN_2'(s(sK12,sK17('const.pred_set.EMPTY_0',sK13,sK12)),s(fun(sK12,bool),sK13)),'T_0') | ($spl0 | $spl3 | $spl5)),
  inference(unit_resulting_resolution,[],[f11948,f484,f481,f12154,f388])).
fof(f12528,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f12187,f481])).
fof(f12187,plain,(
  ~sP4('T_0','F_0','const.bool.IN_2'(s(sK12,sK17('const.pred_set.EMPTY_0',sK13,sK12)),s(fun(sK12,bool),sK13))) | ($spl0 | $spl3 | $spl5)),
  inference(unit_resulting_resolution,[],[f11948,f484,f488,f12154,f388])).
fof(f12527,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f12188,f481])).
fof(f12188,plain,(
  ~sP4('T_0','F_0','const.bool.IN_2'(s(sK12,sK17('const.pred_set.EMPTY_0',sK13,sK12)),s(fun(sK12,bool),sK13))) | ($spl0 | $spl3 | $spl5)),
  inference(unit_resulting_resolution,[],[f11948,f485,f488,f12154,f388])).
fof(f12526,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f12189,f484])).
fof(f12189,plain,(
  ~sP3('const.bool.IN_2'(s(sK12,sK17('const.pred_set.EMPTY_0',sK13,sK12)),s(fun(sK12,bool),sK13)),'F_0','T_0') | ($spl0 | $spl3 | $spl5)),
  inference(unit_resulting_resolution,[],[f11948,f481,f488,f12154,f388])).
fof(f12193,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(unit_resulting_resolution,[],[f11948,f484,f481,f488,f12154,f388])).
fof(f12191,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(unit_resulting_resolution,[],[f11948,f485,f481,f488,f12154,f388])).
fof(f9655,plain,(
  ( ! [X10,X11,X9] : (sP1(sK14,happ(s(fun(X10,bool),'_dst_x0_1'(s(bool,X9))),s(X10,X11))) | s(bool,'T_0') = s(bool,X9)) ) | $spl34),
  inference(cnf_transformation,[],[f9655_D])).
fof(f9655_D,plain,(
  ( ! [X10,X11,X9] : (sP1(sK14,happ(s(fun(X10,bool),'_dst_x0_1'(s(bool,X9))),s(X10,X11))) | s(bool,'T_0') = s(bool,X9)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f9656,plain,(
  ( ! [X14,X12,X15,X13] : (sP1(sK14,happ(s(fun(X14,bool),'_dst_x0_1'(s(bool,X13))),s(X14,X15))) | p(s(bool,X12)) | s(bool,'T_0') = s(bool,X13)) ) | $spl2),
  inference(subsumption_resolution,[],[f9633,f286])).
fof(f9633,plain,(
  ( ! [X14,X12,X15,X13] : (sP1(sK14,happ(s(fun(X14,bool),'_dst_x0_1'(s(bool,X13))),s(X14,X15))) | p(s(bool,X12)) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X12)) ) | $spl2),
  inference(superposition,[],[f8115,f747])).
fof(f8115,plain,(
  ( ! [X14,X15,X16] : (sP1(sK14,happ(s(fun(X14,bool),'_dst_x0_1'(s(bool,X15))),s(X14,X16))) | p(s(bool,X15))) ) | $spl2),
  inference(superposition,[],[f7871,f381])).
fof(f7871,plain,(
  ( ! [X17] : (p(s(bool,X17)) | sP1(sK14,X17)) ) | $spl2),
  inference(subsumption_resolution,[],[f7786,f268])).
fof(f7786,plain,(
  ( ! [X17] : (~p(s(bool,'T_0')) | p(s(bool,X17)) | sP1(sK14,X17)) ) | $spl2),
  inference(superposition,[],[f373,f6790])).
fof(f6790,plain,(
  s(bool,'T_0') = s(bool,sK14) | $spl2),
  inference(unit_resulting_resolution,[],[f268,f469,f355])).
fof(f9654,plain,(
  ( ! [X10,X8,X11,X9] : (sP1(sK14,happ(s(fun(X10,bool),'_dst_x0_1'(s(bool,X9))),s(X10,X11))) | p(s(bool,X8)) | s(bool,'T_0') = s(bool,X9)) ) | $spl2),
  inference(subsumption_resolution,[],[f9632,f286])).
fof(f9632,plain,(
  ( ! [X10,X8,X11,X9] : (sP1(sK14,happ(s(fun(X10,bool),'_dst_x0_1'(s(bool,X9))),s(X10,X11))) | p(s(bool,X8)) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X9)) ) | $spl2),
  inference(superposition,[],[f8115,f747])).
fof(f6785,plain,(
  $false | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f6773,f2002])).
fof(f6773,plain,(
  ~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),'const.pred_set.EMPTY_0'),s(fun(sK12,bool),'_dst_x0_1'(s(bool,sK14)))))) | ($spl1 | $spl4)),
  inference(backward_demodulation,[],[f471,f468])).
fof(f6770,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5228,f5245])).
fof(f5245,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f5183,f5183,f355])).
fof(f5228,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f5183,f290])).
fof(f6769,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5232,f268])).
fof(f5232,plain,(
  ~p(s(bool,'T_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f515,f5183,f355])).
fof(f6768,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5234,f268])).
fof(f5234,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f532,f5183,f355])).
fof(f532,plain,(
  s(bool,'T_0') != s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0'))))) | ($spl1 | $spl2)),
  inference(backward_demodulation,[],[f518,f516])).
fof(f516,plain,(
  s(bool,'T_0') != s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,sK14))))) | $spl1),
  inference(unit_resulting_resolution,[],[f468,f286])).
fof(f518,plain,(
  s(bool,'T_0') = s(bool,sK14) | $spl2),
  inference(unit_resulting_resolution,[],[f469,f287])).
fof(f6767,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5236,f268])).
fof(f5236,plain,(
  ~p(s(bool,'T_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f591,f5183,f355])).
fof(f6766,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5240,f5183])).
fof(f5240,plain,(
  ~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))))) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f532,f5183,f355])).
fof(f6765,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5242,f5183])).
fof(f5242,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f591,f5183,f355])).
fof(f6764,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5246,f268])).
fof(f5246,plain,(
  ~p(s(bool,'T_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f515,f5183,f355])).
fof(f6763,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5248,f268])).
fof(f5248,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f532,f5183,f355])).
fof(f6762,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5250,f268])).
fof(f5250,plain,(
  ~p(s(bool,'T_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f591,f5183,f355])).
fof(f6761,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5254,f5183])).
fof(f5254,plain,(
  ~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))))) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f532,f5183,f355])).
fof(f6760,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5256,f5183])).
fof(f5256,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f591,f5183,f355])).
fof(f6759,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5262,f268])).
fof(f5262,plain,(
  ~p(s(bool,'T_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f573,f5183,f372])).
fof(f6758,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5263,f268])).
fof(f5263,plain,(
  ~p(s(bool,'T_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f598,f5183,f372])).
fof(f6757,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5266,f5183])).
fof(f5266,plain,(
  ~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))))) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f660,f5183,f372])).
fof(f660,plain,(
  sP1(sK14,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0'))))) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f655,f473])).
fof(f655,plain,(
  p(s(bool,'F_0')) | sP1(sK14,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0'))))) | ($spl1 | $spl2)),
  inference(superposition,[],[f566,f556])).
fof(f556,plain,(
  s(bool,'F_0') = s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0'))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f521,f291])).
fof(f521,plain,(
  ~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))))) | ($spl1 | $spl2)),
  inference(backward_demodulation,[],[f518,f468])).
fof(f566,plain,(
  ( ! [X0] : (p(s(bool,X0)) | sP1(sK14,X0)) ) | $spl2),
  inference(subsumption_resolution,[],[f565,f268])).
fof(f565,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP1(sK14,X0)) ) | $spl2),
  inference(superposition,[],[f373,f518])).
fof(f6756,plain,(
  $false | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5267,f5183])).
fof(f5267,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f616,f5183,f372])).
fof(f616,plain,(
  ( ! [X0,X1] : (sP1(sK14,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f332,f566])).
fof(f6755,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5268,f5183])).
fof(f5268,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f597,f5183,f372])).
fof(f6754,plain,(
  $false | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5269,f5183])).
fof(f5269,plain,(
  ~p(s(bool,sK14)) | ($spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f589,f5183,f372])).
fof(f589,plain,(
  sP1('F_0',sK14) | $spl2),
  inference(unit_resulting_resolution,[],[f473,f575])).
fof(f575,plain,(
  ( ! [X0] : (p(s(bool,X0)) | sP1(X0,sK14)) ) | $spl2),
  inference(subsumption_resolution,[],[f574,f268])).
fof(f574,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP1(X0,sK14)) ) | $spl2),
  inference(superposition,[],[f376,f518])).
fof(f6753,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5270,f5183])).
fof(f5270,plain,(
  ~p(s(bool,sK14)) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f661,f5183,f372])).
fof(f661,plain,(
  sP1('const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),sK14) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f656,f473])).
fof(f656,plain,(
  p(s(bool,'F_0')) | sP1('const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),sK14) | ($spl1 | $spl2)),
  inference(superposition,[],[f575,f556])).
fof(f6752,plain,(
  $false | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5271,f5183])).
fof(f5271,plain,(
  ~p(s(bool,sK14)) | ($spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f617,f5183,f372])).
fof(f617,plain,(
  ( ! [X0,X1] : (sP1('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),sK14)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f332,f575])).
fof(f6751,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5277,f5183])).
fof(f5277,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f598,f5183,f372])).
fof(f6750,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5278,f268])).
fof(f5278,plain,(
  ~p(s(bool,'T_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f564,f5183,f372])).
fof(f6749,plain,(
  $false | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5279,f5183])).
fof(f5279,plain,(
  ~p(s(bool,sK14)) | ($spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f569,f5183,f372])).
fof(f569,plain,(
  sP1(sK14,'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f473,f566])).
fof(f6748,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5280,f5183])).
fof(f5280,plain,(
  ~p(s(bool,sK14)) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f660,f5183,f372])).
fof(f6747,plain,(
  $false | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5281,f5183])).
fof(f5281,plain,(
  ~p(s(bool,sK14)) | ($spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f616,f5183,f372])).
fof(f6746,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5282,f268])).
fof(f5282,plain,(
  ~p(s(bool,'T_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f597,f5183,f372])).
fof(f6745,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5284,f5183])).
fof(f5284,plain,(
  ~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))))) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f661,f5183,f372])).
fof(f6744,plain,(
  $false | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5285,f5183])).
fof(f5285,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f617,f5183,f372])).
fof(f6743,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5305,f5183])).
fof(f5305,plain,(
  ~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))))) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f863,f473,f5183,f386])).
fof(f863,plain,(
  sP5('T_0','const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),'F_0') | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f268,f556,f403])).
fof(f6742,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5306,f5183])).
fof(f5306,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X0,X1))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1207,f473,f5183,f386])).
fof(f6741,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5307,f5183])).
fof(f5307,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1915,f473,f5183,f386])).
fof(f6740,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5308,f268])).
fof(f5308,plain,(
  ~p(s(bool,'T_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f854,f473,f5183,f386])).
fof(f6739,plain,(
  $false | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5309,f5183])).
fof(f5309,plain,(
  ~p(s(bool,sK14)) | ($spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f1438,f473,f5183,f386])).
fof(f1438,plain,(
  sP5('F_0',sK14,'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f485,f481,f489,f473,f389])).
fof(f489,plain,(
  ( ! [X0,X1] : (sP2(X0,sK14,X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f469,f400])).
fof(f6738,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5404,f268])).
fof(f5404,plain,(
  ~p(s(bool,'T_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f1381,f473,f5183,f386])).
fof(f6737,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5405,f268])).
fof(f5405,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f863,f473,f5183,f386])).
fof(f6736,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5406,f268])).
fof(f5406,plain,(
  ~p(s(bool,'T_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f1207,f473,f5183,f386])).
fof(f6735,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5407,f268])).
fof(f5407,plain,(
  ~p(s(bool,'T_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f1915,f473,f5183,f386])).
fof(f6734,plain,(
  $false | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5410,f5183])).
fof(f5410,plain,(
  ~p(s(bool,sK14)) | ($spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f1372,f473,f5183,f386])).
fof(f1372,plain,(
  sP5(sK14,'F_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f486,f481,f487,f473,f388])).
fof(f487,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,sK14)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f469,f399])).
fof(f6733,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5415,f5760])).
fof(f5760,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f5183,f401])).
fof(f5415,plain,(
  ~sP2('F_0','T_0','T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f878,f484,f480,f5183,f387])).
fof(f6732,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5416,f5760])).
fof(f5416,plain,(
  ~sP2('const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),'T_0','T_0') | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f879,f484,f480,f5183,f387])).
fof(f879,plain,(
  ~sP5('T_0','T_0','const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0'))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f268,f532,f404])).
fof(f6731,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5417,f5760])).
fof(f5417,plain,(
  ( ! [X0,X1] : (~sP2('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0','T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1856,f484,f480,f5183,f387])).
fof(f6730,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5418,f5760])).
fof(f5418,plain,(
  ~sP2('F_0','T_0','T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f878,f485,f480,f5183,f387])).
fof(f6729,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5419,f5760])).
fof(f5419,plain,(
  ~sP2('const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),'T_0','T_0') | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f879,f525,f480,f5183,f387])).
fof(f525,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),X1)) ) | ($spl1 | $spl2)),
  inference(backward_demodulation,[],[f518,f508])).
fof(f508,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,sK14)))),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f468,f396])).
fof(f6728,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5420,f5760])).
fof(f5420,plain,(
  ( ! [X0,X1] : (~sP2('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0','T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1856,f601,f480,f5183,f387])).
fof(f6727,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5421,f5760])).
fof(f5421,plain,(
  ~sP2('F_0','T_0','T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f878,f485,f481,f5183,f387])).
fof(f6726,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5426,f5760])).
fof(f5426,plain,(
  ~sP2('F_0','T_0','T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f878,f484,f481,f5183,f387])).
fof(f6725,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5430,f5760])).
fof(f5430,plain,(
  ~sP2('const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),'T_0','T_0') | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f879,f525,f523,f5183,f387])).
fof(f523,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),X1)) ) | ($spl1 | $spl2)),
  inference(backward_demodulation,[],[f518,f506])).
fof(f506,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,sK14)))),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f468,f392])).
fof(f6724,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5432,f5760])).
fof(f5432,plain,(
  ( ! [X0,X1] : (~sP2('const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0'))))))),s(X0,X1)),'F_0')) ) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f1206,f525,f523,f5183,f387])).
fof(f6723,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5433,f5760])).
fof(f5433,plain,(
  ~sP2('const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),'T_0','T_0') | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f879,f484,f523,f5183,f387])).
fof(f6722,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5435,f5760])).
fof(f5435,plain,(
  ( ! [X0,X1] : (~sP2('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0','T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1856,f601,f599,f5183,f387])).
fof(f6721,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5437,f5760])).
fof(f5437,plain,(
  ( ! [X2,X0,X3,X1] : (~sP2('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),s(X2,X3)),'F_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1206,f601,f599,f5183,f387])).
fof(f6720,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5438,f5760])).
fof(f5438,plain,(
  ( ! [X0,X1] : (~sP2('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0','T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1856,f484,f599,f5183,f387])).
fof(f6719,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5445,f5760])).
fof(f5445,plain,(
  ( ! [X0,X1] : (~sP2('const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0'))))))),s(X0,X1)),'F_0')) ) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f1206,f525,f482,f5183,f387])).
fof(f6718,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5447,f5760])).
fof(f5447,plain,(
  ( ! [X2,X0,X3,X1] : (~sP2('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),s(X2,X3)),'F_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1206,f601,f482,f5183,f387])).
fof(f6717,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5453,f5760])).
fof(f5453,plain,(
  ( ! [X0,X1] : (~sP2(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0'))))))),s(X0,X1)),'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),'F_0')) ) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f1208,f526,f482,f5183,f387])).
fof(f526,plain,(
  ( ! [X0,X1] : (sP3('const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),X0,X1)) ) | ($spl1 | $spl2)),
  inference(backward_demodulation,[],[f518,f509])).
fof(f509,plain,(
  ( ! [X0,X1] : (sP3('const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,sK14)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f468,f397])).
fof(f6716,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5455,f5760])).
fof(f5455,plain,(
  ( ! [X2,X0,X3,X1] : (~sP2(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0'))))),s(X0,X3)),'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')),'F_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1208,f602,f482,f5183,f387])).
fof(f6715,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5479,f5706])).
fof(f5706,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f5183,f391])).
fof(f5479,plain,(
  ~sP4('T_0','F_0','T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f878,f484,f488,f5183,f387])).
fof(f6714,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5480,f5706])).
fof(f5480,plain,(
  ~sP4('T_0','const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),'T_0') | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f879,f484,f488,f5183,f387])).
fof(f6713,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5481,f5706])).
fof(f5481,plain,(
  ( ! [X0,X1] : (~sP4('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1856,f484,f488,f5183,f387])).
fof(f6712,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5486,f5706])).
fof(f5486,plain,(
  ~sP4('T_0','F_0','T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f878,f485,f488,f5183,f387])).
fof(f6711,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5487,f5706])).
fof(f5487,plain,(
  ~sP4('T_0','const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),'T_0') | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f879,f525,f488,f5183,f387])).
fof(f6710,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5488,f5706])).
fof(f5488,plain,(
  ( ! [X0,X1] : (~sP4('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1856,f601,f488,f5183,f387])).
fof(f6709,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5498,f5754])).
fof(f5754,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f5183,f395])).
fof(f5498,plain,(
  ~sP3('T_0','F_0','T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f878,f480,f488,f5183,f387])).
fof(f6708,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5499,f5754])).
fof(f5499,plain,(
  ~sP3('T_0','const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),'T_0') | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f879,f480,f488,f5183,f387])).
fof(f6707,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5500,f5754])).
fof(f5500,plain,(
  ( ! [X0,X1] : (~sP3('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1856,f480,f488,f5183,f387])).
fof(f6706,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5509,f5754])).
fof(f5509,plain,(
  ~sP3('T_0','F_0','T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f878,f481,f488,f5183,f387])).
fof(f6705,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5517,f5754])).
fof(f5517,plain,(
  ~sP3('T_0','const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),'T_0') | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f879,f523,f488,f5183,f387])).
fof(f6704,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5525,f5754])).
fof(f5525,plain,(
  ( ! [X0,X1] : (~sP3('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1856,f599,f488,f5183,f387])).
fof(f6703,plain,(
  $false | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5535,f5754])).
fof(f5535,plain,(
  ( ! [X0,X1] : (~sP3(sK14,happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,sK14))),s(X0,X1)),'F_0')) ) | ($spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f1208,f479,f489,f5183,f387])).
fof(f479,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,sK14)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f469,f391])).
fof(f6702,plain,(
  $false | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5552,f5754])).
fof(f5552,plain,(
  ( ! [X0,X1] : (~sP3(sK14,happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,sK14))),s(X0,X1)),'F_0')) ) | ($spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f1208,f482,f489,f5183,f387])).
fof(f6701,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5562,f5706])).
fof(f5562,plain,(
  ~sP4('T_0','F_0','T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f878,f484,f490,f5183,f387])).
fof(f6700,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5563,f5706])).
fof(f5563,plain,(
  ~sP4('T_0','const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),'T_0') | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f879,f484,f490,f5183,f387])).
fof(f6699,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5564,f5706])).
fof(f5564,plain,(
  ( ! [X0,X1] : (~sP4('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1856,f484,f490,f5183,f387])).
fof(f6698,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5565,f5706])).
fof(f5565,plain,(
  ~sP4('T_0','F_0','T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f878,f485,f490,f5183,f387])).
fof(f6697,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5566,f5706])).
fof(f5566,plain,(
  ~sP4('T_0','const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),'T_0') | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f879,f525,f490,f5183,f387])).
fof(f6696,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5567,f5706])).
fof(f5567,plain,(
  ( ! [X0,X1] : (~sP4('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1856,f601,f490,f5183,f387])).
fof(f6695,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5568,f5754])).
fof(f5568,plain,(
  ~sP3('T_0','F_0','T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f878,f480,f490,f5183,f387])).
fof(f6694,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5569,f5754])).
fof(f5569,plain,(
  ~sP3('T_0','const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),'T_0') | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f879,f480,f490,f5183,f387])).
fof(f6693,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5570,f5754])).
fof(f5570,plain,(
  ( ! [X0,X1] : (~sP3('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1856,f480,f490,f5183,f387])).
fof(f6692,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5577,f5824])).
fof(f5824,plain,(
  ( ! [X0] : (~sP5(X0,'T_0','F_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f515,f5183,f404])).
fof(f5577,plain,(
  ( ! [X0] : (sP5(X0,'T_0','F_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f485,f480,f490,f5183,f387])).
fof(f6691,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5579,f5826])).
fof(f5826,plain,(
  ( ! [X0] : (~sP5(X0,'T_0','const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))))) ) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f532,f5183,f404])).
fof(f5579,plain,(
  ( ! [X0] : (sP5(X0,'T_0','const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))))) ) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f525,f480,f490,f5183,f387])).
fof(f6690,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5581,f5828])).
fof(f5828,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,'T_0','const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f591,f5183,f404])).
fof(f5581,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,'T_0','const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f601,f480,f490,f5183,f387])).
fof(f6689,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5583,f5754])).
fof(f5583,plain,(
  ~sP3('T_0','F_0','T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f878,f481,f490,f5183,f387])).
fof(f6688,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5584,f5824])).
fof(f5584,plain,(
  ( ! [X0] : (sP5(X0,'T_0','F_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f485,f481,f490,f5183,f387])).
fof(f6687,plain,(
  $false | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5586,f5824])).
fof(f5586,plain,(
  sP5(sK14,'T_0','F_0') | ($spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f483,f481,f490,f5183,f387])).
fof(f483,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,sK14)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f469,f395])).
fof(f6686,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5587,f878])).
fof(f5587,plain,(
  sP5('T_0','T_0','F_0') | $spl30),
  inference(unit_resulting_resolution,[],[f484,f481,f490,f5183,f387])).
fof(f6685,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5589,f5754])).
fof(f5589,plain,(
  ~sP3('T_0','const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),'T_0') | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f879,f523,f490,f5183,f387])).
fof(f6684,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5590,f5826])).
fof(f5590,plain,(
  ( ! [X0] : (sP5(X0,'T_0','const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))))) ) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f525,f523,f490,f5183,f387])).
fof(f6683,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5592,f5826])).
fof(f5592,plain,(
  sP5(sK14,'T_0','const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0'))))) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f483,f523,f490,f5183,f387])).
fof(f6682,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5593,f879])).
fof(f5593,plain,(
  sP5('T_0','T_0','const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0'))))) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f484,f523,f490,f5183,f387])).
fof(f6681,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5595,f5754])).
fof(f5595,plain,(
  ( ! [X0,X1] : (~sP3('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1856,f599,f490,f5183,f387])).
fof(f6680,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5596,f5828])).
fof(f5596,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,'T_0','const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f601,f599,f490,f5183,f387])).
fof(f6679,plain,(
  $false | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5598,f5828])).
fof(f5598,plain,(
  ( ! [X0,X1] : (sP5(sK14,'T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | ($spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f483,f599,f490,f5183,f387])).
fof(f6678,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5599,f1856])).
fof(f5599,plain,(
  ( ! [X0,X1] : (sP5('T_0','T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f484,f599,f490,f5183,f387])).
fof(f6677,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5607,f5824])).
fof(f5607,plain,(
  sP5('const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),'T_0','F_0') | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f485,f524,f490,f5183,f387])).
fof(f524,plain,(
  ( ! [X0,X1] : (sP4('const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),X0,X1)) ) | ($spl1 | $spl2)),
  inference(backward_demodulation,[],[f518,f507])).
fof(f507,plain,(
  ( ! [X0,X1] : (sP4('const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,sK14)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f468,f393])).
fof(f6676,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5608,f5826])).
fof(f5608,plain,(
  sP5('const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),'T_0','const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0'))))) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f525,f524,f490,f5183,f387])).
fof(f6675,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5609,f5828])).
fof(f5609,plain,(
  ( ! [X0,X1] : (sP5('const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),'T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f601,f524,f490,f5183,f387])).
fof(f6674,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5610,f5824])).
fof(f5610,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0','F_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f485,f600,f490,f5183,f387])).
fof(f6673,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5611,f5826])).
fof(f5611,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0','const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))))) ) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f525,f600,f490,f5183,f387])).
fof(f6672,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5612,f5828])).
fof(f5612,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0','const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.EMPTY_0')))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f601,f600,f490,f5183,f387])).
fof(f6671,plain,(
  $false | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5623,f5754])).
fof(f5623,plain,(
  ( ! [X0,X1] : (~sP3(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,sK14))),s(X0,X1)),sK14,'F_0')) ) | ($spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f1206,f482,f491,f5183,f387])).
fof(f491,plain,(
  ( ! [X0,X1] : (sP2(sK14,X0,X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f469,f401])).
fof(f6670,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5658,f5183])).
fof(f5658,plain,(
  ~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))))) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f523,f473,f5183,f390])).
fof(f6669,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5659,f5183])).
fof(f5659,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f599,f473,f5183,f390])).
fof(f6668,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5660,f5183])).
fof(f5660,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f482,f473,f5183,f390])).
fof(f6667,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5661,f5183])).
fof(f5661,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f524,f473,f5183,f390])).
fof(f6666,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5662,f5183])).
fof(f5662,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f600,f473,f5183,f390])).
fof(f6665,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5668,f5706])).
fof(f5668,plain,(
  ( ! [X0] : (~sP4(X0,'T_0','F_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f473,f268,f5183,f390])).
fof(f6664,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5679,f5706])).
fof(f5679,plain,(
  ( ! [X0,X1] : (~sP4(X0,X1,'F_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f473,f5183,f5183,f390])).
fof(f6663,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5681,f5183])).
fof(f5681,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f481,f473,f5183,f390])).
fof(f6662,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5682,f5183])).
fof(f5682,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f523,f473,f5183,f390])).
fof(f6661,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5683,f5183])).
fof(f5683,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f599,f473,f5183,f390])).
fof(f6660,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5685,f5183])).
fof(f5685,plain,(
  ~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))))) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f524,f473,f5183,f390])).
fof(f6659,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5686,f5183])).
fof(f5686,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f600,f473,f5183,f390])).
fof(f6658,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5693,f5706])).
fof(f5693,plain,(
  ( ! [X0] : (~sP4('T_0',X0,'F_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f473,f268,f5183,f390])).
fof(f6657,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5704,f5706])).
fof(f5704,plain,(
  ( ! [X0,X1] : (~sP4(X0,X1,'F_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f473,f5183,f5183,f390])).
fof(f6656,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5709,f5183])).
fof(f5709,plain,(
  ~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))))) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f525,f473,f5183,f394])).
fof(f6655,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5710,f5183])).
fof(f5710,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f601,f473,f5183,f394])).
fof(f6654,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5711,f5183])).
fof(f5711,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f486,f473,f5183,f394])).
fof(f6653,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5712,f5183])).
fof(f5712,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f526,f473,f5183,f394])).
fof(f6652,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5713,f5183])).
fof(f5713,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f602,f473,f5183,f394])).
fof(f6651,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5719,f5754])).
fof(f5719,plain,(
  ( ! [X0] : (~sP3(X0,'T_0','F_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f473,f268,f5183,f394])).
fof(f6650,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5729,f5754])).
fof(f5729,plain,(
  ( ! [X0,X1] : (~sP3(X0,X1,'F_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f473,f5183,f5183,f394])).
fof(f6649,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5744,f5183])).
fof(f5744,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f485,f473,f5183,f394])).
fof(f6648,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5745,f5183])).
fof(f5745,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f525,f473,f5183,f394])).
fof(f6647,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5746,f5183])).
fof(f5746,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f601,f473,f5183,f394])).
fof(f6646,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5748,f5183])).
fof(f5748,plain,(
  ~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))))) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f526,f473,f5183,f394])).
fof(f6645,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5749,f5183])).
fof(f5749,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f602,f473,f5183,f394])).
fof(f6644,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5750,f5754])).
fof(f5750,plain,(
  ( ! [X0] : (~sP3('T_0',X0,'F_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f268,f473,f5183,f394])).
fof(f6643,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5752,f5754])).
fof(f5752,plain,(
  ( ! [X0,X1] : (~sP3(X0,X1,'F_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f5183,f473,f5183,f394])).
fof(f6641,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5769,f381])).
fof(f5769,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,X0))),s(X1,X2)))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1208,f5183,f403])).
fof(f6640,plain,(
  $false | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5770,f518])).
fof(f5770,plain,(
  s(bool,'T_0') != s(bool,sK14) | ($spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f889,f5183,f403])).
fof(f889,plain,(
  ~sP5('F_0','T_0',sK14) | $spl2),
  inference(unit_resulting_resolution,[],[f473,f518,f405])).
fof(f6639,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5773,f556])).
fof(f5773,plain,(
  s(bool,'F_0') != s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0'))))) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f891,f5183,f403])).
fof(f891,plain,(
  ~sP5('F_0','F_0','const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0'))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f473,f556,f405])).
fof(f6638,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5774,f592])).
fof(f5774,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1916,f5183,f403])).
fof(f6637,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5775,f556])).
fof(f5775,plain,(
  s(bool,'F_0') != s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0'))))) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f888,f5183,f403])).
fof(f888,plain,(
  ~sP5('F_0','const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),'F_0') | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f473,f556,f405])).
fof(f6636,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5776,f381])).
fof(f5776,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,X0))),s(X1,X2)))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1206,f5183,f403])).
fof(f6635,plain,(
  $false | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f6634,f518])).
fof(f6634,plain,(
  s(bool,'T_0') != s(bool,sK14) | ($spl2 | $spl30)),
  inference(forward_demodulation,[],[f5777,f381])).
fof(f5777,plain,(
  ( ! [X0,X1] : (s(bool,sK14) != s(bool,happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)))) ) | ($spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f2047,f5183,f403])).
fof(f2047,plain,(
  ( ! [X12,X13] : (~sP5('F_0',happ(s(fun(X12,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X12,X13)),sK14)) ) | $spl2),
  inference(superposition,[],[f1206,f518])).
fof(f6633,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5778,f592])).
fof(f5778,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1914,f5183,f403])).
fof(f6632,plain,(
  $false | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5779,f518])).
fof(f5779,plain,(
  s(bool,'T_0') != s(bool,sK14) | ($spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f886,f5183,f403])).
fof(f886,plain,(
  ~sP5('F_0',sK14,'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f473,f518,f405])).
fof(f6631,plain,(
  $false | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f6630,f518])).
fof(f6630,plain,(
  s(bool,'T_0') != s(bool,sK14) | ($spl2 | $spl30)),
  inference(forward_demodulation,[],[f5780,f381])).
fof(f5780,plain,(
  ( ! [X0,X1] : (s(bool,sK14) != s(bool,happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)))) ) | ($spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f2133,f5183,f403])).
fof(f2133,plain,(
  ( ! [X12,X13] : (~sP5('F_0',sK14,happ(s(fun(X12,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X12,X13)))) ) | $spl2),
  inference(superposition,[],[f1208,f518])).
fof(f6627,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5801,f515])).
fof(f5801,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl30),
  inference(unit_resulting_resolution,[],[f854,f5183,f404])).
fof(f6626,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5802,f532])).
fof(f5802,plain,(
  s(bool,'T_0') = s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0'))))) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f855,f5183,f404])).
fof(f855,plain,(
  sP5('F_0','T_0','const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0'))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f473,f532,f402])).
fof(f6625,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5803,f591])).
fof(f5803,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1453,f5183,f404])).
fof(f6624,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5804,f515])).
fof(f5804,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl30),
  inference(unit_resulting_resolution,[],[f852,f5183,f404])).
fof(f6623,plain,(
  $false | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f6622,f515])).
fof(f6622,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl2 | $spl30)),
  inference(forward_demodulation,[],[f5805,f518])).
fof(f5805,plain,(
  s(bool,'F_0') = s(bool,sK14) | ($spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f1388,f5183,f404])).
fof(f1388,plain,(
  sP5('F_0','F_0',sK14) | $spl2),
  inference(unit_resulting_resolution,[],[f486,f482,f491,f473,f388])).
fof(f6621,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5806,f532])).
fof(f5806,plain,(
  s(bool,'T_0') = s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0'))))) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f853,f5183,f404])).
fof(f853,plain,(
  sP5('F_0','const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),'T_0') | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f473,f532,f402])).
fof(f6620,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f6619,f515])).
fof(f6619,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl2 | $spl30)),
  inference(forward_demodulation,[],[f6618,f518])).
fof(f6618,plain,(
  s(bool,'F_0') = s(bool,sK14) | ($spl1 | $spl2 | $spl30)),
  inference(forward_demodulation,[],[f5807,f556])).
fof(f5807,plain,(
  s(bool,sK14) = s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0'))))) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f1456,f5183,f404])).
fof(f1456,plain,(
  sP5('F_0','const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),sK14) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f526,f482,f491,f473,f389])).
fof(f6617,plain,(
  $false | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f6616,f515])).
fof(f6616,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl2 | $spl30)),
  inference(forward_demodulation,[],[f6615,f518])).
fof(f6615,plain,(
  s(bool,'F_0') = s(bool,sK14) | ($spl2 | $spl30)),
  inference(forward_demodulation,[],[f5808,f592])).
fof(f5808,plain,(
  ( ! [X0,X1] : (s(bool,sK14) = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | ($spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f1457,f5183,f404])).
fof(f1457,plain,(
  ( ! [X0,X1] : (sP5('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),sK14)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f602,f482,f491,f473,f389])).
fof(f6614,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5809,f591])).
fof(f5809,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1311,f5183,f404])).
fof(f6613,plain,(
  $false | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f6612,f515])).
fof(f6612,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl2 | $spl30)),
  inference(forward_demodulation,[],[f5810,f518])).
fof(f5810,plain,(
  s(bool,'F_0') = s(bool,sK14) | ($spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f1438,f5183,f404])).
fof(f6611,plain,(
  $false | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f6610,f515])).
fof(f6610,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl2 | $spl30)),
  inference(forward_demodulation,[],[f6609,f518])).
fof(f6609,plain,(
  s(bool,'F_0') = s(bool,sK14) | ($spl2 | $spl30)),
  inference(forward_demodulation,[],[f5811,f592])).
fof(f5811,plain,(
  ( ! [X0,X1] : (s(bool,sK14) = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | ($spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f1440,f5183,f404])).
fof(f1440,plain,(
  ( ! [X0,X1] : (sP5('F_0',sK14,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f601,f599,f489,f473,f389])).
fof(f6608,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f6607,f515])).
fof(f6607,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl2 | $spl30)),
  inference(forward_demodulation,[],[f6606,f518])).
fof(f6606,plain,(
  s(bool,'F_0') = s(bool,sK14) | ($spl1 | $spl2 | $spl30)),
  inference(forward_demodulation,[],[f5812,f556])).
fof(f5812,plain,(
  s(bool,sK14) = s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0'))))) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f1439,f5183,f404])).
fof(f1439,plain,(
  sP5('F_0',sK14,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0'))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f525,f523,f489,f473,f389])).
fof(f6605,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5813,f515])).
fof(f5813,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f1312,f5183,f404])).
fof(f1312,plain,(
  sP5('const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),'F_0','T_0') | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f486,f524,f492,f268,f387])).
fof(f6604,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f6603,f515])).
fof(f6603,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl2 | $spl30)),
  inference(forward_demodulation,[],[f5814,f518])).
fof(f5814,plain,(
  s(bool,'F_0') = s(bool,sK14) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f1389,f5183,f404])).
fof(f1389,plain,(
  sP5('const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),'F_0',sK14) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f486,f524,f491,f473,f388])).
fof(f6602,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5815,f591])).
fof(f5815,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f1314,f5183,f404])).
fof(f1314,plain,(
  ( ! [X0,X1] : (sP5('const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f602,f524,f492,f268,f387])).
fof(f6601,plain,(
  $false | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f6600,f515])).
fof(f6600,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl2 | $spl30)),
  inference(forward_demodulation,[],[f5816,f518])).
fof(f5816,plain,(
  s(bool,'F_0') = s(bool,sK14) | ($spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f1390,f5183,f404])).
fof(f1390,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0',sK14)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f486,f600,f491,f473,f388])).
fof(f6599,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5817,f515])).
fof(f5817,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl30),
  inference(unit_resulting_resolution,[],[f1315,f5183,f404])).
fof(f6598,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5818,f532])).
fof(f5818,plain,(
  s(bool,'T_0') = s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0'))))) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f1316,f5183,f404])).
fof(f1316,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),'T_0')) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f526,f600,f492,f268,f387])).
fof(f6597,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5844,f6021])).
fof(f6021,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f5183,f5183,f5183,f419])).
fof(f5844,plain,(
  ( ! [X0,X1] : (~sP7(X0,X1,'F_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f473,f5183,f407])).
fof(f6596,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5854,f6021])).
fof(f5854,plain,(
  ( ! [X0,X1] : (~sP7(X0,'F_0',X1)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f473,f5183,f408])).
fof(f6595,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5857,f5183])).
fof(f5857,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f496,f628,f5183,f409])).
fof(f6594,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5858,f5183])).
fof(f5858,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f494,f628,f5183,f409])).
fof(f6593,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5860,f5183])).
fof(f5860,plain,(
  ~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))))) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f527,f628,f5183,f409])).
fof(f527,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),X1)) ) | ($spl1 | $spl2)),
  inference(backward_demodulation,[],[f518,f510])).
fof(f510,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,sK14)))),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f468,f415])).
fof(f6592,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5861,f5183])).
fof(f5861,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f603,f628,f5183,f409])).
fof(f6591,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5867,f5959])).
fof(f5959,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f5183,f414])).
fof(f5867,plain,(
  ~sP6('F_0','T_0','T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f628,f268,f5183,f409])).
fof(f6590,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5878,f5959])).
fof(f5878,plain,(
  ( ! [X0] : (~sP6(X0,'F_0','T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f626,f5183,f5183,f409])).
fof(f6589,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5880,f5959])).
fof(f5880,plain,(
  ( ! [X0] : (~sP6('F_0',X0,'T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f628,f5183,f5183,f409])).
fof(f6588,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5884,f5183])).
fof(f5884,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f495,f626,f5183,f409])).
fof(f6587,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5885,f5183])).
fof(f5885,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f494,f626,f5183,f409])).
fof(f6586,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5887,f5183])).
fof(f5887,plain,(
  ~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))))) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f528,f626,f5183,f409])).
fof(f528,plain,(
  ( ! [X0,X1] : (sP6('const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),X0,X1)) ) | ($spl1 | $spl2)),
  inference(backward_demodulation,[],[f518,f511])).
fof(f511,plain,(
  ( ! [X0,X1] : (sP6('const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,sK14)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f468,f416])).
fof(f6585,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5888,f5183])).
fof(f5888,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f604,f626,f5183,f409])).
fof(f6584,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5895,f5959])).
fof(f5895,plain,(
  ~sP6('T_0','F_0','T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f626,f268,f5183,f409])).
fof(f6583,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5907,f5959])).
fof(f5907,plain,(
  ( ! [X0] : (~sP6(X0,'F_0','T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f626,f5183,f5183,f409])).
fof(f6582,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5909,f5959])).
fof(f5909,plain,(
  ( ! [X0] : (~sP6('F_0',X0,'T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f628,f5183,f5183,f409])).
fof(f6581,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5914,f5183])).
fof(f5914,plain,(
  ~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))))) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f527,f473,f5183,f413])).
fof(f6580,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5915,f5183])).
fof(f5915,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f603,f473,f5183,f413])).
fof(f6579,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5916,f5183])).
fof(f5916,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f496,f473,f5183,f413])).
fof(f6578,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5917,f5183])).
fof(f5917,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f528,f473,f5183,f413])).
fof(f6577,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5918,f5183])).
fof(f5918,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f604,f473,f5183,f413])).
fof(f6576,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5924,f5959])).
fof(f5924,plain,(
  ( ! [X0] : (~sP6(X0,'T_0','F_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f473,f268,f5183,f413])).
fof(f6575,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5934,f5959])).
fof(f5934,plain,(
  ( ! [X0,X1] : (~sP6(X0,X1,'F_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f473,f5183,f5183,f413])).
fof(f6574,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5949,f5183])).
fof(f5949,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f495,f473,f5183,f413])).
fof(f6573,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5950,f5183])).
fof(f5950,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f527,f473,f5183,f413])).
fof(f6572,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5951,f5183])).
fof(f5951,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f603,f473,f5183,f413])).
fof(f6571,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5953,f5183])).
fof(f5953,plain,(
  ~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))))) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f528,f473,f5183,f413])).
fof(f6570,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5954,f5183])).
fof(f5954,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f604,f473,f5183,f413])).
fof(f6569,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5955,f5959])).
fof(f5955,plain,(
  ( ! [X0] : (~sP6('T_0',X0,'F_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f268,f473,f5183,f413])).
fof(f6568,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5957,f5959])).
fof(f5957,plain,(
  ( ! [X0,X1] : (~sP6(X0,X1,'F_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f5183,f473,f5183,f413])).
fof(f6567,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5962,f5183])).
fof(f5962,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f628,f268,f5183,f419])).
fof(f6566,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5964,f5183])).
fof(f5964,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f628,f5183,f5183,f419])).
fof(f6565,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5965,f268])).
fof(f5965,plain,(
  ~p(s(bool,'T_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f996,f5183,f5183,f419])).
fof(f6564,plain,(
  $false | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5966,f5183])).
fof(f5966,plain,(
  ~p(s(bool,sK14)) | ($spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f1024,f5183,f5183,f419])).
fof(f1024,plain,(
  ~sP7('F_0','T_0',sK14) | $spl2),
  inference(unit_resulting_resolution,[],[f268,f473,f1023])).
fof(f1023,plain,(
  ( ! [X8,X9] : (~sP7(X9,X8,sK14) | p(s(bool,X9)) | ~p(s(bool,X8))) ) | $spl2),
  inference(subsumption_resolution,[],[f1019,f268])).
fof(f1019,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X8)) | p(s(bool,X9)) | ~sP7(X9,X8,sK14)) ) | $spl2),
  inference(superposition,[],[f422,f518])).
fof(f6563,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5967,f268])).
fof(f5967,plain,(
  ~p(s(bool,'T_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f628,f268,f5183,f419])).
fof(f6562,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5973,f268])).
fof(f5973,plain,(
  ~p(s(bool,'T_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f626,f5183,f5183,f419])).
fof(f6561,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5974,f268])).
fof(f5974,plain,(
  ~p(s(bool,'T_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f628,f5183,f5183,f419])).
fof(f6560,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5981,f5183])).
fof(f5981,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f626,f268,f5183,f419])).
fof(f6559,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5983,f5183])).
fof(f5983,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f626,f5183,f5183,f419])).
fof(f6558,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5985,f268])).
fof(f5985,plain,(
  ~p(s(bool,'T_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f996,f5183,f5183,f419])).
fof(f6557,plain,(
  $false | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f5986,f268])).
fof(f5986,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f1024,f5183,f5183,f419])).
fof(f6556,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5987,f268])).
fof(f5987,plain,(
  ~p(s(bool,'T_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f626,f268,f5183,f419])).
fof(f6555,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5994,f268])).
fof(f5994,plain,(
  ~p(s(bool,'T_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f626,f5183,f5183,f419])).
fof(f6554,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f5995,f268])).
fof(f5995,plain,(
  ~p(s(bool,'T_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f628,f5183,f5183,f419])).
fof(f6553,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6003,f268])).
fof(f6003,plain,(
  ~p(s(bool,'T_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f996,f268,f5183,f419])).
fof(f6552,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6004,f5183])).
fof(f6004,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f626,f5183,f5183,f419])).
fof(f6551,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6006,f268])).
fof(f6006,plain,(
  ~p(s(bool,'T_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f996,f5183,f5183,f419])).
fof(f6550,plain,(
  $false | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f6007,f268])).
fof(f6007,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f1024,f5183,f5183,f419])).
fof(f6549,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6009,f268])).
fof(f6009,plain,(
  ~p(s(bool,'T_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f996,f268,f5183,f419])).
fof(f6548,plain,(
  $false | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f6010,f5183])).
fof(f6010,plain,(
  ~p(s(bool,sK14)) | ($spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f1024,f268,f5183,f419])).
fof(f6547,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6016,f5183])).
fof(f6016,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f628,f5183,f5183,f419])).
fof(f6546,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6017,f268])).
fof(f6017,plain,(
  ~p(s(bool,'T_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f996,f5183,f5183,f419])).
fof(f6545,plain,(
  $false | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f6018,f5183])).
fof(f6018,plain,(
  ~p(s(bool,sK14)) | ($spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f1024,f5183,f5183,f419])).
fof(f6544,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f6024,f5183])).
fof(f6024,plain,(
  ~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))))) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f737,f473,f5183,f422])).
fof(f737,plain,(
  ( ! [X0] : (sP7('F_0',X0,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f527,f473,f459])).
fof(f6543,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6025,f5183])).
fof(f6025,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1170,f473,f5183,f422])).
fof(f6542,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6026,f5183])).
fof(f6026,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f637,f473,f5183,f422])).
fof(f6541,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f6027,f5183])).
fof(f6027,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f739,f473,f5183,f422])).
fof(f739,plain,(
  ( ! [X0] : (sP7('F_0','const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),X0)) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f528,f473,f459])).
fof(f6540,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6028,f5183])).
fof(f6028,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1178,f473,f5183,f422])).
fof(f6539,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6046,f5183])).
fof(f6046,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f631,f473,f5183,f422])).
fof(f6538,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f6047,f5183])).
fof(f6047,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f737,f473,f5183,f422])).
fof(f6537,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6048,f5183])).
fof(f6048,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1170,f473,f5183,f422])).
fof(f6536,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f6050,f5183])).
fof(f6050,plain,(
  ~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))))) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f739,f473,f5183,f422])).
fof(f6535,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6051,f5183])).
fof(f6051,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1178,f473,f5183,f422])).
fof(f6534,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6076,f6440])).
fof(f6440,plain,(
  ( ! [X0,X1] : (sP9('F_0',X0,X1)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f501,f5183,f458])).
fof(f6076,plain,(
  ( ! [X0,X1] : (~sP9('F_0',X0,X1)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f473,f5183,f423])).
fof(f6533,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6085,f6440])).
fof(f6085,plain,(
  ( ! [X0,X1] : (~sP9('F_0',X0,X1)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f473,f5183,f424])).
fof(f6532,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6097,f6132])).
fof(f6132,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f5183,f432])).
fof(f6097,plain,(
  ~sP8('F_0','T_0','F_0') | $spl30),
  inference(unit_resulting_resolution,[],[f644,f473,f5183,f427])).
fof(f6531,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6108,f6132])).
fof(f6108,plain,(
  ~sP8('F_0','F_0','T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f641,f473,f5183,f428])).
fof(f6530,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6128,f6132])).
fof(f6128,plain,(
  ( ! [X0] : (~sP8(X0,'F_0','F_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f473,f473,f5183,f430])).
fof(f6529,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6136,f268])).
fof(f6136,plain,(
  ~p(s(bool,'T_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f1079,f5183,f435])).
fof(f6528,plain,(
  $false | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f6137,f5183])).
fof(f6137,plain,(
  ~p(s(bool,sK14)) | ($spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f1107,f5183,f435])).
fof(f1107,plain,(
  ~sP9(sK14,'F_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f473,f473,f1106])).
fof(f1106,plain,(
  ( ! [X8,X9] : (~sP9(sK14,X9,X8) | p(s(bool,X9)) | p(s(bool,X8))) ) | $spl2),
  inference(subsumption_resolution,[],[f1103,f268])).
fof(f1103,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | p(s(bool,X8)) | p(s(bool,X9)) | ~sP9(sK14,X9,X8)) ) | $spl2),
  inference(superposition,[],[f437,f518])).
fof(f6527,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6142,f268])).
fof(f6142,plain,(
  ~p(s(bool,'T_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f641,f5183,f435])).
fof(f6526,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6143,f5183])).
fof(f6143,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f644,f5183,f435])).
fof(f6525,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6152,f268])).
fof(f6152,plain,(
  ~p(s(bool,'T_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f1079,f5183,f436])).
fof(f6524,plain,(
  $false | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f6153,f5183])).
fof(f6153,plain,(
  ~p(s(bool,sK14)) | ($spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f1107,f5183,f436])).
fof(f6523,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6158,f5183])).
fof(f6158,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f641,f5183,f436])).
fof(f6522,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6159,f268])).
fof(f6159,plain,(
  ~p(s(bool,'T_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f644,f5183,f436])).
fof(f6521,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6186,f6424])).
fof(f6424,plain,(
  ( ! [X0,X1] : (sP11('F_0',X0,X1)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f505,f5183,f457])).
fof(f6186,plain,(
  ( ! [X0,X1] : (~sP11('F_0',X0,X1)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f473,f5183,f441])).
fof(f6520,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6189,f5183])).
fof(f6189,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f505,f679,f5183,f444])).
fof(f6519,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6190,f5183])).
fof(f6190,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f504,f679,f5183,f444])).
fof(f6518,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f6192,f5183])).
fof(f6192,plain,(
  ~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))))) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f530,f679,f5183,f444])).
fof(f530,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))))) ) | ($spl1 | $spl2)),
  inference(backward_demodulation,[],[f518,f513])).
fof(f513,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,sK14)))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f468,f448])).
fof(f6517,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6193,f5183])).
fof(f6193,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f606,f679,f5183,f444])).
fof(f6516,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6199,f6333])).
fof(f6333,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f5183,f449])).
fof(f6199,plain,(
  ~sP10('F_0','T_0','T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f679,f268,f5183,f444])).
fof(f6515,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6211,f6333])).
fof(f6211,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'F_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f674,f5183,f5183,f444])).
fof(f6514,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6213,f6333])).
fof(f6213,plain,(
  ( ! [X0] : (~sP10('F_0','T_0',X0)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f679,f5183,f5183,f444])).
fof(f6513,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6231,f6333])).
fof(f6231,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'F_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f674,f5183,f5183,f444])).
fof(f6512,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6232,f6333])).
fof(f6232,plain,(
  ( ! [X0] : (~sP10('F_0','T_0',X0)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f679,f5183,f5183,f444])).
fof(f6511,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6257,f6333])).
fof(f6257,plain,(
  ~sP10('F_0','F_0','F_0') | $spl30),
  inference(unit_resulting_resolution,[],[f674,f473,f5183,f445])).
fof(f6510,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6276,f6333])).
fof(f6276,plain,(
  ~sP10('F_0','F_0','F_0') | $spl30),
  inference(unit_resulting_resolution,[],[f674,f473,f5183,f446])).
fof(f6509,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f6301,f5183])).
fof(f6301,plain,(
  ~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))))) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f530,f473,f5183,f447])).
fof(f6508,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6302,f5183])).
fof(f6302,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f606,f473,f5183,f447])).
fof(f6507,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6303,f5183])).
fof(f6303,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f505,f473,f5183,f447])).
fof(f6506,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f6304,f5183])).
fof(f6304,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f531,f473,f5183,f447])).
fof(f531,plain,(
  ( ! [X0,X1] : (sP10('const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),X0,X1)) ) | ($spl1 | $spl2)),
  inference(backward_demodulation,[],[f518,f514])).
fof(f514,plain,(
  ( ! [X0,X1] : (sP10('const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,sK14)))),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f468,f450])).
fof(f6505,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6305,f5183])).
fof(f6305,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f607,f473,f5183,f447])).
fof(f6504,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6306,f6333])).
fof(f6306,plain,(
  ( ! [X0] : (~sP10(X0,'F_0','T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f268,f473,f5183,f447])).
fof(f6503,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6308,f6333])).
fof(f6308,plain,(
  ( ! [X0,X1] : (~sP10(X0,'F_0',X1)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f5183,f473,f5183,f447])).
fof(f6502,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6323,f5183])).
fof(f6323,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f502,f473,f5183,f447])).
fof(f6501,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f6324,f5183])).
fof(f6324,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f530,f473,f5183,f447])).
fof(f6500,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6325,f5183])).
fof(f6325,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f606,f473,f5183,f447])).
fof(f6499,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f6327,f5183])).
fof(f6327,plain,(
  ~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))))) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f531,f473,f5183,f447])).
fof(f6498,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6328,f5183])).
fof(f6328,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f607,f473,f5183,f447])).
fof(f6497,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6329,f6333])).
fof(f6329,plain,(
  ( ! [X0] : (~sP10('T_0','F_0',X0)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f268,f473,f5183,f447])).
fof(f6496,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6331,f6333])).
fof(f6331,plain,(
  ( ! [X0,X1] : (~sP10(X0,'F_0',X1)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f5183,f473,f5183,f447])).
fof(f6495,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6350,f268])).
fof(f6350,plain,(
  ~p(s(bool,'T_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f1158,f5183,f453])).
fof(f6494,plain,(
  $false | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f6351,f5183])).
fof(f6351,plain,(
  ~p(s(bool,sK14)) | ($spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f1200,f5183,f453])).
fof(f1200,plain,(
  ~sP11(sK14,'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f268,f473,f1199])).
fof(f1199,plain,(
  ( ! [X8,X9] : (~sP11(sK14,X9,X8) | p(s(bool,X9)) | ~p(s(bool,X8))) ) | $spl2),
  inference(subsumption_resolution,[],[f1195,f268])).
fof(f1195,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X8)) | p(s(bool,X9)) | ~sP11(sK14,X9,X8)) ) | $spl2),
  inference(superposition,[],[f454,f518])).
fof(f6493,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6356,f5183])).
fof(f6356,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f674,f5183,f453])).
fof(f6492,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6357,f268])).
fof(f6357,plain,(
  ~p(s(bool,'T_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f679,f5183,f453])).
fof(f6491,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6364,f268])).
fof(f6364,plain,(
  ~p(s(bool,'T_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f686,f473,f5183,f454])).
fof(f6490,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f6365,f268])).
fof(f6365,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f712,f473,f5183,f454])).
fof(f712,plain,(
  ( ! [X0] : (sP11('T_0',X0,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f530,f268,f457])).
fof(f6489,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6366,f268])).
fof(f6366,plain,(
  ~p(s(bool,'T_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f1700,f473,f5183,f454])).
fof(f6488,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f6368,f5183])).
fof(f6368,plain,(
  ~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))))) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f1140,f473,f5183,f454])).
fof(f1140,plain,(
  sP11('const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),'F_0','T_0') | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f531,f268,f473,f446])).
fof(f6487,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6369,f5183])).
fof(f6369,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1717,f473,f5183,f454])).
fof(f6486,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f6388,f5183])).
fof(f6388,plain,(
  ~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))))) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f712,f473,f5183,f454])).
fof(f6485,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6389,f5183])).
fof(f6389,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f1700,f473,f5183,f454])).
fof(f6484,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6390,f268])).
fof(f6390,plain,(
  ~p(s(bool,'T_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f1139,f473,f5183,f454])).
fof(f6483,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(subsumption_resolution,[],[f6391,f268])).
fof(f6391,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f1140,f473,f5183,f454])).
fof(f6482,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6392,f268])).
fof(f6392,plain,(
  ~p(s(bool,'T_0')) | $spl30),
  inference(unit_resulting_resolution,[],[f1717,f473,f5183,f454])).
fof(f6481,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6410,f502])).
fof(f6410,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'F_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f674,f5183,f457])).
fof(f6480,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6411,f504])).
fof(f6411,plain,(
  ( ! [X0] : (~sP10('F_0','T_0',X0)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f679,f5183,f457])).
fof(f6479,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6428,f498])).
fof(f6428,plain,(
  ( ! [X0] : (~sP8('F_0',X0,'T_0')) ) | $spl30),
  inference(unit_resulting_resolution,[],[f641,f5183,f458])).
fof(f6478,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f6429,f500])).
fof(f6429,plain,(
  ( ! [X0] : (~sP8('F_0','T_0',X0)) ) | $spl30),
  inference(unit_resulting_resolution,[],[f644,f5183,f458])).
fof(f6456,plain,(
  $false | ($spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f637,f473,f5183,f1023])).
fof(f6451,plain,(
  $false | ($spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f482,f473,f5183,f928])).
fof(f928,plain,(
  ( ! [X8,X9] : (~sP4(X8,sK14,X9) | p(s(bool,X9)) | ~p(s(bool,X8))) ) | $spl2),
  inference(subsumption_resolution,[],[f922,f268])).
fof(f922,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X8)) | p(s(bool,X9)) | ~sP4(X8,sK14,X9)) ) | $spl2),
  inference(superposition,[],[f390,f518])).
fof(f6441,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f641,f501,f5183,f458])).
fof(f6439,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f644,f500,f5183,f458])).
fof(f6437,plain,(
  $false | ($spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f641,f499,f5183,f458])).
fof(f499,plain,(
  ( ! [X0,X1] : (sP8(X0,sK14,X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f469,f432])).
fof(f6435,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f641,f498,f5183,f458])).
fof(f6433,plain,(
  $false | ($spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f644,f497,f5183,f458])).
fof(f497,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,sK14)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f469,f431])).
fof(f6425,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f674,f505,f5183,f457])).
fof(f6423,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f679,f504,f5183,f457])).
fof(f6421,plain,(
  $false | ($spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f674,f503,f5183,f457])).
fof(f503,plain,(
  ( ! [X0,X1] : (sP10(X0,sK14,X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f469,f449])).
fof(f6419,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f679,f606,f5183,f457])).
fof(f6417,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f679,f530,f5183,f457])).
fof(f6415,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f674,f502,f5183,f457])).
fof(f6409,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f686,f473,f5183,f5183,f454])).
fof(f6399,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1139,f473,f268,f5183,f454])).
fof(f6386,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f686,f473,f5183,f5183,f454])).
fof(f6376,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f686,f473,f268,f5183,f454])).
fof(f6363,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f674,f5183,f5183,f453])).
fof(f6361,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f679,f268,f5183,f453])).
fof(f6355,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f674,f5183,f5183,f453])).
fof(f6353,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1158,f268,f5183,f453])).
fof(f6347,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f674,f473,f5183,f452])).
fof(f6341,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f674,f473,f473,f5183,f451])).
fof(f6334,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1149,f5183,f449])).
fof(f6332,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f502,f5183,f473,f5183,f447])).
fof(f6330,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f502,f268,f473,f5183,f447])).
fof(f6309,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f502,f5183,f473,f5183,f447])).
fof(f6307,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f505,f268,f473,f5183,f447])).
fof(f6284,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f674,f505,f473,f5183,f446])).
fof(f6280,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f674,f502,f473,f5183,f446])).
fof(f6265,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f674,f505,f473,f5183,f445])).
fof(f6261,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f674,f502,f473,f5183,f445])).
fof(f6246,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f674,f505,f5183,f5183,f444])).
fof(f6244,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f679,f504,f5183,f5183,f444])).
fof(f6242,plain,(
  $false | ($spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f674,f503,f5183,f5183,f444])).
fof(f6240,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f679,f606,f5183,f5183,f444])).
fof(f6238,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f679,f530,f5183,f5183,f444])).
fof(f6236,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f674,f502,f5183,f5183,f444])).
fof(f6214,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f505,f679,f5183,f5183,f444])).
fof(f6212,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f505,f674,f5183,f5183,f444])).
fof(f6200,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f505,f679,f268,f5183,f444])).
fof(f6187,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1139,f473,f5183,f441])).
fof(f6177,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1059,f473,f473,f5183,f437])).
fof(f6165,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f641,f5183,f5183,f436])).
fof(f6163,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f644,f268,f5183,f436])).
fof(f6157,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f641,f5183,f5183,f436])).
fof(f6155,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1079,f268,f5183,f436])).
fof(f6149,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f641,f5183,f5183,f435])).
fof(f6147,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f641,f268,f5183,f435])).
fof(f6141,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f641,f5183,f5183,f435])).
fof(f6139,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1079,f268,f5183,f435])).
fof(f6133,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1070,f5183,f432])).
fof(f6131,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1070,f5183,f431])).
fof(f6129,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f501,f473,f473,f5183,f430])).
fof(f6115,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f641,f501,f473,f5183,f428])).
fof(f6113,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f641,f498,f473,f5183,f428])).
fof(f6098,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f501,f644,f473,f5183,f427])).
fof(f6086,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1059,f473,f5183,f424])).
fof(f6077,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1059,f473,f5183,f423])).
fof(f6068,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f631,f473,f5183,f5183,f422])).
fof(f6058,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f631,f473,f268,f5183,f422])).
fof(f6045,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f631,f473,f5183,f5183,f422])).
fof(f6035,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f637,f473,f268,f5183,f422])).
fof(f6022,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f626,f5183,f5183,f5183,f419])).
fof(f6020,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f628,f268,f5183,f5183,f419])).
fof(f6014,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f626,f5183,f268,f5183,f419])).
fof(f6012,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f996,f268,f268,f5183,f419])).
fof(f6001,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f626,f5183,f5183,f5183,f419])).
fof(f5999,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f626,f268,f5183,f5183,f419])).
fof(f5993,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f626,f5183,f268,f5183,f419])).
fof(f5991,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f626,f268,f268,f5183,f419])).
fof(f5980,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f626,f5183,f5183,f5183,f419])).
fof(f5978,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f626,f268,f5183,f5183,f419])).
fof(f5972,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f628,f5183,f268,f5183,f419])).
fof(f5970,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f628,f268,f268,f5183,f419])).
fof(f5960,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f987,f5183,f414])).
fof(f5958,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f495,f5183,f473,f5183,f413])).
fof(f5956,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f495,f268,f473,f5183,f413])).
fof(f5935,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f495,f473,f5183,f5183,f413])).
fof(f5925,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f496,f473,f268,f5183,f413])).
fof(f5910,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f496,f628,f5183,f5183,f409])).
fof(f5908,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f495,f626,f5183,f5183,f409])).
fof(f5896,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f495,f626,f268,f5183,f409])).
fof(f5881,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f496,f628,f5183,f5183,f409])).
fof(f5879,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f495,f626,f5183,f5183,f409])).
fof(f5868,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f496,f628,f268,f5183,f409])).
fof(f5855,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f637,f473,f5183,f408])).
fof(f5845,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f631,f473,f5183,f407])).
fof(f5835,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1311,f591,f5183,f404])).
fof(f5833,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f853,f532,f5183,f404])).
fof(f5831,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f852,f515,f5183,f404])).
fof(f5829,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1453,f591,f5183,f404])).
fof(f5827,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f855,f532,f5183,f404])).
fof(f5825,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f854,f515,f5183,f404])).
fof(f5788,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1208,f381,f5183,f403])).
fof(f5786,plain,(
  $false | ($spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f889,f518,f5183,f403])).
fof(f5784,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1206,f381,f5183,f403])).
fof(f5782,plain,(
  $false | ($spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f886,f518,f5183,f403])).
fof(f6477,plain,(
  $false | $spl30),
  inference(trivial_inequality_removal,[],[f5772])).
fof(f5772,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl30),
  inference(unit_resulting_resolution,[],[f976,f5183,f403])).
fof(f6476,plain,(
  $false | $spl30),
  inference(trivial_inequality_removal,[],[f5771])).
fof(f5771,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl30),
  inference(unit_resulting_resolution,[],[f972,f5183,f403])).
fof(f5761,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f963,f5183,f401])).
fof(f5759,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f963,f5183,f400])).
fof(f5757,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f963,f5183,f399])).
fof(f5755,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f946,f5183,f395])).
fof(f5753,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f485,f5183,f473,f5183,f394])).
fof(f5751,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f485,f268,f473,f5183,f394])).
fof(f5730,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f485,f473,f5183,f5183,f394])).
fof(f5720,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f486,f473,f268,f5183,f394])).
fof(f5707,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f918,f5183,f391])).
fof(f5705,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f481,f473,f5183,f5183,f390])).
fof(f5694,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f481,f473,f268,f5183,f390])).
fof(f5680,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f481,f473,f5183,f5183,f390])).
fof(f5669,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f482,f473,f268,f5183,f390])).
fof(f5600,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1856,f484,f599,f490,f5183,f387])).
fof(f5597,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1856,f601,f599,f490,f5183,f387])).
fof(f5594,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f879,f484,f523,f490,f5183,f387])).
fof(f5591,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f879,f525,f523,f490,f5183,f387])).
fof(f5588,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f878,f484,f481,f490,f5183,f387])).
fof(f5585,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f878,f485,f481,f490,f5183,f387])).
fof(f5582,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1856,f601,f480,f490,f5183,f387])).
fof(f5580,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f879,f525,f480,f490,f5183,f387])).
fof(f5578,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f878,f485,f480,f490,f5183,f387])).
fof(f5576,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f878,f484,f480,f490,f5183,f387])).
fof(f5529,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1856,f484,f599,f488,f5183,f387])).
fof(f5527,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1856,f601,f599,f488,f5183,f387])).
fof(f5521,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f879,f484,f523,f488,f5183,f387])).
fof(f5519,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f879,f525,f523,f488,f5183,f387])).
fof(f5513,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f878,f484,f481,f488,f5183,f387])).
fof(f5511,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f878,f485,f481,f488,f5183,f387])).
fof(f5508,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1856,f601,f480,f488,f5183,f387])).
fof(f5506,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f879,f525,f480,f488,f5183,f387])).
fof(f5504,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f878,f485,f480,f488,f5183,f387])).
fof(f5502,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f878,f484,f480,f488,f5183,f387])).
fof(f5414,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1381,f5183,f473,f5183,f386])).
fof(f5412,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1381,f268,f473,f5183,f386])).
fof(f5356,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f1381,f473,f5183,f5183,f386])).
fof(f5319,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f854,f473,f268,f5183,f386])).
fof(f5303,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f560,f473,f5183,f376])).
fof(f5296,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f560,f473,f5183,f373])).
fof(f5289,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f573,f5183,f5183,f372])).
fof(f5287,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f564,f268,f5183,f372])).
fof(f5275,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f573,f5183,f5183,f372])).
fof(f5273,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f573,f268,f5183,f372])).
fof(f5261,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f584,f5183,f358])).
fof(f5259,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f584,f5183,f357])).
fof(f5257,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f5183,f591,f5183,f355])).
fof(f5255,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f5183,f532,f5183,f355])).
fof(f5253,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f5183,f515,f5183,f355])).
fof(f5251,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f268,f591,f5183,f355])).
fof(f5249,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f268,f532,f5183,f355])).
fof(f5247,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f268,f515,f5183,f355])).
fof(f5243,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f5183,f591,f5183,f355])).
fof(f5241,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f5183,f532,f5183,f355])).
fof(f5239,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f5183,f515,f5183,f355])).
fof(f5237,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f268,f591,f5183,f355])).
fof(f5235,plain,(
  $false | ($spl1 | $spl2 | $spl30)),
  inference(unit_resulting_resolution,[],[f268,f532,f5183,f355])).
fof(f5233,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f268,f515,f5183,f355])).
fof(f5229,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f381,f5183,f290])).
fof(f5227,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f515,f5183,f287])).
fof(f5225,plain,(
  $false | $spl30),
  inference(unit_resulting_resolution,[],[f473,f5183])).
fof(f5224,plain,(
  sP1('const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),'T_0') | $spl32),
  inference(cnf_transformation,[],[f5224_D])).
fof(f5224_D,plain,(
  sP1('const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),'T_0') <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f5222,plain,(
  ( ! [X15] : (p(s(bool,X15)) | sP1('const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),'T_0')) ) | ($spl1 | $spl2 | $spl18)),
  inference(subsumption_resolution,[],[f5216,f286])).
fof(f5216,plain,(
  ( ! [X15] : (p(s(bool,X15)) | sP1('const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),'T_0') | s(bool,'T_0') = s(bool,X15)) ) | ($spl1 | $spl2 | $spl18)),
  inference(superposition,[],[f4434,f2227])).
fof(f2227,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,X4)) ) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2222,f473])).
fof(f2222,plain,(
  ( ! [X4] : (p(s(bool,'F_0')) | s(bool,X4) = s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,X4)) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f810,f556])).
fof(f5182,plain,(
  sP1('T_0','const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0'))))) | $spl28),
  inference(cnf_transformation,[],[f5182_D])).
fof(f5182_D,plain,(
  sP1('T_0','const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0'))))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f5180,plain,(
  ( ! [X15] : (p(s(bool,X15)) | sP1('T_0','const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))))) ) | ($spl1 | $spl2 | $spl16)),
  inference(subsumption_resolution,[],[f5174,f286])).
fof(f5174,plain,(
  ( ! [X15] : (p(s(bool,X15)) | sP1('T_0','const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,X15)) ) | ($spl1 | $spl2 | $spl16)),
  inference(superposition,[],[f4431,f2227])).
fof(f4697,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f267,f266,f265,f268,f473,f275,f287,f286,f291,f290,f329,f328,f332,f336,f335,f339,f338,f337,f343,f341,f346,f344,f351,f350,f349,f348,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f592,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1220,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1480,f1477,f1496,f1493,f1511,f1514,f1515,f1539,f1536,f1538,f1562,f1578,f1593,f1610,f1622,f1623,f1691,f1700,f1714,f1709,f1710,f1715,f1712,f1717,f1741,f1771,f1800,f1827,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2002,f1928,f1929,f2010,f2040,f2055,f2037,f2043,f2044,f2046,f2048,f2084,f2087,f2088,f2090,f2092,f2126,f2125,f2142,f2129,f2130,f2132,f2170,f2173,f2174,f2176,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2968,f2970,f2971,f2972,f2973,f2974,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4450,f4451,f4452,f4453,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f4463,f4464,f4465,f4466,f4467,f4472,f4473,f4474,f3420,f3421,f3499,f3501,f3502,f3503,f3504,f3505,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f4415,f4617,f4622,f4624,f4518,f4521,f4522,f4526,f4527,f4530,f4531,f4535,f4692,f4694,f4696])).
fof(f4696,plain,(
  ( ! [X4,X2,X0] : (s(fun(X0,bool),X2) = s(fun(X0,bool),X4)) ) | $spl8),
  inference(subsumption_resolution,[],[f4602,f4415])).
fof(f4602,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK17(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X4,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X2) = s(fun(X0,bool),X4)) ) | $spl8),
  inference(backward_demodulation,[],[f4415,f1715])).
fof(f4694,plain,(
  ( ! [X4,X2,X0] : (s(fun(X0,bool),X2) = s(fun(X0,bool),X4)) ) | $spl8),
  inference(subsumption_resolution,[],[f4601,f4415])).
fof(f4601,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK17(X4,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X2) = s(fun(X0,bool),X4)) ) | $spl8),
  inference(backward_demodulation,[],[f4415,f1714])).
fof(f4692,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | $spl8),
  inference(subsumption_resolution,[],[f4600,f4415])).
fof(f4600,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK17(X2,X1,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | $spl8),
  inference(backward_demodulation,[],[f4415,f336])).
fof(f4535,plain,(
  ( ! [X10,X8,X11,X9] : (sP5('T_0','const.bool.IN_2'(s(X8,X9),s(fun(X8,bool),'const.pred_set.EMPTY_0')),happ(s(fun(X10,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X10,X11)))) ) | $spl8),
  inference(backward_demodulation,[],[f4415,f2176])).
fof(f4531,plain,(
  ( ! [X2,X0,X3,X1] : (~sP2(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X3)),'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')),'F_0')) ) | $spl8),
  inference(backward_demodulation,[],[f4415,f2142])).
fof(f4530,plain,(
  ( ! [X10,X8,X11,X9] : (~sP5('F_0','const.bool.IN_2'(s(X8,X9),s(fun(X8,bool),'const.pred_set.EMPTY_0')),happ(s(fun(X10,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X10,X11)))) ) | $spl8),
  inference(backward_demodulation,[],[f4415,f2132])).
fof(f4527,plain,(
  ( ! [X0,X1] : (~sP2(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)),'F_0','F_0')) ) | $spl8),
  inference(backward_demodulation,[],[f4415,f2126])).
fof(f4526,plain,(
  ( ! [X10,X8,X11,X9] : (sP5('T_0',happ(s(fun(X10,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X10,X11)),'const.bool.IN_2'(s(X8,X9),s(fun(X8,bool),'const.pred_set.EMPTY_0')))) ) | $spl8),
  inference(backward_demodulation,[],[f4415,f2090])).
fof(f4522,plain,(
  ( ! [X2,X0,X3,X1] : (~sP2('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),'F_0')) ) | $spl8),
  inference(backward_demodulation,[],[f4415,f2055])).
fof(f4521,plain,(
  ( ! [X10,X8,X11,X9] : (~sP5('F_0',happ(s(fun(X10,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X10,X11)),'const.bool.IN_2'(s(X8,X9),s(fun(X8,bool),'const.pred_set.EMPTY_0')))) ) | $spl8),
  inference(backward_demodulation,[],[f4415,f2046])).
fof(f4518,plain,(
  ( ! [X0,X1] : (~sP2('F_0',happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)),'F_0')) ) | $spl8),
  inference(backward_demodulation,[],[f4415,f2037])).
fof(f4624,plain,(
  ( ! [X6,X7] : (s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),X7)) ) | $spl8),
  inference(subsumption_resolution,[],[f4517,f4415])).
fof(f4517,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X6,sK17('const.pred_set.EMPTY_0',X7,X6)),s(fun(X6,bool),X7))) | s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),X7)) ) | $spl8),
  inference(backward_demodulation,[],[f4415,f1929])).
fof(f4622,plain,(
  ( ! [X4,X5] : (s(fun(X4,bool),'const.pred_set.EMPTY_0') = s(fun(X4,bool),X5)) ) | $spl8),
  inference(subsumption_resolution,[],[f4516,f4415])).
fof(f4516,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X4,sK17(X5,'const.pred_set.EMPTY_0',X4)),s(fun(X4,bool),X5))) | s(fun(X4,bool),'const.pred_set.EMPTY_0') = s(fun(X4,bool),X5)) ) | $spl8),
  inference(backward_demodulation,[],[f4415,f1928])).
fof(f4617,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(subsumption_resolution,[],[f4499,f4415])).
fof(f4499,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl8),
  inference(backward_demodulation,[],[f4415,f290])).
fof(f4695,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f267,f266,f265,f268,f473,f275,f287,f286,f291,f290,f329,f328,f332,f336,f335,f339,f338,f337,f343,f341,f346,f344,f351,f350,f349,f348,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f592,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1220,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1480,f1477,f1496,f1493,f1511,f1514,f1515,f1539,f1536,f1538,f1562,f1578,f1593,f1610,f1622,f1623,f1691,f1700,f1714,f1709,f1710,f1715,f1712,f1717,f1741,f1771,f1800,f1827,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2002,f1928,f1929,f2010,f2040,f2055,f2037,f2043,f2044,f2046,f2048,f2084,f2087,f2088,f2090,f2092,f2126,f2125,f2142,f2129,f2130,f2132,f2170,f2173,f2174,f2176,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2968,f2970,f2971,f2972,f2973,f2974,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4450,f4451,f4452,f4453,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f4463,f4464,f4465,f4466,f4467,f4472,f4473,f4474,f3420,f3421,f3499,f3501,f3502,f3503,f3504,f3505,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f4415,f4617,f4622,f4624,f4518,f4521,f4522,f4526,f4527,f4530,f4531,f4535,f4692,f4694])).
fof(f4693,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f267,f266,f265,f268,f473,f275,f287,f286,f291,f290,f329,f328,f332,f336,f335,f339,f338,f337,f343,f341,f346,f344,f351,f350,f349,f348,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f592,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1220,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1480,f1477,f1496,f1493,f1511,f1514,f1515,f1539,f1536,f1538,f1562,f1578,f1593,f1610,f1622,f1623,f1691,f1700,f1714,f1709,f1710,f1715,f1712,f1717,f1741,f1771,f1800,f1827,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2002,f1928,f1929,f2010,f2040,f2055,f2037,f2043,f2044,f2046,f2048,f2084,f2087,f2088,f2090,f2092,f2126,f2125,f2142,f2129,f2130,f2132,f2170,f2173,f2174,f2176,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2968,f2970,f2971,f2972,f2973,f2974,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4450,f4451,f4452,f4453,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f4463,f4464,f4465,f4466,f4467,f4472,f4473,f4474,f3420,f3421,f3499,f3501,f3502,f3503,f3504,f3505,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f4415,f4617,f4622,f4624,f4518,f4521,f4522,f4526,f4527,f4530,f4531,f4535,f4692])).
fof(f4639,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f267,f266,f265,f268,f473,f275,f287,f286,f291,f290,f329,f328,f332,f336,f335,f339,f338,f337,f343,f341,f346,f344,f351,f350,f349,f348,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f592,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1220,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1480,f1477,f1496,f1493,f1511,f1514,f1515,f1539,f1536,f1538,f1562,f1578,f1593,f1610,f1622,f1623,f1691,f1700,f1714,f1709,f1710,f1715,f1712,f1717,f1741,f1771,f1800,f1827,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2002,f1928,f1929,f2010,f2040,f2055,f2037,f2043,f2044,f2046,f2048,f2084,f2087,f2088,f2090,f2092,f2126,f2125,f2142,f2129,f2130,f2132,f2170,f2173,f2174,f2176,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2968,f2970,f2971,f2972,f2973,f2974,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4450,f4451,f4452,f4453,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f4463,f4464,f4465,f4466,f4467,f4472,f4473,f4474,f3420,f3421,f3499,f3501,f3502,f3503,f3504,f3505,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f4415,f4617,f4622,f4624,f4518,f4521,f4522,f4526,f4527,f4530,f4531,f4535])).
fof(f4638,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(global_subsumption,[],[f4534,f267,f266,f265,f268,f473,f275,f287,f286,f291,f290,f329,f328,f332,f336,f335,f339,f338,f337,f343,f341,f346,f344,f351,f350,f349,f348,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f592,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1220,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1480,f1477,f1496,f1493,f1511,f1514,f1515,f1539,f1536,f1538,f1562,f1578,f1593,f1610,f1622,f1623,f1691,f1700,f1714,f1709,f1710,f1715,f1712,f1717,f1741,f1771,f1800,f1827,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2002,f1928,f1929,f2010,f2040,f2055,f2037,f2043,f2044,f2046,f2048,f2084,f2087,f2088,f2090,f2092,f2126,f2125,f2142,f2129,f2130,f2132,f2170,f2173,f2174,f2176,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2968,f2970,f2971,f2972,f2973,f2974,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4450,f4451,f4452,f4453,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f4463,f4464,f4465,f4466,f4467,f4472,f4473,f4474,f3420,f3421,f3499,f3501,f3502,f3503,f3504,f3505,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f4415,f4617,f4622,f4624,f4518,f4521,f4522,f4526,f4527,f4530,f4531])).
fof(f4534,plain,(
  ( ! [X6,X7] : (sP5('T_0','const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),happ(s(fun(X6,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X6,X7)))) ) | ($spl1 | $spl2 | $spl8)),
  inference(backward_demodulation,[],[f4415,f2175])).
fof(f2175,plain,(
  ( ! [X6,X7] : (sP5('T_0','const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),happ(s(fun(X6,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X6,X7)))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f1209,f556])).
fof(f4637,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(global_subsumption,[],[f4532,f267,f266,f265,f268,f473,f275,f287,f286,f291,f290,f329,f328,f332,f336,f335,f339,f338,f337,f343,f341,f346,f344,f351,f350,f349,f348,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f592,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1220,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1480,f1477,f1496,f1493,f1511,f1514,f1515,f1539,f1536,f1538,f1562,f1578,f1593,f1610,f1622,f1623,f1691,f1700,f1714,f1709,f1710,f1715,f1712,f1717,f1741,f1771,f1800,f1827,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2002,f1928,f1929,f2010,f2040,f2055,f2037,f2043,f2044,f2046,f2048,f2084,f2087,f2088,f2090,f2092,f2126,f2125,f2142,f2129,f2130,f2132,f2170,f2173,f2174,f2176,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2968,f2970,f2971,f2972,f2973,f2974,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4450,f4451,f4452,f4453,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f4463,f4464,f4465,f4466,f4467,f4472,f4473,f4474,f3420,f3421,f3499,f3501,f3502,f3503,f3504,f3505,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f4415,f4617,f4622,f4624,f4518,f4521,f4522,f4526,f4527,f4530,f4531])).
fof(f4532,plain,(
  ( ! [X0,X1] : (~sP2(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)),'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),'F_0')) ) | ($spl1 | $spl2 | $spl8)),
  inference(backward_demodulation,[],[f4415,f2143])).
fof(f2143,plain,(
  ( ! [X0,X1] : (~sP2(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X0,X1)),'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),'F_0')) ) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f2122,f556])).
fof(f2122,plain,(
  ( ! [X0,X1] : (~sP2(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0'))))))),s(X0,X1)),'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),'F_0')) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f473,f526,f482,f1208,f389])).
fof(f4636,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f267,f266,f265,f268,f473,f275,f287,f286,f291,f290,f329,f328,f332,f336,f335,f339,f338,f337,f343,f341,f346,f344,f351,f350,f349,f348,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f592,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1220,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1480,f1477,f1496,f1493,f1511,f1514,f1515,f1539,f1536,f1538,f1562,f1578,f1593,f1610,f1622,f1623,f1691,f1700,f1714,f1709,f1710,f1715,f1712,f1717,f1741,f1771,f1800,f1827,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2002,f1928,f1929,f2010,f2040,f2055,f2037,f2043,f2044,f2046,f2048,f2084,f2087,f2088,f2090,f2092,f2126,f2125,f2142,f2129,f2130,f2132,f2170,f2173,f2174,f2176,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2968,f2970,f2971,f2972,f2973,f2974,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4450,f4451,f4452,f4453,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f4463,f4464,f4465,f4466,f4467,f4472,f4473,f4474,f3420,f3421,f3499,f3501,f3502,f3503,f3504,f3505,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f4415,f4617,f4622,f4624,f4518,f4521,f4522,f4526,f4527,f4530,f4531])).
fof(f4635,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f267,f266,f265,f268,f473,f275,f287,f286,f291,f290,f329,f328,f332,f336,f335,f339,f338,f337,f343,f341,f346,f344,f351,f350,f349,f348,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f592,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1220,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1480,f1477,f1496,f1493,f1511,f1514,f1515,f1539,f1536,f1538,f1562,f1578,f1593,f1610,f1622,f1623,f1691,f1700,f1714,f1709,f1710,f1715,f1712,f1717,f1741,f1771,f1800,f1827,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2002,f1928,f1929,f2010,f2040,f2055,f2037,f2043,f2044,f2046,f2048,f2084,f2087,f2088,f2090,f2092,f2126,f2125,f2142,f2129,f2130,f2132,f2170,f2173,f2174,f2176,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2968,f2970,f2971,f2972,f2973,f2974,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4450,f4451,f4452,f4453,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f4463,f4464,f4465,f4466,f4467,f4472,f4473,f4474,f3420,f3421,f3499,f3501,f3502,f3503,f3504,f3505,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f4415,f4617,f4622,f4624,f4518,f4521,f4522,f4526,f4527,f4530])).
fof(f4634,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(global_subsumption,[],[f4529,f267,f266,f265,f268,f473,f275,f287,f286,f291,f290,f329,f328,f332,f336,f335,f339,f338,f337,f343,f341,f346,f344,f351,f350,f349,f348,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f592,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1220,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1480,f1477,f1496,f1493,f1511,f1514,f1515,f1539,f1536,f1538,f1562,f1578,f1593,f1610,f1622,f1623,f1691,f1700,f1714,f1709,f1710,f1715,f1712,f1717,f1741,f1771,f1800,f1827,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2002,f1928,f1929,f2010,f2040,f2055,f2037,f2043,f2044,f2046,f2048,f2084,f2087,f2088,f2090,f2092,f2126,f2125,f2142,f2129,f2130,f2132,f2170,f2173,f2174,f2176,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2968,f2970,f2971,f2972,f2973,f2974,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4450,f4451,f4452,f4453,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f4463,f4464,f4465,f4466,f4467,f4472,f4473,f4474,f3420,f3421,f3499,f3501,f3502,f3503,f3504,f3505,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f4415,f4617,f4622,f4624,f4518,f4521,f4522,f4526,f4527])).
fof(f4529,plain,(
  ( ! [X6,X7] : (~sP5('F_0','const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),happ(s(fun(X6,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X6,X7)))) ) | ($spl1 | $spl2 | $spl8)),
  inference(backward_demodulation,[],[f4415,f2131])).
fof(f2131,plain,(
  ( ! [X6,X7] : (~sP5('F_0','const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),happ(s(fun(X6,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X6,X7)))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f1208,f556])).
fof(f4633,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f267,f266,f265,f268,f473,f275,f287,f286,f291,f290,f329,f328,f332,f336,f335,f339,f338,f337,f343,f341,f346,f344,f351,f350,f349,f348,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f592,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1220,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1480,f1477,f1496,f1493,f1511,f1514,f1515,f1539,f1536,f1538,f1562,f1578,f1593,f1610,f1622,f1623,f1691,f1700,f1714,f1709,f1710,f1715,f1712,f1717,f1741,f1771,f1800,f1827,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2002,f1928,f1929,f2010,f2040,f2055,f2037,f2043,f2044,f2046,f2048,f2084,f2087,f2088,f2090,f2092,f2126,f2125,f2142,f2129,f2130,f2132,f2170,f2173,f2174,f2176,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2968,f2970,f2971,f2972,f2973,f2974,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4450,f4451,f4452,f4453,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f4463,f4464,f4465,f4466,f4467,f4472,f4473,f4474,f3420,f3421,f3499,f3501,f3502,f3503,f3504,f3505,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f4415,f4617,f4622,f4624,f4518,f4521,f4522,f4526,f4527])).
fof(f4632,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f267,f266,f265,f268,f473,f275,f287,f286,f291,f290,f329,f328,f332,f336,f335,f339,f338,f337,f343,f341,f346,f344,f351,f350,f349,f348,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f592,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1220,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1480,f1477,f1496,f1493,f1511,f1514,f1515,f1539,f1536,f1538,f1562,f1578,f1593,f1610,f1622,f1623,f1691,f1700,f1714,f1709,f1710,f1715,f1712,f1717,f1741,f1771,f1800,f1827,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2002,f1928,f1929,f2010,f2040,f2055,f2037,f2043,f2044,f2046,f2048,f2084,f2087,f2088,f2090,f2092,f2126,f2125,f2142,f2129,f2130,f2132,f2170,f2173,f2174,f2176,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2968,f2970,f2971,f2972,f2973,f2974,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4450,f4451,f4452,f4453,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f4463,f4464,f4465,f4466,f4467,f4472,f4473,f4474,f3420,f3421,f3499,f3501,f3502,f3503,f3504,f3505,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f4415,f4617,f4622,f4624,f4518,f4521,f4522,f4526])).
fof(f4631,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(global_subsumption,[],[f4525,f267,f266,f265,f268,f473,f275,f287,f286,f291,f290,f329,f328,f332,f336,f335,f339,f338,f337,f343,f341,f346,f344,f351,f350,f349,f348,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f592,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1220,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1480,f1477,f1496,f1493,f1511,f1514,f1515,f1539,f1536,f1538,f1562,f1578,f1593,f1610,f1622,f1623,f1691,f1700,f1714,f1709,f1710,f1715,f1712,f1717,f1741,f1771,f1800,f1827,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2002,f1928,f1929,f2010,f2040,f2055,f2037,f2043,f2044,f2046,f2048,f2084,f2087,f2088,f2090,f2092,f2126,f2125,f2142,f2129,f2130,f2132,f2170,f2173,f2174,f2176,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2968,f2970,f2971,f2972,f2973,f2974,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4450,f4451,f4452,f4453,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f4463,f4464,f4465,f4466,f4467,f4472,f4473,f4474,f3420,f3421,f3499,f3501,f3502,f3503,f3504,f3505,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f4415,f4617,f4622,f4624,f4518,f4521,f4522])).
fof(f4525,plain,(
  ( ! [X6,X7] : (sP5('T_0',happ(s(fun(X6,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X6,X7)),'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))))) ) | ($spl1 | $spl2 | $spl8)),
  inference(backward_demodulation,[],[f4415,f2089])).
fof(f2089,plain,(
  ( ! [X6,X7] : (sP5('T_0',happ(s(fun(X6,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X6,X7)),'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f1207,f556])).
fof(f4630,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(global_subsumption,[],[f4523,f267,f266,f265,f268,f473,f275,f287,f286,f291,f290,f329,f328,f332,f336,f335,f339,f338,f337,f343,f341,f346,f344,f351,f350,f349,f348,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f592,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1220,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1480,f1477,f1496,f1493,f1511,f1514,f1515,f1539,f1536,f1538,f1562,f1578,f1593,f1610,f1622,f1623,f1691,f1700,f1714,f1709,f1710,f1715,f1712,f1717,f1741,f1771,f1800,f1827,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2002,f1928,f1929,f2010,f2040,f2055,f2037,f2043,f2044,f2046,f2048,f2084,f2087,f2088,f2090,f2092,f2126,f2125,f2142,f2129,f2130,f2132,f2170,f2173,f2174,f2176,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2968,f2970,f2971,f2972,f2973,f2974,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4450,f4451,f4452,f4453,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f4463,f4464,f4465,f4466,f4467,f4472,f4473,f4474,f3420,f3421,f3499,f3501,f3502,f3503,f3504,f3505,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f4415,f4617,f4622,f4624,f4518,f4521,f4522])).
fof(f4523,plain,(
  ( ! [X0,X1] : (~sP2('const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)),'F_0')) ) | ($spl1 | $spl2 | $spl8)),
  inference(backward_demodulation,[],[f4415,f2056])).
fof(f2056,plain,(
  ( ! [X0,X1] : (~sP2('const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X0,X1)),'F_0')) ) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f2038,f556])).
fof(f2038,plain,(
  ( ! [X0,X1] : (~sP2('const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0'))))))),s(X0,X1)),'F_0')) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f473,f525,f523,f1206,f389])).
fof(f4629,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f267,f266,f265,f268,f473,f275,f287,f286,f291,f290,f329,f328,f332,f336,f335,f339,f338,f337,f343,f341,f346,f344,f351,f350,f349,f348,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f592,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1220,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1480,f1477,f1496,f1493,f1511,f1514,f1515,f1539,f1536,f1538,f1562,f1578,f1593,f1610,f1622,f1623,f1691,f1700,f1714,f1709,f1710,f1715,f1712,f1717,f1741,f1771,f1800,f1827,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2002,f1928,f1929,f2010,f2040,f2055,f2037,f2043,f2044,f2046,f2048,f2084,f2087,f2088,f2090,f2092,f2126,f2125,f2142,f2129,f2130,f2132,f2170,f2173,f2174,f2176,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2968,f2970,f2971,f2972,f2973,f2974,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4450,f4451,f4452,f4453,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f4463,f4464,f4465,f4466,f4467,f4472,f4473,f4474,f3420,f3421,f3499,f3501,f3502,f3503,f3504,f3505,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f4415,f4617,f4622,f4624,f4518,f4521,f4522])).
fof(f4628,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f267,f266,f265,f268,f473,f275,f287,f286,f291,f290,f329,f328,f332,f336,f335,f339,f338,f337,f343,f341,f346,f344,f351,f350,f349,f348,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f592,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1220,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1480,f1477,f1496,f1493,f1511,f1514,f1515,f1539,f1536,f1538,f1562,f1578,f1593,f1610,f1622,f1623,f1691,f1700,f1714,f1709,f1710,f1715,f1712,f1717,f1741,f1771,f1800,f1827,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2002,f1928,f1929,f2010,f2040,f2055,f2037,f2043,f2044,f2046,f2048,f2084,f2087,f2088,f2090,f2092,f2126,f2125,f2142,f2129,f2130,f2132,f2170,f2173,f2174,f2176,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2968,f2970,f2971,f2972,f2973,f2974,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4450,f4451,f4452,f4453,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f4463,f4464,f4465,f4466,f4467,f4472,f4473,f4474,f3420,f3421,f3499,f3501,f3502,f3503,f3504,f3505,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f4415,f4617,f4622,f4624,f4518,f4521])).
fof(f4627,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(global_subsumption,[],[f4520,f267,f266,f265,f268,f473,f275,f287,f286,f291,f290,f329,f328,f332,f336,f335,f339,f338,f337,f343,f341,f346,f344,f351,f350,f349,f348,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f592,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1220,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1480,f1477,f1496,f1493,f1511,f1514,f1515,f1539,f1536,f1538,f1562,f1578,f1593,f1610,f1622,f1623,f1691,f1700,f1714,f1709,f1710,f1715,f1712,f1717,f1741,f1771,f1800,f1827,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2002,f1928,f1929,f2010,f2040,f2055,f2037,f2043,f2044,f2046,f2048,f2084,f2087,f2088,f2090,f2092,f2126,f2125,f2142,f2129,f2130,f2132,f2170,f2173,f2174,f2176,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2968,f2970,f2971,f2972,f2973,f2974,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4450,f4451,f4452,f4453,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f4463,f4464,f4465,f4466,f4467,f4472,f4473,f4474,f3420,f3421,f3499,f3501,f3502,f3503,f3504,f3505,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f4415,f4617,f4622,f4624,f4518])).
fof(f4520,plain,(
  ( ! [X6,X7] : (~sP5('F_0',happ(s(fun(X6,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X6,X7)),'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))))) ) | ($spl1 | $spl2 | $spl8)),
  inference(backward_demodulation,[],[f4415,f2045])).
fof(f2045,plain,(
  ( ! [X6,X7] : (~sP5('F_0',happ(s(fun(X6,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X6,X7)),'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,'T_0')))))) ) | ($spl1 | $spl2)),
  inference(superposition,[],[f1206,f556])).
fof(f4626,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f267,f266,f265,f268,f473,f275,f287,f286,f291,f290,f329,f328,f332,f336,f335,f339,f338,f337,f343,f341,f346,f344,f351,f350,f349,f348,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f592,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1220,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1480,f1477,f1496,f1493,f1511,f1514,f1515,f1539,f1536,f1538,f1562,f1578,f1593,f1610,f1622,f1623,f1691,f1700,f1714,f1709,f1710,f1715,f1712,f1717,f1741,f1771,f1800,f1827,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2002,f1928,f1929,f2010,f2040,f2055,f2037,f2043,f2044,f2046,f2048,f2084,f2087,f2088,f2090,f2092,f2126,f2125,f2142,f2129,f2130,f2132,f2170,f2173,f2174,f2176,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2968,f2970,f2971,f2972,f2973,f2974,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4450,f4451,f4452,f4453,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f4463,f4464,f4465,f4466,f4467,f4472,f4473,f4474,f3420,f3421,f3499,f3501,f3502,f3503,f3504,f3505,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f4415,f4617,f4622,f4624,f4518])).
fof(f4625,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f267,f266,f265,f268,f473,f275,f287,f286,f291,f290,f329,f328,f332,f336,f335,f339,f338,f337,f343,f341,f346,f344,f351,f350,f349,f348,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f592,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1220,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1480,f1477,f1496,f1493,f1511,f1514,f1515,f1539,f1536,f1538,f1562,f1578,f1593,f1610,f1622,f1623,f1691,f1700,f1714,f1709,f1710,f1715,f1712,f1717,f1741,f1771,f1800,f1827,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2002,f1928,f1929,f2010,f2040,f2055,f2037,f2043,f2044,f2046,f2048,f2084,f2087,f2088,f2090,f2092,f2126,f2125,f2142,f2129,f2130,f2132,f2170,f2173,f2174,f2176,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2968,f2970,f2971,f2972,f2973,f2974,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4450,f4451,f4452,f4453,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f4463,f4464,f4465,f4466,f4467,f4472,f4473,f4474,f3420,f3421,f3499,f3501,f3502,f3503,f3504,f3505,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f4415,f4617,f4622,f4624])).
fof(f4623,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f267,f266,f265,f268,f473,f275,f287,f286,f291,f290,f329,f328,f332,f336,f335,f339,f338,f337,f343,f341,f346,f344,f351,f350,f349,f348,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f592,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1220,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1480,f1477,f1496,f1493,f1511,f1514,f1515,f1539,f1536,f1538,f1562,f1578,f1593,f1610,f1622,f1623,f1691,f1700,f1714,f1709,f1710,f1715,f1712,f1717,f1741,f1771,f1800,f1827,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2002,f1928,f1929,f2010,f2040,f2055,f2037,f2043,f2044,f2046,f2048,f2084,f2087,f2088,f2090,f2092,f2126,f2125,f2142,f2129,f2130,f2132,f2170,f2173,f2174,f2176,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2968,f2970,f2971,f2972,f2973,f2974,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4450,f4451,f4452,f4453,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f4463,f4464,f4465,f4466,f4467,f4472,f4473,f4474,f3420,f3421,f3499,f3501,f3502,f3503,f3504,f3505,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f4415,f4617,f4622])).
fof(f4621,plain,(
  $false | ($spl2 | $spl8)),
  inference(global_subsumption,[],[f4512,f267,f266,f265,f268,f473,f275,f287,f286,f291,f290,f329,f328,f332,f336,f335,f339,f338,f337,f343,f341,f346,f344,f351,f350,f349,f348,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f592,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1220,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1480,f1477,f1496,f1493,f1511,f1514,f1515,f1539,f1536,f1538,f1562,f1578,f1593,f1610,f1622,f1623,f1691,f1700,f1714,f1709,f1710,f1715,f1712,f1717,f1741,f1771,f1800,f1827,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2002,f1928,f1929,f2010,f2040,f2055,f2037,f2043,f2044,f2046,f2048,f2084,f2087,f2088,f2090,f2092,f2126,f2125,f2142,f2129,f2130,f2132,f2170,f2173,f2174,f2176,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2968,f2970,f2971,f2972,f2973,f2974,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4450,f4451,f4452,f4453,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f4463,f4464,f4465,f4466,f4467,f4472,f4473,f4474,f3420,f3421,f3499,f3501,f3502,f3503,f3504,f3505,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f4415,f4617])).
fof(f4512,plain,(
  ( ! [X0,X1] : (sP1(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)),sK14)) ) | ($spl2 | $spl8)),
  inference(backward_demodulation,[],[f4415,f1624])).
fof(f1624,plain,(
  ( ! [X0,X1] : (sP1(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X0,X1)),sK14)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f473,f1238])).
fof(f1238,plain,(
  ( ! [X68,X66,X67] : (sP1(happ(s(fun(X66,bool),'_dst_x0_1'(s(bool,X67))),s(X66,X68)),sK14) | p(s(bool,X67))) ) | $spl2),
  inference(superposition,[],[f575,f381])).
fof(f4620,plain,(
  $false | ($spl2 | $spl8)),
  inference(global_subsumption,[],[f4505,f267,f266,f265,f268,f473,f275,f287,f286,f291,f290,f329,f328,f332,f336,f335,f339,f338,f337,f343,f341,f346,f344,f351,f350,f349,f348,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f592,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1220,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1480,f1477,f1496,f1493,f1511,f1514,f1515,f1539,f1536,f1538,f1562,f1578,f1593,f1610,f1622,f1623,f1691,f1700,f1714,f1709,f1710,f1715,f1712,f1717,f1741,f1771,f1800,f1827,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2002,f1928,f1929,f2010,f2040,f2055,f2037,f2043,f2044,f2046,f2048,f2084,f2087,f2088,f2090,f2092,f2126,f2125,f2142,f2129,f2130,f2132,f2170,f2173,f2174,f2176,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2968,f2970,f2971,f2972,f2973,f2974,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4450,f4451,f4452,f4453,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f4463,f4464,f4465,f4466,f4467,f4472,f4473,f4474,f3420,f3421,f3499,f3501,f3502,f3503,f3504,f3505,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f4415,f4617])).
fof(f4505,plain,(
  ( ! [X0,X1] : (sP1(sK14,happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)))) ) | ($spl2 | $spl8)),
  inference(backward_demodulation,[],[f4415,f1267])).
fof(f1267,plain,(
  ( ! [X0,X1] : (sP1(sK14,happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X0,X1)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f473,f1237])).
fof(f1237,plain,(
  ( ! [X64,X65,X63] : (sP1(sK14,happ(s(fun(X63,bool),'_dst_x0_1'(s(bool,X64))),s(X63,X65))) | p(s(bool,X64))) ) | $spl2),
  inference(superposition,[],[f566,f381])).
fof(f4619,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f4501,f268])).
fof(f4501,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(backward_demodulation,[],[f4415,f473])).
fof(f4618,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f267,f266,f265,f268,f473,f275,f287,f286,f291,f290,f329,f328,f332,f336,f335,f339,f338,f337,f343,f341,f346,f344,f351,f350,f349,f348,f355,f352,f358,f357,f474,f376,f373,f372,f371,f381,f382,f389,f388,f387,f386,f385,f384,f383,f393,f392,f391,f390,f397,f396,f395,f394,f401,f400,f399,f398,f405,f404,f403,f402,f459,f409,f408,f407,f406,f416,f415,f414,f413,f422,f419,f418,f417,f429,f428,f427,f458,f425,f424,f423,f433,f432,f431,f430,f437,f436,f435,f434,f446,f445,f444,f457,f442,f441,f440,f450,f449,f448,f447,f454,f453,f452,f451,f476,f478,f480,f481,f482,f484,f485,f486,f488,f490,f492,f494,f495,f496,f498,f500,f501,f502,f504,f505,f515,f560,f561,f564,f573,f584,f613,f612,f608,f607,f606,f605,f604,f603,f602,f601,f600,f599,f598,f597,f596,f595,f593,f592,f591,f626,f628,f631,f637,f641,f644,f664,f670,f674,f679,f686,f696,f747,f810,f854,f852,f878,f875,f918,f946,f963,f972,f976,f980,f987,f996,f1059,f1070,f1079,f1139,f1149,f1158,f1170,f1178,f1209,f1208,f1207,f1206,f1220,f1229,f1230,f1231,f1232,f1233,f1234,f1235,f1236,f1332,f1317,f1315,f1311,f1385,f1381,f1453,f1480,f1477,f1496,f1493,f1511,f1514,f1515,f1539,f1536,f1538,f1562,f1578,f1593,f1610,f1622,f1623,f1691,f1700,f1714,f1709,f1710,f1715,f1712,f1717,f1741,f1771,f1800,f1827,f1856,f1852,f1874,f1916,f1915,f1914,f1921,f2002,f1928,f1929,f2010,f2040,f2055,f2037,f2043,f2044,f2046,f2048,f2084,f2087,f2088,f2090,f2092,f2126,f2125,f2142,f2129,f2130,f2132,f2170,f2173,f2174,f2176,f2204,f2228,f2252,f2290,f2742,f2757,f2771,f2810,f2811,f2812,f2813,f2888,f2889,f2968,f2970,f2971,f2972,f2973,f2974,f2975,f3166,f4413,f4417,f4420,f4423,f4430,f4433,f4436,f4439,f4442,f4445,f3233,f3234,f3235,f3236,f4447,f4448,f3250,f4449,f4450,f4451,f4452,f4453,f4454,f4455,f4456,f4457,f4458,f4459,f4460,f4461,f4462,f3342,f3343,f3344,f3345,f4463,f4464,f4465,f4466,f4467,f4472,f4473,f4474,f3420,f3421,f3499,f3501,f3502,f3503,f3504,f3505,f3506,f4476,f4478,f4480,f4482,f4488,f4490,f4492,f4494,f4496,f4498,f4415,f4617])).
fof(f4446,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) ) | $spl26),
  inference(cnf_transformation,[],[f4446_D])).
fof(f4446_D,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f4498,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f4497])).
fof(f4497,plain,(
  ( ! [X198,X200,X199,X201] : (~p(s(bool,X199)) | p(s(bool,X200)) | ~sP11('T_0',X200,X199) | s(bool,X198) = s(bool,X201) | s(bool,'T_0') = s(bool,X201)) )),
  inference(subsumption_resolution,[],[f3571,f810])).
fof(f3571,plain,(
  ( ! [X198,X200,X199,X201] : (~p(s(bool,X198)) | ~p(s(bool,X199)) | p(s(bool,X200)) | ~sP11('T_0',X200,X199) | s(bool,X198) = s(bool,X201) | s(bool,'T_0') = s(bool,X201)) )),
  inference(superposition,[],[f454,f747])).
fof(f4443,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) ) | $spl24),
  inference(cnf_transformation,[],[f4443_D])).
fof(f4443_D,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f4496,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f4495])).
fof(f4495,plain,(
  ( ! [X196,X194,X197,X195] : (p(s(bool,X195)) | p(s(bool,X196)) | ~sP9('T_0',X196,X195) | s(bool,X194) = s(bool,X197) | s(bool,'T_0') = s(bool,X197)) )),
  inference(subsumption_resolution,[],[f3570,f810])).
fof(f3570,plain,(
  ( ! [X196,X194,X197,X195] : (~p(s(bool,X194)) | p(s(bool,X195)) | p(s(bool,X196)) | ~sP9('T_0',X196,X195) | s(bool,X194) = s(bool,X197) | s(bool,'T_0') = s(bool,X197)) )),
  inference(superposition,[],[f437,f747])).
fof(f4440,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) ) | $spl22),
  inference(cnf_transformation,[],[f4440_D])).
fof(f4440_D,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f4494,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f4493])).
fof(f4493,plain,(
  ( ! [X191,X192,X190,X193] : (~p(s(bool,X191)) | p(s(bool,X192)) | ~sP7(X192,X191,'T_0') | s(bool,X190) = s(bool,X193) | s(bool,'T_0') = s(bool,X193)) )),
  inference(subsumption_resolution,[],[f3569,f810])).
fof(f3569,plain,(
  ( ! [X191,X192,X190,X193] : (~p(s(bool,X190)) | ~p(s(bool,X191)) | p(s(bool,X192)) | ~sP7(X192,X191,'T_0') | s(bool,X190) = s(bool,X193) | s(bool,'T_0') = s(bool,X193)) )),
  inference(superposition,[],[f422,f747])).
fof(f4437,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | ~sP4(X1,'T_0',X2) | p(s(bool,X2))) ) | $spl20),
  inference(cnf_transformation,[],[f4437_D])).
fof(f4437_D,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | ~sP4(X1,'T_0',X2) | p(s(bool,X2))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f4492,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f4491])).
fof(f4491,plain,(
  ( ! [X189,X187,X188,X186] : (~p(s(bool,X187)) | p(s(bool,X188)) | ~sP4(X187,'T_0',X188) | s(bool,X186) = s(bool,X189) | s(bool,'T_0') = s(bool,X189)) )),
  inference(subsumption_resolution,[],[f3568,f810])).
fof(f3568,plain,(
  ( ! [X189,X187,X188,X186] : (~p(s(bool,X186)) | ~p(s(bool,X187)) | p(s(bool,X188)) | ~sP4(X187,'T_0',X188) | s(bool,X186) = s(bool,X189) | s(bool,'T_0') = s(bool,X189)) )),
  inference(superposition,[],[f390,f747])).
fof(f4434,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) | $spl18),
  inference(cnf_transformation,[],[f4434_D])).
fof(f4434_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f4490,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f4489])).
fof(f4489,plain,(
  ( ! [X185,X184,X183] : (p(s(bool,X184)) | sP1(X184,'T_0') | s(bool,X183) = s(bool,X185) | s(bool,'T_0') = s(bool,X185)) )),
  inference(subsumption_resolution,[],[f3567,f810])).
fof(f3567,plain,(
  ( ! [X185,X184,X183] : (~p(s(bool,X183)) | p(s(bool,X184)) | sP1(X184,'T_0') | s(bool,X183) = s(bool,X185) | s(bool,'T_0') = s(bool,X185)) )),
  inference(superposition,[],[f376,f747])).
fof(f4431,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) | $spl16),
  inference(cnf_transformation,[],[f4431_D])).
fof(f4431_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f4488,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f4487])).
fof(f4487,plain,(
  ( ! [X182,X180,X181] : (p(s(bool,X181)) | sP1('T_0',X181) | s(bool,X180) = s(bool,X182) | s(bool,'T_0') = s(bool,X182)) )),
  inference(subsumption_resolution,[],[f3566,f810])).
fof(f3566,plain,(
  ( ! [X182,X180,X181] : (~p(s(bool,X180)) | p(s(bool,X181)) | sP1('T_0',X181) | s(bool,X180) = s(bool,X182) | s(bool,'T_0') = s(bool,X182)) )),
  inference(superposition,[],[f373,f747])).
fof(f4424,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK22('T_0',X2,X1))))) | ~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3))))) ) | $spl14),
  inference(cnf_transformation,[],[f4424_D])).
fof(f4424_D,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK22('T_0',X2,X1))))) | ~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3))))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f4482,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK22('T_0',X2,X1))))) | ~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3))))) )),
  inference(condensation,[],[f4481])).
fof(f4481,plain,(
  ( ! [X163,X161,X164,X162,X160] : (p(s(bool,happ(s(fun(X161,bool),X162),s(X161,sK22('T_0',X162,X161))))) | ~p(s(bool,happ(s(fun(X161,bool),X162),s(X161,X163)))) | s(bool,X160) = s(bool,X164) | s(bool,'T_0') = s(bool,X164)) )),
  inference(subsumption_resolution,[],[f3560,f810])).
fof(f3560,plain,(
  ( ! [X163,X161,X164,X162,X160] : (~p(s(bool,X160)) | p(s(bool,happ(s(fun(X161,bool),X162),s(X161,sK22('T_0',X162,X161))))) | ~p(s(bool,happ(s(fun(X161,bool),X162),s(X161,X163)))) | s(bool,X160) = s(bool,X164) | s(bool,'T_0') = s(bool,X164)) )),
  inference(superposition,[],[f346,f747])).
fof(f4421,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK21('T_0',X2,X1))))) | ~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3))))) ) | $spl12),
  inference(cnf_transformation,[],[f4421_D])).
fof(f4421_D,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK21('T_0',X2,X1))))) | ~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3))))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f4480,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK21('T_0',X2,X1))))) | ~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3))))) )),
  inference(condensation,[],[f4479])).
fof(f4479,plain,(
  ( ! [X158,X156,X159,X157,X155] : (p(s(bool,happ(s(fun(X156,bool),X157),s(X156,sK21('T_0',X157,X156))))) | ~p(s(bool,happ(s(fun(X156,bool),X157),s(X156,X158)))) | s(bool,X155) = s(bool,X159) | s(bool,'T_0') = s(bool,X159)) )),
  inference(subsumption_resolution,[],[f3559,f810])).
fof(f3559,plain,(
  ( ! [X158,X156,X159,X157,X155] : (~p(s(bool,X155)) | p(s(bool,happ(s(fun(X156,bool),X157),s(X156,sK21('T_0',X157,X156))))) | ~p(s(bool,happ(s(fun(X156,bool),X157),s(X156,X158)))) | s(bool,X155) = s(bool,X159) | s(bool,'T_0') = s(bool,X159)) )),
  inference(superposition,[],[f344,f747])).
fof(f4418,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3)))) | p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK20(X2,'T_0',X1)))))) ) | $spl10),
  inference(cnf_transformation,[],[f4418_D])).
fof(f4418_D,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3)))) | p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK20(X2,'T_0',X1)))))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f4478,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3)))) | p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK20(X2,'T_0',X1)))))) )),
  inference(condensation,[],[f4477])).
fof(f4477,plain,(
  ( ! [X154,X152,X151,X153,X150] : (~p(s(bool,happ(s(fun(X151,bool),X152),s(X151,X153)))) | p(s(bool,happ(s(fun(X151,bool),X152),s(X151,sK20(X152,'T_0',X151))))) | s(bool,X150) = s(bool,X154) | s(bool,'T_0') = s(bool,X154)) )),
  inference(subsumption_resolution,[],[f3558,f810])).
fof(f3558,plain,(
  ( ! [X154,X152,X151,X153,X150] : (~p(s(bool,X150)) | ~p(s(bool,happ(s(fun(X151,bool),X152),s(X151,X153)))) | p(s(bool,happ(s(fun(X151,bool),X152),s(X151,sK20(X152,'T_0',X151))))) | s(bool,X150) = s(bool,X154) | s(bool,'T_0') = s(bool,X154)) )),
  inference(superposition,[],[f343,f747])).
fof(f4414,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3)))) | p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK19(X2,'T_0',X1)))))) ) | $spl6),
  inference(cnf_transformation,[],[f4414_D])).
fof(f4414_D,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3)))) | p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK19(X2,'T_0',X1)))))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f4476,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3)))) | p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK19(X2,'T_0',X1)))))) )),
  inference(condensation,[],[f4475])).
fof(f4475,plain,(
  ( ! [X146,X149,X147,X145,X148] : (~p(s(bool,happ(s(fun(X146,bool),X147),s(X146,X148)))) | p(s(bool,happ(s(fun(X146,bool),X147),s(X146,sK19(X147,'T_0',X146))))) | s(bool,X145) = s(bool,X149) | s(bool,'T_0') = s(bool,X149)) )),
  inference(subsumption_resolution,[],[f3557,f810])).
fof(f3557,plain,(
  ( ! [X146,X149,X147,X145,X148] : (~p(s(bool,X145)) | ~p(s(bool,happ(s(fun(X146,bool),X147),s(X146,X148)))) | p(s(bool,happ(s(fun(X146,bool),X147),s(X146,sK19(X147,'T_0',X146))))) | s(bool,X145) = s(bool,X149) | s(bool,'T_0') = s(bool,X149)) )),
  inference(superposition,[],[f341,f747])).
fof(f4445,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f4444])).
fof(f4444,plain,(
  ( ! [X198,X200,X199,X201] : (~p(s(bool,X199)) | p(s(bool,X200)) | ~sP11('T_0',X200,X199) | s(bool,X198) = s(bool,X201) | s(bool,'T_0') = s(bool,X201)) )),
  inference(subsumption_resolution,[],[f3229,f810])).
fof(f3229,plain,(
  ( ! [X198,X200,X199,X201] : (~p(s(bool,X198)) | ~p(s(bool,X199)) | p(s(bool,X200)) | ~sP11('T_0',X200,X199) | s(bool,X198) = s(bool,X201) | s(bool,'T_0') = s(bool,X201)) )),
  inference(superposition,[],[f454,f747])).
fof(f4442,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f4441])).
fof(f4441,plain,(
  ( ! [X196,X194,X197,X195] : (p(s(bool,X195)) | p(s(bool,X196)) | ~sP9('T_0',X196,X195) | s(bool,X194) = s(bool,X197) | s(bool,'T_0') = s(bool,X197)) )),
  inference(subsumption_resolution,[],[f3228,f810])).
fof(f3228,plain,(
  ( ! [X196,X194,X197,X195] : (~p(s(bool,X194)) | p(s(bool,X195)) | p(s(bool,X196)) | ~sP9('T_0',X196,X195) | s(bool,X194) = s(bool,X197) | s(bool,'T_0') = s(bool,X197)) )),
  inference(superposition,[],[f437,f747])).
fof(f4439,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f4438])).
fof(f4438,plain,(
  ( ! [X191,X192,X190,X193] : (~p(s(bool,X191)) | p(s(bool,X192)) | ~sP7(X192,X191,'T_0') | s(bool,X190) = s(bool,X193) | s(bool,'T_0') = s(bool,X193)) )),
  inference(subsumption_resolution,[],[f3227,f810])).
fof(f3227,plain,(
  ( ! [X191,X192,X190,X193] : (~p(s(bool,X190)) | ~p(s(bool,X191)) | p(s(bool,X192)) | ~sP7(X192,X191,'T_0') | s(bool,X190) = s(bool,X193) | s(bool,'T_0') = s(bool,X193)) )),
  inference(superposition,[],[f422,f747])).
fof(f4436,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f4435])).
fof(f4435,plain,(
  ( ! [X189,X187,X188,X186] : (~p(s(bool,X187)) | p(s(bool,X188)) | ~sP4(X187,'T_0',X188) | s(bool,X186) = s(bool,X189) | s(bool,'T_0') = s(bool,X189)) )),
  inference(subsumption_resolution,[],[f3226,f810])).
fof(f3226,plain,(
  ( ! [X189,X187,X188,X186] : (~p(s(bool,X186)) | ~p(s(bool,X187)) | p(s(bool,X188)) | ~sP4(X187,'T_0',X188) | s(bool,X186) = s(bool,X189) | s(bool,'T_0') = s(bool,X189)) )),
  inference(superposition,[],[f390,f747])).
fof(f4433,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f4432])).
fof(f4432,plain,(
  ( ! [X185,X184,X183] : (p(s(bool,X184)) | sP1(X184,'T_0') | s(bool,X183) = s(bool,X185) | s(bool,'T_0') = s(bool,X185)) )),
  inference(subsumption_resolution,[],[f3225,f810])).
fof(f3225,plain,(
  ( ! [X185,X184,X183] : (~p(s(bool,X183)) | p(s(bool,X184)) | sP1(X184,'T_0') | s(bool,X183) = s(bool,X185) | s(bool,'T_0') = s(bool,X185)) )),
  inference(superposition,[],[f376,f747])).
fof(f4430,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f4429])).
fof(f4429,plain,(
  ( ! [X182,X180,X181] : (p(s(bool,X181)) | sP1('T_0',X181) | s(bool,X180) = s(bool,X182) | s(bool,'T_0') = s(bool,X182)) )),
  inference(subsumption_resolution,[],[f3224,f810])).
fof(f3224,plain,(
  ( ! [X182,X180,X181] : (~p(s(bool,X180)) | p(s(bool,X181)) | sP1('T_0',X181) | s(bool,X180) = s(bool,X182) | s(bool,'T_0') = s(bool,X182)) )),
  inference(superposition,[],[f373,f747])).
fof(f4423,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK22('T_0',X2,X1))))) | ~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3))))) )),
  inference(condensation,[],[f4422])).
fof(f4422,plain,(
  ( ! [X163,X161,X164,X162,X160] : (p(s(bool,happ(s(fun(X161,bool),X162),s(X161,sK22('T_0',X162,X161))))) | ~p(s(bool,happ(s(fun(X161,bool),X162),s(X161,X163)))) | s(bool,X160) = s(bool,X164) | s(bool,'T_0') = s(bool,X164)) )),
  inference(subsumption_resolution,[],[f3218,f810])).
fof(f3218,plain,(
  ( ! [X163,X161,X164,X162,X160] : (~p(s(bool,X160)) | p(s(bool,happ(s(fun(X161,bool),X162),s(X161,sK22('T_0',X162,X161))))) | ~p(s(bool,happ(s(fun(X161,bool),X162),s(X161,X163)))) | s(bool,X160) = s(bool,X164) | s(bool,'T_0') = s(bool,X164)) )),
  inference(superposition,[],[f346,f747])).
fof(f4420,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK21('T_0',X2,X1))))) | ~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3))))) )),
  inference(condensation,[],[f4419])).
fof(f4419,plain,(
  ( ! [X158,X156,X159,X157,X155] : (p(s(bool,happ(s(fun(X156,bool),X157),s(X156,sK21('T_0',X157,X156))))) | ~p(s(bool,happ(s(fun(X156,bool),X157),s(X156,X158)))) | s(bool,X155) = s(bool,X159) | s(bool,'T_0') = s(bool,X159)) )),
  inference(subsumption_resolution,[],[f3217,f810])).
fof(f3217,plain,(
  ( ! [X158,X156,X159,X157,X155] : (~p(s(bool,X155)) | p(s(bool,happ(s(fun(X156,bool),X157),s(X156,sK21('T_0',X157,X156))))) | ~p(s(bool,happ(s(fun(X156,bool),X157),s(X156,X158)))) | s(bool,X155) = s(bool,X159) | s(bool,'T_0') = s(bool,X159)) )),
  inference(superposition,[],[f344,f747])).
fof(f4417,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3)))) | p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK20(X2,'T_0',X1)))))) )),
  inference(condensation,[],[f4416])).
fof(f4416,plain,(
  ( ! [X154,X152,X151,X153,X150] : (~p(s(bool,happ(s(fun(X151,bool),X152),s(X151,X153)))) | p(s(bool,happ(s(fun(X151,bool),X152),s(X151,sK20(X152,'T_0',X151))))) | s(bool,X150) = s(bool,X154) | s(bool,'T_0') = s(bool,X154)) )),
  inference(subsumption_resolution,[],[f3216,f810])).
fof(f3216,plain,(
  ( ! [X154,X152,X151,X153,X150] : (~p(s(bool,X150)) | ~p(s(bool,happ(s(fun(X151,bool),X152),s(X151,X153)))) | p(s(bool,happ(s(fun(X151,bool),X152),s(X151,sK20(X152,'T_0',X151))))) | s(bool,X150) = s(bool,X154) | s(bool,'T_0') = s(bool,X154)) )),
  inference(superposition,[],[f343,f747])).
fof(f4413,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3)))) | p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK19(X2,'T_0',X1)))))) )),
  inference(condensation,[],[f4412])).
fof(f4412,plain,(
  ( ! [X146,X149,X147,X145,X148] : (~p(s(bool,happ(s(fun(X146,bool),X147),s(X146,X148)))) | p(s(bool,happ(s(fun(X146,bool),X147),s(X146,sK19(X147,'T_0',X146))))) | s(bool,X145) = s(bool,X149) | s(bool,'T_0') = s(bool,X149)) )),
  inference(subsumption_resolution,[],[f3215,f810])).
fof(f3215,plain,(
  ( ! [X146,X149,X147,X145,X148] : (~p(s(bool,X145)) | ~p(s(bool,happ(s(fun(X146,bool),X147),s(X146,X148)))) | p(s(bool,happ(s(fun(X146,bool),X147),s(X146,sK19(X147,'T_0',X146))))) | s(bool,X145) = s(bool,X149) | s(bool,'T_0') = s(bool,X149)) )),
  inference(superposition,[],[f341,f747])).
fof(f467,plain,(
  p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,sK14)))))) | $spl0),
  inference(cnf_transformation,[],[f467_D])).
fof(f467_D,plain,(
  p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,sK14)))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f471,plain,(
  s(fun(sK12,bool),'const.pred_set.EMPTY_0') = s(fun(sK12,bool),sK13) | $spl4),
  inference(cnf_transformation,[],[f471_D])).
fof(f471_D,plain,(
  s(fun(sK12,bool),'const.pred_set.EMPTY_0') = s(fun(sK12,bool),sK13) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f469,plain,(
  p(s(bool,sK14)) | $spl2),
  inference(cnf_transformation,[],[f469_D])).
fof(f469_D,plain,(
  p(s(bool,sK14)) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f265,plain,(
  p(s(bool,sK14)) | s(fun(sK12,bool),'const.pred_set.EMPTY_0') = s(fun(sK12,bool),sK13) | p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,sK14))))))),
  inference(cnf_transformation,[],[f166])).
fof(f166,plain,(
  (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,sK14)))))) | s(fun(sK12,bool),'const.pred_set.EMPTY_0') = s(fun(sK12,bool),sK13) | p(s(bool,sK14))) & (~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,sK14)))))) | (s(fun(sK12,bool),'const.pred_set.EMPTY_0') != s(fun(sK12,bool),sK13) & ~p(s(bool,sK14))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14])],[f165])).
fof(f165,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'_dst_x0_1'(s(bool,X2)))))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1) | p(s(bool,X2))) & (~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'_dst_x0_1'(s(bool,X2)))))) | (s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),X1) & ~p(s(bool,X2)))))),
  inference(flattening,[],[f164])).
fof(f164,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'_dst_x0_1'(s(bool,X2)))))) | (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1) | p(s(bool,X2)))) & (~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'_dst_x0_1'(s(bool,X2)))))) | (s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),X1) & ~p(s(bool,X2)))))),
  inference(nnf_transformation,[],[f122])).
fof(f122,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'_dst_x0_1'(s(bool,X2)))))) <~> (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1) | p(s(bool,X2))))),
  inference(ennf_transformation,[],[f51])).
fof(f51,plain,(
  ~! [X0,X1,X2] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'_dst_x0_1'(s(bool,X2)))))) <=> (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1) | p(s(bool,X2))))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ~! [X0] : ! [X1,X2] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'_dst_x0_1'(s(bool,X2)))))) <=> (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1) | p(s(bool,X2))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X6,bool),X7),s(fun(X6,bool),'_dst_x0_1'(s(bool,X8)))))) <=> (s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),X7) | p(s(bool,X8))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : (p(s(bool,'const.bool.RES_FORALL_2'(s(fun(X6,bool),X7),s(fun(X6,bool),'_dst_x0_1'(s(bool,X8)))))) <=> (s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),X7) | p(s(bool,X8))))),
  file('thm.res_quan.RES_FORALL_NULL','thm.res_quan.RES_FORALL_NULL')).
fof(f468,plain,(
  ~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,sK14)))))) | $spl1),
  inference(cnf_transformation,[],[f468_D])).
fof(f468_D,plain,(
  ~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,sK14)))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f472,plain,(
  s(fun(sK12,bool),'const.pred_set.EMPTY_0') != s(fun(sK12,bool),sK13) | $spl5),
  inference(cnf_transformation,[],[f472_D])).
fof(f472_D,plain,(
  s(fun(sK12,bool),'const.pred_set.EMPTY_0') != s(fun(sK12,bool),sK13) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f266,plain,(
  s(fun(sK12,bool),'const.pred_set.EMPTY_0') != s(fun(sK12,bool),sK13) | ~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,sK14))))))),
  inference(cnf_transformation,[],[f166])).
fof(f470,plain,(
  ~p(s(bool,sK14)) | $spl3),
  inference(cnf_transformation,[],[f470_D])).
fof(f470_D,plain,(
  ~p(s(bool,sK14)) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f267,plain,(
  ~p(s(bool,sK14)) | ~p(s(bool,'const.bool.RES_FORALL_2'(s(fun(sK12,bool),sK13),s(fun(sK12,bool),'_dst_x0_1'(s(bool,sK14))))))),
  inference(cnf_transformation,[],[f166])).
% SZS output end Proof for thm.res_quan
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 819
% Passive clauses: 2028
% Generated clauses: 15315
% Final active clauses: 311
% Final passive clauses: 268
% Input formulas: 48
% Initial clauses: 192
% 
% Duplicate literals: 383
% Trivial inequalities: 36
% Fw subsumption resolutions: 680
% Fw demodulations: 1428
% Bw demodulations: 449
% Condensations: 753
% Global subsumptions: 62
% 
% Simple tautologies: 110
% Equational tautologies: 110
% Forward subsumptions: 9137
% Fw demodulations to eq. taut.: 4
% Bw demodulations to eq. taut.: 2
% 
% Unit resulting resolution: 5976
% Factoring: 1
% Forward superposition: 2709
% Backward superposition: 2122
% Self superposition: 32
% Equality factoring: 30
% Equality resolution: 5
% 
% Split clauses: 56
% Split components: 118
% SAT solver clauses: 2732
% SAT solver unit clauses: 492
% SAT solver binary clauses: 825
% 
% Sat splits: 56
% Sat splitting refutations: 528
% 
% Memory used [KB]: 6268
% Time elapsed: 0.406 s
% ------------------------------
% ------------------------------
% Success in time 1.744 s
