Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.transc
% (43444)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1077
% Passive clauses: 5054
% Generated clauses: 36994
% Final active clauses: 1054
% Final passive clauses: 3505
% Input formulas: 260
% Initial clauses: 423
% 
% Duplicate literals: 272
% Trivial inequalities: 133
% Fw subsumption resolutions: 1744
% Fw demodulations: 2270
% Bw demodulations: 15
% 
% Simple tautologies: 218
% Equational tautologies: 74
% Forward subsumptions: 26634
% Fw demodulations to eq. taut.: 5
% 
% Binary resolution: 7732
% Factoring: 1340
% Forward superposition: 10521
% Backward superposition: 12268
% Self superposition: 81
% Equality factoring: 2
% Equality resolution: 24
% 
% Split clauses: 638
% Split components: 1278
% SAT solver clauses: 1028
% SAT solver unit clauses: 16
% SAT solver binary clauses: 96
% 
% Sat splits: 638
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 1028
% TWLsolver calls for satisfiability: 517
% 
% Memory used [KB]: 19573
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.transc
% (43465)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 432
% Passive clauses: 1817
% Generated clauses: 30761
% Final active clauses: 404
% Final passive clauses: 1132
% Input formulas: 260
% Initial clauses: 423
% 
% Duplicate literals: 84
% Trivial inequalities: 1008
% Fw subsumption resolutions: 1815
% Fw demodulations: 10086
% Bw demodulations: 270
% Condensations: 59
% Global subsumptions: 22
% 
% Simple tautologies: 230
% Equational tautologies: 7
% Forward subsumptions: 12910
% Fw demodulations to eq. taut.: 5
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 576
% Factoring: 4
% Forward superposition: 8577
% Backward superposition: 5830
% Self superposition: 53
% Equality factoring: 2
% Equality resolution: 66
% 
% Split clauses: 91
% Split components: 183
% SAT solver clauses: 2952
% SAT solver unit clauses: 590
% SAT solver binary clauses: 121
% SAT solver learnt clauses: 44
% SAT solver learnt literals: 10
% 
% Sat splits: 91
% Sat splitting refutations: 30
% 
% TWLsolver clauses: 2864
% TWLsolver calls for satisfiability: 2401
% 
% Memory used [KB]: 22643
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.transc
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.transc
% SZS output start Proof for thm.transc
fof(f42793,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f762,f1199,f1197,f1195,f761,f1198,f760,f1196,f1194,f5444,f5447,f5433,f5453,f5467,f5468,f5470,f5473,f5472,f5474,f5476,f5478,f5479,f5482,f5485,f5484,f5487,f5488,f5483,f5489,f5490,f5491,f5492,f5493,f5494,f5496,f5497,f5499,f5500,f5502,f5503,f5504,f5506,f5508,f5509,f5511,f5445,f5455,f5536,f5537,f5538,f5539,f5540,f5541,f5542,f5543,f5544,f5545,f5612,f6046,f6047,f6048,f6049,f6050,f6052,f6053,f6054,f6055,f6056,f6057,f6058,f6059,f6060,f6061,f6062,f6063,f6064,f6065,f6066,f6067,f6068,f6069,f6071,f6072,f6075,f6076,f6077,f6078,f6079,f6080,f6081,f6082,f6083,f6084,f6085,f6086,f6087,f6088,f6089,f6090,f6091,f6092,f6093,f6094,f6096,f6097,f6100,f6106,f6107,f6108,f6109,f6110,f6111,f6112,f6113,f6114,f6115,f6116,f6117,f6118,f6119,f6120,f6121,f6122,f6123,f6124,f6125,f6126,f6127,f6128,f6129,f6130,f6131,f6132,f6133,f6134,f6135,f6136,f6137,f6138,f6139,f6140,f6141,f6144,f6145,f6146,f6147,f6251,f6261,f6264,f6270,f6275,f6277,f6279,f6293,f6294,f6295,f6296,f6299,f6300,f6301,f6302,f6303,f6304,f6305,f6306,f6307,f6308,f6309,f6310,f6311,f6312,f6328,f6332,f6330,f6329,f6346,f6347,f6348,f6349,f6350,f6351,f6352,f6353,f6354,f6355,f6356,f6357,f6358,f6359,f6360,f6361,f6362,f6363,f6364,f6365,f6366,f6367,f6368,f6369,f6370,f6371,f6372,f6373,f6374,f6375,f6376,f6377,f6378,f6379,f6380,f6381,f6382,f6385,f6388,f6389,f6390,f6391,f6402,f6403,f6404,f6405,f6406,f6410,f6411,f6412,f6413,f6414,f6415,f6416,f6417,f6418,f10929,f10930,f10931,f10932,f10933,f10934,f10935,f10936,f12180,f12168,f12182,f12184,f12542,f12545,f12543,f15282,f15329,f15327,f15283,f15328,f15326,f19966,f19968,f24146,f24242,f24261,f31056,f31059,f31058,f42769])).
fof(f42769,plain,(
  $false | ($spl1 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f42639,f25015])).
fof(f25015,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f1195,f1196])).
fof(f42639,plain,(
  s('type.realax.real',sK10) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))) | ($spl2 | $spl5)),
  inference(superposition,[],[f24273,f20003])).
fof(f20003,plain,(
  s('type.num.num',sK11) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK12(sK11)))) | $spl5),
  inference(unit_resulting_resolution,[],[f1198,f856])).
fof(f856,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK12(X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f652])).
fof(f652,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK12(X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12])],[f344])).
fof(f344,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | ? [X1] : s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))),
  inference(rectify,[],[f156])).
fof(f156,axiom,(
  ! [X21] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X21) | ? [X7] : s('type.num.num',X21) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.arithmetic.num_CASES')).
fof(f24273,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) ) | $spl2),
  inference(backward_demodulation,[],[f1196,f833])).
fof(f833,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f321])).
fof(f321,plain,(
  ! [X0] : s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7] : s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7)))))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.real.POW_|48|')).
fof(f31058,plain,(
  p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.tan_0'),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10)))))) | $spl62),
  inference(cnf_transformation,[],[f31058_D])).
fof(f31058_D,plain,(
  p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.tan_0'),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10)))))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f31059,plain,(
  ( ! [X29] : (s('type.realax.real',sK10) = s('type.realax.real','const.transc.cos_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X29),s('type.realax.real',X29))))) | s('type.realax.real',sK10) = s('type.realax.real',X29)) ) | $spl64),
  inference(cnf_transformation,[],[f31059_D])).
fof(f31059_D,plain,(
  ( ! [X29] : (s('type.realax.real',sK10) = s('type.realax.real','const.transc.cos_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X29),s('type.realax.real',X29))))) | s('type.realax.real',sK10) = s('type.realax.real',X29)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f31056,plain,(
  ( ! [X29] : (s('type.realax.real',sK10) = s('type.realax.real','const.transc.cos_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X29),s('type.realax.real',X29))))) | p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.tan_0'),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10)))))) | s('type.realax.real',sK10) = s('type.realax.real',X29)) ) | $spl2),
  inference(forward_demodulation,[],[f31055,f958])).
fof(f958,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1)))))) )),
  inference(cnf_transformation,[],[f414])).
fof(f414,plain,(
  ! [X0,X1] : s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1)))))),
  inference(rectify,[],[f194])).
fof(f194,axiom,(
  ! [X6,X15] : s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X6),s('type.realax.real',X15))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X15)))))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.real.real_div')).
fof(f31055,plain,(
  ( ! [X29] : (s('type.realax.real',sK10) = s('type.realax.real','const.transc.cos_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X29),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X29))))))) | p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.tan_0'),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10)))))) | s('type.realax.real',sK10) = s('type.realax.real',X29)) ) | $spl2),
  inference(forward_demodulation,[],[f31006,f952])).
fof(f952,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) )),
  inference(cnf_transformation,[],[f408])).
fof(f408,plain,(
  ! [X0,X1] : s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0)))),
  inference(rectify,[],[f97])).
fof(f97,axiom,(
  ! [X6,X15] : s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real',X15))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X15),s('type.realax.real',X6)))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.realax.REAL_MUL_SYM')).
fof(f31006,plain,(
  ( ! [X29] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.tan_0'),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10)))))) | s('type.realax.real',sK10) = s('type.realax.real','const.transc.cos_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X29))),s('type.realax.real',X29))))) | s('type.realax.real',sK10) = s('type.realax.real',X29)) ) | $spl2),
  inference(superposition,[],[f24306,f24849])).
fof(f24849,plain,(
  ( ! [X0] : (s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.realax.real',X0))) | s('type.realax.real',sK10) = s('type.realax.real',X0)) ) | $spl2),
  inference(forward_demodulation,[],[f24312,f24301])).
fof(f24301,plain,(
  s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl2),
  inference(backward_demodulation,[],[f1196,f1202])).
fof(f1202,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.transc.cos_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f1201,f1200])).
fof(f1200,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f768,f765])).
fof(f765,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.ZERO_0')),
  inference(cnf_transformation,[],[f79])).
fof(f79,axiom,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.ZERO_0')),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.arithmetic.ALT_ZERO')).
fof(f768,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f254])).
fof(f254,axiom,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.numeral.numeral_distrib|split|13')).
fof(f1201,plain,(
  s('type.realax.real','const.transc.cos_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(forward_demodulation,[],[f772,f765])).
fof(f772,plain,(
  s('type.realax.real','const.transc.cos_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(cnf_transformation,[],[f205])).
fof(f205,axiom,(
  s('type.realax.real','const.transc.cos_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.transc.COS_|48|')).
fof(f24312,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) = s('type.realax.real',X0) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.realax.real',X0)))) ) | $spl2),
  inference(backward_demodulation,[],[f1196,f10947])).
fof(f10947,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.realax.real',X0))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f1260,f10496])).
fof(f10496,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f1200,f781])).
fof(f781,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f73])).
fof(f73,axiom,(
  ! [X6] : s('type.num.num',X6) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6)))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.arithmetic.NUMERAL_DEF')).
fof(f1260,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.realax.real',X0))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f1259,f765])).
fof(f1259,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.realax.real',X0))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f866,f781])).
fof(f866,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.realax.real',X0))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f563])).
fof(f563,plain,(
  ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.realax.real',X0))))),
  inference(ennf_transformation,[],[f353])).
fof(f353,plain,(
  ! [X0] : (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.realax.real',X0))))),
  inference(flattening,[],[f352])).
fof(f352,plain,(
  ! [X0] : (~s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.realax.real',X0))))),
  inference(rectify,[],[f198])).
fof(f198,axiom,(
  ! [X6] : (~s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X6))),s('type.realax.real',X6))))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.real.REAL_MUL_LINV')).
fof(f24306,plain,(
  ( ! [X0] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.tan_0'),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real',X0)))) | s('type.realax.real',sK10) = s('type.realax.real','const.transc.cos_1'(s('type.realax.real',X0)))) ) | $spl2),
  inference(backward_demodulation,[],[f1196,f1265])).
fof(f1265,plain,(
  ( ! [X0] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.tan_0'),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real',X0)))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.cos_1'(s('type.realax.real',X0)))) )),
  inference(forward_demodulation,[],[f1264,f765])).
fof(f1264,plain,(
  ( ! [X0] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.tan_0'),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X0)))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.cos_1'(s('type.realax.real',X0)))) )),
  inference(forward_demodulation,[],[f868,f781])).
fof(f868,plain,(
  ( ! [X0] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.tan_0'),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real',X0)))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.cos_1'(s('type.realax.real',X0)))) )),
  inference(cnf_transformation,[],[f565])).
fof(f565,plain,(
  ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.cos_1'(s('type.realax.real',X0))) | p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.tan_0'),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real',X0)))))),
  inference(ennf_transformation,[],[f357])).
fof(f357,plain,(
  ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.transc.cos_1'(s('type.realax.real',X0))) => p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.tan_0'),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real',X0)))))),
  inference(flattening,[],[f356])).
fof(f356,plain,(
  ! [X0] : (~s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.cos_1'(s('type.realax.real',X0))) => p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.tan_0'),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real',X0)))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X6] : (~s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.cos_1'(s('type.realax.real',X6))) => p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'const.transc.tan_0'),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',X6))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.realax.real',X6)))))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.transc.DIFF_TAN')).
fof(f24261,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f24241,f1197])).
fof(f24241,plain,(
  s('type.realax.real',sK10) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl0),
  inference(trivial_inequality_removal,[],[f24240])).
fof(f24240,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',sK10) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl0),
  inference(superposition,[],[f1011,f1194])).
fof(f1011,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X1),s('type.num.num',X0))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f576])).
fof(f576,plain,(
  ! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X1),s('type.num.num',X0))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(ennf_transformation,[],[f467])).
fof(f467,plain,(
  ! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',X1),s('type.num.num',X0))) => s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X7,X6] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',X6),s('type.num.num',X7))) => s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.real.POW_ZERO')).
fof(f24242,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f24145,f1197])).
fof(f24145,plain,(
  s('type.realax.real',sK10) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl0),
  inference(unit_resulting_resolution,[],[f1194,f1011])).
fof(f24146,plain,(
  $false | ($spl0 | $spl3)),
  inference(unit_resulting_resolution,[],[f1197,f1194,f1011])).
fof(f19968,plain,(
  $false | ($spl0 | $spl4)),
  inference(subsumption_resolution,[],[f19942,f19853])).
fof(f19853,plain,(
  ( ! [X6] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X6),s('type.num.num','const.num.|48|_0')))))) )),
  inference(subsumption_resolution,[],[f19813,f1208])).
fof(f1208,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f1207,f1200])).
fof(f1207,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(forward_demodulation,[],[f780,f765])).
fof(f780,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(cnf_transformation,[],[f268])).
fof(f268,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(flattening,[],[f27])).
fof(f27,axiom,(
  ~s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.real.REAL_|49||48|')).
fof(f19813,plain,(
  ( ! [X6] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X6),s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f886,f10546])).
fof(f10546,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(backward_demodulation,[],[f10496,f1236])).
fof(f1236,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f1235,f765])).
fof(f1235,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f839,f781])).
fof(f839,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(cnf_transformation,[],[f327])).
fof(f327,plain,(
  ! [X0] : s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6] : s('type.realax.real','const.real.pow_2'(s('type.realax.real',X6),s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.real.pow|split|0')).
fof(f886,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f665])).
fof(f665,plain,(
  ! [X0] : ((s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0)))))),
  inference(nnf_transformation,[],[f370])).
fof(f370,plain,(
  ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0))) <=> s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f183])).
fof(f183,axiom,(
  ! [X6] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X6))) <=> s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.real.ABS_ZERO')).
fof(f19942,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl4)),
  inference(trivial_inequality_removal,[],[f19890])).
fof(f19890,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl4)),
  inference(superposition,[],[f887,f15333])).
fof(f15333,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl4)),
  inference(backward_demodulation,[],[f1199,f1194])).
fof(f887,plain,(
  ( ! [X0] : (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0)))) )),
  inference(cnf_transformation,[],[f665])).
fof(f19966,plain,(
  $false | ($spl0 | $spl4)),
  inference(subsumption_resolution,[],[f19864,f19853])).
fof(f19864,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl4)),
  inference(unit_resulting_resolution,[],[f15333,f887])).
fof(f15326,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num',sK12(sK11))))) | $spl59),
  inference(cnf_transformation,[],[f15326_D])).
fof(f15326_D,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num',sK12(sK11))))) <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f15328,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num',sK11)))) | $spl61),
  inference(cnf_transformation,[],[f15328_D])).
fof(f15328_D,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num',sK11)))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f15283,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num',sK11)))) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num',sK12(sK11))))) | $spl5),
  inference(superposition,[],[f884,f13508])).
fof(f13508,plain,(
  s('type.num.num',sK11) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK12(sK11)))) | $spl5),
  inference(unit_resulting_resolution,[],[f1198,f856])).
fof(f884,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f664])).
fof(f664,plain,(
  ! [X0] : ((~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X0))))) & (p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X0)))) | p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))))),
  inference(nnf_transformation,[],[f369])).
fof(f369,plain,(
  ! [X0] : (p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) <=> ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X0)))))),
  inference(flattening,[],[f368])).
fof(f368,plain,(
  ! [X0] : (p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) <=> ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f75])).
fof(f75,axiom,(
  ! [X7] : (p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7)))))) <=> ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X7)))))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.arithmetic.ODD|split|1')).
fof(f15327,plain,(
  p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num',sK12(sK11))))) | $spl58),
  inference(cnf_transformation,[],[f15327_D])).
fof(f15327_D,plain,(
  p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num',sK12(sK11))))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f15329,plain,(
  p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num',sK11)))) | $spl60),
  inference(cnf_transformation,[],[f15329_D])).
fof(f15329_D,plain,(
  p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num',sK11)))) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f15282,plain,(
  p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num',sK11)))) | p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num',sK12(sK11))))) | $spl5),
  inference(superposition,[],[f885,f13508])).
fof(f885,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) | p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f664])).
fof(f12543,plain,(
  ( ! [X7] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X7),s('type.num.num','const.num.|48|_0')))) ) | $spl54),
  inference(cnf_transformation,[],[f12543_D])).
fof(f12543_D,plain,(
  ( ! [X7] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X7),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f12545,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl57),
  inference(cnf_transformation,[],[f12545_D])).
fof(f12545_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl57),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl57])])).
fof(f12542,plain,(
  ( ! [X7] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X7),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f12527,f781])).
fof(f12527,plain,(
  ( ! [X7] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X7),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1267,f10541])).
fof(f10541,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(backward_demodulation,[],[f10496,f1215])).
fof(f1215,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f1214,f765])).
fof(f1214,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f819,f781])).
fof(f819,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f307])).
fof(f307,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f76])).
fof(f76,axiom,(
  ! [X21] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X21),s('type.num.num','const.num.|48|_0')))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.arithmetic.EXP|split|0')).
fof(f1267,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(forward_demodulation,[],[f878,f765])).
fof(f878,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f661])).
fof(f661,plain,(
  ! [X0] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) & (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(nnf_transformation,[],[f365])).
fof(f365,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0))),
  inference(rectify,[],[f245])).
fof(f245,axiom,(
  ! [X7] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X7))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.numeral.numeral_distrib|split|17')).
fof(f12184,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | $spl53),
  inference(cnf_transformation,[],[f12184_D])).
fof(f12184_D,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f12182,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))) | $spl51),
  inference(cnf_transformation,[],[f12182_D])).
fof(f12182_D,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f12168,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(superposition,[],[f884,f1209])).
fof(f1209,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(backward_demodulation,[],[f781,f1204])).
fof(f1204,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f1203,f1200])).
fof(f1203,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(forward_demodulation,[],[f773,f765])).
fof(f773,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(cnf_transformation,[],[f123])).
fof(f123,axiom,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.arithmetic.TWO')).
fof(f12180,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(forward_demodulation,[],[f12179,f10541])).
fof(f12179,plain,(
  ( ! [X7] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X7),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f12158,f1209])).
fof(f12158,plain,(
  ( ! [X7] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X7),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f884,f10541])).
fof(f10936,plain,(
  $false | ($spl2 | $spl17)),
  inference(subsumption_resolution,[],[f10861,f10756])).
fof(f10756,plain,(
  ( ! [X0] : (~sP9('const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10)))),'T_0',X0)) ) | ($spl2 | $spl17)),
  inference(forward_demodulation,[],[f10624,f1374])).
fof(f1374,plain,(
  s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl2),
  inference(backward_demodulation,[],[f1196,f1202])).
fof(f10624,plain,(
  ( ! [X0] : (~sP9('const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),'T_0',X0)) ) | $spl17),
  inference(backward_demodulation,[],[f10496,f8290])).
fof(f8290,plain,(
  ( ! [X0] : (~sP9('const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))),'T_0',X0)) ) | $spl17),
  inference(unit_resulting_resolution,[],[f763,f5472,f1148])).
fof(f1148,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f753])).
fof(f753,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP8(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f752])).
fof(f752,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f751])).
fof(f751,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1)) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f645])).
fof(f645,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f763,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.transc.transc_2|49|4_prove_ltb',aHL_TRUTH)).
fof(f10861,plain,(
  ( ! [X22] : (sP9('const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10)))),'T_0',X22)) ) | $spl2),
  inference(superposition,[],[f3644,f10680])).
fof(f10680,plain,(
  ( ! [X0] : (s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num','const.num.|48|_0')))) ) | $spl2),
  inference(forward_demodulation,[],[f10546,f1374])).
fof(f3644,plain,(
  ( ! [X0,X1] : (sP9('const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',X0)))),'T_0',X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f1692,f3571,f1183])).
fof(f1183,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f1150])).
fof(f1150,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f753])).
fof(f3571,plain,(
  ( ! [X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',X1))))))) ) | $spl2),
  inference(forward_demodulation,[],[f3570,f1343])).
fof(f1343,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) ) | $spl2),
  inference(backward_demodulation,[],[f1196,f833])).
fof(f3570,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))),s('type.num.num',X1))))))) ) | $spl2),
  inference(subsumption_resolution,[],[f3569,f763])).
fof(f3569,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))),s('type.num.num',X1))))))) ) | $spl2),
  inference(forward_demodulation,[],[f3551,f2408])).
fof(f2408,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)))) )),
  inference(unit_resulting_resolution,[],[f782,f871])).
fof(f871,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f657])).
fof(f657,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f359])).
fof(f359,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.bool.EQ_CLAUSES|split|0')).
fof(f782,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0))))) )),
  inference(cnf_transformation,[],[f270])).
fof(f270,plain,(
  ! [X0] : p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0))))),
  inference(rectify,[],[f153])).
fof(f153,axiom,(
  ! [X6] : p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X6))))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.real.REAL_LE_REFL')).
fof(f3551,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real',sK10)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))),s('type.num.num',X1))))))) ) | $spl2),
  inference(superposition,[],[f1387,f1343])).
fof(f1387,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real',X0)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num',X1))))))) ) | $spl2),
  inference(forward_demodulation,[],[f1344,f1196])).
fof(f1344,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num',X1)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))))) ) | $spl2),
  inference(backward_demodulation,[],[f1196,f859])).
fof(f859,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num',X1)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))))) )),
  inference(cnf_transformation,[],[f558])).
fof(f558,plain,(
  ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | ! [X1] : p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num',X1)))))))),
  inference(ennf_transformation,[],[f348])).
fof(f348,plain,(
  ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) => ! [X1] : p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num',X1)))))))),
  inference(rectify,[],[f139])).
fof(f139,axiom,(
  ! [X6] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X6)))) => ! [X7] : p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real',X6),s('type.num.num',X7)))))))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.real.POW_POS')).
fof(f1692,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f763,f1156])).
fof(f1156,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f756])).
fof(f756,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f755])).
fof(f755,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f754])).
fof(f754,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f644])).
fof(f644,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f10935,plain,(
  $false | ($spl2 | $spl17)),
  inference(subsumption_resolution,[],[f10860,f10753])).
fof(f10753,plain,(
  ( ! [X0] : (~sP9('const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10)))),X0,'F_0')) ) | ($spl2 | $spl17)),
  inference(forward_demodulation,[],[f10621,f1374])).
fof(f10621,plain,(
  ( ! [X0] : (~sP9('const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X0,'F_0')) ) | $spl17),
  inference(backward_demodulation,[],[f10496,f8280])).
fof(f8280,plain,(
  ( ! [X0] : (~sP9('const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))),X0,'F_0')) ) | $spl17),
  inference(unit_resulting_resolution,[],[f1205,f5472,f1147])).
fof(f1147,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f753])).
fof(f1205,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f775,f763])).
fof(f775,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f650])).
fof(f650,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f264])).
fof(f264,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f30])).
fof(f30,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.bool.NOT_CLAUSES|split|1')).
fof(f10860,plain,(
  ( ! [X21] : (sP9('const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10)))),X21,'F_0')) ) | $spl2),
  inference(superposition,[],[f3643,f10680])).
fof(f3643,plain,(
  ( ! [X0,X1] : (sP9('const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',X0)))),X1,'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f1691,f3571,f1183])).
fof(f1691,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1205,f1155])).
fof(f1155,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f756])).
fof(f10934,plain,(
  $false | ($spl2 | $spl17)),
  inference(subsumption_resolution,[],[f10855,f10761])).
fof(f10761,plain,(
  ( ! [X0] : (sP7('const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10)))),'F_0',X0)) ) | ($spl2 | $spl17)),
  inference(forward_demodulation,[],[f10629,f1374])).
fof(f10629,plain,(
  ( ! [X0] : (sP7('const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),'F_0',X0)) ) | $spl17),
  inference(backward_demodulation,[],[f10496,f8297])).
fof(f8297,plain,(
  ( ! [X0] : (sP7('const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))),'F_0',X0)) ) | $spl17),
  inference(unit_resulting_resolution,[],[f1690,f5472,f1184])).
fof(f1184,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f1136])).
fof(f1136,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f746])).
fof(f746,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP6(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f745])).
fof(f745,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f744])).
fof(f744,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f642])).
fof(f642,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f1690,plain,(
  ( ! [X0,X1] : (sP6('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1205,f1140])).
fof(f1140,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f748])).
fof(f748,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f747])).
fof(f747,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f641])).
fof(f641,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f10855,plain,(
  ( ! [X14] : (~sP7('const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10)))),'F_0',X14)) ) | $spl2),
  inference(superposition,[],[f3630,f10680])).
fof(f3630,plain,(
  ( ! [X0,X1] : (~sP7('const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',X0)))),'F_0',X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f1205,f3571,f1132])).
fof(f1132,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f746])).
fof(f10933,plain,(
  $false | ($spl2 | $spl17)),
  inference(subsumption_resolution,[],[f10854,f10760])).
fof(f10760,plain,(
  ( ! [X0] : (sP7('const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10)))),X0,'F_0')) ) | ($spl2 | $spl17)),
  inference(forward_demodulation,[],[f10628,f1374])).
fof(f10628,plain,(
  ( ! [X0] : (sP7('const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X0,'F_0')) ) | $spl17),
  inference(backward_demodulation,[],[f10496,f8296])).
fof(f8296,plain,(
  ( ! [X0] : (sP7('const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))),X0,'F_0')) ) | $spl17),
  inference(unit_resulting_resolution,[],[f1673,f5472,f1184])).
fof(f1673,plain,(
  ( ! [X0,X1] : (sP6(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f1205,f1139])).
fof(f1139,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f748])).
fof(f10854,plain,(
  ( ! [X13] : (~sP7('const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10)))),X13,'F_0')) ) | $spl2),
  inference(superposition,[],[f3629,f10680])).
fof(f3629,plain,(
  ( ! [X0,X1] : (~sP7('const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',X0)))),X1,'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f1205,f3571,f1131])).
fof(f1131,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f746])).
fof(f10932,plain,(
  $false | ($spl2 | $spl17)),
  inference(subsumption_resolution,[],[f10848,f10737])).
fof(f10737,plain,(
  ~sP1('F_0','const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))) | ($spl2 | $spl17)),
  inference(forward_demodulation,[],[f10605,f1374])).
fof(f10605,plain,(
  ~sP1('F_0','const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl17),
  inference(backward_demodulation,[],[f10496,f8254])).
fof(f8254,plain,(
  ~sP1('F_0','const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl17),
  inference(unit_resulting_resolution,[],[f1205,f5472,f1052])).
fof(f1052,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f719])).
fof(f719,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f718])).
fof(f718,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f717])).
fof(f717,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f634])).
fof(f634,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f10848,plain,(
  sP1('F_0','const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))) | $spl2),
  inference(superposition,[],[f3623,f10680])).
fof(f3623,plain,(
  ( ! [X0] : (sP1('F_0','const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',X0)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f1205,f3571,f1057])).
fof(f1057,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f719])).
fof(f10931,plain,(
  $false | ($spl2 | $spl17)),
  inference(subsumption_resolution,[],[f10847,f10736])).
fof(f10736,plain,(
  ~sP1('const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10)))),'F_0') | ($spl2 | $spl17)),
  inference(forward_demodulation,[],[f10604,f1374])).
fof(f10604,plain,(
  ~sP1('const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),'F_0') | $spl17),
  inference(backward_demodulation,[],[f10496,f8252])).
fof(f8252,plain,(
  ~sP1('const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))),'F_0') | $spl17),
  inference(unit_resulting_resolution,[],[f1205,f5472,f1052])).
fof(f10847,plain,(
  sP1('const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10)))),'F_0') | $spl2),
  inference(superposition,[],[f3622,f10680])).
fof(f3622,plain,(
  ( ! [X0] : (sP1('const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',X0)))),'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f1205,f3571,f1054])).
fof(f1054,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f719])).
fof(f10930,plain,(
  $false | ($spl2 | $spl17)),
  inference(subsumption_resolution,[],[f10846,f10739])).
fof(f10739,plain,(
  sP1('T_0','const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))) | ($spl2 | $spl17)),
  inference(forward_demodulation,[],[f10607,f1374])).
fof(f10607,plain,(
  sP1('T_0','const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl17),
  inference(backward_demodulation,[],[f10496,f8256])).
fof(f8256,plain,(
  sP1('T_0','const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl17),
  inference(unit_resulting_resolution,[],[f763,f5472,f1054])).
fof(f10846,plain,(
  ~sP1('T_0','const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))) | $spl2),
  inference(superposition,[],[f3620,f10680])).
fof(f3620,plain,(
  ( ! [X0] : (~sP1('T_0','const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',X0)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f763,f3571,f1053])).
fof(f1053,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f719])).
fof(f10929,plain,(
  $false | ($spl2 | $spl17)),
  inference(subsumption_resolution,[],[f10845,f10740])).
fof(f10740,plain,(
  sP1('const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10)))),'T_0') | ($spl2 | $spl17)),
  inference(forward_demodulation,[],[f10608,f1374])).
fof(f10608,plain,(
  sP1('const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),'T_0') | $spl17),
  inference(backward_demodulation,[],[f10496,f8257])).
fof(f8257,plain,(
  sP1('const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))),'T_0') | $spl17),
  inference(unit_resulting_resolution,[],[f763,f5472,f1057])).
fof(f10845,plain,(
  ~sP1('const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10)))),'T_0') | $spl2),
  inference(superposition,[],[f3618,f10680])).
fof(f3618,plain,(
  ( ! [X0] : (~sP1('const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',X0)))),'T_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f763,f3571,f1053])).
fof(f6418,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5959,f1340])).
fof(f1340,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',sK10)))) ) | $spl2),
  inference(backward_demodulation,[],[f1196,f823])).
fof(f823,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(cnf_transformation,[],[f311])).
fof(f311,plain,(
  ! [X0] : s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f151])).
fof(f151,axiom,(
  ! [X6] : s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.real.REAL_MUL_RZERO')).
fof(f5959,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real',sK10))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4957])).
fof(f4957,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))) | $spl2),
  inference(forward_demodulation,[],[f4692,f1337])).
fof(f1337,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real',sK10)))) ) | $spl2),
  inference(backward_demodulation,[],[f1196,f806])).
fof(f806,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(cnf_transformation,[],[f294])).
fof(f294,plain,(
  ! [X0] : s('type.realax.real',X0) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f184])).
fof(f184,axiom,(
  ! [X6] : s('type.realax.real',X6) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X6),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.real.REAL_SUB_RZERO')).
fof(f4692,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real',sK10))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real',sK10))))) | $spl2),
  inference(unit_resulting_resolution,[],[f1813,f1813,f1394])).
fof(f1394,plain,(
  ( ! [X0,X1] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))) | s('type.realax.real',sK10) = s('type.realax.real',X1) | s('type.realax.real',sK10) = s('type.realax.real',X0)) ) | $spl2),
  inference(forward_demodulation,[],[f1393,f1196])).
fof(f1393,plain,(
  ( ! [X0,X1] : (s('type.realax.real',sK10) = s('type.realax.real',X1) | s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl2),
  inference(forward_demodulation,[],[f1356,f1196])).
fof(f1356,plain,(
  ( ! [X0,X1] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl2),
  inference(backward_demodulation,[],[f1196,f1030])).
fof(f1030,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) )),
  inference(cnf_transformation,[],[f708])).
fof(f708,plain,(
  ! [X0,X1] : ((s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & ((s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.realax.real',X1) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))),
  inference(flattening,[],[f707])).
fof(f707,plain,(
  ! [X0,X1] : ((s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))) | (s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) & ((s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.realax.real',X1) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))),
  inference(nnf_transformation,[],[f479])).
fof(f479,plain,(
  ! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))) <=> (s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6,X15] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real',X15))) <=> (s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X15) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.real.REAL_ENTIRE')).
fof(f1813,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real',sK10))) | $spl2),
  inference(unit_resulting_resolution,[],[f1538,f1352])).
fof(f1352,plain,(
  ( ! [X0,X1] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real',X1))) | s('type.realax.real',X0) = s('type.realax.real',X1)) ) | $spl2),
  inference(backward_demodulation,[],[f1196,f1026])).
fof(f1026,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real',X1))) | s('type.realax.real',X0) = s('type.realax.real',X1)) )),
  inference(cnf_transformation,[],[f705])).
fof(f705,plain,(
  ! [X0,X1] : ((s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real',X1))) | s('type.realax.real',X0) = s('type.realax.real',X1)) & (s('type.realax.real',X0) != s('type.realax.real',X1) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))),
  inference(nnf_transformation,[],[f477])).
fof(f477,plain,(
  ! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real',X1))) <=> s('type.realax.real',X0) = s('type.realax.real',X1))),
  inference(rectify,[],[f190])).
fof(f190,axiom,(
  ! [X6,X15] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X6),s('type.realax.real',X15))) <=> s('type.realax.real',X6) = s('type.realax.real',X15))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.real.REAL_SUB_|48|')).
fof(f1538,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))) | $spl2),
  inference(superposition,[],[f1375,f1374])).
fof(f1375,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl2),
  inference(backward_demodulation,[],[f1196,f1208])).
fof(f5613,plain,(
  ( ! [X39] : (s('type.realax.real',sK10) = s('type.realax.real',X39)) ) | ($spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5488,f5485])).
fof(f6417,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5958,f1340])).
fof(f5958,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))),s('type.realax.real',sK10))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4949])).
fof(f4949,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))) | $spl2),
  inference(forward_demodulation,[],[f4696,f1337])).
fof(f4696,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real',sK10))))) | $spl2),
  inference(unit_resulting_resolution,[],[f1714,f1813,f1394])).
fof(f1714,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))) | $spl2),
  inference(unit_resulting_resolution,[],[f1538,f1388])).
fof(f1388,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0))) | s('type.realax.real',sK10) = s('type.realax.real',X0)) ) | $spl2),
  inference(forward_demodulation,[],[f1346,f1196])).
fof(f1346,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl2),
  inference(backward_demodulation,[],[f1196,f886])).
fof(f6416,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5957,f1340])).
fof(f5957,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))))),s('type.realax.real',sK10))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4948])).
fof(f4948,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))) | $spl2),
  inference(forward_demodulation,[],[f4697,f1337])).
fof(f4697,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real',sK10))))) | $spl2),
  inference(unit_resulting_resolution,[],[f1734,f1813,f1394])).
fof(f1734,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f1714,f1388])).
fof(f6415,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5956,f1342])).
fof(f1342,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real',X0)))) ) | $spl2),
  inference(backward_demodulation,[],[f1196,f830])).
fof(f830,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) )),
  inference(cnf_transformation,[],[f318])).
fof(f318,plain,(
  ! [X0] : s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6] : s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X6)))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.real.REAL_MUL_LZERO')).
fof(f5956,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4844])).
fof(f4844,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | $spl2),
  inference(forward_demodulation,[],[f4764,f1337])).
fof(f4764,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real',sK10))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f1813,f1714,f1394])).
fof(f6414,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5955,f1342])).
fof(f5955,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))))))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4834])).
fof(f4834,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | $spl2),
  inference(forward_demodulation,[],[f4782,f1337])).
fof(f4782,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real',sK10))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f1813,f1734,f1394])).
fof(f6413,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5954,f5613])).
fof(f5954,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))))))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4787])).
fof(f4787,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f1734,f1734,f1394])).
fof(f6412,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5953,f5613])).
fof(f5953,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))))))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4786])).
fof(f4786,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f1714,f1734,f1394])).
fof(f6411,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5952,f5613])).
fof(f5952,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4769])).
fof(f4769,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f1734,f1714,f1394])).
fof(f6410,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5951,f5613])).
fof(f5951,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4768])).
fof(f4768,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f1714,f1714,f1394])).
fof(f6406,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5948,f5613])).
fof(f5948,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f2682])).
fof(f2682,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f2465,f1354])).
fof(f1354,plain,(
  ( ! [X0,X1] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))) | s('type.realax.real',X0) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))) ) | $spl2),
  inference(backward_demodulation,[],[f1196,f1028])).
fof(f1028,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X0) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) )),
  inference(cnf_transformation,[],[f706])).
fof(f706,plain,(
  ! [X0,X1] : ((s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))) | s('type.realax.real',X0) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))) & (s('type.realax.real',X0) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))),
  inference(nnf_transformation,[],[f478])).
fof(f478,plain,(
  ! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))) <=> s('type.realax.real',X0) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))),
  inference(rectify,[],[f192])).
fof(f192,axiom,(
  ! [X6,X15] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X6),s('type.realax.real',X15))) <=> s('type.realax.real',X6) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X15))))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.real.REAL_LNEG_UNIQ')).
fof(f2465,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f1714,f1850])).
fof(f1850,plain,(
  ( ! [X2] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))) | s('type.realax.real',sK10) = s('type.realax.real',X2)) ) | $spl2),
  inference(superposition,[],[f1352,f1341])).
fof(f1341,plain,(
  ( ! [X0] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real',X0)))) ) | $spl2),
  inference(backward_demodulation,[],[f1196,f824])).
fof(f824,plain,(
  ( ! [X0] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) )),
  inference(cnf_transformation,[],[f312])).
fof(f312,plain,(
  ! [X0] : s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))),
  inference(rectify,[],[f204])).
fof(f204,axiom,(
  ! [X6] : s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X6))) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X6)))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.real.REAL_SUB_LZERO')).
fof(f6405,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5947,f5613])).
fof(f5947,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK10),s('type.realax.real',sK10))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f2681])).
fof(f2681,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))) | $spl2),
  inference(unit_resulting_resolution,[],[f1875,f1354])).
fof(f1875,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))) | $spl2),
  inference(superposition,[],[f1808,f1341])).
fof(f1808,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))) | $spl2),
  inference(unit_resulting_resolution,[],[f1538,f1352])).
fof(f6404,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5946,f5613])).
fof(f5946,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))))))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f2466])).
fof(f2466,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f1734,f1850])).
fof(f6403,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5945,f5613])).
fof(f5945,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f2465])).
fof(f6402,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f6401,f2403])).
fof(f2403,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X0)))) )),
  inference(unit_resulting_resolution,[],[f937,f870])).
fof(f870,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f657])).
fof(f937,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X0))))) )),
  inference(cnf_transformation,[],[f392])).
fof(f392,plain,(
  ! [X0] : ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X0))))),
  inference(flattening,[],[f391])).
fof(f391,plain,(
  ! [X0] : ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X0))))),
  inference(rectify,[],[f195])).
fof(f195,axiom,(
  ! [X6] : ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X6),s('type.realax.real',X6))))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.realax.REAL_LT_REFL')).
fof(f6401,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real',sK10))) | ($spl2 | $spl26 | $spl28)),
  inference(forward_demodulation,[],[f5944,f1334])).
fof(f1334,plain,(
  s('type.realax.real',sK10) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))) | $spl2),
  inference(backward_demodulation,[],[f1196,f771])).
fof(f771,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(cnf_transformation,[],[f143])).
fof(f143,axiom,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.real.ABS_|48|')).
fof(f5944,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f2410])).
fof(f2410,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f2044,f871])).
fof(f2044,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10)))))))) | $spl2),
  inference(forward_demodulation,[],[f2020,f1337])).
fof(f2020,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real',sK10)))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f1813,f1390])).
fof(f1390,plain,(
  ( ! [X0] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0)))))) | s('type.realax.real',sK10) = s('type.realax.real',X0)) ) | $spl2),
  inference(forward_demodulation,[],[f1348,f1196])).
fof(f1348,plain,(
  ( ! [X0] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0)))))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl2),
  inference(backward_demodulation,[],[f1196,f933])).
fof(f933,plain,(
  ( ! [X0] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0)))))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f696])).
fof(f696,plain,(
  ! [X0] : ((s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0))))))) & (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0)))))) | s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(nnf_transformation,[],[f388])).
fof(f388,plain,(
  ! [X0] : (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0)))))))),
  inference(flattening,[],[f387])).
fof(f387,plain,(
  ! [X0] : (~s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0)))))))),
  inference(rectify,[],[f122])).
fof(f122,axiom,(
  ! [X6] : (~s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X6)))))))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.real.ABS_NZ')).
fof(f6391,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5934,f5613])).
fof(f5934,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',X0)))) ) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f2193])).
fof(f2193,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.num.num',X0)))) ) | $spl2),
  inference(forward_demodulation,[],[f2163,f1337])).
fof(f2163,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real',sK10))),s('type.num.num',X0)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f1813,f1392])).
fof(f1392,plain,(
  ( ! [X0,X1] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X1),s('type.num.num',X0))) | s('type.realax.real',sK10) = s('type.realax.real',X1)) ) | $spl2),
  inference(forward_demodulation,[],[f1351,f1196])).
fof(f1351,plain,(
  ( ! [X0,X1] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X1),s('type.num.num',X0))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl2),
  inference(backward_demodulation,[],[f1196,f1011])).
fof(f6390,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5933,f5613])).
fof(f5933,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))))),s('type.num.num',X0)))) ) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f2167])).
fof(f2167,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.num.num',X0)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f1734,f1392])).
fof(f6389,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5932,f5613])).
fof(f5932,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))),s('type.num.num',X0)))) ) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f2166])).
fof(f2166,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.num.num',X0)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f1714,f1392])).
fof(f6388,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f6387,f1205])).
fof(f6387,plain,(
  p(s(bool,'F_0')) | ($spl2 | $spl26 | $spl28)),
  inference(forward_demodulation,[],[f6386,f2451])).
fof(f2451,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X0)))) )),
  inference(unit_resulting_resolution,[],[f937,f875])).
fof(f875,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f659])).
fof(f659,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f362])).
fof(f362,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f361])).
fof(f361,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X11] : (s(bool,'F_0') = s(bool,X11) <=> ~p(s(bool,X11)))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.bool.EQ_CLAUSES|split|2')).
fof(f6386,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real',sK10)))) | ($spl2 | $spl26 | $spl28)),
  inference(forward_demodulation,[],[f5931,f5613])).
fof(f5931,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10)))))))))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f2023])).
fof(f2023,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10)))))))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f1734,f1390])).
fof(f6385,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f6384,f1205])).
fof(f6384,plain,(
  p(s(bool,'F_0')) | ($spl2 | $spl26 | $spl28)),
  inference(forward_demodulation,[],[f6383,f2451])).
fof(f6383,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real',sK10)))) | ($spl2 | $spl26 | $spl28)),
  inference(forward_demodulation,[],[f5930,f5613])).
fof(f5930,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10)))))))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f2021])).
fof(f2021,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10)))))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f1714,f1390])).
fof(f6382,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5929,f5613])).
fof(f5929,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))))),s('type.realax.real',sK10))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f2014])).
fof(f2014,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real',sK10))) | $spl2),
  inference(unit_resulting_resolution,[],[f1734,f1352])).
fof(f6381,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5928,f5613])).
fof(f5928,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))))))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f2013])).
fof(f2013,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f1734,f1352])).
fof(f6380,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5927,f5613])).
fof(f5927,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))))))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f2012])).
fof(f2012,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f1734,f1388])).
fof(f6379,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5926,f1332])).
fof(f1332,plain,(
  s('type.realax.real',sK10) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK10))) | $spl2),
  inference(backward_demodulation,[],[f1196,f769])).
fof(f769,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(cnf_transformation,[],[f206])).
fof(f206,axiom,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.real.REAL_NEG_|48|')).
fof(f5926,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK10))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f1875])).
fof(f6378,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5925,f5613])).
fof(f5925,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))),s('type.realax.real',sK10))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f1814])).
fof(f1814,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real',sK10))) | $spl2),
  inference(unit_resulting_resolution,[],[f1714,f1352])).
fof(f6377,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5924,f1335])).
fof(f1335,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real',X0)))) ) | $spl2),
  inference(backward_demodulation,[],[f1196,f800])).
fof(f800,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real',X0)))) )),
  inference(cnf_transformation,[],[f288])).
fof(f288,plain,(
  ! [X0] : s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real',X0)))),
  inference(rectify,[],[f207])).
fof(f207,axiom,(
  ! [X6] : s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X6),s('type.realax.real',X6)))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.real.REAL_SUB_REFL')).
fof(f5924,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real',sK10))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f1813])).
fof(f6376,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5923,f5613])).
fof(f5923,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f1809])).
fof(f1809,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f1714,f1352])).
fof(f6375,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5922,f1335])).
fof(f5922,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real',sK10))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f1808])).
fof(f6374,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5921,f5613])).
fof(f5921,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f1734])).
fof(f6373,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5920,f1334])).
fof(f5920,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f1714])).
fof(f6372,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5913,f1340])).
fof(f5913,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real',sK10))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5070])).
fof(f5070,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | $spl2),
  inference(forward_demodulation,[],[f4602,f1337])).
fof(f4602,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real',sK10))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f1813,f2465,f1394])).
fof(f6371,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5912,f1342])).
fof(f5912,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4965])).
fof(f4965,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))) | $spl2),
  inference(forward_demodulation,[],[f4687,f1337])).
fof(f4687,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real',sK10))))) | $spl2),
  inference(unit_resulting_resolution,[],[f2465,f1813,f1394])).
fof(f6370,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5911,f1342])).
fof(f5911,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4777])).
fof(f4777,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f2465,f1734,f1394])).
fof(f6369,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5910,f1342])).
fof(f5910,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4759])).
fof(f4759,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f2465,f1714,f1394])).
fof(f6368,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5909,f1340])).
fof(f5909,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real',sK10))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4607])).
fof(f4607,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f1734,f2465,f1394])).
fof(f6367,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5908,f1340])).
fof(f5908,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real',sK10))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4606])).
fof(f4606,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f1714,f2465,f1394])).
fof(f6366,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5907,f1340])).
fof(f5907,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real',sK10))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4597])).
fof(f4597,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f2465,f2465,f1394])).
fof(f6365,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5906,f1335])).
fof(f5906,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real',sK10))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f2507])).
fof(f2507,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real',sK10))) | $spl2),
  inference(unit_resulting_resolution,[],[f2465,f1352])).
fof(f6364,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5905,f1335])).
fof(f5905,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real',sK10))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f2506])).
fof(f2506,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f2465,f1352])).
fof(f6363,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5904,f1334])).
fof(f5904,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f2505])).
fof(f2505,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f2465,f1388])).
fof(f6362,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5903,f5613])).
fof(f5903,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',X0)))) ) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f2503])).
fof(f2503,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.num.num',X0)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f2465,f1392])).
fof(f6361,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5902,f1332])).
fof(f5902,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK10))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f2502])).
fof(f2502,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f2465,f1850])).
fof(f6360,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5900,f1342])).
fof(f5900,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5072])).
fof(f5072,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | $spl2),
  inference(forward_demodulation,[],[f4600,f1341])).
fof(f4600,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f1808,f2465,f1394])).
fof(f6359,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5899,f1340])).
fof(f5899,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real',sK10))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5019])).
fof(f5019,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | $spl2),
  inference(forward_demodulation,[],[f4651,f1341])).
fof(f4651,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f2465,f1808,f1394])).
fof(f6358,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5898,f1340])).
fof(f5898,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real',sK10))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5015])).
fof(f5015,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | $spl2),
  inference(forward_demodulation,[],[f4654,f1341])).
fof(f4654,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f1808,f1808,f1394])).
fof(f6357,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5897,f1340])).
fof(f5897,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real',sK10))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5012])).
fof(f5012,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | $spl2),
  inference(forward_demodulation,[],[f5011,f1337])).
fof(f5011,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real',sK10))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | $spl2),
  inference(forward_demodulation,[],[f4656,f1341])).
fof(f4656,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real',sK10))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f1813,f1808,f1394])).
fof(f6356,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5896,f1340])).
fof(f5896,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real',sK10))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5003])).
fof(f5003,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | $spl2),
  inference(forward_demodulation,[],[f4660,f1341])).
fof(f4660,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f1714,f1808,f1394])).
fof(f6355,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5895,f1340])).
fof(f5895,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real',sK10))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5002])).
fof(f5002,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | $spl2),
  inference(forward_demodulation,[],[f4661,f1341])).
fof(f4661,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f1734,f1808,f1394])).
fof(f6354,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5894,f1342])).
fof(f5894,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4961])).
fof(f4961,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))) | $spl2),
  inference(forward_demodulation,[],[f4960,f1341])).
fof(f4960,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))) | $spl2),
  inference(forward_demodulation,[],[f4690,f1337])).
fof(f4690,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real',sK10))))) | $spl2),
  inference(unit_resulting_resolution,[],[f1808,f1813,f1394])).
fof(f6353,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5893,f1342])).
fof(f5893,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4846])).
fof(f4846,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | $spl2),
  inference(forward_demodulation,[],[f4762,f1341])).
fof(f4762,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f1808,f1714,f1394])).
fof(f6352,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5892,f1342])).
fof(f5892,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4836])).
fof(f4836,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | $spl2),
  inference(forward_demodulation,[],[f4780,f1341])).
fof(f4780,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f1808,f1734,f1394])).
fof(f6351,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5891,f5613])).
fof(f5891,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK10),s('type.realax.real',sK10))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f2680])).
fof(f2680,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK10),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f2488,f1354])).
fof(f2488,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | $spl2),
  inference(forward_demodulation,[],[f2461,f1341])).
fof(f2461,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f1808,f1850])).
fof(f6350,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5890,f1332])).
fof(f5890,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK10))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f2488])).
fof(f6349,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5889,f5613])).
fof(f5889,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',X0)))) ) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f2194])).
fof(f2194,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.num.num',X0)))) ) | $spl2),
  inference(forward_demodulation,[],[f2162,f1341])).
fof(f2162,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.num.num',X0)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f1808,f1392])).
fof(f6348,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5888,f1334])).
fof(f5888,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f1878])).
fof(f1878,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | $spl2),
  inference(forward_demodulation,[],[f1872,f1341])).
fof(f1872,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f1808,f1388])).
fof(f6347,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5887,f1335])).
fof(f5887,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real',sK10))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f1877])).
fof(f1877,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | $spl2),
  inference(forward_demodulation,[],[f1873,f1341])).
fof(f1873,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f1808,f1352])).
fof(f6346,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5886,f1335])).
fof(f5886,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real',sK10))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f1876])).
fof(f1876,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real',sK10))) | $spl2),
  inference(forward_demodulation,[],[f1874,f1341])).
fof(f1874,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real',sK10))) | $spl2),
  inference(unit_resulting_resolution,[],[f1808,f1352])).
fof(f6329,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real',X0))))) ) | $spl46),
  inference(cnf_transformation,[],[f6329_D])).
fof(f6329_D,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real',X0))))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f6330,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) ) | $spl48),
  inference(cnf_transformation,[],[f6330_D])).
fof(f6330_D,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f6332,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real',X2))))) ) | ($spl2 | $spl26 | $spl28)),
  inference(forward_demodulation,[],[f6331,f2408])).
fof(f6331,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real',sK10))) = s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real',X2))))) ) | ($spl2 | $spl26 | $spl28)),
  inference(forward_demodulation,[],[f5842,f5613])).
fof(f5842,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real',sK10))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real',X2))))) ) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f1362])).
fof(f1362,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real',X2)))) | s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2)))))) ) | $spl2),
  inference(backward_demodulation,[],[f1196,f1075])).
fof(f1075,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2))))) )),
  inference(cnf_transformation,[],[f590])).
fof(f590,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) | s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))))))),
  inference(ennf_transformation,[],[f505])).
fof(f505,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) => s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))))))),
  inference(rectify,[],[f147])).
fof(f147,axiom,(
  ! [X6,X15,X18] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X18)))) => s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X15))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real',X18))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X15),s('type.realax.real',X18))))))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.real.REAL_LE_RMUL')).
fof(f6328,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real',X0))))) ) | ($spl2 | $spl26 | $spl28)),
  inference(forward_demodulation,[],[f6327,f2408])).
fof(f6327,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real',sK10))) = s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real',X0))))) ) | ($spl2 | $spl26 | $spl28)),
  inference(forward_demodulation,[],[f5841,f5613])).
fof(f5841,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2))) = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X2))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real',X0))))) ) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f1361])).
fof(f1361,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real',X0)))) | s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))) ) | $spl2),
  inference(backward_demodulation,[],[f1196,f1074])).
fof(f1074,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X2))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))))) )),
  inference(cnf_transformation,[],[f589])).
fof(f589,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X2))))))),
  inference(ennf_transformation,[],[f504])).
fof(f504,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) => s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X2))))))),
  inference(rectify,[],[f146])).
fof(f146,axiom,(
  ! [X6,X15,X18] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X6)))) => s(bool,'const.real.real_lte_2'(s('type.realax.real',X15),s('type.realax.real',X18))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real',X15))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real',X18))))))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.real.REAL_LE_LMUL')).
fof(f6312,plain,(
  $false | ($spl2 | $spl13 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5812,f1342])).
fof(f5812,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | ($spl2 | $spl13 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5608])).
fof(f5608,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | ($spl2 | $spl13)),
  inference(forward_demodulation,[],[f5560,f1341])).
fof(f5560,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | ($spl2 | $spl13)),
  inference(unit_resulting_resolution,[],[f1808,f5455,f1394])).
fof(f6311,plain,(
  $false | ($spl2 | $spl13 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5811,f1342])).
fof(f5811,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))) | ($spl2 | $spl13 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5606])).
fof(f5606,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))) | ($spl2 | $spl13)),
  inference(forward_demodulation,[],[f5562,f1337])).
fof(f5562,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real',sK10))))) | ($spl2 | $spl13)),
  inference(unit_resulting_resolution,[],[f1813,f5455,f1394])).
fof(f6310,plain,(
  $false | ($spl2 | $spl13 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5810,f1340])).
fof(f5810,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real',sK10))) | ($spl2 | $spl13 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5598])).
fof(f5598,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl2 | $spl13)),
  inference(forward_demodulation,[],[f5580,f1341])).
fof(f5580,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl2 | $spl13)),
  inference(unit_resulting_resolution,[],[f1808,f5455,f1394])).
fof(f6309,plain,(
  $false | ($spl2 | $spl13 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5809,f1340])).
fof(f5809,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real',sK10))) | ($spl2 | $spl13 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5596])).
fof(f5596,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl2 | $spl13)),
  inference(forward_demodulation,[],[f5582,f1337])).
fof(f5582,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real',sK10))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl2 | $spl13)),
  inference(unit_resulting_resolution,[],[f1813,f5455,f1394])).
fof(f6308,plain,(
  $false | ($spl2 | $spl13 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5808,f1335])).
fof(f5808,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real',sK10))) | ($spl2 | $spl13 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5591])).
fof(f5591,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real',sK10))) | ($spl2 | $spl13)),
  inference(unit_resulting_resolution,[],[f5455,f1352])).
fof(f6307,plain,(
  $false | ($spl2 | $spl13 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5807,f1335])).
fof(f5807,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real',sK10))) | ($spl2 | $spl13 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5590])).
fof(f5590,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl2 | $spl13)),
  inference(unit_resulting_resolution,[],[f5455,f1352])).
fof(f6306,plain,(
  $false | ($spl2 | $spl13 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5806,f1340])).
fof(f5806,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real',sK10))) | ($spl2 | $spl13 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5587])).
fof(f5587,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl2 | $spl13)),
  inference(unit_resulting_resolution,[],[f1734,f5455,f1394])).
fof(f6305,plain,(
  $false | ($spl2 | $spl13 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5805,f1340])).
fof(f5805,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real',sK10))) | ($spl2 | $spl13 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5586])).
fof(f5586,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl2 | $spl13)),
  inference(unit_resulting_resolution,[],[f1714,f5455,f1394])).
fof(f6304,plain,(
  $false | ($spl2 | $spl13 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5804,f1340])).
fof(f5804,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real',sK10))) | ($spl2 | $spl13 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5577])).
fof(f5577,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl2 | $spl13)),
  inference(unit_resulting_resolution,[],[f2465,f5455,f1394])).
fof(f6303,plain,(
  $false | ($spl2 | $spl13 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5803,f1340])).
fof(f5803,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real',sK10))) | ($spl2 | $spl13 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5573])).
fof(f5573,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl2 | $spl13)),
  inference(unit_resulting_resolution,[],[f5455,f5455,f1394])).
fof(f6302,plain,(
  $false | ($spl2 | $spl13 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5802,f1342])).
fof(f5802,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | ($spl2 | $spl13 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5567])).
fof(f5567,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | ($spl2 | $spl13)),
  inference(unit_resulting_resolution,[],[f1734,f5455,f1394])).
fof(f6301,plain,(
  $false | ($spl2 | $spl13 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5801,f1342])).
fof(f5801,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | ($spl2 | $spl13 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5566])).
fof(f5566,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | ($spl2 | $spl13)),
  inference(unit_resulting_resolution,[],[f1714,f5455,f1394])).
fof(f6300,plain,(
  $false | ($spl2 | $spl13 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5800,f1342])).
fof(f5800,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | ($spl2 | $spl13 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5557])).
fof(f5557,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | ($spl2 | $spl13)),
  inference(unit_resulting_resolution,[],[f2465,f5455,f1394])).
fof(f6299,plain,(
  $false | ($spl2 | $spl13 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f6298,f1205])).
fof(f6298,plain,(
  p(s(bool,'F_0')) | ($spl2 | $spl13 | $spl26 | $spl28)),
  inference(forward_demodulation,[],[f6297,f2451])).
fof(f6297,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real',sK10)))) | ($spl2 | $spl13 | $spl26 | $spl28)),
  inference(forward_demodulation,[],[f5799,f1334])).
fof(f5799,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10)))))) | ($spl2 | $spl13 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5547])).
fof(f5547,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))))) | ($spl2 | $spl13)),
  inference(unit_resulting_resolution,[],[f5455,f1390])).
fof(f6296,plain,(
  $false | ($spl2 | $spl26 | $spl28 | $spl43)),
  inference(subsumption_resolution,[],[f5793,f1332])).
fof(f5793,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK10))) | ($spl26 | $spl28 | $spl43)),
  inference(backward_demodulation,[],[f5613,f5506])).
fof(f6295,plain,(
  $false | ($spl26 | $spl28 | $spl30)),
  inference(subsumption_resolution,[],[f5790,f5613])).
fof(f5790,plain,(
  ( ! [X41] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',X41)))) ) | ($spl26 | $spl28 | $spl30)),
  inference(backward_demodulation,[],[f5613,f5490])).
fof(f6294,plain,(
  $false | ($spl2 | $spl21 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5789,f1334])).
fof(f5789,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))) | ($spl21 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5476])).
fof(f6293,plain,(
  $false | ($spl17 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f6292,f763])).
fof(f6292,plain,(
  ~p(s(bool,'T_0')) | ($spl17 | $spl26 | $spl28)),
  inference(forward_demodulation,[],[f5788,f2408])).
fof(f5788,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real',sK10)))) | ($spl17 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5472])).
fof(f6279,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f6243,f5613])).
fof(f6243,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num','const.num.|48|_0'))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f6157,f1403])).
fof(f1403,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f1195,f1196])).
fof(f6157,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) ) | ($spl26 | $spl28)),
  inference(subsumption_resolution,[],[f6155,f763])).
fof(f6155,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) ) | ($spl26 | $spl28)),
  inference(backward_demodulation,[],[f6149,f1277])).
fof(f1277,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(backward_demodulation,[],[f1276,f882])).
fof(f882,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f663])).
fof(f663,plain,(
  ! [X0] : ((~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))))),
  inference(nnf_transformation,[],[f367])).
fof(f367,plain,(
  ! [X0] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X0))),
  inference(rectify,[],[f231])).
fof(f231,axiom,(
  ! [X7] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7)))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X7))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.numeral.numeral_distrib|split|30')).
fof(f1276,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(forward_demodulation,[],[f1275,f781])).
fof(f1275,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num',X1)))) )),
  inference(forward_demodulation,[],[f973,f781])).
fof(f973,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f429])).
fof(f429,plain,(
  ! [X0,X1] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f233])).
fof(f233,axiom,(
  ! [X7,X21] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X21),s('type.num.num',X7))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X21)))))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.numeral.numeral_distrib|split|31')).
fof(f6149,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) ) | ($spl26 | $spl28)),
  inference(forward_demodulation,[],[f6148,f2408])).
fof(f6148,plain,(
  ( ! [X0,X1] : (s(bool,'const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real',sK10))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) ) | ($spl26 | $spl28)),
  inference(forward_demodulation,[],[f5766,f5613])).
fof(f5766,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))) ) | ($spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f963])).
fof(f963,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f419])).
fof(f419,plain,(
  ! [X0,X1] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f152])).
fof(f152,axiom,(
  ! [X21,X7] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X21),s('type.num.num',X7))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X21))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X7)))))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.real.REAL_LE')).
fof(f6277,plain,(
  $false | ($spl2 | $spl13 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f6229,f1196])).
fof(f6229,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl13 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f6157,f5455])).
fof(f6275,plain,(
  $false | ($spl26 | $spl28)),
  inference(subsumption_resolution,[],[f6274,f763])).
fof(f6274,plain,(
  ~p(s(bool,'T_0')) | ($spl26 | $spl28)),
  inference(forward_demodulation,[],[f6273,f6266])).
fof(f6266,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) ) | ($spl26 | $spl28)),
  inference(forward_demodulation,[],[f6265,f6175])).
fof(f6175,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) ) | ($spl26 | $spl28)),
  inference(backward_demodulation,[],[f6157,f4328])).
fof(f4328,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f798,f871])).
fof(f798,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))) )),
  inference(cnf_transformation,[],[f286])).
fof(f286,plain,(
  ! [X0] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))),
  inference(rectify,[],[f170])).
fof(f170,axiom,(
  ! [X7] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.prim_rec.LESS_SUC_REFL')).
fof(f6265,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) ) | ($spl26 | $spl28)),
  inference(forward_demodulation,[],[f6206,f6157])).
fof(f6206,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) ) | ($spl26 | $spl28)),
  inference(backward_demodulation,[],[f6157,f968])).
fof(f968,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f424])).
fof(f424,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f220])).
fof(f220,axiom,(
  ! [X7,X21] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',X21))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X7))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X21)))))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.numeral.numeral_lt|split|6')).
fof(f6273,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) ) | ($spl26 | $spl28)),
  inference(subsumption_resolution,[],[f6272,f763])).
fof(f6272,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) ) | ($spl26 | $spl28)),
  inference(forward_demodulation,[],[f6214,f6175])).
fof(f6214,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) ) | ($spl26 | $spl28)),
  inference(backward_demodulation,[],[f6157,f1033])).
fof(f1033,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f709])).
fof(f709,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))))))))),
  inference(nnf_transformation,[],[f481])).
fof(f481,plain,(
  ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) <=> ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f480])).
fof(f480,plain,(
  ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) <=> ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f222])).
fof(f222,axiom,(
  ! [X7,X21] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X7))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X21)))))) <=> ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X21),s('type.num.num',X7)))))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.numeral.numeral_lt|split|5')).
fof(f6270,plain,(
  $false | ($spl26 | $spl28)),
  inference(subsumption_resolution,[],[f6269,f2402])).
fof(f2402,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f1205,f870])).
fof(f6269,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl26 | $spl28)),
  inference(backward_demodulation,[],[f6267,f786])).
fof(f786,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f274])).
fof(f274,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f238])).
fof(f238,axiom,(
  ! [X7] : s(bool,'F_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7)))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.numeral.numeral_distrib|split|23')).
fof(f6267,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num',X0),s('type.num.num',X1)))) ) | ($spl26 | $spl28)),
  inference(backward_demodulation,[],[f6266,f1280])).
fof(f1280,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.arithmetic.>_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(forward_demodulation,[],[f1279,f781])).
fof(f1279,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num',X1)))) )),
  inference(forward_demodulation,[],[f977,f781])).
fof(f977,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f433])).
fof(f433,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f237])).
fof(f237,axiom,(
  ! [X7,X21] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X21),s('type.num.num',X7))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X21)))))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.numeral.numeral_distrib|split|25')).
fof(f6264,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f6202,f1196])).
fof(f6202,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f6157,f1375])).
fof(f6261,plain,(
  $false | ($spl26 | $spl28)),
  inference(subsumption_resolution,[],[f6260,f2402])).
fof(f6260,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl26 | $spl28)),
  inference(backward_demodulation,[],[f6175,f1210])).
fof(f1210,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f789,f765])).
fof(f789,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(cnf_transformation,[],[f277])).
fof(f277,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(rectify,[],[f227])).
fof(f227,axiom,(
  ! [X7] : s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.ZERO_0')))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.numeral.numeral_lt|split|2')).
fof(f6251,plain,(
  $false | ($spl26 | $spl28)),
  inference(subsumption_resolution,[],[f6250,f1205])).
fof(f6250,plain,(
  p(s(bool,'F_0')) | ($spl26 | $spl28)),
  inference(forward_demodulation,[],[f6247,f767])).
fof(f767,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))),
  inference(cnf_transformation,[],[f77])).
fof(f77,axiom,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.arithmetic.ODD|split|0')).
fof(f6247,plain,(
  p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) | ($spl26 | $spl28)),
  inference(condensation,[],[f6159])).
fof(f6159,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X0))))) ) | ($spl26 | $spl28)),
  inference(backward_demodulation,[],[f6157,f885])).
fof(f6147,plain,(
  $false | ($spl1 | $spl2 | $spl13 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5765,f5613])).
fof(f5765,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK10))))) | ($spl1 | $spl2 | $spl13 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5607])).
fof(f5607,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))))) | ($spl1 | $spl2 | $spl13)),
  inference(forward_demodulation,[],[f5561,f1341])).
fof(f5561,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))))) | ($spl1 | $spl2 | $spl13)),
  inference(unit_resulting_resolution,[],[f1855,f5455,f1394])).
fof(f1855,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f1854,f1334])).
fof(f1854,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))),s('type.num.num',sK11))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f1810,f974])).
fof(f974,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0))),s('type.num.num',X1))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f430])).
fof(f430,plain,(
  ! [X0,X1] : s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0))),s('type.num.num',X1))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num',X1)))))),
  inference(rectify,[],[f142])).
fof(f142,axiom,(
  ! [X8,X7] : s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',X8))),s('type.num.num',X7))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X8),s('type.num.num',X7)))))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.real.POW_ABS')).
fof(f1810,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1713,f1352])).
fof(f1713,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1403,f1388])).
fof(f6146,plain,(
  $false | ($spl1 | $spl2 | $spl13 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5764,f1340])).
fof(f5764,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real',sK10))) | ($spl1 | $spl2 | $spl13 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5604])).
fof(f5604,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))) | ($spl1 | $spl2 | $spl13)),
  inference(forward_demodulation,[],[f5603,f1334])).
fof(f5603,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))),s('type.num.num',sK11))))) | ($spl1 | $spl2 | $spl13)),
  inference(forward_demodulation,[],[f5564,f974])).
fof(f5564,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))))) | ($spl1 | $spl2 | $spl13)),
  inference(unit_resulting_resolution,[],[f1713,f5455,f1394])).
fof(f6145,plain,(
  $false | ($spl1 | $spl2 | $spl13 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5763,f5613])).
fof(f5763,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK10))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl1 | $spl2 | $spl13 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5597])).
fof(f5597,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl1 | $spl2 | $spl13)),
  inference(forward_demodulation,[],[f5581,f1341])).
fof(f5581,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl1 | $spl2 | $spl13)),
  inference(unit_resulting_resolution,[],[f1855,f5455,f1394])).
fof(f6144,plain,(
  $false | ($spl1 | $spl2 | $spl13 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5762,f1342])).
fof(f5762,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl1 | $spl2 | $spl13 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5594])).
fof(f5594,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl1 | $spl2 | $spl13)),
  inference(forward_demodulation,[],[f5593,f1334])).
fof(f5593,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))),s('type.num.num',sK11))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl1 | $spl2 | $spl13)),
  inference(forward_demodulation,[],[f5584,f974])).
fof(f5584,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl1 | $spl2 | $spl13)),
  inference(unit_resulting_resolution,[],[f1713,f5455,f1394])).
fof(f6141,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5759,f5613])).
fof(f5759,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK10))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5071])).
fof(f5071,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4601,f1341])).
fof(f4601,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1855,f2465,f1394])).
fof(f6140,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5758,f1342])).
fof(f5758,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5068])).
fof(f5068,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f5067,f1334])).
fof(f5067,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))),s('type.num.num',sK11))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4604,f974])).
fof(f4604,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1713,f2465,f1394])).
fof(f6139,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5757,f5613])).
fof(f5757,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK10))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5014])).
fof(f5014,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f5013,f1341])).
fof(f5013,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4655,f1341])).
fof(f4655,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1855,f1808,f1394])).
fof(f6138,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5756,f1342])).
fof(f5756,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5008])).
fof(f5008,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f5007,f1334])).
fof(f5007,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))),s('type.num.num',sK11))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f5006,f974])).
fof(f5006,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4658,f1341])).
fof(f4658,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1713,f1808,f1394])).
fof(f6137,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5755,f5613])).
fof(f5755,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK10))))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4992])).
fof(f4992,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4669,f1341])).
fof(f4669,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f2465,f1855,f1394])).
fof(f6136,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5754,f5613])).
fof(f5754,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK10))))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4988])).
fof(f4988,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4987,f1341])).
fof(f4987,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4672,f1341])).
fof(f4672,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1808,f1855,f1394])).
fof(f6135,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5753,f5613])).
fof(f5753,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK10))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK10))))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4986])).
fof(f4986,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4673,f1341])).
fof(f4673,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1855,f1855,f1394])).
fof(f6134,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5752,f5613])).
fof(f5752,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK10))))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4985])).
fof(f4985,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4984,f1337])).
fof(f4984,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real',sK10))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4674,f1341])).
fof(f4674,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real',sK10))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1813,f1855,f1394])).
fof(f6133,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5751,f1342])).
fof(f5751,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK10))))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4981])).
fof(f4981,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4980,f1334])).
fof(f4980,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))),s('type.num.num',sK11))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4979,f974])).
fof(f4979,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4676,f1341])).
fof(f4676,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1713,f1855,f1394])).
fof(f6132,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5750,f5613])).
fof(f5750,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK10))))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4976])).
fof(f4976,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4678,f1341])).
fof(f4678,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1714,f1855,f1394])).
fof(f6131,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5749,f5613])).
fof(f5749,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK10))))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4975])).
fof(f4975,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4679,f1341])).
fof(f4679,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1734,f1855,f1394])).
fof(f6130,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5748,f5613])).
fof(f5748,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK10))),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4959])).
fof(f4959,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4958,f1341])).
fof(f4958,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4691,f1337])).
fof(f4691,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real',sK10))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1855,f1813,f1394])).
fof(f6129,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5747,f1342])).
fof(f5747,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4954])).
fof(f4954,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4953,f1334])).
fof(f4953,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))),s('type.num.num',sK11))),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4952,f974])).
fof(f4952,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4694,f1337])).
fof(f4694,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real',sK10))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1713,f1813,f1394])).
fof(f6128,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5746,f1340])).
fof(f5746,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real',sK10))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4904])).
fof(f4904,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4903,f1334])).
fof(f4903,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))),s('type.num.num',sK11))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4723,f974])).
fof(f4723,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f2465,f1713,f1394])).
fof(f6127,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5745,f1340])).
fof(f5745,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real',sK10))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4897])).
fof(f4897,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4896,f1341])).
fof(f4896,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4895,f1334])).
fof(f4895,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))),s('type.num.num',sK11))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4726,f974])).
fof(f4726,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1808,f1713,f1394])).
fof(f6126,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5744,f1340])).
fof(f5744,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK10))),s('type.realax.real',sK10))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4894])).
fof(f4894,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4893,f1341])).
fof(f4893,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4892,f1334])).
fof(f4892,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))),s('type.num.num',sK11))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4727,f974])).
fof(f4727,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1855,f1713,f1394])).
fof(f6125,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5743,f1340])).
fof(f5743,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real',sK10))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4891])).
fof(f4891,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4890,f1337])).
fof(f4890,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real',sK10))),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4889,f1334])).
fof(f4889,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real',sK10))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))),s('type.num.num',sK11))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4728,f974])).
fof(f4728,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real',sK10))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1813,f1713,f1394])).
fof(f6124,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5742,f1340])).
fof(f5742,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real',sK10))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4885])).
fof(f4885,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4884,f1334])).
fof(f4884,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))),s('type.num.num',sK11))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))),s('type.num.num',sK11))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4730,f974])).
fof(f4730,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1713,f1713,f1394])).
fof(f6123,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5741,f1340])).
fof(f5741,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real',sK10))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4880])).
fof(f4880,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4879,f1334])).
fof(f4879,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))),s('type.num.num',sK11))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4732,f974])).
fof(f4732,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1714,f1713,f1394])).
fof(f6122,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5740,f1340])).
fof(f5740,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real',sK10))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4878])).
fof(f4878,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4877,f1334])).
fof(f4877,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))),s('type.num.num',sK11))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4733,f974])).
fof(f4733,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1734,f1713,f1394])).
fof(f6121,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5739,f5613])).
fof(f5739,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK10))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4845])).
fof(f4845,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4763,f1341])).
fof(f4763,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1855,f1714,f1394])).
fof(f6120,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5738,f1342])).
fof(f5738,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4842])).
fof(f4842,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4841,f1334])).
fof(f4841,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))),s('type.num.num',sK11))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4766,f974])).
fof(f4766,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1713,f1714,f1394])).
fof(f6119,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5737,f5613])).
fof(f5737,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK10))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4835])).
fof(f4835,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4781,f1341])).
fof(f4781,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1855,f1734,f1394])).
fof(f6118,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5736,f1342])).
fof(f5736,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4832])).
fof(f4832,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4831,f1334])).
fof(f4831,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))),s('type.num.num',sK11))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4784,f974])).
fof(f4784,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1713,f1734,f1394])).
fof(f6117,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5733,f5613])).
fof(f5733,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK10),s('type.realax.real',sK10))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f2679])).
fof(f2679,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f2486,f1354])).
fof(f2486,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f2485,f1334])).
fof(f2485,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))),s('type.num.num',sK11))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f2463,f974])).
fof(f2463,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1713,f1850])).
fof(f6116,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5732,f5613])).
fof(f5732,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK10))),s('type.realax.real',sK10))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f2496])).
fof(f2496,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.realax.real',sK10))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f2486,f1352])).
fof(f6115,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5731,f5613])).
fof(f5731,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK10))))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f2495])).
fof(f2495,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f2486,f1352])).
fof(f6114,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5730,f5613])).
fof(f5730,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK10))))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f2494])).
fof(f2494,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f2486,f1388])).
fof(f6113,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5729,f5613])).
fof(f5729,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK10))),s('type.num.num',X0)))) ) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f2492])).
fof(f2492,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.num.num',X0)))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f2486,f1392])).
fof(f6112,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5728,f5613])).
fof(f5728,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK10))))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f2491])).
fof(f2491,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f2486,f1850])).
fof(f6111,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5727,f1332])).
fof(f5727,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK10))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f2486])).
fof(f6110,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5726,f2403])).
fof(f5726,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real',sK10))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f2409])).
fof(f2409,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f2043,f871])).
fof(f2043,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11)))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f2042,f1334])).
fof(f2042,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))),s('type.num.num',sK11)))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f2041,f1334])).
fof(f2041,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))))),s('type.num.num',sK11)))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f2040,f974])).
fof(f2040,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))),s('type.num.num',sK11)))))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f2022,f974])).
fof(f2022,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11)))))))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1713,f1390])).
fof(f6109,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5725,f5613])).
fof(f5725,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',X0)))) ) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f2192])).
fof(f2192,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.num.num',X0)))) ) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f2191,f1334])).
fof(f2191,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))),s('type.num.num',sK11))),s('type.num.num',X0)))) ) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f2164,f974])).
fof(f2164,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.num.num',X0)))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1713,f1392])).
fof(f6108,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5715,f1335])).
fof(f5715,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real',sK10))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f1855])).
fof(f6107,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5714,f1335])).
fof(f5714,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real',sK10))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f1852])).
fof(f1852,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.realax.real',sK10))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f1851,f1334])).
fof(f1851,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))),s('type.num.num',sK11))),s('type.realax.real',sK10))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f1815,f974])).
fof(f1815,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.realax.real',sK10))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1713,f1352])).
fof(f6106,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5713,f1334])).
fof(f5713,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f1713])).
fof(f6100,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f6099,f1205])).
fof(f6099,plain,(
  p(s(bool,'F_0')) | ($spl2 | $spl26 | $spl28)),
  inference(forward_demodulation,[],[f5685,f2451])).
fof(f5685,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real',sK10)))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f2378])).
fof(f2378,plain,(
  ( ! [X0] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.num.num',X0))))))) ) | $spl2),
  inference(forward_demodulation,[],[f2372,f974])).
fof(f2372,plain,(
  ( ! [X0] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.num.num',X0))))))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f2193,f1390])).
fof(f6097,plain,(
  $false | ($spl2 | $spl13 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5682,f1342])).
fof(f5682,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl2 | $spl13 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5570])).
fof(f5570,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.num.num',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))) ) | ($spl2 | $spl13)),
  inference(unit_resulting_resolution,[],[f2193,f5455,f1394])).
fof(f6096,plain,(
  $false | ($spl2 | $spl13 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5681,f1340])).
fof(f5681,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real',sK10))) | ($spl2 | $spl13 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5550])).
fof(f5550,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.num.num',X0)))))) ) | ($spl2 | $spl13)),
  inference(unit_resulting_resolution,[],[f2193,f5455,f1394])).
fof(f6094,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5679,f1340])).
fof(f5679,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real',sK10))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5166])).
fof(f5166,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.num.num',X0)))))) ) | $spl2),
  inference(forward_demodulation,[],[f4492,f1341])).
fof(f4492,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.num.num',X0)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f1808,f2193,f1394])).
fof(f6093,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5678,f1340])).
fof(f5678,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.realax.real',sK10))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5165])).
fof(f5165,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.num.num',X0)))))) ) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4493,f1341])).
fof(f4493,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.num.num',X0)))))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1855,f2193,f1394])).
fof(f6092,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5677,f1340])).
fof(f5677,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real',sK10))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5164])).
fof(f5164,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.num.num',X0)))))) ) | $spl2),
  inference(forward_demodulation,[],[f4494,f1337])).
fof(f4494,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real',sK10))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.num.num',X0)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f1813,f2193,f1394])).
fof(f6091,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5676,f1340])).
fof(f5676,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.realax.real',sK10))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5162])).
fof(f5162,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.num.num',X0)))))) ) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f5161,f1334])).
fof(f5161,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))),s('type.num.num',sK11))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.num.num',X0)))))) ) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4496,f974])).
fof(f4496,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.num.num',X0)))))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1713,f2193,f1394])).
fof(f6090,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5675,f1342])).
fof(f5675,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5027])).
fof(f5027,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.num.num',X0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10)))))))) ) | $spl2),
  inference(forward_demodulation,[],[f4645,f1341])).
fof(f4645,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.num.num',X0))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10)))))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f2193,f1808,f1394])).
fof(f6089,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5674,f1342])).
fof(f5674,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5000])).
fof(f5000,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.num.num',X0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11)))))))) ) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4663,f1341])).
fof(f4663,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.num.num',X0))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11)))))))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f2193,f1855,f1394])).
fof(f6088,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5673,f1342])).
fof(f5673,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4973])).
fof(f4973,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.num.num',X0))),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10)))))) ) | $spl2),
  inference(forward_demodulation,[],[f4681,f1337])).
fof(f4681,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.num.num',X0))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real',sK10)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f2193,f1813,f1394])).
fof(f6087,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5672,f1342])).
fof(f5672,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4918])).
fof(f4918,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.num.num',X0))),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11)))))) ) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4917,f1334])).
fof(f4917,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.num.num',X0))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))),s('type.num.num',sK11)))))) ) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4717,f974])).
fof(f4717,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.num.num',X0))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11)))))))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f2193,f1713,f1394])).
fof(f6086,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5671,f1342])).
fof(f5671,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4771])).
fof(f4771,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.num.num',X0))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10)))))))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f2193,f1734,f1394])).
fof(f6085,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5670,f1342])).
fof(f5670,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4753])).
fof(f4753,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.num.num',X0))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10)))))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f2193,f1714,f1394])).
fof(f6084,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5669,f1342])).
fof(f5669,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4591])).
fof(f4591,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.num.num',X0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10)))))))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f2193,f2465,f1394])).
fof(f6083,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5668,f1340])).
fof(f5668,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real',sK10))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4499])).
fof(f4499,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.num.num',X0)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f1734,f2193,f1394])).
fof(f6082,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5667,f1340])).
fof(f5667,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real',sK10))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4498])).
fof(f4498,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.num.num',X0)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f1714,f2193,f1394])).
fof(f6081,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5666,f1340])).
fof(f5666,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real',sK10))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4489])).
fof(f4489,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.num.num',X0)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f2465,f2193,f1394])).
fof(f6080,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5665,f1342])).
fof(f5665,plain,(
  ( ! [X1] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.num.num',X1)))))) ) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4483])).
fof(f4483,plain,(
  ( ! [X0,X1] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.num.num',X0))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.num.num',X1)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f2193,f2193,f1394])).
fof(f6079,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5664,f1332])).
fof(f5664,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK10))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f2457])).
fof(f2457,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.num.num',X0)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f2193,f1850])).
fof(f6078,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5663,f1335])).
fof(f5663,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real',sK10))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f2375])).
fof(f2375,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.num.num',X0))),s('type.realax.real',sK10)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f2193,f1352])).
fof(f6077,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5662,f1335])).
fof(f5662,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real',sK10))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f2374])).
fof(f2374,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.num.num',X0)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f2193,f1352])).
fof(f6076,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5661,f1334])).
fof(f5661,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f2373])).
fof(f2373,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.num.num',X0)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f2193,f1388])).
fof(f6075,plain,(
  $false | ($spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5660,f5613])).
fof(f5660,plain,(
  ( ! [X1] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',X1)))) ) | ($spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f2371])).
fof(f2371,plain,(
  ( ! [X0,X1] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.num.num',X0))),s('type.num.num',X1)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f2193,f1392])).
fof(f6072,plain,(
  $false | ($spl1 | $spl2 | $spl13 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5652,f1342])).
fof(f5652,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl1 | $spl2 | $spl13 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5571])).
fof(f5571,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.num.num',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))) ) | ($spl1 | $spl2 | $spl13)),
  inference(unit_resulting_resolution,[],[f2192,f5455,f1394])).
fof(f6071,plain,(
  $false | ($spl1 | $spl2 | $spl13 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5651,f1340])).
fof(f5651,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real',sK10))) | ($spl1 | $spl2 | $spl13 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5551])).
fof(f5551,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.num.num',X0)))))) ) | ($spl1 | $spl2 | $spl13)),
  inference(unit_resulting_resolution,[],[f2192,f5455,f1394])).
fof(f6069,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5649,f1340])).
fof(f5649,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real',sK10))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5156])).
fof(f5156,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.num.num',X0)))))) ) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4510,f1341])).
fof(f4510,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.num.num',X0)))))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1808,f2192,f1394])).
fof(f6068,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5648,f1340])).
fof(f5648,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.realax.real',sK10))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5155])).
fof(f5155,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.num.num',X0)))))) ) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4511,f1341])).
fof(f4511,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.num.num',X0)))))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1855,f2192,f1394])).
fof(f6067,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5647,f1340])).
fof(f5647,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real',sK10))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5154])).
fof(f5154,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.num.num',X0)))))) ) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4512,f1337])).
fof(f4512,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real',sK10))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.num.num',X0)))))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1813,f2192,f1394])).
fof(f6066,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5646,f1340])).
fof(f5646,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.realax.real',sK10))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5152])).
fof(f5152,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.num.num',X0)))))) ) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f5151,f1334])).
fof(f5151,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))),s('type.num.num',sK11))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.num.num',X0)))))) ) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4514,f974])).
fof(f4514,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.num.num',X0)))))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1713,f2192,f1394])).
fof(f6065,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5645,f1342])).
fof(f5645,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f5026])).
fof(f5026,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.num.num',X0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10)))))))) ) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4646,f1341])).
fof(f4646,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.num.num',X0))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10)))))))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f2192,f1808,f1394])).
fof(f6064,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5644,f1342])).
fof(f5644,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4999])).
fof(f4999,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.num.num',X0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11)))))))) ) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4664,f1341])).
fof(f4664,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.num.num',X0))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11)))))))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f2192,f1855,f1394])).
fof(f6063,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5643,f1342])).
fof(f5643,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4972])).
fof(f4972,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.num.num',X0))),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10)))))) ) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4682,f1337])).
fof(f4682,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.num.num',X0))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real',sK10)))))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f2192,f1813,f1394])).
fof(f6062,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5642,f1342])).
fof(f5642,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4916])).
fof(f4916,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.num.num',X0))),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11)))))) ) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4915,f1334])).
fof(f4915,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.num.num',X0))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))),s('type.num.num',sK11)))))) ) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f4718,f974])).
fof(f4718,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.num.num',X0))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11)))))))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f2192,f1713,f1394])).
fof(f6061,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5641,f1342])).
fof(f5641,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4772])).
fof(f4772,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.num.num',X0))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10)))))))))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f2192,f1734,f1394])).
fof(f6060,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5640,f1342])).
fof(f5640,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4754])).
fof(f4754,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.num.num',X0))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10)))))))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f2192,f1714,f1394])).
fof(f6059,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5639,f1342])).
fof(f5639,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4592])).
fof(f4592,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.num.num',X0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10)))))))))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f2192,f2465,f1394])).
fof(f6058,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5638,f1340])).
fof(f5638,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real',sK10))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4517])).
fof(f4517,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.num.num',X0)))))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1734,f2192,f1394])).
fof(f6057,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5637,f1340])).
fof(f5637,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real',sK10))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4516])).
fof(f4516,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.num.num',X0)))))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1714,f2192,f1394])).
fof(f6056,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5636,f1340])).
fof(f5636,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real',sK10))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4507])).
fof(f4507,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.num.num',X0)))))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f2465,f2192,f1394])).
fof(f6055,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5635,f1340])).
fof(f5635,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.num.num',X0))),s('type.realax.real',sK10)))) ) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4502])).
fof(f4502,plain,(
  ( ! [X0,X1] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.num.num',X0))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.num.num',X1)))))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f2192,f2192,f1394])).
fof(f6054,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5634,f1340])).
fof(f5634,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.num.num',X0))),s('type.realax.real',sK10)))) ) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4501])).
fof(f4501,plain,(
  ( ! [X0,X1] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.num.num',X0))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.num.num',X1)))))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f2193,f2192,f1394])).
fof(f6053,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5633,f1342])).
fof(f5633,plain,(
  ( ! [X1] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.num.num',X1)))))) ) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f4484])).
fof(f4484,plain,(
  ( ! [X0,X1] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.num.num',X0))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.num.num',X1)))))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f2192,f2193,f1394])).
fof(f6052,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f6051,f1205])).
fof(f6051,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(forward_demodulation,[],[f5632,f2451])).
fof(f5632,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real',sK10)))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f3326])).
fof(f3326,plain,(
  ( ! [X0] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.num.num',X0))))))) ) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f3325,f1334])).
fof(f3325,plain,(
  ( ! [X0] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))),s('type.num.num',sK11))),s('type.num.num',X0))))))) ) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f3324,f974])).
fof(f3324,plain,(
  ( ! [X0] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.num.num',X0))))))) ) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f3318,f974])).
fof(f3318,plain,(
  ( ! [X0] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.num.num',X0))))))))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f2192,f1390])).
fof(f6050,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5631,f1335])).
fof(f5631,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real',sK10))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f3321])).
fof(f3321,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.num.num',X0))),s('type.realax.real',sK10)))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f2192,f1352])).
fof(f6049,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5630,f1335])).
fof(f5630,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real',sK10))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f3320])).
fof(f3320,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.num.num',X0)))))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f2192,f1352])).
fof(f6048,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5629,f1334])).
fof(f5629,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f3319])).
fof(f3319,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.num.num',X0)))))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f2192,f1388])).
fof(f6047,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5628,f5613])).
fof(f5628,plain,(
  ( ! [X1] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',X1)))) ) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f3317])).
fof(f3317,plain,(
  ( ! [X0,X1] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.num.num',X0))),s('type.num.num',X1)))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f2192,f1392])).
fof(f6046,plain,(
  $false | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(subsumption_resolution,[],[f5627,f1332])).
fof(f5627,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK10))) | ($spl1 | $spl2 | $spl26 | $spl28)),
  inference(backward_demodulation,[],[f5613,f3316])).
fof(f3316,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.num.num',X0)))))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f2192,f1850])).
fof(f5612,plain,(
  $false | ($spl2 | $spl13 | $spl25)),
  inference(subsumption_resolution,[],[f5547,f5483])).
fof(f5545,plain,(
  $false | ($spl2 | $spl6 | $spl13)),
  inference(subsumption_resolution,[],[f5533,f5455])).
fof(f5533,plain,(
  s('type.realax.real',sK10) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl2 | $spl6)),
  inference(backward_demodulation,[],[f5445,f5348])).
fof(f5348,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f1714,f1378])).
fof(f1378,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.realax.real',X0))) | s('type.realax.real',sK10) = s('type.realax.real',X0)) ) | $spl2),
  inference(backward_demodulation,[],[f1196,f1260])).
fof(f5544,plain,(
  $false | ($spl2 | $spl6 | $spl13)),
  inference(subsumption_resolution,[],[f5532,f5455])).
fof(f5532,plain,(
  s('type.realax.real',sK10) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl2 | $spl6)),
  inference(backward_demodulation,[],[f5445,f5349])).
fof(f5349,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f1734,f1378])).
fof(f5543,plain,(
  $false | ($spl2 | $spl6 | $spl13)),
  inference(subsumption_resolution,[],[f5531,f5455])).
fof(f5531,plain,(
  s('type.realax.real',sK10) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl2 | $spl6)),
  inference(backward_demodulation,[],[f5445,f5438])).
fof(f5438,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))) | $spl2),
  inference(forward_demodulation,[],[f5344,f1337])).
fof(f5344,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real',sK10))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.realax.real',sK10))))) | $spl2),
  inference(unit_resulting_resolution,[],[f1813,f1378])).
fof(f5542,plain,(
  $false | ($spl2 | $spl6 | $spl13)),
  inference(subsumption_resolution,[],[f5530,f5455])).
fof(f5530,plain,(
  s('type.realax.real',sK10) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl2 | $spl6)),
  inference(backward_demodulation,[],[f5445,f5339])).
fof(f5339,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f2465,f1378])).
fof(f5541,plain,(
  $false | ($spl2 | $spl6 | $spl13)),
  inference(subsumption_resolution,[],[f5529,f5455])).
fof(f5529,plain,(
  s('type.realax.real',sK10) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl2 | $spl6)),
  inference(backward_demodulation,[],[f5445,f5440])).
fof(f5440,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | $spl2),
  inference(forward_demodulation,[],[f5342,f1341])).
fof(f5342,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))))))) | $spl2),
  inference(unit_resulting_resolution,[],[f1808,f1378])).
fof(f5540,plain,(
  $false | ($spl1 | $spl2 | $spl6 | $spl13)),
  inference(subsumption_resolution,[],[f5528,f5455])).
fof(f5528,plain,(
  s('type.realax.real',sK10) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl1 | $spl2 | $spl6)),
  inference(backward_demodulation,[],[f5445,f5439])).
fof(f5439,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f5343,f1341])).
fof(f5343,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1855,f1378])).
fof(f5539,plain,(
  $false | ($spl2 | $spl6 | $spl13)),
  inference(subsumption_resolution,[],[f5527,f5455])).
fof(f5527,plain,(
  s('type.realax.real',sK10) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl2 | $spl6)),
  inference(backward_demodulation,[],[f5445,f5333])).
fof(f5333,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.num.num',X0))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.cos_1'(s('type.realax.real',sK10))),s('type.num.num',X0)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f2193,f1378])).
fof(f5538,plain,(
  $false | ($spl1 | $spl2 | $spl6 | $spl13)),
  inference(subsumption_resolution,[],[f5526,f5455])).
fof(f5526,plain,(
  s('type.realax.real',sK10) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl1 | $spl2 | $spl6)),
  inference(backward_demodulation,[],[f5445,f5334])).
fof(f5334,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.num.num',X0))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))),s('type.num.num',X0)))))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f2192,f1378])).
fof(f5537,plain,(
  $false | ($spl1 | $spl2 | $spl6 | $spl13)),
  inference(subsumption_resolution,[],[f5525,f5455])).
fof(f5525,plain,(
  s('type.realax.real',sK10) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl1 | $spl2 | $spl6)),
  inference(backward_demodulation,[],[f5445,f5436])).
fof(f5436,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f5435,f1334])).
fof(f5435,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))),s('type.num.num',sK11))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',sK10))),s('type.num.num',sK11))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f5346,f974])).
fof(f5346,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))))),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))))))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f1713,f1378])).
fof(f5536,plain,(
  $false | ($spl2 | $spl6 | $spl13)),
  inference(subsumption_resolution,[],[f5534,f5455])).
fof(f5534,plain,(
  s('type.realax.real',sK10) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl2 | $spl6)),
  inference(condensation,[],[f5524])).
fof(f5524,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',sK10) = s('type.realax.real',X0)) ) | ($spl2 | $spl6)),
  inference(backward_demodulation,[],[f5445,f1378])).
fof(f5455,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | $spl13),
  inference(cnf_transformation,[],[f5455_D])).
fof(f5455_D,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f5445,plain,(
  ( ! [X5] : (s('type.realax.real',sK10) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X5))),s('type.realax.real',X5)))) ) | $spl6),
  inference(cnf_transformation,[],[f5445_D])).
fof(f5445_D,plain,(
  ( ! [X5] : (s('type.realax.real',sK10) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X5))),s('type.realax.real',X5)))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f5511,plain,(
  ( ! [X90] : (s('type.realax.real',sK10) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X90))),s('type.realax.real',X90))) | s('type.realax.real',sK10) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl2),
  inference(subsumption_resolution,[],[f5510,f1396])).
fof(f1396,plain,(
  ( ! [X0,X1] : (s('type.realax.real',sK10) != s('type.realax.real',X1) | s('type.realax.real',sK10) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) ) | $spl2),
  inference(forward_demodulation,[],[f1358,f1196])).
fof(f1358,plain,(
  ( ! [X0,X1] : (s('type.realax.real',sK10) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))) | s('type.realax.real',X1) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl2),
  inference(backward_demodulation,[],[f1196,f1032])).
fof(f1032,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))) | s('type.realax.real',X1) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f708])).
fof(f5510,plain,(
  ( ! [X90] : (s('type.realax.real',sK10) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X90))),s('type.realax.real',X90))) | s('type.realax.real',sK10) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',sK10) = s('type.realax.real',X90)) ) | $spl2),
  inference(forward_demodulation,[],[f5421,f1332])).
fof(f5421,plain,(
  ( ! [X90] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK10))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X90))),s('type.realax.real',X90))) | s('type.realax.real',sK10) = s('type.realax.real',X90)) ) | $spl2),
  inference(superposition,[],[f2718,f1378])).
fof(f2718,plain,(
  ( ! [X1] : (s('type.realax.real',sK10) != s('type.realax.real',X1) | s('type.realax.real',X1) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK10)))) ) | $spl2),
  inference(superposition,[],[f1354,f1338])).
fof(f1338,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',sK10)))) ) | $spl2),
  inference(backward_demodulation,[],[f1196,f809])).
fof(f809,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(cnf_transformation,[],[f297])).
fof(f297,plain,(
  ! [X0] : s('type.realax.real',X0) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f193])).
fof(f193,axiom,(
  ! [X6] : s('type.realax.real',X6) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X6),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.real.REAL_ADD_RID')).
fof(f5509,plain,(
  ( ! [X89] : (s('type.realax.real',sK10) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X89),s('type.realax.real',X89))))) | s('type.realax.real',sK10) = s('type.realax.real',X89)) ) | $spl44),
  inference(cnf_transformation,[],[f5509_D])).
fof(f5509_D,plain,(
  ( ! [X89] : (s('type.realax.real',sK10) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X89),s('type.realax.real',X89))))) | s('type.realax.real',sK10) = s('type.realax.real',X89)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f5508,plain,(
  ( ! [X89] : (s('type.realax.real',sK10) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X89),s('type.realax.real',X89))))) | s('type.realax.real',sK10) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',sK10) = s('type.realax.real',X89)) ) | $spl2),
  inference(forward_demodulation,[],[f5507,f958])).
fof(f5507,plain,(
  ( ! [X89] : (s('type.realax.real',sK10) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X89),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X89))))))) | s('type.realax.real',sK10) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',sK10) = s('type.realax.real',X89)) ) | $spl2),
  inference(forward_demodulation,[],[f5420,f952])).
fof(f5420,plain,(
  ( ! [X89] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',sK10) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X89))),s('type.realax.real',X89))))) | s('type.realax.real',sK10) = s('type.realax.real',X89)) ) | $spl2),
  inference(superposition,[],[f2717,f1378])).
fof(f2717,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real',X0) | s('type.realax.real',sK10) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))) ) | $spl2),
  inference(superposition,[],[f1354,f1336])).
fof(f1336,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',sK10),s('type.realax.real',X0)))) ) | $spl2),
  inference(backward_demodulation,[],[f1196,f803])).
fof(f803,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) )),
  inference(cnf_transformation,[],[f291])).
fof(f291,plain,(
  ! [X0] : s('type.realax.real',X0) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))),
  inference(rectify,[],[f209])).
fof(f209,axiom,(
  ! [X6] : s('type.realax.real',X6) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X6)))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.real.REAL_ADD_LID')).
fof(f5506,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl43),
  inference(cnf_transformation,[],[f5506_D])).
fof(f5506_D,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f5504,plain,(
  ( ! [X77] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.realax.real',sK10) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X77))),s('type.realax.real',X77)))) ) | $spl2),
  inference(subsumption_resolution,[],[f5412,f1396])).
fof(f5412,plain,(
  ( ! [X77] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.realax.real',sK10) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X77))),s('type.realax.real',X77))) | s('type.realax.real',sK10) = s('type.realax.real',X77)) ) | $spl2),
  inference(superposition,[],[f1850,f1378])).
fof(f5503,plain,(
  ( ! [X50,X51] : (s('type.realax.real',sK10) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X50),s('type.realax.real',X50))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X51))))) | s('type.realax.real',sK10) = s('type.realax.real',X50)) ) | $spl40),
  inference(cnf_transformation,[],[f5503_D])).
fof(f5503_D,plain,(
  ( ! [X50,X51] : (s('type.realax.real',sK10) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X50),s('type.realax.real',X50))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X51))))) | s('type.realax.real',sK10) = s('type.realax.real',X50)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f5502,plain,(
  ( ! [X50,X51] : (s('type.realax.real',sK10) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X50),s('type.realax.real',X50))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X51))))) | s('type.realax.real',sK10) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',sK10) = s('type.realax.real',X50)) ) | $spl2),
  inference(forward_demodulation,[],[f5501,f958])).
fof(f5501,plain,(
  ( ! [X50,X51] : (s('type.realax.real',sK10) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X50),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X50))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X51))))) | s('type.realax.real',sK10) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',sK10) = s('type.realax.real',X50)) ) | $spl2),
  inference(forward_demodulation,[],[f5402,f952])).
fof(f5402,plain,(
  ( ! [X50,X51] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',sK10) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X50))),s('type.realax.real',X50))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X51))))) | s('type.realax.real',sK10) = s('type.realax.real',X50)) ) | $spl2),
  inference(superposition,[],[f1397,f1378])).
fof(f1397,plain,(
  ( ! [X0,X1] : (s('type.realax.real',sK10) != s('type.realax.real',X1) | s('type.realax.real',sK10) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) ) | $spl2),
  inference(forward_demodulation,[],[f1359,f1196])).
fof(f1359,plain,(
  ( ! [X0,X1] : (s('type.realax.real',sK10) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) | s('type.realax.real',X1) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl2),
  inference(backward_demodulation,[],[f1196,f1036])).
fof(f1036,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) | s('type.realax.real',X1) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f710])).
fof(f710,plain,(
  ! [X0,X1] : ((s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.realax.real',X1) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))))),
  inference(nnf_transformation,[],[f482])).
fof(f482,plain,(
  ! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) <=> s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X7,X6] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))) <=> s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.real.POW_ZERO_EQ')).
fof(f5500,plain,(
  ( ! [X48,X49] : (s('type.realax.real',sK10) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X49),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X48),s('type.realax.real',X48))))) | s('type.realax.real',sK10) = s('type.realax.real',X48)) ) | $spl38),
  inference(cnf_transformation,[],[f5500_D])).
fof(f5500_D,plain,(
  ( ! [X48,X49] : (s('type.realax.real',sK10) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X49),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X48),s('type.realax.real',X48))))) | s('type.realax.real',sK10) = s('type.realax.real',X48)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f5499,plain,(
  ( ! [X48,X49] : (s('type.realax.real',sK10) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X49),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X48),s('type.realax.real',X48))))) | s('type.realax.real',sK10) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',sK10) = s('type.realax.real',X48)) ) | $spl2),
  inference(forward_demodulation,[],[f5498,f958])).
fof(f5498,plain,(
  ( ! [X48,X49] : (s('type.realax.real',sK10) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X49),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X48),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X48))))))) | s('type.realax.real',sK10) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',sK10) = s('type.realax.real',X48)) ) | $spl2),
  inference(forward_demodulation,[],[f5401,f952])).
fof(f5401,plain,(
  ( ! [X48,X49] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',sK10) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X49),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X48))),s('type.realax.real',X48))))) | s('type.realax.real',sK10) = s('type.realax.real',X48)) ) | $spl2),
  inference(superposition,[],[f1396,f1378])).
fof(f5497,plain,(
  ( ! [X47,X46] : (s('type.realax.real',sK10) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X46),s('type.realax.real',X46))),s('type.realax.real',X47))) | s('type.realax.real',sK10) = s('type.realax.real',X46)) ) | $spl36),
  inference(cnf_transformation,[],[f5497_D])).
fof(f5497_D,plain,(
  ( ! [X47,X46] : (s('type.realax.real',sK10) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X46),s('type.realax.real',X46))),s('type.realax.real',X47))) | s('type.realax.real',sK10) = s('type.realax.real',X46)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f5496,plain,(
  ( ! [X47,X46] : (s('type.realax.real',sK10) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X46),s('type.realax.real',X46))),s('type.realax.real',X47))) | s('type.realax.real',sK10) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',sK10) = s('type.realax.real',X46)) ) | $spl2),
  inference(forward_demodulation,[],[f5495,f958])).
fof(f5495,plain,(
  ( ! [X47,X46] : (s('type.realax.real',sK10) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X46),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X46))))),s('type.realax.real',X47))) | s('type.realax.real',sK10) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',sK10) = s('type.realax.real',X46)) ) | $spl2),
  inference(forward_demodulation,[],[f5400,f952])).
fof(f5400,plain,(
  ( ! [X47,X46] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',sK10) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X46))),s('type.realax.real',X46))),s('type.realax.real',X47))) | s('type.realax.real',sK10) = s('type.realax.real',X46)) ) | $spl2),
  inference(superposition,[],[f1395,f1378])).
fof(f1395,plain,(
  ( ! [X0,X1] : (s('type.realax.real',sK10) != s('type.realax.real',X0) | s('type.realax.real',sK10) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) ) | $spl2),
  inference(forward_demodulation,[],[f1357,f1196])).
fof(f1357,plain,(
  ( ! [X0,X1] : (s('type.realax.real',sK10) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))) | s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl2),
  inference(backward_demodulation,[],[f1196,f1031])).
fof(f1031,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))) | s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f708])).
fof(f5494,plain,(
  ( ! [X45] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X45),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.realax.real',sK10) = s('type.realax.real',X45)) ) | $spl34),
  inference(cnf_transformation,[],[f5494_D])).
fof(f5494_D,plain,(
  ( ! [X45] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X45),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.realax.real',sK10) = s('type.realax.real',X45)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f5493,plain,(
  ( ! [X45,X44] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X45),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.realax.real',sK10) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X44))),s('type.realax.real',X44))) | s('type.realax.real',sK10) = s('type.realax.real',X45)) ) | $spl2),
  inference(subsumption_resolution,[],[f5399,f1396])).
fof(f5399,plain,(
  ( ! [X45,X44] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X45),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.realax.real',sK10) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X44))),s('type.realax.real',X44))) | s('type.realax.real',sK10) = s('type.realax.real',X45) | s('type.realax.real',sK10) = s('type.realax.real',X44)) ) | $spl2),
  inference(superposition,[],[f1394,f1378])).
fof(f5492,plain,(
  ( ! [X43] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real',X43))) | s('type.realax.real',sK10) = s('type.realax.real',X43)) ) | $spl32),
  inference(cnf_transformation,[],[f5492_D])).
fof(f5492_D,plain,(
  ( ! [X43] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real',X43))) | s('type.realax.real',sK10) = s('type.realax.real',X43)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f5491,plain,(
  ( ! [X43,X42] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real',X43))) | s('type.realax.real',sK10) = s('type.realax.real',X43) | s('type.realax.real',sK10) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X42))),s('type.realax.real',X42)))) ) | $spl2),
  inference(subsumption_resolution,[],[f5398,f1396])).
fof(f5398,plain,(
  ( ! [X43,X42] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real',X43))) | s('type.realax.real',sK10) = s('type.realax.real',X43) | s('type.realax.real',sK10) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X42))),s('type.realax.real',X42))) | s('type.realax.real',sK10) = s('type.realax.real',X42)) ) | $spl2),
  inference(superposition,[],[f1394,f1378])).
fof(f5490,plain,(
  ( ! [X41] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X41)))) ) | $spl30),
  inference(cnf_transformation,[],[f5490_D])).
fof(f5490_D,plain,(
  ( ! [X41] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X41)))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f5489,plain,(
  ( ! [X41,X40] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X41))) | s('type.realax.real',sK10) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X40))),s('type.realax.real',X40)))) ) | $spl2),
  inference(subsumption_resolution,[],[f5397,f1396])).
fof(f5397,plain,(
  ( ! [X41,X40] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X41))) | s('type.realax.real',sK10) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X40))),s('type.realax.real',X40))) | s('type.realax.real',sK10) = s('type.realax.real',X40)) ) | $spl2),
  inference(superposition,[],[f1392,f1378])).
fof(f5483,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))))) | $spl25),
  inference(cnf_transformation,[],[f5483_D])).
fof(f5483_D,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f5488,plain,(
  ( ! [X39] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X39),s('type.realax.real',X39))) | s('type.realax.real',sK10) = s('type.realax.real',X39)) ) | $spl28),
  inference(cnf_transformation,[],[f5488_D])).
fof(f5488_D,plain,(
  ( ! [X39] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X39),s('type.realax.real',X39))) | s('type.realax.real',sK10) = s('type.realax.real',X39)) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f5487,plain,(
  ( ! [X39] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X39),s('type.realax.real',X39))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))))) | s('type.realax.real',sK10) = s('type.realax.real',X39)) ) | $spl2),
  inference(forward_demodulation,[],[f5486,f958])).
fof(f5486,plain,(
  ( ! [X39] : (s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X39),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X39))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))))) | s('type.realax.real',sK10) = s('type.realax.real',X39)) ) | $spl2),
  inference(forward_demodulation,[],[f5396,f952])).
fof(f5396,plain,(
  ( ! [X39] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))))) | s('type.realax.real',sK10) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X39))),s('type.realax.real',X39))) | s('type.realax.real',sK10) = s('type.realax.real',X39)) ) | $spl2),
  inference(superposition,[],[f1391,f1378])).
fof(f1391,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0)))))) | s('type.realax.real',sK10) != s('type.realax.real',X0)) ) | $spl2),
  inference(forward_demodulation,[],[f1349,f1196])).
fof(f1349,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0)))))) | s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl2),
  inference(backward_demodulation,[],[f1196,f934])).
fof(f934,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0)))))) | s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f696])).
fof(f5484,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))))) | $spl24),
  inference(cnf_transformation,[],[f5484_D])).
fof(f5484_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f5485,plain,(
  ( ! [X38] : (s('type.realax.real',sK10) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X38),s('type.realax.real',X38)))) ) | $spl26),
  inference(cnf_transformation,[],[f5485_D])).
fof(f5485_D,plain,(
  ( ! [X38] : (s('type.realax.real',sK10) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X38),s('type.realax.real',X38)))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f5482,plain,(
  ( ! [X38] : (s('type.realax.real',sK10) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X38),s('type.realax.real',X38))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))))))) ) | $spl2),
  inference(forward_demodulation,[],[f5481,f958])).
fof(f5481,plain,(
  ( ! [X38] : (s('type.realax.real',sK10) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X38),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X38))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))))))) ) | $spl2),
  inference(forward_demodulation,[],[f5480,f952])).
fof(f5480,plain,(
  ( ! [X38] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))))) | s('type.realax.real',sK10) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X38))),s('type.realax.real',X38)))) ) | $spl2),
  inference(subsumption_resolution,[],[f5395,f1396])).
fof(f5395,plain,(
  ( ! [X38] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))))) | s('type.realax.real',sK10) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X38))),s('type.realax.real',X38))) | s('type.realax.real',sK10) = s('type.realax.real',X38)) ) | $spl2),
  inference(superposition,[],[f1390,f1378])).
fof(f5479,plain,(
  ( ! [X37] : (s('type.realax.real',sK10) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X37),s('type.realax.real',X37))))) | s('type.realax.real',sK10) = s('type.realax.real',X37)) ) | $spl22),
  inference(cnf_transformation,[],[f5479_D])).
fof(f5479_D,plain,(
  ( ! [X37] : (s('type.realax.real',sK10) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X37),s('type.realax.real',X37))))) | s('type.realax.real',sK10) = s('type.realax.real',X37)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f5478,plain,(
  ( ! [X37] : (s('type.realax.real',sK10) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X37),s('type.realax.real',X37))))) | s('type.realax.real',sK10) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',sK10) = s('type.realax.real',X37)) ) | $spl2),
  inference(forward_demodulation,[],[f5477,f958])).
fof(f5477,plain,(
  ( ! [X37] : (s('type.realax.real',sK10) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X37),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X37))))))) | s('type.realax.real',sK10) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',sK10) = s('type.realax.real',X37)) ) | $spl2),
  inference(forward_demodulation,[],[f5394,f952])).
fof(f5394,plain,(
  ( ! [X37] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',sK10) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X37))),s('type.realax.real',X37))))) | s('type.realax.real',sK10) = s('type.realax.real',X37)) ) | $spl2),
  inference(superposition,[],[f1389,f1378])).
fof(f1389,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real',X0) | s('type.realax.real',sK10) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0)))) ) | $spl2),
  inference(forward_demodulation,[],[f1347,f1196])).
fof(f1347,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) = s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0))) | s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl2),
  inference(backward_demodulation,[],[f1196,f887])).
fof(f5476,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl21),
  inference(cnf_transformation,[],[f5476_D])).
fof(f5476_D,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f5474,plain,(
  ( ! [X36] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.realax.real',sK10) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X36))),s('type.realax.real',X36)))) ) | $spl2),
  inference(subsumption_resolution,[],[f5393,f1396])).
fof(f5393,plain,(
  ( ! [X36] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.realax.real',sK10) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X36))),s('type.realax.real',X36))) | s('type.realax.real',sK10) = s('type.realax.real',X36)) ) | $spl2),
  inference(superposition,[],[f1388,f1378])).
fof(f5472,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl17),
  inference(cnf_transformation,[],[f5472_D])).
fof(f5472_D,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f5473,plain,(
  ( ! [X35,X34] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X34),s('type.realax.real',X34))),s('type.num.num',X35)))))) | s('type.realax.real',sK10) = s('type.realax.real',X34)) ) | $spl18),
  inference(cnf_transformation,[],[f5473_D])).
fof(f5473_D,plain,(
  ( ! [X35,X34] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X34),s('type.realax.real',X34))),s('type.num.num',X35)))))) | s('type.realax.real',sK10) = s('type.realax.real',X34)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f5470,plain,(
  ( ! [X35,X34] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X34),s('type.realax.real',X34))),s('type.num.num',X35)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.realax.real',sK10) = s('type.realax.real',X34)) ) | $spl2),
  inference(forward_demodulation,[],[f5469,f958])).
fof(f5469,plain,(
  ( ! [X35,X34] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X34),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X34))))),s('type.num.num',X35)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.realax.real',sK10) = s('type.realax.real',X34)) ) | $spl2),
  inference(forward_demodulation,[],[f5392,f952])).
fof(f5392,plain,(
  ( ! [X35,X34] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK10),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X34))),s('type.realax.real',X34))),s('type.num.num',X35)))))) | s('type.realax.real',sK10) = s('type.realax.real',X34)) ) | $spl2),
  inference(superposition,[],[f1387,f1378])).
fof(f5468,plain,(
  ( ! [X33,X31,X32] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X31),s('type.realax.real',X31))),s('type.realax.real',X32))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X31),s('type.realax.real',X31))),s('type.realax.real',X33))) | s('type.realax.real',sK10) = s('type.realax.real',X31)) ) | $spl14),
  inference(cnf_transformation,[],[f5468_D])).
fof(f5468_D,plain,(
  ( ! [X33,X31,X32] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X31),s('type.realax.real',X31))),s('type.realax.real',X32))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X31),s('type.realax.real',X31))),s('type.realax.real',X33))) | s('type.realax.real',sK10) = s('type.realax.real',X31)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f5467,plain,(
  ( ! [X33,X31,X32] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X31),s('type.realax.real',X31))),s('type.realax.real',X32))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X31),s('type.realax.real',X31))),s('type.realax.real',X33))) | s('type.realax.real',sK10) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',sK10) = s('type.realax.real',X31)) ) | $spl2),
  inference(forward_demodulation,[],[f5466,f958])).
fof(f5466,plain,(
  ( ! [X33,X31,X32] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X31),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X31))))),s('type.realax.real',X32))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X31),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X31))))),s('type.realax.real',X33))) | s('type.realax.real',sK10) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',sK10) = s('type.realax.real',X31)) ) | $spl2),
  inference(forward_demodulation,[],[f5391,f952])).
fof(f5391,plain,(
  ( ! [X33,X31,X32] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X31))),s('type.realax.real',X31))),s('type.realax.real',X32))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X31))),s('type.realax.real',X31))),s('type.realax.real',X33))) | s('type.realax.real',sK10) = s('type.realax.real',X31)) ) | $spl2),
  inference(superposition,[],[f1371,f1378])).
fof(f1371,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real',sK10) != s('type.realax.real',X0) | s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X2)))) ) | $spl2),
  inference(backward_demodulation,[],[f1196,f1093])).
fof(f1093,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X2))) | s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f727])).
fof(f727,plain,(
  ! [X0,X1,X2] : ((s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X2))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X1) = s('type.realax.real',X2)) & ((s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.realax.real',X1) != s('type.realax.real',X2)) | s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))),
  inference(flattening,[],[f726])).
fof(f726,plain,(
  ! [X0,X1,X2] : ((s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X2))) | (s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X1) = s('type.realax.real',X2))) & ((s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.realax.real',X1) != s('type.realax.real',X2)) | s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))),
  inference(nnf_transformation,[],[f520])).
fof(f520,plain,(
  ! [X0,X1,X2] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X2))) <=> (s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X1) = s('type.realax.real',X2)))),
  inference(rectify,[],[f185])).
fof(f185,axiom,(
  ! [X6,X15,X18] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real',X15))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real',X18))) <=> (s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X15) = s('type.realax.real',X18)))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.real.REAL_EQ_LMUL')).
fof(f5453,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) = s('type.realax.real',X0) | s('type.realax.real',sK10) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0)))) ) | $spl10),
  inference(cnf_transformation,[],[f5453_D])).
fof(f5453_D,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) = s('type.realax.real',X0) | s('type.realax.real',sK10) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0)))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f5433,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',sK10) = s('type.realax.real',X0) | s('type.realax.real',sK10) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0)))) ) | $spl2),
  inference(duplicate_literal_removal,[],[f5372])).
fof(f5372,plain,(
  ( ! [X0] : (s('type.realax.real',sK10) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real',sK10) = s('type.realax.real',X0) | s('type.realax.real',sK10) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))) | s('type.realax.real',sK10) = s('type.realax.real',X0)) ) | $spl2),
  inference(superposition,[],[f1394,f1378])).
fof(f5447,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl8),
  inference(cnf_transformation,[],[f5447_D])).
fof(f5447_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f5444,plain,(
  ( ! [X5] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.realax.real',sK10) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X5))),s('type.realax.real',X5)))) ) | $spl2),
  inference(subsumption_resolution,[],[f5357,f1396])).
fof(f5357,plain,(
  ( ! [X5] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.realax.real',sK10) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X5))),s('type.realax.real',X5))) | s('type.realax.real',sK10) = s('type.realax.real',X5)) ) | $spl2),
  inference(superposition,[],[f1378,f1378])).
fof(f1194,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))) | $spl0),
  inference(cnf_transformation,[],[f1194_D])).
fof(f1194_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f1196,plain,(
  s('type.realax.real',sK10) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl2),
  inference(cnf_transformation,[],[f1196_D])).
fof(f1196_D,plain,(
  s('type.realax.real',sK10) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f760,plain,(
  s('type.realax.real',sK10) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11)))),
  inference(cnf_transformation,[],[f649])).
fof(f649,plain,(
  (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))) | (s('type.realax.real',sK10) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.num.num','const.num.|48|_0') != s('type.num.num',sK11))) & (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))) | s('type.realax.real',sK10) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK11))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11])],[f648])).
fof(f648,plain,(
  ? [X0,X1] : ((s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num',X1))) | (s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.num.num','const.num.|48|_0') != s('type.num.num',X1))) & (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num',X1))) | s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)))),
  inference(flattening,[],[f647])).
fof(f647,plain,(
  ? [X0,X1] : ((s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num',X1))) | (s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.num.num','const.num.|48|_0') != s('type.num.num',X1))) & (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num',X1))) | (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1))))),
  inference(nnf_transformation,[],[f554])).
fof(f554,plain,(
  ? [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num',X1))) <~> (s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)))),
  inference(ennf_transformation,[],[f263])).
fof(f263,plain,(
  ~! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num',X1))) <=> (s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)))),
  inference(flattening,[],[f262])).
fof(f262,plain,(
  ~! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num',X1))) <=> (s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) & ~s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',X6),s('type.num.num',X7))) <=> (s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) & ~s('type.num.num','const.num.|48|_0') = s('type.num.num',X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',X6),s('type.num.num',X7))) <=> (s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) & ~s('type.num.num','const.num.|48|_0') = s('type.num.num',X7)))),
  file('thm.transc.transc_2|49|4_prove_ltb','thm.transc.transc_2|49|4_prove_ltb')).
fof(f1198,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK11) | $spl5),
  inference(cnf_transformation,[],[f1198_D])).
fof(f1198_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK11) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f761,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK11) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11)))),
  inference(cnf_transformation,[],[f649])).
fof(f1195,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))) | $spl1),
  inference(cnf_transformation,[],[f1195_D])).
fof(f1195_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f1197,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl3),
  inference(cnf_transformation,[],[f1197_D])).
fof(f1197_D,plain,(
  s('type.realax.real',sK10) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f1199,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK11) | $spl4),
  inference(cnf_transformation,[],[f1199_D])).
fof(f1199_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK11) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f762,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK11) | s('type.realax.real',sK10) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK10),s('type.num.num',sK11)))),
  inference(cnf_transformation,[],[f649])).
% SZS output end Proof for thm.transc
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 969
% Passive clauses: 5643
% Generated clauses: 42607
% Final active clauses: 529
% Final passive clauses: 2698
% Input formulas: 260
% Initial clauses: 423
% 
% Duplicate literals: 11
% Trivial inequalities: 256
% Fw subsumption resolutions: 517
% Fw demodulations: 9982
% Bw demodulations: 1487
% Condensations: 38
% 
% Simple tautologies: 89
% Equational tautologies: 99
% Forward subsumptions: 25443
% Fw demodulations to eq. taut.: 307
% Bw demodulations to eq. taut.: 5
% 
% Unit resulting resolution: 7200
% Factoring: 1
% Forward superposition: 12873
% Backward superposition: 8953
% Self superposition: 173
% Equality factoring: 2
% Equality resolution: 29
% 
% Split clauses: 28
% Split components: 57
% SAT solver clauses: 5895
% SAT solver unit clauses: 1239
% SAT solver binary clauses: 2232
% 
% Sat splits: 28
% Sat splitting refutations: 194
% 
% Memory used [KB]: 30703
% Time elapsed: 1.322 s
% ------------------------------
% ------------------------------
% Success in time 2.66 s
