Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 44
dis+11_7_41 on thm.integer
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.integer
% SZS output start Proof for thm.integer
fof(f116284,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f116281,f1884,f2641,f2639,f1696,f2640,f116017,f53911])).
fof(f53911,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|49|_0'),s('type.integer.int','const.integer.int_|48|_0')))) | $spl554),
  inference(cnf_transformation,[],[f53911_D])).
fof(f53911_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|49|_0'),s('type.integer.int','const.integer.int_|48|_0')))) <=> ~$spl554),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl554])])).
fof(f116017,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|49|_0'),s('type.integer.int','const.integer.int_|48|_0')))) | $spl73),
  inference(subsumption_resolution,[],[f116016,f1697])).
fof(f1697,plain,(
  s('type.integer.int','const.integer.int_|49|_0') != s('type.integer.int','const.integer.int_|48|_0')),
  inference(cnf_transformation,[],[f539])).
fof(f539,plain,(
  s('type.integer.int','const.integer.int_|49|_0') != s('type.integer.int','const.integer.int_|48|_0')),
  inference(flattening,[],[f229])).
fof(f229,axiom,(
  ~s('type.integer.int','const.integer.int_|49|_0') = s('type.integer.int','const.integer.int_|48|_0')),
  file('thm.integer.INT_LT_|48||49|','thm.integer.INT_|49||48|')).
fof(f116016,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|49|_0'),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int','const.integer.int_|49|_0') = s('type.integer.int','const.integer.int_|48|_0') | $spl73),
  inference(subsumption_resolution,[],[f115732,f2640])).
fof(f115732,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|49|_0'),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int','const.integer.int_|49|_0') = s('type.integer.int','const.integer.int_|48|_0')),
  inference(superposition,[],[f2091,f56197])).
fof(f56197,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|49|_0')))),
  inference(resolution,[],[f56196,f1838])).
fof(f1838,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f1366])).
fof(f1366,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f680])).
fof(f680,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f679])).
fof(f679,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.integer.INT_LT_|48||49|','thm.bool.EQ_CLAUSES|split|2')).
fof(f56196,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|49|_0'))))),
  inference(forward_demodulation,[],[f56195,f1649])).
fof(f1649,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),
  inference(cnf_transformation,[],[f67])).
fof(f67,axiom,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),
  file('thm.integer.INT_LT_|48||49|','thm.integer.INT_|48|')).
fof(f56195,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|49|_0'))))),
  inference(forward_demodulation,[],[f56194,f17581])).
fof(f17581,plain,(
  s('type.integer.int','const.integer.int_|49|_0') = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(backward_demodulation,[],[f10519,f10383])).
fof(f10383,plain,(
  s('type.integer.int','const.integer.int_|49|_0') = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(forward_demodulation,[],[f1656,f5686])).
fof(f5686,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.ZERO_0')),
  inference(superposition,[],[f1719,f1717])).
fof(f1717,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f559])).
fof(f559,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f317])).
fof(f317,axiom,(
  ! [X47] : s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X47)))),
  file('thm.integer.INT_LT_|48||49|','thm.numeral.numeral_mult|split|0')).
fof(f1719,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f561])).
fof(f561,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f513])).
fof(f513,axiom,(
  ! [X76] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X76),s('type.num.num','const.num.|48|_0')))),
  file('thm.integer.INT_LT_|48||49|','thm.arithmetic.MULT_|48|')).
fof(f1656,plain,(
  s('type.integer.int','const.integer.int_|49|_0') = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(cnf_transformation,[],[f13])).
fof(f13,axiom,(
  s('type.integer.int','const.integer.int_|49|_0') = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  file('thm.integer.INT_LT_|48||49|','thm.integer.INT_|49|')).
fof(f10519,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f10518,f5779])).
fof(f5779,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),
  inference(backward_demodulation,[],[f5686,f1654])).
fof(f1654,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(cnf_transformation,[],[f499])).
fof(f499,axiom,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  file('thm.integer.INT_LT_|48||49|','thm.numeral.numeral_suc|split|0')).
fof(f10518,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f1657,f5686])).
fof(f1657,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f154])).
fof(f154,axiom,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  file('thm.integer.INT_LT_|48||49|','thm.arithmetic.ONE')).
fof(f56194,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))))),
  inference(forward_demodulation,[],[f56193,f10519])).
fof(f56193,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))))))),
  inference(forward_demodulation,[],[f1642,f5686])).
fof(f1642,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(cnf_transformation,[],[f519])).
fof(f519,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(flattening,[],[f6])).
fof(f6,negated_conjecture,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  file('thm.integer.INT_LT_|48||49|','thm.integer.INT_LT_|48||49|')).
fof(f2091,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1)) )),
  inference(cnf_transformation,[],[f874])).
fof(f874,plain,(
  ! [X0,X1] : (s('type.integer.int',X0) = s('type.integer.int',X1) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f111])).
fof(f111,axiom,(
  ! [X7,X8] : (s('type.integer.int',X7) = s('type.integer.int',X8) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int',X8)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X8),s('type.integer.int',X7)))))),
  file('thm.integer.INT_LT_|48||49|','thm.integer.INT_LT_TOTAL')).
fof(f2640,plain,(
  ~p(s(bool,'F_0')) | $spl73),
  inference(cnf_transformation,[],[f2640_D])).
fof(f2640_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl73),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl73])])).
fof(f1696,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f538])).
fof(f538,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.integer.INT_LT_|48||49|',aHL_FALSITY)).
fof(f2639,plain,(
  p(s(bool,'F_0')) | $spl72),
  inference(cnf_transformation,[],[f2639_D])).
fof(f2639_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f2641,plain,(
  p(s(bool,'T_0')) | $spl74),
  inference(cnf_transformation,[],[f2641_D])).
fof(f2641_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f1884,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f1396])).
fof(f1396,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f1395])).
fof(f1395,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f1186])).
fof(f1186,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f700])).
fof(f700,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f98])).
fof(f98,axiom,(
  ! [X10] : ((p(s(bool,'F_0')) => p(s(bool,X10))) <=> p(s(bool,'T_0')))),
  file('thm.integer.INT_LT_|48||49|','thm.bool.IMP_CLAUSES|split|2')).
fof(f116281,plain,(
  $false | ($spl74 | $spl554)),
  inference(subsumption_resolution,[],[f116280,f2641])).
fof(f116280,plain,(
  ~p(s(bool,'T_0')) | $spl554),
  inference(forward_demodulation,[],[f116206,f56904])).
fof(f56904,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|49|_0')))),
  inference(resolution,[],[f56903,f1834])).
fof(f1834,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f1364])).
fof(f1364,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f677])).
fof(f677,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.integer.INT_LT_|48||49|','thm.bool.EQ_CLAUSES|split|1')).
fof(f56903,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|49|_0'))))),
  inference(forward_demodulation,[],[f56902,f1649])).
fof(f56902,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|49|_0'))))),
  inference(forward_demodulation,[],[f56901,f17581])).
fof(f56901,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))))),
  inference(forward_demodulation,[],[f56900,f10519])).
fof(f56900,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))))))),
  inference(forward_demodulation,[],[f1661,f5686])).
fof(f1661,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(cnf_transformation,[],[f14])).
fof(f14,axiom,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  file('thm.integer.INT_LT_|48||49|','thm.integer.INT_LE_|48||49|')).
fof(f116206,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|49|_0')))) | $spl554),
  inference(resolution,[],[f53911,f2138])).
fof(f2138,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f1477])).
fof(f1477,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) & (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f915])).
fof(f915,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(flattening,[],[f914])).
fof(f914,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X7,X8] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X7),s('type.integer.int',X8)))) <=> ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X8),s('type.integer.int',X7)))))),
  file('thm.integer.INT_LT_|48||49|','thm.integer.int_le')).
% SZS output end Proof for thm.integer
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 2371
% Passive clauses: 21483
% Generated clauses: 114600
% Final active clauses: 2109
% Final passive clauses: 17936
% Input formulas: 517
% Initial clauses: 914
% 
% Duplicate literals: 474
% Trivial inequalities: 227
% Fw subsumption resolutions: 4537
% Fw demodulations: 15011
% Bw demodulations: 253
% 
% Simple tautologies: 387
% Equational tautologies: 43
% Forward subsumptions: 72274
% Fw demodulations to eq. taut.: 14
% 
% Binary resolution: 17275
% Factoring: 2777
% Forward superposition: 37949
% Backward superposition: 34531
% Self superposition: 117
% Equality factoring: 2
% Equality resolution: 42
% 
% Split clauses: 748
% Split components: 1498
% SAT solver clauses: 1173
% SAT solver unit clauses: 35
% SAT solver binary clauses: 154
% SAT solver learnt clauses: 1
% 
% Sat splits: 748
% Sat splitting refutations: 14
% 
% TWLsolver clauses: 1168
% TWLsolver calls for satisfiability: 818
% 
% Memory used [KB]: 59743
% Time elapsed: 1.945 s
% ------------------------------
% ------------------------------
% Success in time 1.99 s
