Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.combin
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.combin
% SZS output start Proof for thm.combin
fof(f3627,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f3626,f52,f78,f76,f3207,f51,f77,f75])).
fof(f75,plain,(
  s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK4),s(fun(sK0,sK1),sK2))) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) | $spl0),
  inference(cnf_transformation,[],[f75_D])).
fof(f75_D,plain,(
  s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK4),s(fun(sK0,sK1),sK2))) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f77,plain,(
  s(sK1,sK4) = s(sK1,sK5) | $spl2),
  inference(cnf_transformation,[],[f77_D])).
fof(f77_D,plain,(
  s(sK1,sK4) = s(sK1,sK5) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f51,plain,(
  s(sK1,sK4) = s(sK1,sK5) | s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK4),s(fun(sK0,sK1),sK2))) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2)))),
  inference(cnf_transformation,[],[f40])).
fof(f40,plain,(
  (s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK4),s(fun(sK0,sK1),sK2))) = s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) | s(sK1,sK4) = s(sK1,sK5)) & (s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK4),s(fun(sK0,sK1),sK2))) != s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) | s(sK1,sK4) != s(sK1,sK5))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5])],[f39])).
fof(f39,plain,(
  ? [X0,X1,X2,X3,X4,X5] : ((s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X4),s(fun(X0,X1),X2))) = s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X5),s(fun(X0,X1),X2))) | s(X1,X4) = s(X1,X5)) & (s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X4),s(fun(X0,X1),X2))) != s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X5),s(fun(X0,X1),X2))) | s(X1,X4) != s(X1,X5)))),
  inference(nnf_transformation,[],[f35])).
fof(f35,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X4),s(fun(X0,X1),X2))) = s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X5),s(fun(X0,X1),X2))) <~> s(X1,X4) = s(X1,X5))),
  inference(ennf_transformation,[],[f17])).
fof(f17,plain,(
  ~! [X0,X1,X2,X3,X4,X5] : (s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X4),s(fun(X0,X1),X2))) = s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X5),s(fun(X0,X1),X2))) <=> s(X1,X4) = s(X1,X5))),
  inference(flattening,[],[f16])).
fof(f16,plain,(
  ~! [X0,X1] : ! [X2,X3,X4,X5] : (s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X4),s(fun(X0,X1),X2))) = s(fun(X0,X1),'const.combin.UPDATE_3'(s(X0,X3),s(X1,X5),s(fun(X0,X1),X2))) <=> s(X1,X4) = s(X1,X5))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10,X11] : (s(fun(X6,X7),'const.combin.UPDATE_3'(s(X6,X9),s(X7,X10),s(fun(X6,X7),X8))) = s(fun(X6,X7),'const.combin.UPDATE_3'(s(X6,X9),s(X7,X11),s(fun(X6,X7),X8))) <=> s(X7,X10) = s(X7,X11))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10,X11] : (s(fun(X6,X7),'const.combin.UPDATE_3'(s(X6,X9),s(X7,X10),s(fun(X6,X7),X8))) = s(fun(X6,X7),'const.combin.UPDATE_3'(s(X6,X9),s(X7,X11),s(fun(X6,X7),X8))) <=> s(X7,X10) = s(X7,X11))),
  file('thm.combin.UPD|49||49|_SAME_KEY_AND_BASE','thm.combin.UPD|49||49|_SAME_KEY_AND_BASE')).
fof(f3207,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f3206,f78])).
fof(f3206,plain,(
  s(sK1,sK4) = s(sK1,sK5) | $spl0),
  inference(forward_demodulation,[],[f3187,f2721])).
fof(f2721,plain,(
  ( ! [X6,X10,X8,X7,X9] : (s(X10,X8) = s(X10,'const.combin.UPDATE_4'(s(X9,X6),s(X10,X8),s(fun(X9,X10),X7),s(X9,X6)))) )),
  inference(forward_demodulation,[],[f2588,f61])).
fof(f61,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f23])).
fof(f23,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f22])).
fof(f22,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : ! [X12,X13] : s(X6,X12) = s(X6,'const.bool.COND_3'(s(bool,'T_0'),s(X6,X12),s(X6,X13)))),
  file('thm.combin.UPD|49||49|_SAME_KEY_AND_BASE','thm.bool.COND_CLAUSES|split|0')).
fof(f2588,plain,(
  ( ! [X6,X10,X8,X7,X9] : (s(X10,'const.combin.UPDATE_4'(s(X9,X6),s(X10,X8),s(fun(X9,X10),X7),s(X9,X6))) = s(X10,'const.bool.COND_3'(s(bool,'T_0'),s(X10,X8),s(X10,happ(s(fun(X9,X10),X7),s(X9,X6)))))) )),
  inference(superposition,[],[f73,f214])).
fof(f214,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,sK8(X0,X1,X2,X0,X3,X4))) )),
  inference(equality_resolution,[],[f181])).
fof(f181,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X4,X0) != s(X4,X3) | s(bool,'T_0') = s(bool,sK8(X0,X1,X2,X3,X4,X5))) )),
  inference(subsumption_resolution,[],[f180,f55])).
fof(f55,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f18])).
fof(f18,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.combin.UPD|49||49|_SAME_KEY_AND_BASE',aHL_FALSITY)).
fof(f180,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'F_0')) | s(X4,X0) != s(X4,X3) | s(bool,'T_0') = s(bool,sK8(X0,X1,X2,X3,X4,X5))) )),
  inference(superposition,[],[f72,f56])).
fof(f56,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.combin.UPD|49||49|_SAME_KEY_AND_BASE',aHL_BOOL_CASES)).
fof(f72,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,sK8(X5,X4,X3,X2,X1,X0))) | s(X1,X2) != s(X1,X5)) )),
  inference(cnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,sK8(X5,X4,X3,X2,X1,X0))) | s(X1,X2) = s(X1,X5)) & (s(X1,X2) != s(X1,X5) | p(s(bool,sK8(X5,X4,X3,X2,X1,X0)))) & s(X0,'const.combin.UPDATE_4'(s(X1,X2),s(X0,X3),s(fun(X1,X0),X4),s(X1,X5))) = s(X0,'const.bool.COND_3'(s(bool,sK8(X5,X4,X3,X2,X1,X0)),s(X0,X3),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8])],[f49])).
fof(f49,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ? [X6] : ((~p(s(bool,X6)) | s(X1,X2) = s(X1,X5)) & (s(X1,X2) != s(X1,X5) | p(s(bool,X6))) & s(X0,'const.combin.UPDATE_4'(s(X1,X2),s(X0,X3),s(fun(X1,X0),X4),s(X1,X5))) = s(X0,'const.bool.COND_3'(s(bool,X6),s(X0,X3),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ? [X6] : (((~p(s(bool,X6)) | s(X1,X2) = s(X1,X5)) & (s(X1,X2) != s(X1,X5) | p(s(bool,X6)))) & s(X0,'const.combin.UPDATE_4'(s(X1,X2),s(X0,X3),s(fun(X1,X0),X4),s(X1,X5))) = s(X0,'const.bool.COND_3'(s(bool,X6),s(X0,X3),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))))),
  inference(nnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ? [X6] : ((p(s(bool,X6)) <=> s(X1,X2) = s(X1,X5)) & s(X0,'const.combin.UPDATE_4'(s(X1,X2),s(X0,X3),s(fun(X1,X0),X4),s(X1,X5))) = s(X0,'const.bool.COND_3'(s(bool,X6),s(X0,X3),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : ? [X6] : ((p(s(bool,X6)) <=> s(X1,X2) = s(X1,X5)) & s(X0,'const.combin.UPDATE_4'(s(X1,X2),s(X0,X3),s(fun(X1,X0),X4),s(X1,X5))) = s(X0,'const.bool.COND_3'(s(bool,X6),s(X0,X3),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X7,X6] : ! [X9,X10,X5,X17] : ? [X18] : ((p(s(bool,X18)) <=> s(X6,X9) = s(X6,X17)) & s(X7,'const.combin.UPDATE_4'(s(X6,X9),s(X7,X10),s(fun(X6,X7),X5),s(X6,X17))) = s(X7,'const.bool.COND_3'(s(bool,X18),s(X7,X10),s(X7,happ(s(fun(X6,X7),X5),s(X6,X17))))))),
  file('thm.combin.UPD|49||49|_SAME_KEY_AND_BASE','thm.combin.UPDATE_def')).
fof(f73,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.combin.UPDATE_4'(s(X1,X2),s(X0,X3),s(fun(X1,X0),X4),s(X1,X5))) = s(X0,'const.bool.COND_3'(s(bool,sK8(X5,X4,X3,X2,X1,X0)),s(X0,X3),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5)))))) )),
  inference(cnf_transformation,[],[f50])).
fof(f3187,plain,(
  s(sK1,sK5) = s(sK1,'const.combin.UPDATE_4'(s(sK0,sK3),s(sK1,sK4),s(fun(sK0,sK1),sK2),s(sK0,sK3))) | $spl0),
  inference(superposition,[],[f1564,f2721])).
fof(f1564,plain,(
  ( ! [X0] : (s(sK1,'const.combin.UPDATE_4'(s(sK0,sK3),s(sK1,sK4),s(fun(sK0,sK1),sK2),s(sK0,X0))) = s(sK1,'const.combin.UPDATE_4'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2),s(sK0,X0)))) ) | $spl0),
  inference(forward_demodulation,[],[f1494,f70])).
fof(f70,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.combin.UPDATE_4'(s(X1,X2),s(X0,X3),s(fun(X1,X0),X4),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),'const.combin.UPDATE_3'(s(X1,X2),s(X0,X3),s(fun(X1,X0),X4))),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,'const.combin.UPDATE_4'(s(X1,X2),s(X0,X3),s(fun(X1,X0),X4),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),'const.combin.UPDATE_3'(s(X1,X2),s(X0,X3),s(fun(X1,X0),X4))),s(X1,X5)))),
  inference(flattening,[],[f31])).
fof(f31,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : s(X0,'const.combin.UPDATE_4'(s(X1,X2),s(X0,X3),s(fun(X1,X0),X4),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),'const.combin.UPDATE_3'(s(X1,X2),s(X0,X3),s(fun(X1,X0),X4))),s(X1,X5)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X7,X6] : ! [X19,X20,X21,X22] : s(X7,'const.combin.UPDATE_4'(s(X6,X19),s(X7,X20),s(fun(X6,X7),X21),s(X6,X22))) = s(X7,happ(s(fun(X6,X7),'const.combin.UPDATE_3'(s(X6,X19),s(X7,X20),s(fun(X6,X7),X21))),s(X6,X22)))),
  file('thm.combin.UPD|49||49|_SAME_KEY_AND_BASE','ARITY_3_4_const.combin.UPDATE')).
fof(f1494,plain,(
  ( ! [X0] : (s(sK1,'const.combin.UPDATE_4'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2),s(sK0,X0))) = s(sK1,happ(s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK4),s(fun(sK0,sK1),sK2))),s(sK0,X0)))) ) | $spl0),
  inference(superposition,[],[f70,f75])).
fof(f76,plain,(
  s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK4),s(fun(sK0,sK1),sK2))) != s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) | $spl1),
  inference(cnf_transformation,[],[f76_D])).
fof(f76_D,plain,(
  s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK4),s(fun(sK0,sK1),sK2))) != s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f78,plain,(
  s(sK1,sK4) != s(sK1,sK5) | $spl3),
  inference(cnf_transformation,[],[f78_D])).
fof(f78_D,plain,(
  s(sK1,sK4) != s(sK1,sK5) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f52,plain,(
  s(sK1,sK4) != s(sK1,sK5) | s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK4),s(fun(sK0,sK1),sK2))) != s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK5),s(fun(sK0,sK1),sK2)))),
  inference(cnf_transformation,[],[f40])).
fof(f3626,plain,(
  $false | ($spl1 | $spl2)),
  inference(trivial_inequality_removal,[],[f3625])).
fof(f3625,plain,(
  s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK4),s(fun(sK0,sK1),sK2))) != s(fun(sK0,sK1),'const.combin.UPDATE_3'(s(sK0,sK3),s(sK1,sK4),s(fun(sK0,sK1),sK2))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f76,f77])).
% SZS output end Proof for thm.combin
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 59
% Passive clauses: 894
% Generated clauses: 3482
% Final active clauses: 41
% Final passive clauses: 769
% Input formulas: 14
% Initial clauses: 23
% 
% Duplicate literals: 211
% Trivial inequalities: 24
% Fw subsumption resolutions: 228
% Fw demodulations: 215
% Bw demodulations: 12
% 
% Simple tautologies: 39
% Equational tautologies: 35
% Forward subsumptions: 1818
% 
% Binary resolution: 81
% Factoring: 1
% Forward superposition: 1314
% Backward superposition: 1193
% Self superposition: 91
% Equality factoring: 56
% Equality resolution: 6
% 
% Split clauses: 19
% Split components: 38
% SAT solver clauses: 59
% SAT solver unit clauses: 17
% SAT solver binary clauses: 24
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 1
% 
% Sat splits: 19
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 57
% TWLsolver calls for satisfiability: 26
% 
% Memory used [KB]: 2942
% Time elapsed: 0.104 s
% ------------------------------
% ------------------------------
% Success in time 0.109 s
