Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.pred_set
% (104775)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 131
% Passive clauses: 1209
% Generated clauses: 26486
% Final active clauses: 126
% Final passive clauses: 991
% Input formulas: 17
% Initial clauses: 34
% 
% Duplicate literals: 2248
% Trivial inequalities: 28
% Fw subsumption resolutions: 3792
% Fw demodulations: 4586
% Bw demodulations: 2
% 
% Simple tautologies: 244
% Equational tautologies: 22
% Forward subsumptions: 12182
% 
% Binary resolution: 156
% Forward superposition: 9543
% Backward superposition: 5828
% Self superposition: 139
% Equality factoring: 34
% Equality resolution: 7
% 
% Split clauses: 1325
% Split components: 2655
% SAT solver clauses: 2251
% SAT solver unit clauses: 7
% SAT solver binary clauses: 38
% 
% Sat splits: 1325
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 2251
% TWLsolver calls for satisfiability: 121
% 
% Memory used [KB]: 12025
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (104797)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 98
% Passive clauses: 806
% Generated clauses: 13881
% Final active clauses: 89
% Final passive clauses: 353
% Input formulas: 17
% Initial clauses: 34
% 
% Duplicate literals: 1241
% Trivial inequalities: 100
% Fw subsumption resolutions: 953
% Fw demodulations: 918
% Bw demodulations: 298
% Condensations: 1517
% Global subsumptions: 4
% 
% Simple tautologies: 230
% Equational tautologies: 45
% Forward subsumptions: 5404
% Bw demodulations to eq. taut.: 65
% 
% Binary resolution: 168
% Forward superposition: 3476
% Backward superposition: 4838
% Self superposition: 193
% Equality factoring: 141
% Equality resolution: 12
% 
% Split clauses: 145
% Split components: 295
% SAT solver clauses: 1283
% SAT solver unit clauses: 38
% SAT solver binary clauses: 30
% SAT solver learnt clauses: 8
% 
% Sat splits: 145
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1267
% TWLsolver calls for satisfiability: 188
% 
% Memory used [KB]: 10106
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% (104824)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 681
% Passive clauses: 3732
% Generated clauses: 206566
% Final active clauses: 155
% Final passive clauses: 542
% Input formulas: 17
% Initial clauses: 34
% 
% Duplicate literals: 8511
% Trivial inequalities: 1319
% Fw subsumption resolutions: 20600
% Fw demodulations: 27598
% Bw demodulations: 2405
% Condensations: 19007
% Global subsumptions: 55
% 
% Simple tautologies: 2293
% Equational tautologies: 838
% Forward subsumptions: 121056
% Fw demodulations to eq. taut.: 99
% Bw demodulations to eq. taut.: 45
% 
% Unit resulting resolution: 3963
% Forward superposition: 71308
% Backward superposition: 34146
% Self superposition: 727
% Equality factoring: 331
% Equality resolution: 245
% 
% Split clauses: 294
% Split components: 594
% SAT solver clauses: 4594
% SAT solver unit clauses: 64
% SAT solver binary clauses: 283
% 
% Sat splits: 294
% Sat splitting refutations: 217
% 
% Memory used [KB]: 78420
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f39428,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f161,f164,f162,f1321,f1326,f1324,f1322,f1330,f1328,f1377,f1398,f1403,f1401,f1399,f1407,f1405,f1587,f16446,f16731,f16447,f16732,f16730,f16728,f16726,f16470,f16739,f16471,f16740,f16472,f16741,f16742,f16738,f16744,f16736,f16734,f16494,f16751,f16495,f16752,f16496,f16753,f16754,f16750,f16756,f16748,f16746,f18401,f18406,f18404,f18402,f18469,f18914,f18919,f18917,f18915,f18923,f18921,f19456,f33836,f33803,f33842,f33840,f33719,f33846,f33718,f33850,f33717,f33852,f33844,f33716,f33854,f33848,f33755,f33856,f33713,f33858,f33712,f33862,f33711,f33860,f33866,f33710,f33864,f33868,f33709,f33870,f33708,f33874,f33707,f33872,f33878,f33706,f33876,f33880,f33705,f33884,f33882,f33704,f33888,f33886,f34189,f34326,f18410,f18408,f36411,f34508,f36415,f36413,f36564,f39427])).
fof(f39427,plain,(
  $false | ($spl62 | $spl127)),
  inference(subsumption_resolution,[],[f39426,f100])).
fof(f100,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.INTER_UNION_COMPL','thm.bool.TRUTH')).
fof(f39426,plain,(
  ~p(s(bool,'T_0')) | ($spl62 | $spl127)),
  inference(forward_demodulation,[],[f34970,f38088])).
fof(f38088,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),X5),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))))) ) | $spl127),
  inference(subsumption_resolution,[],[f38086,f94])).
fof(f94,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f40])).
fof(f40,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pred_set.INTER_UNION_COMPL',aHL_FALSITY)).
fof(f38086,plain,(
  ( ! [X5] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),X5),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))))) ) | $spl127),
  inference(superposition,[],[f36806,f73])).
fof(f73,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.INTER_UNION_COMPL',aHL_BOOL_CASES)).
fof(f36806,plain,(
  ( ! [X3] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),X3),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))))) ) | $spl127),
  inference(resolution,[],[f36598,f76])).
fof(f76,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(nnf_transformation,[],[f28])).
fof(f28,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f27])).
fof(f27,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6] : ! [X7,X8,X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X7)))) | p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X8))))))),
  file('thm.pred_set.INTER_UNION_COMPL','thm.pred_set.IN_UNION')).
fof(f36598,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))) | $spl127),
  inference(resolution,[],[f36415,f78])).
fof(f78,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))))) & (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X2))))))))),
  inference(nnf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X2)))))) <=> ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))))),
  inference(flattening,[],[f29])).
fof(f29,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X2)))))) <=> ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6] : ! [X9,X7] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),'const.pred_set.COMPL_1'(s(fun(X6,bool),X7)))))) <=> ~p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X7)))))),
  file('thm.pred_set.INTER_UNION_COMPL','thm.pred_set.IN_COMPL')).
fof(f34970,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))))) | $spl62),
  inference(subsumption_resolution,[],[f34935,f101])).
fof(f101,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f71])).
fof(f71,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f26])).
fof(f26,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f25])).
fof(f25,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : ! [X9] : (s(X6,X9) = s(X6,X9) <=> p(s(bool,'T_0')))),
  file('thm.pred_set.INTER_UNION_COMPL','thm.bool.REFL_CLAUSE')).
fof(f34935,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))))) | $spl62),
  inference(superposition,[],[f77,f18408])).
fof(f77,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X2))))))) )),
  inference(cnf_transformation,[],[f55])).
fof(f36564,plain,(
  $false | ($spl75 | $spl125)),
  inference(subsumption_resolution,[],[f36563,f94])).
fof(f36563,plain,(
  p(s(bool,'F_0')) | ($spl75 | $spl125)),
  inference(forward_demodulation,[],[f36533,f36447])).
fof(f36447,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl125),
  inference(subsumption_resolution,[],[f36445,f101])).
fof(f36445,plain,(
  ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl125),
  inference(superposition,[],[f36413,f73])).
fof(f36533,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl75),
  inference(resolution,[],[f33839,f78])).
fof(f33839,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) | $spl75),
  inference(cnf_transformation,[],[f33839_D])).
fof(f33839_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) <=> ~$spl75),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl75])])).
fof(f36413,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl125),
  inference(cnf_transformation,[],[f36413_D])).
fof(f36413_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) <=> ~$spl125),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl125])])).
fof(f36415,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | $spl127),
  inference(cnf_transformation,[],[f36415_D])).
fof(f36415_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) <=> ~$spl127),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl127])])).
fof(f34508,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl64),
  inference(subsumption_resolution,[],[f34475,f94])).
fof(f34475,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl64),
  inference(superposition,[],[f81,f18409])).
fof(f18409,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl64),
  inference(cnf_transformation,[],[f18409_D])).
fof(f18409_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f81,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(nnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f31])).
fof(f31,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X6] : ! [X7,X8,X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X7)))) & p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X8))))))),
  file('thm.pred_set.INTER_UNION_COMPL','thm.pred_set.IN_INTER')).
fof(f36411,plain,(
  $false | ($spl62 | $spl74)),
  inference(subsumption_resolution,[],[f36410,f100])).
fof(f36410,plain,(
  ~p(s(bool,'T_0')) | ($spl62 | $spl74)),
  inference(forward_demodulation,[],[f34970,f35404])).
fof(f35404,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),X5)))))) ) | $spl74),
  inference(subsumption_resolution,[],[f35402,f94])).
fof(f35402,plain,(
  ( ! [X5] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),X5)))))) ) | $spl74),
  inference(superposition,[],[f34015,f73])).
fof(f34015,plain,(
  ( ! [X2] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),X2))))))) ) | $spl74),
  inference(resolution,[],[f33840,f75])).
fof(f75,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f54])).
fof(f18408,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))))) | $spl62),
  inference(cnf_transformation,[],[f18408_D])).
fof(f18408_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f18410,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl65),
  inference(cnf_transformation,[],[f18410_D])).
fof(f18410_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl65),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl65])])).
fof(f34326,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))))))),
  inference(superposition,[],[f1229,f73])).
fof(f1229,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))))))),
  inference(equality_resolution,[],[f124])).
fof(f124,plain,(
  ( ! [X2] : (s(fun(sK0,bool),X2) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),X2,sK0)),s(fun(sK0,bool),X2))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),X2,sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))))))) )),
  inference(superposition,[],[f67,f70])).
fof(f70,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X2)))) )),
  inference(cnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f50])).
fof(f50,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f49])).
fof(f49,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f24])).
fof(f24,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f23])).
fof(f23,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6] : ! [X7,X8] : (s(fun(X6,bool),X7) = s(fun(X6,bool),X8) <=> ! [X9] : s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X7))) = s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X8))))),
  file('thm.pred_set.INTER_UNION_COMPL','thm.pred_set.EXTENSION')).
fof(f67,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))))),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f43])).
fof(f43,plain,(
  ? [X0,X1,X2] : s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) != s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X1))),s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X2)))))))),
  inference(ennf_transformation,[],[f20])).
fof(f20,plain,(
  ~! [X0,X1,X2] : s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) = s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X1))),s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X2)))))))),
  inference(flattening,[],[f19])).
fof(f19,plain,(
  ~! [X0] : ! [X1,X2] : s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) = s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X1))),s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X2)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8))) = s(fun(X6,bool),'const.pred_set.COMPL_1'(s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),'const.pred_set.COMPL_1'(s(fun(X6,bool),X7))),s(fun(X6,bool),'const.pred_set.COMPL_1'(s(fun(X6,bool),X8)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8))) = s(fun(X6,bool),'const.pred_set.COMPL_1'(s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),'const.pred_set.COMPL_1'(s(fun(X6,bool),X7))),s(fun(X6,bool),'const.pred_set.COMPL_1'(s(fun(X6,bool),X8)))))))),
  file('thm.pred_set.INTER_UNION_COMPL','thm.pred_set.INTER_UNION_COMPL')).
fof(f34189,plain,(
  $false | ($spl58 | $spl61)),
  inference(subsumption_resolution,[],[f34188,f33889])).
fof(f33889,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl58),
  inference(forward_demodulation,[],[f1229,f18404])).
fof(f34188,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl61),
  inference(trivial_inequality_removal,[],[f34186])).
fof(f34186,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl61),
  inference(superposition,[],[f18406,f73])).
fof(f33886,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK1))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))))) | $spl121),
  inference(cnf_transformation,[],[f33886_D])).
fof(f33886_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK1))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))))) <=> ~$spl121),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl121])])).
fof(f33888,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) | $spl123),
  inference(cnf_transformation,[],[f33888_D])).
fof(f33888_D,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) <=> ~$spl123),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl123])])).
fof(f33704,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) | s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),sK1))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))))) | $spl58),
  inference(superposition,[],[f217,f18404])).
fof(f217,plain,(
  ( ! [X6,X7] : (s(bool,'const.bool.IN_2'(s(sK0,sK4(X7,X6,sK0)),s(fun(sK0,bool),X6))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X7,X6,sK0)),s(fun(sK0,bool),X7))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X6),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),X6))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) | s(bool,'const.bool.IN_2'(s(sK0,sK4(X6,sK1,sK0)),s(fun(sK0,bool),sK1))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X6,sK1,sK0)),s(fun(sK0,bool),X7)))) )),
  inference(superposition,[],[f105,f70])).
fof(f105,plain,(
  ( ! [X3] : (s(bool,'const.bool.IN_2'(s(sK0,sK4(X3,sK1,sK0)),s(fun(sK0,bool),sK1))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X3,sK1,sK0)),s(fun(sK0,bool),X3))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X3),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),X3))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))))) )),
  inference(superposition,[],[f67,f70])).
fof(f33882,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))))) | $spl117),
  inference(cnf_transformation,[],[f33882_D])).
fof(f33882_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))))) <=> ~$spl117),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl117])])).
fof(f33884,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))) | $spl119),
  inference(cnf_transformation,[],[f33884_D])).
fof(f33884_D,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))) <=> ~$spl119),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl119])])).
fof(f33705,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))))) | s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))))) | $spl58),
  inference(superposition,[],[f443,f18404])).
fof(f443,plain,(
  ( ! [X6,X7] : (s(bool,'const.bool.IN_2'(s(sK0,sK4(X7,X6,sK0)),s(fun(sK0,bool),X6))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X7,X6,sK0)),s(fun(sK0,bool),X7))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X6))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),X6))))))) | s(bool,'const.bool.IN_2'(s(sK0,sK4(X6,sK2,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X6,sK2,sK0)),s(fun(sK0,bool),X7)))) )),
  inference(superposition,[],[f113,f70])).
fof(f113,plain,(
  ( ! [X3] : (s(bool,'const.bool.IN_2'(s(sK0,sK4(X3,sK2,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X3,sK2,sK0)),s(fun(sK0,bool),X3))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X3))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),X3)))))))) )),
  inference(superposition,[],[f67,f70])).
fof(f33880,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl115),
  inference(cnf_transformation,[],[f33880_D])).
fof(f33880_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl115),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl115])])).
fof(f33876,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) | $spl111),
  inference(cnf_transformation,[],[f33876_D])).
fof(f33876_D,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) <=> ~$spl111),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl111])])).
fof(f33706,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) | s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl58),
  inference(superposition,[],[f458,f18404])).
fof(f458,plain,(
  ( ! [X4,X5] : (s(bool,'const.bool.IN_2'(s(sK0,sK4(X4,X5,sK0)),s(fun(sK0,bool),X4))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X4,X5,sK0)),s(fun(sK0,bool),X5))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),X4),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) | s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X4,sK0)),s(fun(sK0,bool),X5))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X4,sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) )),
  inference(superposition,[],[f108,f70])).
fof(f108,plain,(
  ( ! [X2] : (s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X2,sK0)),s(fun(sK0,bool),X2))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X2,sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),X2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))))) )),
  inference(superposition,[],[f67,f70])).
fof(f33878,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))))) | $spl113),
  inference(cnf_transformation,[],[f33878_D])).
fof(f33878_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))))) <=> ~$spl113),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl113])])).
fof(f33872,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) | $spl107),
  inference(cnf_transformation,[],[f33872_D])).
fof(f33872_D,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) <=> ~$spl107),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl107])])).
fof(f33707,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) | s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))))) | $spl58),
  inference(superposition,[],[f459,f18404])).
fof(f459,plain,(
  ( ! [X6,X7] : (s(bool,'const.bool.IN_2'(s(sK0,sK4(X7,X6,sK0)),s(fun(sK0,bool),X6))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X7,X6,sK0)),s(fun(sK0,bool),X7))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),X6),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) | s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X6,sK0)),s(fun(sK0,bool),X7))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X6,sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) )),
  inference(superposition,[],[f108,f70])).
fof(f33874,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl109),
  inference(cnf_transformation,[],[f33874_D])).
fof(f33874_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl109),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl109])])).
fof(f33708,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) | s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl58),
  inference(superposition,[],[f514,f18404])).
fof(f514,plain,(
  ( ! [X4,X5] : (s(bool,'const.bool.IN_2'(s(sK0,sK4(X4,X5,sK0)),s(fun(sK0,bool),X4))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X4,X5,sK0)),s(fun(sK0,bool),X5))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),X4),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) | s(bool,'const.bool.IN_2'(s(sK0,sK4(X4,'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),X5))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X4,'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) )),
  inference(superposition,[],[f109,f70])).
fof(f109,plain,(
  ( ! [X3] : (s(bool,'const.bool.IN_2'(s(sK0,sK4(X3,'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),X3))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X3,'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),X3),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))))) )),
  inference(superposition,[],[f67,f70])).
fof(f33870,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))))) | $spl105),
  inference(cnf_transformation,[],[f33870_D])).
fof(f33870_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))))) <=> ~$spl105),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl105])])).
fof(f33709,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) | s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))))) | $spl58),
  inference(superposition,[],[f515,f18404])).
fof(f515,plain,(
  ( ! [X6,X7] : (s(bool,'const.bool.IN_2'(s(sK0,sK4(X7,X6,sK0)),s(fun(sK0,bool),X6))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X7,X6,sK0)),s(fun(sK0,bool),X7))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),X6),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) | s(bool,'const.bool.IN_2'(s(sK0,sK4(X6,'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),X7))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X6,'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) )),
  inference(superposition,[],[f109,f70])).
fof(f33868,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl103),
  inference(cnf_transformation,[],[f33868_D])).
fof(f33868_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl103),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl103])])).
fof(f33864,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))))))) | $spl99),
  inference(cnf_transformation,[],[f33864_D])).
fof(f33864_D,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))))))) <=> ~$spl99),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl99])])).
fof(f33710,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))))))) | s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl58),
  inference(superposition,[],[f609,f18404])).
fof(f609,plain,(
  ( ! [X4,X5] : (s(bool,'const.bool.IN_2'(s(sK0,sK4(X4,X5,sK0)),s(fun(sK0,bool),X4))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X4,X5,sK0)),s(fun(sK0,bool),X5))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),X4))))) | s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),X4,sK0)),s(fun(sK0,bool),X5))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),X4,sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))) )),
  inference(superposition,[],[f116,f70])).
fof(f116,plain,(
  ( ! [X2] : (s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),X2,sK0)),s(fun(sK0,bool),X2))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),X2,sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),X2)))))) )),
  inference(superposition,[],[f67,f70])).
fof(f33866,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))))) | $spl101),
  inference(cnf_transformation,[],[f33866_D])).
fof(f33866_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))))) <=> ~$spl101),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl101])])).
fof(f33860,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))) | $spl95),
  inference(cnf_transformation,[],[f33860_D])).
fof(f33860_D,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))) <=> ~$spl95),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl95])])).
fof(f33711,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))) | s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))))) | $spl58),
  inference(superposition,[],[f610,f18404])).
fof(f610,plain,(
  ( ! [X6,X7] : (s(bool,'const.bool.IN_2'(s(sK0,sK4(X7,X6,sK0)),s(fun(sK0,bool),X6))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X7,X6,sK0)),s(fun(sK0,bool),X7))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),X6))))) | s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),X6,sK0)),s(fun(sK0,bool),X7))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),X6,sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))) )),
  inference(superposition,[],[f116,f70])).
fof(f33862,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl97),
  inference(cnf_transformation,[],[f33862_D])).
fof(f33862_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl97),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl97])])).
fof(f33712,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))))))) | s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl58),
  inference(superposition,[],[f730,f18404])).
fof(f730,plain,(
  ( ! [X4,X5] : (s(bool,'const.bool.IN_2'(s(sK0,sK4(X4,X5,sK0)),s(fun(sK0,bool),X4))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X4,X5,sK0)),s(fun(sK0,bool),X5))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),X4))))) | s(bool,'const.bool.IN_2'(s(sK0,sK4(X4,'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),X5))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X4,'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))) )),
  inference(superposition,[],[f117,f70])).
fof(f117,plain,(
  ( ! [X3] : (s(bool,'const.bool.IN_2'(s(sK0,sK4(X3,'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),X3))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X3,'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),X3)))))) )),
  inference(superposition,[],[f67,f70])).
fof(f33858,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))))) | $spl93),
  inference(cnf_transformation,[],[f33858_D])).
fof(f33858_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))))) <=> ~$spl93),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl93])])).
fof(f33713,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))))) | s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))))) | $spl58),
  inference(superposition,[],[f731,f18404])).
fof(f731,plain,(
  ( ! [X6,X7] : (s(bool,'const.bool.IN_2'(s(sK0,sK4(X7,X6,sK0)),s(fun(sK0,bool),X6))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X7,X6,sK0)),s(fun(sK0,bool),X7))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),X6))))) | s(bool,'const.bool.IN_2'(s(sK0,sK4(X6,'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),X7))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X6,'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))) )),
  inference(superposition,[],[f117,f70])).
fof(f33856,plain,(
  s(bool,f(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),bool,sK0)))) != s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),bool,sK0)))) | $spl91),
  inference(cnf_transformation,[],[f33856_D])).
fof(f33856_D,plain,(
  s(bool,f(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),bool,sK0)))) != s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),bool,sK0)))) <=> ~$spl91),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl91])])).
fof(f33755,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(bool,f(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),bool,sK0)))) != s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),bool,sK0)))) | $spl58),
  inference(trivial_inequality_removal,[],[f33714])).
fof(f33714,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(bool,f(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),bool,sK0)))) != s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),bool,sK0)))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | $spl58),
  inference(superposition,[],[f922,f18404])).
fof(f922,plain,(
  ( ! [X4,X5] : (s(bool,'const.bool.IN_2'(s(sK0,sK4(X4,X5,sK0)),s(fun(sK0,bool),X4))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X4,X5,sK0)),s(fun(sK0,bool),X5))) | s(bool,f(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),X4,bool,sK0)))) != s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),X4,bool,sK0)))) | s(fun(sK0,bool),X5) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) )),
  inference(superposition,[],[f122,f70])).
fof(f122,plain,(
  ( ! [X0] : (s(fun(sK0,bool),X0) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | s(bool,f(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),X0,bool,sK0)))) != s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),X0,bool,sK0))))) )),
  inference(superposition,[],[f67,f68])).
fof(f68,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,X1),X2) = s(fun(X0,X1),X3) | s(X1,f(s(X0,sK3(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK3(X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK3(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK3(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f44])).
fof(f44,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f22])).
fof(f22,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f21])).
fof(f21,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.pred_set.INTER_UNION_COMPL',aHL_EXT)).
fof(f33848,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl83),
  inference(cnf_transformation,[],[f33848_D])).
fof(f33848_D,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl83),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl83])])).
fof(f33854,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))))) | $spl89),
  inference(cnf_transformation,[],[f33854_D])).
fof(f33854_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))))) <=> ~$spl89),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl89])])).
fof(f33716,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl58),
  inference(superposition,[],[f1085,f18404])).
fof(f1085,plain,(
  ( ! [X4,X5] : (s(bool,'const.bool.IN_2'(s(sK0,sK4(X4,X5,sK0)),s(fun(sK0,bool),X4))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X4,X5,sK0)),s(fun(sK0,bool),X5))) | s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))),X4,sK0)),s(fun(sK0,bool),X4))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))),X4,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) | s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),X5))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) )),
  inference(superposition,[],[f120,f70])).
fof(f120,plain,(
  ( ! [X2] : (s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),X2))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))),X2,sK0)),s(fun(sK0,bool),X2))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))),X2,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))))) )),
  inference(superposition,[],[f67,f70])).
fof(f33844,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))))) | $spl79),
  inference(cnf_transformation,[],[f33844_D])).
fof(f33844_D,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))))) <=> ~$spl79),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl79])])).
fof(f33852,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) | $spl87),
  inference(cnf_transformation,[],[f33852_D])).
fof(f33852_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) <=> ~$spl87),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl87])])).
fof(f33717,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))))) | $spl58),
  inference(superposition,[],[f1086,f18404])).
fof(f1086,plain,(
  ( ! [X6,X7] : (s(bool,'const.bool.IN_2'(s(sK0,sK4(X7,X6,sK0)),s(fun(sK0,bool),X6))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X7,X6,sK0)),s(fun(sK0,bool),X7))) | s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))),X6,sK0)),s(fun(sK0,bool),X6))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))),X6,sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) | s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),X7))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) )),
  inference(superposition,[],[f120,f70])).
fof(f33850,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))))) | $spl85),
  inference(cnf_transformation,[],[f33850_D])).
fof(f33850_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))))) <=> ~$spl85),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl85])])).
fof(f33718,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl58),
  inference(superposition,[],[f1177,f18404])).
fof(f1177,plain,(
  ( ! [X4,X5] : (s(bool,'const.bool.IN_2'(s(sK0,sK4(X4,X5,sK0)),s(fun(sK0,bool),X4))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X4,X5,sK0)),s(fun(sK0,bool),X5))) | s(bool,'const.bool.IN_2'(s(sK0,sK4(X4,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),X4))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X4,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) | s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),X5))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) )),
  inference(superposition,[],[f121,f70])).
fof(f121,plain,(
  ( ! [X3] : (s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),X3))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | s(bool,'const.bool.IN_2'(s(sK0,sK4(X3,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),X3))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X3,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))))) )),
  inference(superposition,[],[f67,f70])).
fof(f33846,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) | $spl81),
  inference(cnf_transformation,[],[f33846_D])).
fof(f33846_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) <=> ~$spl81),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl81])])).
fof(f33719,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))))) | $spl58),
  inference(superposition,[],[f1178,f18404])).
fof(f1178,plain,(
  ( ! [X6,X7] : (s(bool,'const.bool.IN_2'(s(sK0,sK4(X7,X6,sK0)),s(fun(sK0,bool),X6))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X7,X6,sK0)),s(fun(sK0,bool),X7))) | s(bool,'const.bool.IN_2'(s(sK0,sK4(X6,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),X6))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X6,'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) | s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),X7))) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) )),
  inference(superposition,[],[f121,f70])).
fof(f33840,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) | $spl74),
  inference(cnf_transformation,[],[f33840_D])).
fof(f33840_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f33842,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))) | $spl76),
  inference(cnf_transformation,[],[f33842_D])).
fof(f33842_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f33803,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) | $spl58),
  inference(resolution,[],[f33773,f74])).
fof(f74,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))) )),
  inference(cnf_transformation,[],[f54])).
fof(f33773,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))))) | $spl58),
  inference(subsumption_resolution,[],[f33722,f94])).
fof(f33722,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))))) | $spl58),
  inference(superposition,[],[f78,f18404])).
fof(f33836,plain,(
  $false | ($spl58 | $spl60)),
  inference(subsumption_resolution,[],[f33835,f94])).
fof(f33835,plain,(
  p(s(bool,'F_0')) | ($spl58 | $spl60)),
  inference(forward_demodulation,[],[f33834,f18855])).
fof(f18855,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | $spl60),
  inference(subsumption_resolution,[],[f18852,f101])).
fof(f18852,plain,(
  ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | $spl60),
  inference(superposition,[],[f18567,f73])).
fof(f18567,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) | $spl60),
  inference(resolution,[],[f18534,f77])).
fof(f18534,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl60),
  inference(subsumption_resolution,[],[f18499,f101])).
fof(f18499,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl60),
  inference(superposition,[],[f79,f18405])).
fof(f18405,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl60),
  inference(cnf_transformation,[],[f18405_D])).
fof(f18405_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f79,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))) )),
  inference(cnf_transformation,[],[f57])).
fof(f33834,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) | ($spl58 | $spl60)),
  inference(subsumption_resolution,[],[f33833,f94])).
fof(f33833,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) | ($spl58 | $spl60)),
  inference(forward_demodulation,[],[f33803,f18887])).
fof(f18887,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))) | $spl60),
  inference(subsumption_resolution,[],[f18884,f101])).
fof(f18884,plain,(
  ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))) | $spl60),
  inference(superposition,[],[f18597,f73])).
fof(f18597,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))) | $spl60),
  inference(resolution,[],[f18535,f77])).
fof(f18535,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | $spl60),
  inference(subsumption_resolution,[],[f18500,f101])).
fof(f18500,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | $spl60),
  inference(superposition,[],[f80,f18405])).
fof(f80,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))) )),
  inference(cnf_transformation,[],[f57])).
fof(f19456,plain,(
  $false | ($spl69 | $spl73)),
  inference(subsumption_resolution,[],[f19455,f18919])).
fof(f19455,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl73),
  inference(trivial_inequality_removal,[],[f19454])).
fof(f19454,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl73),
  inference(superposition,[],[f18923,f73])).
fof(f18921,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))))) | $spl70),
  inference(cnf_transformation,[],[f18921_D])).
fof(f18921_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))))) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f18923,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl73),
  inference(cnf_transformation,[],[f18923_D])).
fof(f18923_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl73),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl73])])).
fof(f18915,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))))))),
  inference(superposition,[],[f1264,f73])).
fof(f1264,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))))))),
  inference(equality_resolution,[],[f125])).
fof(f125,plain,(
  ( ! [X3] : (s(fun(sK0,bool),X3) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | s(bool,'const.bool.IN_2'(s(sK0,sK4(X3,'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),sK0)),s(fun(sK0,bool),X3))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X3,'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))))))) )),
  inference(superposition,[],[f67,f70])).
fof(f18917,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))))) | $spl66),
  inference(cnf_transformation,[],[f18917_D])).
fof(f18917_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))))) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f18919,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl69),
  inference(cnf_transformation,[],[f18919_D])).
fof(f18919_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl69),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl69])])).
fof(f18914,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))))))),
  inference(superposition,[],[f1264,f73])).
fof(f18469,plain,(
  $false | ($spl61 | $spl65)),
  inference(subsumption_resolution,[],[f18468,f18406])).
fof(f18468,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl65),
  inference(trivial_inequality_removal,[],[f18467])).
fof(f18467,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl65),
  inference(superposition,[],[f18410,f73])).
fof(f18402,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))))))),
  inference(superposition,[],[f1229,f73])).
fof(f18404,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))))) | $spl58),
  inference(cnf_transformation,[],[f18404_D])).
fof(f18404_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f18406,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | $spl61),
  inference(cnf_transformation,[],[f18406_D])).
fof(f18406_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f18401,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))))))),
  inference(superposition,[],[f1229,f73])).
fof(f16746,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | $spl44),
  inference(cnf_transformation,[],[f16746_D])).
fof(f16746_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f16748,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl46),
  inference(cnf_transformation,[],[f16748_D])).
fof(f16748_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f16756,plain,(
  s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))) | $spl57),
  inference(cnf_transformation,[],[f16756_D])).
fof(f16756_D,plain,(
  s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))) <=> ~$spl57),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl57])])).
fof(f16750,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))))))) | $spl49),
  inference(cnf_transformation,[],[f16750_D])).
fof(f16750_D,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))))))) <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f16754,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))))))) | s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | $spl4),
  inference(subsumption_resolution,[],[f16565,f446])).
fof(f446,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(X4,sK2,sK0)),s(fun(sK0,bool),sK2))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X4))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),X4))))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(X4,sK2,sK0)),s(fun(sK0,bool),X4)))) )),
  inference(superposition,[],[f113,f73])).
fof(f16565,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK2,sK0)),s(fun(sK0,bool),sK2))) | s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | $spl4),
  inference(trivial_inequality_removal,[],[f16515])).
fof(f16515,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK2,sK0)),s(fun(sK0,bool),sK2))) | s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | $spl4),
  inference(superposition,[],[f906,f14397])).
fof(f14397,plain,(
  ( ! [X4] : (s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),X4),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X4,sK0)),s(fun(sK0,bool),X4))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X4,sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) ) | $spl4),
  inference(trivial_inequality_removal,[],[f14395])).
fof(f14395,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),X4),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X4,sK0)),s(fun(sK0,bool),X4))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X4,sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) ) | $spl4),
  inference(superposition,[],[f11709,f73])).
fof(f11709,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X4,sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),X4),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X4,sK0)),s(fun(sK0,bool),X4)))) ) | $spl4),
  inference(superposition,[],[f1411,f73])).
fof(f1411,plain,(
  ( ! [X2] : (s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X2,sK0)),s(fun(sK0,bool),X2))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X2,sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),X2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0))))) ) | $spl4),
  inference(superposition,[],[f1324,f70])).
fof(f906,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(X0,sK2,sK0)),s(fun(sK0,bool),X1))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X0))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),X0))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(X0,sK2,sK0)),s(fun(sK0,bool),sK2))) | s(fun(sK0,bool),X0) != s(fun(sK0,bool),X1)) )),
  inference(superposition,[],[f439,f69])).
fof(f69,plain,(
  ( ! [X4,X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) != s(fun(X0,bool),X2)) )),
  inference(cnf_transformation,[],[f51])).
fof(f439,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(X5,sK2,sK0)),s(fun(sK0,bool),X5))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X5))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),X5))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(X5,sK2,sK0)),s(fun(sK0,bool),sK2)))) )),
  inference(superposition,[],[f113,f73])).
fof(f16753,plain,(
  ( ! [X46] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK2,sK0)),s(fun(sK0,bool),X46))) | s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X46,sK0)),s(fun(sK0,bool),X46))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X46,sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) ) | $spl54),
  inference(cnf_transformation,[],[f16753_D])).
fof(f16753_D,plain,(
  ( ! [X46] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK2,sK0)),s(fun(sK0,bool),X46))) | s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X46,sK0)),s(fun(sK0,bool),X46))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X46,sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f16496,plain,(
  ( ! [X46] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK2,sK0)),s(fun(sK0,bool),X46))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))))))) | s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X46,sK0)),s(fun(sK0,bool),X46))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X46,sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) ) | $spl4),
  inference(superposition,[],[f442,f14397])).
fof(f442,plain,(
  ( ! [X4,X5] : (s(bool,'const.bool.IN_2'(s(sK0,sK4(X4,sK2,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X4,sK2,sK0)),s(fun(sK0,bool),X5))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X4))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),X4))))))) | s(bool,'const.bool.IN_2'(s(sK0,sK4(X4,X5,sK0)),s(fun(sK0,bool),X4))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X4,X5,sK0)),s(fun(sK0,bool),X5)))) )),
  inference(superposition,[],[f113,f70])).
fof(f16752,plain,(
  ( ! [X45] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK2,sK0)),s(fun(sK0,bool),X45))) | s(bool,f(s(sK0,sK3(X45,'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),bool,sK0)))) != s(bool,g(s(sK0,sK3(X45,'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),bool,sK0))))) ) | $spl52),
  inference(cnf_transformation,[],[f16752_D])).
fof(f16752_D,plain,(
  ( ! [X45] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK2,sK0)),s(fun(sK0,bool),X45))) | s(bool,f(s(sK0,sK3(X45,'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),bool,sK0)))) != s(bool,g(s(sK0,sK3(X45,'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),bool,sK0))))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f16495,plain,(
  ( ! [X45] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK2,sK0)),s(fun(sK0,bool),X45))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))))))) | s(bool,f(s(sK0,sK3(X45,'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),bool,sK0)))) != s(bool,g(s(sK0,sK3(X45,'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) ) | $spl4),
  inference(superposition,[],[f441,f14397])).
fof(f441,plain,(
  ( ! [X2,X3] : (s(bool,'const.bool.IN_2'(s(sK0,sK4(X2,sK2,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X2,sK2,sK0)),s(fun(sK0,bool),X3))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),X2))))))) | s(bool,f(s(sK0,sK3(X3,X2,bool,sK0)))) != s(bool,g(s(sK0,sK3(X3,X2,bool,sK0))))) )),
  inference(superposition,[],[f113,f68])).
fof(f16751,plain,(
  ( ! [X44] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK2,sK0)),s(fun(sK0,bool),X44))) | s(bool,f(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X44,bool,sK0)))) != s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X44,bool,sK0))))) ) | $spl50),
  inference(cnf_transformation,[],[f16751_D])).
fof(f16751_D,plain,(
  ( ! [X44] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK2,sK0)),s(fun(sK0,bool),X44))) | s(bool,f(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X44,bool,sK0)))) != s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X44,bool,sK0))))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f16494,plain,(
  ( ! [X44] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK2,sK0)),s(fun(sK0,bool),X44))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))))))) | s(bool,f(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X44,bool,sK0)))) != s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X44,bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) ) | $spl4),
  inference(superposition,[],[f440,f14397])).
fof(f440,plain,(
  ( ! [X0,X1] : (s(bool,'const.bool.IN_2'(s(sK0,sK4(X0,sK2,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X0,sK2,sK0)),s(fun(sK0,bool),X1))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X0))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),X0))))))) | s(bool,f(s(sK0,sK3(X0,X1,bool,sK0)))) != s(bool,g(s(sK0,sK3(X0,X1,bool,sK0))))) )),
  inference(superposition,[],[f113,f68])).
fof(f16734,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | $spl30),
  inference(cnf_transformation,[],[f16734_D])).
fof(f16734_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f16736,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl32),
  inference(cnf_transformation,[],[f16736_D])).
fof(f16736_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f16744,plain,(
  s(fun(sK0,bool),sK1) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))) | $spl43),
  inference(cnf_transformation,[],[f16744_D])).
fof(f16744_D,plain,(
  s(fun(sK0,bool),sK1) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f16738,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) | $spl35),
  inference(cnf_transformation,[],[f16738_D])).
fof(f16738_D,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f16742,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) | s(fun(sK0,bool),sK1) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | $spl4),
  inference(subsumption_resolution,[],[f16586,f220])).
fof(f220,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(X4,sK1,sK0)),s(fun(sK0,bool),sK1))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X4),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),X4))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(X4,sK1,sK0)),s(fun(sK0,bool),X4)))) )),
  inference(superposition,[],[f105,f73])).
fof(f16586,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(fun(sK0,bool),sK1) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | $spl4),
  inference(trivial_inequality_removal,[],[f16491])).
fof(f16491,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(fun(sK0,bool),sK1) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | $spl4),
  inference(superposition,[],[f816,f14397])).
fof(f816,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(X0,sK1,sK0)),s(fun(sK0,bool),X1))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),X0))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(X0,sK1,sK0)),s(fun(sK0,bool),sK1))) | s(fun(sK0,bool),X0) != s(fun(sK0,bool),X1)) )),
  inference(superposition,[],[f213,f69])).
fof(f213,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(X5,sK1,sK0)),s(fun(sK0,bool),X5))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X5),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),X5))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(X5,sK1,sK0)),s(fun(sK0,bool),sK1)))) )),
  inference(superposition,[],[f105,f73])).
fof(f16741,plain,(
  ( ! [X43] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),X43))) | s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X43,sK0)),s(fun(sK0,bool),X43))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X43,sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) ) | $spl40),
  inference(cnf_transformation,[],[f16741_D])).
fof(f16741_D,plain,(
  ( ! [X43] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),X43))) | s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X43,sK0)),s(fun(sK0,bool),X43))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X43,sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f16472,plain,(
  ( ! [X43] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),X43))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) | s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X43,sK0)),s(fun(sK0,bool),X43))) != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X43,sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) ) | $spl4),
  inference(superposition,[],[f216,f14397])).
fof(f216,plain,(
  ( ! [X4,X5] : (s(bool,'const.bool.IN_2'(s(sK0,sK4(X4,sK1,sK0)),s(fun(sK0,bool),sK1))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X4,sK1,sK0)),s(fun(sK0,bool),X5))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X4),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),X4))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) | s(bool,'const.bool.IN_2'(s(sK0,sK4(X4,X5,sK0)),s(fun(sK0,bool),X4))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X4,X5,sK0)),s(fun(sK0,bool),X5)))) )),
  inference(superposition,[],[f105,f70])).
fof(f16740,plain,(
  ( ! [X42] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),X42))) | s(bool,f(s(sK0,sK3(X42,'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),bool,sK0)))) != s(bool,g(s(sK0,sK3(X42,'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),bool,sK0))))) ) | $spl38),
  inference(cnf_transformation,[],[f16740_D])).
fof(f16740_D,plain,(
  ( ! [X42] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),X42))) | s(bool,f(s(sK0,sK3(X42,'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),bool,sK0)))) != s(bool,g(s(sK0,sK3(X42,'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),bool,sK0))))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f16471,plain,(
  ( ! [X42] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),X42))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) | s(bool,f(s(sK0,sK3(X42,'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),bool,sK0)))) != s(bool,g(s(sK0,sK3(X42,'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) ) | $spl4),
  inference(superposition,[],[f215,f14397])).
fof(f215,plain,(
  ( ! [X2,X3] : (s(bool,'const.bool.IN_2'(s(sK0,sK4(X2,sK1,sK0)),s(fun(sK0,bool),sK1))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X2,sK1,sK0)),s(fun(sK0,bool),X3))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X2),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),X2))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) | s(bool,f(s(sK0,sK3(X3,X2,bool,sK0)))) != s(bool,g(s(sK0,sK3(X3,X2,bool,sK0))))) )),
  inference(superposition,[],[f105,f68])).
fof(f16739,plain,(
  ( ! [X41] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),X41))) | s(bool,f(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X41,bool,sK0)))) != s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X41,bool,sK0))))) ) | $spl36),
  inference(cnf_transformation,[],[f16739_D])).
fof(f16739_D,plain,(
  ( ! [X41] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),X41))) | s(bool,f(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X41,bool,sK0)))) != s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X41,bool,sK0))))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f16470,plain,(
  ( ! [X41] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),X41))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) | s(bool,f(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X41,bool,sK0)))) != s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X41,bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) ) | $spl4),
  inference(superposition,[],[f214,f14397])).
fof(f214,plain,(
  ( ! [X0,X1] : (s(bool,'const.bool.IN_2'(s(sK0,sK4(X0,sK1,sK0)),s(fun(sK0,bool),sK1))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X0,sK1,sK0)),s(fun(sK0,bool),X1))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),X0))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) | s(bool,f(s(sK0,sK3(X0,X1,bool,sK0)))) != s(bool,g(s(sK0,sK3(X0,X1,bool,sK0))))) )),
  inference(superposition,[],[f105,f68])).
fof(f16726,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | $spl20),
  inference(cnf_transformation,[],[f16726_D])).
fof(f16726_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f16728,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl22),
  inference(cnf_transformation,[],[f16728_D])).
fof(f16728_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f16730,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))))) | $spl25),
  inference(cnf_transformation,[],[f16730_D])).
fof(f16730_D,plain,(
  s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f16732,plain,(
  ( ! [X40] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),X40))) | s(bool,f(s(sK0,sK3(X40,'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),bool,sK0)))) != s(bool,g(s(sK0,sK3(X40,'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),bool,sK0))))) ) | $spl28),
  inference(cnf_transformation,[],[f16732_D])).
fof(f16732_D,plain,(
  ( ! [X40] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),X40))) | s(bool,f(s(sK0,sK3(X40,'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),bool,sK0)))) != s(bool,g(s(sK0,sK3(X40,'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),bool,sK0))))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f16447,plain,(
  ( ! [X40] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),X40))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))))) | s(bool,f(s(sK0,sK3(X40,'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),bool,sK0)))) != s(bool,g(s(sK0,sK3(X40,'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) ) | $spl4),
  inference(superposition,[],[f729,f14397])).
fof(f729,plain,(
  ( ! [X2,X3] : (s(bool,'const.bool.IN_2'(s(sK0,sK4(X2,'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),X3))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X2,'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),X2))))) | s(bool,f(s(sK0,sK3(X3,X2,bool,sK0)))) != s(bool,g(s(sK0,sK3(X3,X2,bool,sK0))))) )),
  inference(superposition,[],[f117,f68])).
fof(f16731,plain,(
  ( ! [X39] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),X39))) | s(bool,f(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X39,bool,sK0)))) != s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X39,bool,sK0))))) ) | $spl26),
  inference(cnf_transformation,[],[f16731_D])).
fof(f16731_D,plain,(
  ( ! [X39] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),X39))) | s(bool,f(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X39,bool,sK0)))) != s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X39,bool,sK0))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f16446,plain,(
  ( ! [X39] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),X39))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))))) | s(bool,f(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X39,bool,sK0)))) != s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),X39,bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) ) | $spl4),
  inference(superposition,[],[f728,f14397])).
fof(f728,plain,(
  ( ! [X0,X1] : (s(bool,'const.bool.IN_2'(s(sK0,sK4(X0,'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),X1))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(X0,'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),X0))))) | s(bool,f(s(sK0,sK3(X0,X1,bool,sK0)))) != s(bool,g(s(sK0,sK3(X0,X1,bool,sK0))))) )),
  inference(superposition,[],[f117,f68])).
fof(f1587,plain,(
  $false | ($spl15 | $spl19)),
  inference(subsumption_resolution,[],[f1586,f1403])).
fof(f1586,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),bool,sK0)))) | $spl19),
  inference(trivial_inequality_removal,[],[f1585])).
fof(f1585,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),bool,sK0)))) | $spl19),
  inference(superposition,[],[f1407,f73])).
fof(f1405,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),bool,sK0)))) | $spl16),
  inference(cnf_transformation,[],[f1405_D])).
fof(f1405_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),bool,sK0)))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1407,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),bool,sK0)))) | $spl19),
  inference(cnf_transformation,[],[f1407_D])).
fof(f1407_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),bool,sK0)))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f1399,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),bool,sK0))))),
  inference(superposition,[],[f1010,f73])).
fof(f1010,plain,(
  s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),bool,sK0)))) != s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),bool,sK0))))),
  inference(equality_resolution,[],[f123])).
fof(f123,plain,(
  ( ! [X1] : (s(fun(sK0,bool),X1) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | s(bool,f(s(sK0,sK3(X1,'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),bool,sK0)))) != s(bool,g(s(sK0,sK3(X1,'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),bool,sK0))))) )),
  inference(superposition,[],[f67,f68])).
fof(f1401,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),bool,sK0)))) | $spl12),
  inference(cnf_transformation,[],[f1401_D])).
fof(f1401_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),bool,sK0)))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1403,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),bool,sK0)))) | $spl15),
  inference(cnf_transformation,[],[f1403_D])).
fof(f1403_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),bool,sK0)))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f1398,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),bool,sK0))))),
  inference(superposition,[],[f1010,f73])).
fof(f1377,plain,(
  $false | ($spl7 | $spl11)),
  inference(subsumption_resolution,[],[f1376,f1326])).
fof(f1376,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl11),
  inference(trivial_inequality_removal,[],[f1375])).
fof(f1375,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl11),
  inference(superposition,[],[f1330,f73])).
fof(f1328,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl8),
  inference(cnf_transformation,[],[f1328_D])).
fof(f1328_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1330,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl11),
  inference(cnf_transformation,[],[f1330_D])).
fof(f1330_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f1322,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0))))),
  inference(superposition,[],[f936,f73])).
fof(f936,plain,(
  s(bool,f(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) != s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0))))),
  inference(equality_resolution,[],[f122])).
fof(f1324,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl4),
  inference(cnf_transformation,[],[f1324_D])).
fof(f1324_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1326,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl7),
  inference(cnf_transformation,[],[f1326_D])).
fof(f1326_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f1321,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)),bool,sK0))))),
  inference(superposition,[],[f936,f73])).
fof(f162,plain,(
  ( ! [X1] : (s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),X1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,X1,bool,sK0))))) ) | $spl0),
  inference(cnf_transformation,[],[f162_D])).
fof(f162_D,plain,(
  ( ! [X1] : (s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),X1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,X1,bool,sK0))))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f164,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl3),
  inference(cnf_transformation,[],[f164_D])).
fof(f164_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f161,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),X1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,X1,bool,sK0))))) )),
  inference(subsumption_resolution,[],[f159,f126])).
fof(f126,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK0,sK3(sK1,X0,bool,sK0)))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),X0))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,X0,bool,sK0))))) )),
  inference(superposition,[],[f102,f73])).
fof(f102,plain,(
  ( ! [X0] : (s(bool,f(s(sK0,sK3(sK1,X0,bool,sK0)))) != s(bool,g(s(sK0,sK3(sK1,X0,bool,sK0)))) | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),X0))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2)))))))) )),
  inference(superposition,[],[f67,f68])).
fof(f159,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'F_0') | s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),sK2))) != s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),X1))),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK2))))))) | s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,X1,bool,sK0)))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3(sK1,X1,bool,sK0))))) )),
  inference(superposition,[],[f126,f73])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1555
% Passive clauses: 25413
% Generated clauses: 39327
% Final active clauses: 1316
% Final passive clauses: 19153
% Input formulas: 17
% Initial clauses: 34
% 
% Selected by SInE selection: 17
% SInE iterations: 1
% 
% Duplicate literals: 920
% Trivial inequalities: 595
% Fw subsumption resolutions: 929
% Fw demodulations: 3525
% Bw demodulations: 134
% 
% Simple tautologies: 52
% Equational tautologies: 1
% Forward subsumptions: 8401
% 
% Binary resolution: 754
% Forward superposition: 24423
% Backward superposition: 7757
% Equality factoring: 1
% Equality resolution: 354
% 
% Split clauses: 37
% Split components: 108
% SAT solver clauses: 47
% SAT solver binary clauses: 19
% 
% Sat splits: 37
% Sat splitting refutations: 9
% 
% Memory used [KB]: 20340
% Time elapsed: 1.945 s
% ------------------------------
% ------------------------------
% Success in time 13.66 s
