Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.ringNorm
% (75801)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 53
% Passive clauses: 1617
% Generated clauses: 8888
% Final active clauses: 51
% Final passive clauses: 1485
% Input formulas: 21
% Initial clauses: 22
% 
% Duplicate literals: 619
% Trivial inequalities: 108
% Fw subsumption resolutions: 974
% Fw demodulations: 2
% 
% Simple tautologies: 208
% Equational tautologies: 129
% Forward subsumptions: 3368
% 
% Binary resolution: 25
% Forward superposition: 3294
% Backward superposition: 2914
% Self superposition: 481
% Equality factoring: 406
% Equality resolution: 11
% 
% Split clauses: 379
% Split components: 762
% SAT solver clauses: 1399
% SAT solver unit clauses: 352
% SAT solver binary clauses: 505
% 
% Sat splits: 379
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1395
% TWLsolver calls for satisfiability: 39
% 
% Memory used [KB]: 4605
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ringNorm
% (75820)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 37
% Passive clauses: 94
% Generated clauses: 5315
% Final active clauses: 34
% Final passive clauses: 55
% Input formulas: 21
% Initial clauses: 22
% 
% Duplicate literals: 1173
% Trivial inequalities: 60
% Fw subsumption resolutions: 137
% Fw demodulations: 3
% Bw demodulations: 1
% Condensations: 1624
% 
% Simple tautologies: 122
% Equational tautologies: 77
% Forward subsumptions: 1309
% 
% Binary resolution: 16
% Forward superposition: 990
% Backward superposition: 925
% Self superposition: 249
% Equality factoring: 222
% Equality resolution: 5
% 
% Split clauses: 27
% Split components: 54
% SAT solver clauses: 276
% SAT solver unit clauses: 38
% SAT solver binary clauses: 66
% 
% Sat splits: 27
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 276
% TWLsolver calls for satisfiability: 47
% 
% Memory used [KB]: 3582
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.ringNorm
% (75842)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 108
% Passive clauses: 671
% Generated clauses: 31840
% Final active clauses: 74
% Final passive clauses: 223
% Input formulas: 21
% Initial clauses: 22
% Discarded non-redundant clauses: 31
% 
% Duplicate literals: 3407
% Trivial inequalities: 260
% Fw subsumption resolutions: 1465
% Fw demodulations: 443
% Bw demodulations: 169
% Condensations: 6257
% Global subsumptions: 13
% 
% Simple tautologies: 429
% Equational tautologies: 2079
% Forward subsumptions: 16816
% 
% Unit resulting resolution: 4209
% Forward superposition: 8414
% Backward superposition: 5677
% Self superposition: 501
% Equality factoring: 414
% Equality resolution: 15
% 
% Split clauses: 200
% Split components: 448
% SAT solver clauses: 1094
% SAT solver unit clauses: 38
% SAT solver binary clauses: 101
% 
% Sat splits: 200
% Sat splitting refutations: 8
% 
% Memory used [KB]: 12665
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.ringNorm
% Refutation not found, incomplete strategy
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation not found, incomplete strategy

% Active clauses: 11
% Passive clauses: 1
% Generated clauses: 1
% Final active clauses: 11
% Input formulas: 21
% Initial clauses: 13
% 
% Selected by SInE selection: 12
% SInE iterations: 1
% 
% Simple tautologies: 2
% 
% SAT solver clauses: 1
% 
% Memory used [KB]: 767
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.ringNorm
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ringNorm
% SZS output start Proof for thm.ringNorm
fof(f8634,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f8626])).
fof(f8626,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_prod_3'(s('type.ring.ring'(sK0),sK1),s('type.canonical.canonical_sum'(sK0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),sK2),s('type.canonical.canonical_sum'(sK0),sK3))),s('type.canonical.canonical_sum'(sK0),sK4))) != s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_prod_3'(s('type.ring.ring'(sK0),sK1),s('type.canonical.canonical_sum'(sK0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),sK2),s('type.canonical.canonical_sum'(sK0),sK3))),s('type.canonical.canonical_sum'(sK0),sK4)))),
  inference(superposition,[],[f65,f2511])).
fof(f2511,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_prod_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X4))),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar2_3'(s('type.ring.ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))),s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_prod_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X4),s('type.canonical.canonical_sum'(X0),X3)))))) )),
  inference(forward_demodulation,[],[f2510,f632])).
fof(f632,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_prod_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_prod_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(forward_demodulation,[],[f619,f79])).
fof(f79,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_prod_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_prod_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(cnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_prod_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_prod_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(flattening,[],[f44])).
fof(f44,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_prod_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_prod_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6] : ! [X15,X16,X17] : s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_prod_3'(s('type.ring.ring'(X6),X15),s('type.canonical.canonical_sum'(X6),X16),s('type.canonical.canonical_sum'(X6),X17))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)),happ(s(fun('type.canonical.canonical_sum'(X6),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.ringNorm.r_canonical_sum_prod_1'(s('type.ring.ring'(X6),X15))),s('type.canonical.canonical_sum'(X6),X16))),s('type.canonical.canonical_sum'(X6),X17)))),
  file('thm.ringNorm.canonical_sum_prod_def|split|1','ARITY_1_3_const.ringNorm.r_canonical_sum_prod')).
fof(f619,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_prod_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_prod_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(superposition,[],[f78,f71])).
fof(f71,plain,(
  ( ! [X0,X1] : (s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_prod_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_prod_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1] : s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_prod_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_prod_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f30])).
fof(f30,plain,(
  ! [X0] : ! [X1] : s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_prod_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_prod_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : ! [X7] : s(fun('type.canonical.canonical_sum'(X6),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.ringNorm.r_canonical_sum_prod_1'(s('type.ring.ring'(X6),X7))) = s(fun('type.canonical.canonical_sum'(X6),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.canonical.canonical_sum_prod_1'(s('type.semi_ring.semi_ring'(X6),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X6),X7)))))),
  file('thm.ringNorm.canonical_sum_prod_def|split|1','thm.ringNorm.r_canonical_sum_prod_def')).
fof(f78,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_prod_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_prod_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_prod_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_prod_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(flattening,[],[f42])).
fof(f42,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_prod_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_prod_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6] : ! [X15,X16,X17] : s('type.canonical.canonical_sum'(X6),'const.canonical.canonical_sum_prod_3'(s('type.semi_ring.semi_ring'(X6),X15),s('type.canonical.canonical_sum'(X6),X16),s('type.canonical.canonical_sum'(X6),X17))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)),happ(s(fun('type.canonical.canonical_sum'(X6),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.canonical.canonical_sum_prod_1'(s('type.semi_ring.semi_ring'(X6),X15))),s('type.canonical.canonical_sum'(X6),X16))),s('type.canonical.canonical_sum'(X6),X17)))),
  file('thm.ringNorm.canonical_sum_prod_def|split|1','ARITY_1_3_const.canonical.canonical_sum_prod')).
fof(f2510,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_prod_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X4))),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar2_3'(s('type.ring.ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))),s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_prod_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X4),s('type.canonical.canonical_sum'(X0),X3)))))) )),
  inference(forward_demodulation,[],[f2509,f1045])).
fof(f1045,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar2_3'(s('type.ring.ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar2_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(forward_demodulation,[],[f1027,f76])).
fof(f76,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar2_3'(s('type.ring.ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_scalar2_1'(s('type.ring.ring'(X0),X1))),s('type.list.list'('type.quote.index'),X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(cnf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar2_3'(s('type.ring.ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_scalar2_1'(s('type.ring.ring'(X0),X1))),s('type.list.list'('type.quote.index'),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(flattening,[],[f38])).
fof(f38,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar2_3'(s('type.ring.ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_scalar2_1'(s('type.ring.ring'(X0),X1))),s('type.list.list'('type.quote.index'),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X15,X16,X17] : s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_scalar2_3'(s('type.ring.ring'(X6),X15),s('type.list.list'('type.quote.index'),X16),s('type.canonical.canonical_sum'(X6),X17))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.ringNorm.r_canonical_sum_scalar2_1'(s('type.ring.ring'(X6),X15))),s('type.list.list'('type.quote.index'),X16))),s('type.canonical.canonical_sum'(X6),X17)))),
  file('thm.ringNorm.canonical_sum_prod_def|split|1','ARITY_1_3_const.ringNorm.r_canonical_sum_scalar2')).
fof(f1027,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar2_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_scalar2_1'(s('type.ring.ring'(X0),X1))),s('type.list.list'('type.quote.index'),X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(superposition,[],[f81,f70])).
fof(f70,plain,(
  ( ! [X0,X1] : (s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_scalar2_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_scalar2_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f29])).
fof(f29,plain,(
  ! [X0,X1] : s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_scalar2_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_scalar2_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f28])).
fof(f28,plain,(
  ! [X0] : ! [X1] : s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_scalar2_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_scalar2_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6] : ! [X7] : s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.ringNorm.r_canonical_sum_scalar2_1'(s('type.ring.ring'(X6),X7))) = s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.canonical.canonical_sum_scalar2_1'(s('type.semi_ring.semi_ring'(X6),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X6),X7)))))),
  file('thm.ringNorm.canonical_sum_prod_def|split|1','thm.ringNorm.r_canonical_sum_scalar2_def')).
fof(f81,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar2_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_scalar2_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.list.list'('type.quote.index'),X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar2_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_scalar2_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.list.list'('type.quote.index'),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar2_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_scalar2_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.list.list'('type.quote.index'),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X15,X16,X17] : s('type.canonical.canonical_sum'(X6),'const.canonical.canonical_sum_scalar2_3'(s('type.semi_ring.semi_ring'(X6),X15),s('type.list.list'('type.quote.index'),X16),s('type.canonical.canonical_sum'(X6),X17))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)),happ(s(fun('type.list.list'('type.quote.index'),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.canonical.canonical_sum_scalar2_1'(s('type.semi_ring.semi_ring'(X6),X15))),s('type.list.list'('type.quote.index'),X16))),s('type.canonical.canonical_sum'(X6),X17)))),
  file('thm.ringNorm.canonical_sum_prod_def|split|1','ARITY_1_3_const.canonical.canonical_sum_scalar2')).
fof(f2509,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_prod_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X4))),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar2_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))),s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_prod_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X4),s('type.canonical.canonical_sum'(X0),X3)))))) )),
  inference(forward_demodulation,[],[f2483,f632])).
fof(f2483,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_prod_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X4))),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar2_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))),s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_prod_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X4),s('type.canonical.canonical_sum'(X0),X3)))))) )),
  inference(superposition,[],[f811,f82])).
fof(f82,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_prod_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))),s('type.canonical.canonical_sum'(X0),X4))) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar2_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X4))),s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_prod_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X3),s('type.canonical.canonical_sum'(X0),X4)))))) )),
  inference(cnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0,X1,X2,X3,X4] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_prod_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))),s('type.canonical.canonical_sum'(X0),X4))) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar2_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X4))),s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_prod_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X3),s('type.canonical.canonical_sum'(X0),X4)))))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_prod_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))),s('type.canonical.canonical_sum'(X0),X4))) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_scalar2_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X4))),s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_prod_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X3),s('type.canonical.canonical_sum'(X0),X4)))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : ! [X13,X8,X9,X10] : s('type.canonical.canonical_sum'(X6),'const.canonical.canonical_sum_prod_3'(s('type.semi_ring.semi_ring'(X6),X13),s('type.canonical.canonical_sum'(X6),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X8),s('type.canonical.canonical_sum'(X6),X9))),s('type.canonical.canonical_sum'(X6),X10))) = s('type.canonical.canonical_sum'(X6),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X6),X13),s('type.canonical.canonical_sum'(X6),'const.canonical.canonical_sum_scalar2_3'(s('type.semi_ring.semi_ring'(X6),X13),s('type.list.list'('type.quote.index'),X8),s('type.canonical.canonical_sum'(X6),X10))),s('type.canonical.canonical_sum'(X6),'const.canonical.canonical_sum_prod_3'(s('type.semi_ring.semi_ring'(X6),X13),s('type.canonical.canonical_sum'(X6),X9),s('type.canonical.canonical_sum'(X6),X10)))))),
  file('thm.ringNorm.canonical_sum_prod_def|split|1','thm.canonical.canonical_sum_prod_def|split|1')).
fof(f811,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(forward_demodulation,[],[f790,f77])).
fof(f77,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(cnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(flattening,[],[f40])).
fof(f40,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : ! [X15,X16,X17] : s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X6),X15),s('type.canonical.canonical_sum'(X6),X16),s('type.canonical.canonical_sum'(X6),X17))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)),happ(s(fun('type.canonical.canonical_sum'(X6),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X6),X15))),s('type.canonical.canonical_sum'(X6),X16))),s('type.canonical.canonical_sum'(X6),X17)))),
  file('thm.ringNorm.canonical_sum_prod_def|split|1','ARITY_1_3_const.ringNorm.r_canonical_sum_merge')).
fof(f790,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(superposition,[],[f80,f69])).
fof(f69,plain,(
  ( ! [X0,X1] : (s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0,X1] : s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f26])).
fof(f26,plain,(
  ! [X0] : ! [X1] : s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6] : ! [X7] : s(fun('type.canonical.canonical_sum'(X6),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.ringNorm.r_canonical_sum_merge_1'(s('type.ring.ring'(X6),X7))) = s(fun('type.canonical.canonical_sum'(X6),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X6),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X6),X7)))))),
  file('thm.ringNorm.canonical_sum_prod_def|split|1','thm.ringNorm.r_canonical_sum_merge_def')).
fof(f80,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6] : ! [X15,X16,X17] : s('type.canonical.canonical_sum'(X6),'const.canonical.canonical_sum_merge_3'(s('type.semi_ring.semi_ring'(X6),X15),s('type.canonical.canonical_sum'(X6),X16),s('type.canonical.canonical_sum'(X6),X17))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)),happ(s(fun('type.canonical.canonical_sum'(X6),fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6))),'const.canonical.canonical_sum_merge_1'(s('type.semi_ring.semi_ring'(X6),X15))),s('type.canonical.canonical_sum'(X6),X16))),s('type.canonical.canonical_sum'(X6),X17)))),
  file('thm.ringNorm.canonical_sum_prod_def|split|1','ARITY_1_3_const.canonical.canonical_sum_merge')).
fof(f65,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_prod_3'(s('type.ring.ring'(sK0),sK1),s('type.canonical.canonical_sum'(sK0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),sK2),s('type.canonical.canonical_sum'(sK0),sK3))),s('type.canonical.canonical_sum'(sK0),sK4))) != s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(sK0),sK1),s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_scalar2_3'(s('type.ring.ring'(sK0),sK1),s('type.list.list'('type.quote.index'),sK2),s('type.canonical.canonical_sum'(sK0),sK4))),s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_prod_3'(s('type.ring.ring'(sK0),sK1),s('type.canonical.canonical_sum'(sK0),sK3),s('type.canonical.canonical_sum'(sK0),sK4)))))),
  inference(cnf_transformation,[],[f62])).
fof(f62,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_prod_3'(s('type.ring.ring'(sK0),sK1),s('type.canonical.canonical_sum'(sK0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),sK2),s('type.canonical.canonical_sum'(sK0),sK3))),s('type.canonical.canonical_sum'(sK0),sK4))) != s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(sK0),sK1),s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_scalar2_3'(s('type.ring.ring'(sK0),sK1),s('type.list.list'('type.quote.index'),sK2),s('type.canonical.canonical_sum'(sK0),sK4))),s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_prod_3'(s('type.ring.ring'(sK0),sK1),s('type.canonical.canonical_sum'(sK0),sK3),s('type.canonical.canonical_sum'(sK0),sK4)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f60])).
fof(f60,plain,(
  ? [X0,X1,X2,X3,X4] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_prod_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))),s('type.canonical.canonical_sum'(X0),X4))) != s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar2_3'(s('type.ring.ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X4))),s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_prod_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X3),s('type.canonical.canonical_sum'(X0),X4)))))),
  inference(ennf_transformation,[],[f24])).
fof(f24,plain,(
  ~! [X0,X1,X2,X3,X4] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_prod_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))),s('type.canonical.canonical_sum'(X0),X4))) = s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar2_3'(s('type.ring.ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X4))),s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_prod_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X3),s('type.canonical.canonical_sum'(X0),X4)))))),
  inference(flattening,[],[f23])).
fof(f23,plain,(
  ~! [X0] : ! [X1,X2,X3,X4] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_prod_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X3))),s('type.canonical.canonical_sum'(X0),X4))) = s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_scalar2_3'(s('type.ring.ring'(X0),X1),s('type.list.list'('type.quote.index'),X2),s('type.canonical.canonical_sum'(X0),X4))),s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_prod_3'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X3),s('type.canonical.canonical_sum'(X0),X4)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10] : s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_prod_3'(s('type.ring.ring'(X6),X7),s('type.canonical.canonical_sum'(X6),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X8),s('type.canonical.canonical_sum'(X6),X9))),s('type.canonical.canonical_sum'(X6),X10))) = s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X6),X7),s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_scalar2_3'(s('type.ring.ring'(X6),X7),s('type.list.list'('type.quote.index'),X8),s('type.canonical.canonical_sum'(X6),X10))),s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_prod_3'(s('type.ring.ring'(X6),X7),s('type.canonical.canonical_sum'(X6),X9),s('type.canonical.canonical_sum'(X6),X10)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10] : s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_prod_3'(s('type.ring.ring'(X6),X7),s('type.canonical.canonical_sum'(X6),'const.canonical.Cons_varlist_2'(s('type.list.list'('type.quote.index'),X8),s('type.canonical.canonical_sum'(X6),X9))),s('type.canonical.canonical_sum'(X6),X10))) = s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_merge_3'(s('type.ring.ring'(X6),X7),s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_scalar2_3'(s('type.ring.ring'(X6),X7),s('type.list.list'('type.quote.index'),X8),s('type.canonical.canonical_sum'(X6),X10))),s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_prod_3'(s('type.ring.ring'(X6),X7),s('type.canonical.canonical_sum'(X6),X9),s('type.canonical.canonical_sum'(X6),X10)))))),
  file('thm.ringNorm.canonical_sum_prod_def|split|1','thm.ringNorm.canonical_sum_prod_def|split|1')).
% SZS output end Proof for thm.ringNorm
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 102
% Passive clauses: 662
% Generated clauses: 8625
% Final active clauses: 78
% Final passive clauses: 251
% Input formulas: 21
% Initial clauses: 22
% 
% Duplicate literals: 391
% Trivial inequalities: 192
% Fw subsumption resolutions: 1092
% Fw demodulations: 406
% Bw demodulations: 214
% Global subsumptions: 7
% 
% Simple tautologies: 316
% Equational tautologies: 141
% Forward subsumptions: 5377
% Bw demodulations to eq. taut.: 11
% 
% Unit resulting resolution: 404
% Forward superposition: 3299
% Backward superposition: 2285
% Self superposition: 118
% Equality factoring: 50
% Equality resolution: 11
% 
% Split clauses: 15
% Split components: 31
% SAT solver clauses: 722
% SAT solver unit clauses: 39
% SAT solver binary clauses: 60
% 
% Sat splits: 15
% Sat splitting refutations: 36
% 
% Memory used [KB]: 7291
% Time elapsed: 0.350 s
% ------------------------------
% ------------------------------
% Success in time 12.09 s
