Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.real
% (51314)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 164
% Passive clauses: 2940
% Generated clauses: 24433
% Final active clauses: 156
% Final passive clauses: 2542
% Input formulas: 28
% Initial clauses: 49
% 
% Duplicate literals: 968
% Trivial inequalities: 705
% Fw subsumption resolutions: 3213
% Fw demodulations: 3473
% Bw demodulations: 4
% 
% Simple tautologies: 596
% Equational tautologies: 39
% Forward subsumptions: 11542
% Fw demodulations to eq. taut.: 2
% 
% Binary resolution: 210
% Factoring: 2
% Forward superposition: 8274
% Backward superposition: 7096
% Self superposition: 169
% Equality factoring: 83
% Equality resolution: 45
% 
% Split clauses: 567
% Split components: 1164
% SAT solver clauses: 1071
% SAT solver unit clauses: 50
% SAT solver binary clauses: 225
% 
% Sat splits: 567
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1058
% TWLsolver calls for satisfiability: 153
% 
% Memory used [KB]: 12153
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 294 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.real
% (51351)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 167
% Passive clauses: 860
% Generated clauses: 15095
% Final active clauses: 123
% Final passive clauses: 482
% Input formulas: 28
% Initial clauses: 49
% 
% Duplicate literals: 626
% Trivial inequalities: 192
% Fw subsumption resolutions: 1100
% Fw demodulations: 2016
% Bw demodulations: 196
% Condensations: 962
% Global subsumptions: 5
% 
% Simple tautologies: 238
% Equational tautologies: 35
% Forward subsumptions: 7329
% Fw demodulations to eq. taut.: 34
% Bw demodulations to eq. taut.: 36
% 
% Binary resolution: 143
% Factoring: 1
% Forward superposition: 4672
% Backward superposition: 3793
% Self superposition: 143
% Equality factoring: 60
% Equality resolution: 22
% 
% Split clauses: 141
% Split components: 288
% SAT solver clauses: 1696
% SAT solver unit clauses: 59
% SAT solver binary clauses: 43
% SAT solver learnt clauses: 12
% SAT solver learnt literals: 5
% 
% Sat splits: 141
% Sat splitting refutations: 17
% 
% TWLsolver clauses: 1674
% TWLsolver calls for satisfiability: 362
% 
% Memory used [KB]: 10874
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.real
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.real
% SZS output start Proof for thm.real
fof(f1986,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f91,f142,f141,f90,f143,f140,f1333,f1334,f1335,f1337,f1342,f1344,f1346,f1348,f1349,f1352,f1356,f1362,f1365,f1756,f1739,f1881,f1899,f1900,f1958,f1959,f1960,f1962,f1967,f1969,f1971,f1973,f1981,f1982,f1985])).
fof(f1985,plain,(
  $false | ($spl0 | $spl2)),
  inference(global_subsumption,[],[f91,f90,f92,f94,f101,f103,f105,f104,f114,f117,f124,f123,f122,f125,f126,f127,f131,f128,f132,f147,f167,f199,f245,f390,f406,f425,f426,f440,f450,f451,f522,f770,f771,f778,f821,f962,f988,f990,f1029,f1031,f1203,f1205,f1207,f1209,f1054,f1048,f1046,f1043,f1042,f1206,f1195,f1190,f1129,f1128,f1127,f1105,f1092,f1091,f1060,f1058,f1055,f1045,f1044,f1033,f989,f945,f940,f872,f871,f854,f836,f835,f834,f766,f452,f427,f380,f341,f340,f307,f305,f304,f303,f259,f224,f169,f144,f118,f109,f108,f140,f1369,f1370,f1371,f1372,f1383,f1374,f1375,f1376,f1377,f1466,f1468,f1469,f1508,f1525,f1526,f1527,f1660,f1661,f1662,f1761,f1793,f1794,f1795,f1850,f1852,f1853,f1854,f1903,f1957,f1906,f1961,f1966,f1968,f1944,f1972,f1980,f1984])).
fof(f1984,plain,(
  s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))) | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f1983,f92])).
fof(f1983,plain,(
  ~p(s(bool,'T_0')) | s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f1896,f1903])).
fof(f1896,plain,(
  s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))) | ~p(s(bool,'F_0')) | ($spl0 | $spl2)),
  inference(superposition,[],[f1761,f118])).
fof(f1980,plain,(
  ( ! [X19] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))) != s('type.realax.real','const.bool.COND_3'(s(bool,X19),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0)))))))) ) | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f1979,f92])).
fof(f1979,plain,(
  ( ! [X19] : (~p(s(bool,'T_0')) | s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))) != s('type.realax.real','const.bool.COND_3'(s(bool,X19),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0)))))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f1893,f1903])).
fof(f1893,plain,(
  ( ! [X19,X18] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))) != s('type.realax.real','const.bool.COND_3'(s(bool,X19),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))))) | ~p(s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X18),s(bool,X19))))) ) | ($spl0 | $spl2)),
  inference(superposition,[],[f1761,f118])).
fof(f1972,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f1954,f104])).
fof(f1954,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | ($spl0 | $spl2)),
  inference(condensation,[],[f1953])).
fof(f1953,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X0))) ) | ($spl0 | $spl2)),
  inference(condensation,[],[f1946])).
fof(f1946,plain,(
  ( ! [X47,X50,X49] : (s(bool,'T_0') = s(bool,X50) | p(s(bool,X49)) | p(s(bool,X50)) | s(bool,'T_0') = s(bool,X47)) ) | ($spl0 | $spl2)),
  inference(backward_demodulation,[],[f1903,f1046])).
fof(f1944,plain,(
  ( ! [X10,X11] : (s('type.realax.real',sK0) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',X10),s('type.realax.real',X11))) | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X11))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X11)))))) ) | ($spl0 | $spl2)),
  inference(backward_demodulation,[],[f1903,f1375])).
fof(f1968,plain,(
  ( ! [X7] : (s(bool,'F_0') = s(bool,X7)) ) | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f1942,f1903])).
fof(f1942,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X6),s(bool,X7))) | s(bool,'F_0') = s(bool,X7)) ) | ($spl0 | $spl2)),
  inference(backward_demodulation,[],[f1903,f1190])).
fof(f1966,plain,(
  ( ! [X6] : (s(bool,'F_0') = s(bool,X6)) ) | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f1927,f1903])).
fof(f1927,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X6),s(bool,X7))) | s(bool,'F_0') = s(bool,X6)) ) | ($spl0 | $spl2)),
  inference(backward_demodulation,[],[f1903,f940])).
fof(f1961,plain,(
  ( ! [X3] : (s(bool,'F_0') = s(bool,X3)) ) | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f1913,f1903])).
fof(f1913,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'const.marker.unint_1'(s(bool,X3))) | s(bool,'F_0') = s(bool,X3)) ) | ($spl0 | $spl2)),
  inference(backward_demodulation,[],[f1903,f305])).
fof(f1906,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) ) | ($spl0 | $spl2)),
  inference(backward_demodulation,[],[f1903,f118])).
fof(f1957,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f1904,f1903])).
fof(f1904,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl0 | $spl2)),
  inference(backward_demodulation,[],[f1903,f108])).
fof(f1903,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f1902,f117])).
fof(f1902,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))) != s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0)))))))) ) | ($spl0 | $spl2)),
  inference(condensation,[],[f1901])).
fof(f1901,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))) != s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f1882,f199])).
fof(f1882,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))) != s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))))) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | ($spl0 | $spl2)),
  inference(superposition,[],[f1761,f167])).
fof(f1854,plain,(
  ( ! [X10,X8,X9] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',X8))),s('type.realax.real',X9))),s('type.realax.real',X10))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',X8))),s('type.realax.real',X9))),s('type.realax.real',X10)))))) ) | $spl0),
  inference(superposition,[],[f126,f1469])).
fof(f1853,plain,(
  ( ! [X6,X7,X5] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X7),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',X5))),s('type.realax.real',X6))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X7),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',X5))),s('type.realax.real',X6)))))))) ) | $spl0),
  inference(superposition,[],[f127,f1469])).
fof(f1852,plain,(
  ( ! [X4,X3] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',X3))),s('type.realax.real',X4))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',X3))),s('type.realax.real',X4))))))) | s('type.realax.real',sK0) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',X3))),s('type.realax.real',X4)))) ) | $spl0),
  inference(superposition,[],[f1370,f1469])).
fof(f1850,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',X0))),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',X0))),s('type.realax.real',X1)))))) ) | $spl0),
  inference(superposition,[],[f1469,f125])).
fof(f1795,plain,(
  ( ! [X6,X8,X7] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',X7))))),s('type.realax.real',X8))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',X7))))),s('type.realax.real',X8)))))) ) | $spl0),
  inference(superposition,[],[f126,f1468])).
fof(f1794,plain,(
  ( ! [X4,X5,X3] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',X4))))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X5),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',X4)))))))))) ) | $spl0),
  inference(superposition,[],[f127,f1468])).
fof(f1793,plain,(
  ( ! [X2,X1] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',X2))))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',X2))))))))) | s('type.realax.real',sK0) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',X2)))))) ) | $spl0),
  inference(superposition,[],[f1370,f1468])).
fof(f1761,plain,(
  ( ! [X2] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))) != s('type.realax.real','const.bool.COND_3'(s(bool,X2),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))))) | ~p(s(bool,X2))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f1367,f1369])).
fof(f1367,plain,(
  ( ! [X2] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK0))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X2),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))))) | ~p(s(bool,X2))) ) | $spl2),
  inference(forward_demodulation,[],[f1366,f114])).
fof(f1366,plain,(
  ( ! [X2] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK0))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X2),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))))) | ~p(s(bool,X2))) ) | $spl2),
  inference(forward_demodulation,[],[f142,f406])).
fof(f1662,plain,(
  ( ! [X4,X3] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK0),s('type.realax.real',X3))),s('type.realax.real',X4))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK0),s('type.realax.real',X3))),s('type.realax.real',X4)))))) ) | $spl0),
  inference(superposition,[],[f126,f1466])).
fof(f1661,plain,(
  ( ! [X2,X1] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK0),s('type.realax.real',X1))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK0),s('type.realax.real',X1)))))))) ) | $spl0),
  inference(superposition,[],[f127,f1466])).
fof(f1660,plain,(
  ( ! [X0] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK0),s('type.realax.real',X0))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK0),s('type.realax.real',X0))))))) | s('type.realax.real',sK0) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK0),s('type.realax.real',X0)))) ) | $spl0),
  inference(superposition,[],[f1370,f1466])).
fof(f1527,plain,(
  ( ! [X4,X3] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real',sK0))),s('type.realax.real',X4))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real',sK0))),s('type.realax.real',X4)))))) ) | $spl0),
  inference(superposition,[],[f126,f1372])).
fof(f1526,plain,(
  ( ! [X2,X1] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',sK0))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',sK0)))))))) ) | $spl0),
  inference(superposition,[],[f127,f1372])).
fof(f1525,plain,(
  ( ! [X0] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',sK0))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',sK0))))))) | s('type.realax.real',sK0) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',sK0)))) ) | $spl0),
  inference(superposition,[],[f1370,f1372])).
fof(f1508,plain,(
  ( ! [X0] : (s('type.realax.real',sK0) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',X0))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',X0))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',X0)))))))) ) | $spl0),
  inference(forward_demodulation,[],[f1507,f1371])).
fof(f1507,plain,(
  ( ! [X0] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',X0))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',X0))))))) | s('type.realax.real',sK0) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',X0)))))) ) | $spl0),
  inference(forward_demodulation,[],[f1484,f1371])).
fof(f1484,plain,(
  ( ! [X0] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',X0))))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',X0))))))) | s('type.realax.real',sK0) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',X0)))))) ) | $spl0),
  inference(superposition,[],[f1370,f1371])).
fof(f1469,plain,(
  ( ! [X2,X3] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',X2))),s('type.realax.real',X3))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',X2))),s('type.realax.real',X3)))))) ) | $spl0),
  inference(superposition,[],[f126,f1371])).
fof(f1468,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',X0))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',X0)))))))) ) | $spl0),
  inference(superposition,[],[f127,f1371])).
fof(f1466,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK0),s('type.realax.real',X0))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK0),s('type.realax.real',X0)))))) ) | $spl0),
  inference(superposition,[],[f1371,f125])).
fof(f1377,plain,(
  ( ! [X0,X1] : (s('type.realax.real',sK0) = s('type.realax.real',X0) | s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0)))))))) ) | $spl0),
  inference(backward_demodulation,[],[f140,f1209])).
fof(f1376,plain,(
  ( ! [X14,X12,X13] : (s('type.realax.real',sK0) = s('type.realax.real','const.bool.COND_3'(s(bool,X12),s('type.realax.real',X13),s('type.realax.real',X14))) | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X14))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X14))))) | s(bool,'T_0') = s(bool,X12)) ) | $spl0),
  inference(backward_demodulation,[],[f140,f1207])).
fof(f1375,plain,(
  ( ! [X10,X11] : (s('type.realax.real',sK0) = s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real',X10),s('type.realax.real',X11))) | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X11))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X11)))))) ) | $spl0),
  inference(backward_demodulation,[],[f140,f1206])).
fof(f1374,plain,(
  ( ! [X8,X9] : (s('type.realax.real',sK0) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',X8),s('type.realax.real',X9))) | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X8))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X8)))))) ) | $spl0),
  inference(backward_demodulation,[],[f140,f1205])).
fof(f1383,plain,(
  ( ! [X6] : (s('type.realax.real',sK0) = s('type.realax.real',X6) | s('type.realax.real',sK0) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X6))))) | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X6))))))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X6)))))))))) ) | $spl0),
  inference(forward_demodulation,[],[f1373,f140])).
fof(f1373,plain,(
  ( ! [X6] : (s('type.realax.real',sK0) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X6))))) | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X6))))))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X6))))))))) | s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl0),
  inference(backward_demodulation,[],[f140,f1203])).
fof(f1372,plain,(
  ( ! [X0] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',sK0))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',sK0)))))) ) | $spl0),
  inference(backward_demodulation,[],[f140,f962])).
fof(f1371,plain,(
  ( ! [X0] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',X0))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK0),s('type.realax.real',X0)))))) ) | $spl0),
  inference(backward_demodulation,[],[f140,f390])).
fof(f1370,plain,(
  ( ! [X0] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))) | s('type.realax.real',sK0) = s('type.realax.real',X0)) ) | $spl0),
  inference(backward_demodulation,[],[f140,f103])).
fof(f1369,plain,(
  s('type.realax.real',sK0) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK0))) | $spl0),
  inference(backward_demodulation,[],[f140,f94])).
fof(f108,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f41])).
fof(f41,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X11] : (s(bool,'F_0') = s(bool,X11) <=> ~p(s(bool,X11)))),
  file('thm.real.neg_rat|split|1','thm.bool.EQ_CLAUSES|split|2')).
fof(f109,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f82])).
fof(f118,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X12] : ! [X9,X10] : s(X12,X10) = s(X12,'const.bool.COND_3'(s(bool,'F_0'),s(X12,X9),s(X12,X10)))),
  file('thm.real.neg_rat|split|1','thm.bool.COND_CLAUSES|split|1')).
fof(f144,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f96,f92])).
fof(f96,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f31])).
fof(f31,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f13])).
fof(f13,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.real.neg_rat|split|1','thm.bool.NOT_CLAUSES|split|2')).
fof(f169,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'const.marker.unint_1'(s(bool,X2))) | s(bool,'F_0') = s(bool,X2)) )),
  inference(superposition,[],[f101,f114])).
fof(f224,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X3),s(bool,X4))) | s(bool,'F_0') = s(bool,X3)) )),
  inference(superposition,[],[f117,f101])).
fof(f259,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,X4))) | s(bool,'F_0') = s(bool,X4)) )),
  inference(superposition,[],[f118,f101])).
fof(f303,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') = s(bool,X4) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X4),s(bool,X5)))) )),
  inference(forward_demodulation,[],[f285,f114])).
fof(f285,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') = s(bool,'const.marker.unint_1'(s(bool,X4))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X4),s(bool,X5)))) )),
  inference(superposition,[],[f169,f117])).
fof(f304,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,X7) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X6),s(bool,X7)))) )),
  inference(forward_demodulation,[],[f286,f114])).
fof(f286,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,'const.marker.unint_1'(s(bool,X7))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X6),s(bool,X7)))) )),
  inference(superposition,[],[f169,f118])).
fof(f305,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,'const.marker.unint_1'(s(bool,X3))) | s(bool,'F_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f292,f92])).
fof(f292,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.marker.unint_1'(s(bool,X3))) | s(bool,'F_0') = s(bool,X3)) )),
  inference(superposition,[],[f108,f169])).
fof(f307,plain,(
  ( ! [X10,X9] : (~p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,'F_0') = s(bool,X9)) )),
  inference(forward_demodulation,[],[f306,f114])).
fof(f306,plain,(
  ( ! [X10,X9] : (~p(s(bool,X10)) | s(bool,X10) = s(bool,'const.marker.unint_1'(s(bool,X9))) | s(bool,'F_0') = s(bool,X9)) )),
  inference(subsumption_resolution,[],[f296,f92])).
fof(f296,plain,(
  ( ! [X10,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X10)) | s(bool,X10) = s(bool,'const.marker.unint_1'(s(bool,X9))) | s(bool,'F_0') = s(bool,X9)) )),
  inference(superposition,[],[f131,f169])).
fof(f340,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X4) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X4),s(bool,X5)))) )),
  inference(forward_demodulation,[],[f330,f114])).
fof(f330,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,'const.marker.unint_1'(s(bool,X4))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X4),s(bool,X5)))) )),
  inference(superposition,[],[f305,f117])).
fof(f341,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X7) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X6),s(bool,X7)))) )),
  inference(forward_demodulation,[],[f331,f114])).
fof(f331,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,'const.marker.unint_1'(s(bool,X7))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X6),s(bool,X7)))) )),
  inference(superposition,[],[f305,f118])).
fof(f380,plain,(
  ( ! [X4,X5] : (s(bool,X5) = s(bool,'const.marker.unint_1'(s(bool,X4))) | s(bool,'F_0') = s(bool,X5) | s(bool,'F_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f375,f92])).
fof(f375,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | s(bool,X5) = s(bool,'const.marker.unint_1'(s(bool,X4))) | s(bool,'F_0') = s(bool,X5) | s(bool,'F_0') = s(bool,X4)) )),
  inference(superposition,[],[f307,f169])).
fof(f427,plain,(
  ( ! [X28,X29,X27] : (s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X27),s(bool,X28)))) )),
  inference(superposition,[],[f167,f118])).
fof(f452,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') = s(bool,X21) | s(bool,X22) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X20),s(bool,X21))) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f167,f118])).
fof(f766,plain,(
  ( ! [X30,X31] : (s(bool,'F_0') != s(bool,'const.marker.unint_1'(s(bool,X30))) | s(bool,X30) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(subsumption_resolution,[],[f533,f147])).
fof(f533,plain,(
  ( ! [X30,X31] : (s(bool,'F_0') != s(bool,'const.marker.unint_1'(s(bool,X30))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X30) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f305,f167])).
fof(f834,plain,(
  ( ! [X21,X20] : (s(bool,'F_0') != s(bool,X20) | s(bool,X21) = s(bool,'const.marker.unint_1'(s(bool,X20))) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f770,f114])).
fof(f835,plain,(
  ( ! [X24,X23,X22] : (s(bool,'F_0') != s(bool,X22) | s(bool,X24) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X22),s(bool,X23))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f770,f117])).
fof(f836,plain,(
  ( ! [X26,X27,X25] : (s(bool,'F_0') != s(bool,X26) | s(bool,X27) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X25),s(bool,X26))) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f770,f118])).
fof(f854,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2))) | s(bool,'F_0') = s(bool,X1) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f224,f167])).
fof(f871,plain,(
  ( ! [X30,X31,X29] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X29),s(bool,X31))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X29),s(bool,X30)))) )),
  inference(superposition,[],[f224,f117])).
fof(f872,plain,(
  ( ! [X33,X34,X32] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X33),s(bool,X34))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X32),s(bool,X33)))) )),
  inference(superposition,[],[f224,f118])).
fof(f940,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X6),s(bool,X7))) | s(bool,'F_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f893,f92])).
fof(f893,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X6),s(bool,X7))) | s(bool,'F_0') = s(bool,X6)) )),
  inference(superposition,[],[f108,f224])).
fof(f945,plain,(
  ( ! [X37,X38,X36] : (s(bool,X38) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X36),s(bool,X37))) | s(bool,'F_0') = s(bool,X38) | s(bool,'F_0') = s(bool,X36)) )),
  inference(subsumption_resolution,[],[f903,f92])).
fof(f903,plain,(
  ( ! [X37,X38,X36] : (~p(s(bool,'T_0')) | s(bool,X38) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X36),s(bool,X37))) | s(bool,'F_0') = s(bool,X38) | s(bool,'F_0') = s(bool,X36)) )),
  inference(superposition,[],[f307,f224])).
fof(f989,plain,(
  ( ! [X57,X54,X58,X56,X55] : (s(X56,X57) = s(X56,'const.bool.COND_3'(s(bool,X55),s(X56,X58),s(X56,X57))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X54),s(bool,X55)))) )),
  inference(superposition,[],[f245,f118])).
fof(f1033,plain,(
  ( ! [X30,X28,X29] : (s(bool,'F_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X28) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X28),s(bool,X29),s(bool,X30)))) )),
  inference(superposition,[],[f245,f101])).
fof(f1044,plain,(
  ( ! [X43,X41,X42] : (~p(s(bool,X43)) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,X41),s(bool,X42),s(bool,X43))) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f108,f245])).
fof(f1045,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,X46)) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X44),s(bool,X45),s(bool,X46))) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f109,f245])).
fof(f1055,plain,(
  ( ! [X85,X83,X84,X82] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X84),s(bool,X85))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X82),s(bool,X83),s(bool,X84))) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f224,f245])).
fof(f1058,plain,(
  ( ! [X94,X95,X93,X96] : (~p(s(bool,X95)) | s(bool,X96) = s(bool,'const.bool.COND_3'(s(bool,X93),s(bool,X94),s(bool,X95))) | s(bool,'F_0') = s(bool,X96) | s(bool,'T_0') = s(bool,X93)) )),
  inference(superposition,[],[f307,f245])).
fof(f1060,plain,(
  ( ! [X103,X101,X102,X100] : (s(bool,'F_0') != s(bool,X102) | s(bool,X103) = s(bool,'const.bool.COND_3'(s(bool,X100),s(bool,X101),s(bool,X102))) | s(bool,'T_0') = s(bool,X103) | s(bool,'T_0') = s(bool,X100)) )),
  inference(superposition,[],[f770,f245])).
fof(f1091,plain,(
  ( ! [X76,X77,X75] : (s(bool,'T_0') = s(bool,X77) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X75),s(bool,X76),s(bool,X77))) | s(bool,'T_0') = s(bool,X75)) )),
  inference(forward_demodulation,[],[f1053,f114])).
fof(f1053,plain,(
  ( ! [X76,X77,X75] : (s(bool,'T_0') = s(bool,'const.marker.unint_1'(s(bool,X77))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X75),s(bool,X76),s(bool,X77))) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f169,f245])).
fof(f1092,plain,(
  ( ! [X92,X90,X91] : (s(bool,'F_0') != s(bool,X92) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X90),s(bool,X91),s(bool,X92))) | s(bool,'T_0') = s(bool,X90)) )),
  inference(forward_demodulation,[],[f1057,f114])).
fof(f1057,plain,(
  ( ! [X92,X90,X91] : (s(bool,'F_0') != s(bool,'const.marker.unint_1'(s(bool,X92))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X90),s(bool,X91),s(bool,X92))) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f305,f245])).
fof(f1105,plain,(
  ( ! [X14,X15,X16] : (s(bool,X14) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X14),s(bool,X15))) | s(bool,'F_0') = s(bool,X15) | s(bool,X14) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f259,f167])).
fof(f1127,plain,(
  ( ! [X33,X34,X32] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X34),s(bool,X32))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X32),s(bool,X33)))) )),
  inference(superposition,[],[f259,f117])).
fof(f1128,plain,(
  ( ! [X37,X35,X36] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X37),s(bool,X36))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X35),s(bool,X36)))) )),
  inference(superposition,[],[f259,f118])).
fof(f1129,plain,(
  ( ! [X39,X41,X38,X40] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X41),s(bool,X40))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X38),s(bool,X39),s(bool,X40))) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f259,f245])).
fof(f1190,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X6),s(bool,X7))) | s(bool,'F_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f1138,f92])).
fof(f1138,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X6),s(bool,X7))) | s(bool,'F_0') = s(bool,X7)) )),
  inference(superposition,[],[f108,f259])).
fof(f1195,plain,(
  ( ! [X47,X48,X49] : (s(bool,X49) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X47),s(bool,X48))) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X48)) )),
  inference(subsumption_resolution,[],[f1151,f92])).
fof(f1151,plain,(
  ( ! [X47,X48,X49] : (~p(s(bool,'T_0')) | s(bool,X49) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X47),s(bool,X48))) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X48)) )),
  inference(superposition,[],[f307,f259])).
fof(f1206,plain,(
  ( ! [X10,X11] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X11))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X11))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real',X10),s('type.realax.real',X11)))) )),
  inference(superposition,[],[f103,f118])).
fof(f1042,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X37)) | s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,X35),s(bool,X36),s(bool,X37))) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f104,f245])).
fof(f1043,plain,(
  ( ! [X39,X38,X40] : (~p(s(bool,X40)) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X38),s(bool,X39),s(bool,X40))) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f105,f245])).
fof(f1046,plain,(
  ( ! [X47,X50,X48,X49] : (p(s(bool,X49)) | p(s(bool,X50)) | s(bool,X50) = s(bool,'const.bool.COND_3'(s(bool,X47),s(bool,X48),s(bool,X49))) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f128,f245])).
fof(f1048,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X58) = s(bool,'const.bool.COND_3'(s(bool,X55),s(bool,X56),s(bool,X57))) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f131,f245])).
fof(f1054,plain,(
  ( ! [X80,X78,X81,X79] : (p(s(bool,X80)) | s(bool,X81) = s(bool,'const.bool.COND_3'(s(bool,X78),s(bool,X79),s(bool,X80))) | s(bool,'T_0') = s(bool,X81) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f199,f245])).
fof(f1209,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))))))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f125,f103])).
fof(f1207,plain,(
  ( ! [X14,X12,X13] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X14))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X14))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.bool.COND_3'(s(bool,X12),s('type.realax.real',X13),s('type.realax.real',X14))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f103,f245])).
fof(f1205,plain,(
  ( ! [X8,X9] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X8))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X8))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',X8),s('type.realax.real',X9)))) )),
  inference(superposition,[],[f103,f117])).
fof(f1203,plain,(
  ( ! [X6] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X6))))))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X6))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X6))))) | s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f103,f103])).
fof(f1031,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X22) = s(bool,X23) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X20),s(bool,X21),s(bool,X22))) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f245,f167])).
fof(f1029,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,X12) | s(bool,X15) = s(bool,'const.bool.COND_3'(s(bool,X12),s(bool,X13),s(bool,X14))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f245,f167])).
fof(f990,plain,(
  ( ! [X61,X59,X64,X62,X60,X63] : (s(X62,X63) = s(X62,'const.bool.COND_3'(s(bool,X61),s(X62,X64),s(X62,X63))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X59),s(bool,X60),s(bool,X61))) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f245,f245])).
fof(f988,plain,(
  ( ! [X52,X50,X53,X51,X49] : (s(X51,X52) = s(X51,'const.bool.COND_3'(s(bool,X49),s(X51,X53),s(X51,X52))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X49),s(bool,X50)))) )),
  inference(superposition,[],[f245,f117])).
fof(f962,plain,(
  ( ! [X0] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f127,f94])).
fof(f821,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f770,f101])).
fof(f778,plain,(
  ( ! [X28,X26,X27] : (s(bool,X26) = s(bool,X27) | s(bool,'F_0') = s(bool,X27) | s(bool,X26) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f551,f92])).
fof(f551,plain,(
  ( ! [X28,X26,X27] : (~p(s(bool,'T_0')) | s(bool,X26) = s(bool,X27) | s(bool,'F_0') = s(bool,X27) | s(bool,X26) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f307,f167])).
fof(f771,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f547,f92])).
fof(f547,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f131,f167])).
fof(f770,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f543,f92])).
fof(f543,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f108,f167])).
fof(f522,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.bool.COND_3'(s(bool,X0),s(X1,X2),s(X1,X3))) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f117,f167])).
fof(f451,plain,(
  ( ! [X19,X17,X18] : (s(bool,'T_0') = s(bool,X17) | s(bool,X19) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X17),s(bool,X18))) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f167,f117])).
fof(f450,plain,(
  ( ! [X15,X16] : (s(bool,'T_0') = s(bool,X15) | s(bool,X16) = s(bool,'const.marker.unint_1'(s(bool,X15))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f167,f114])).
fof(f440,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f167,f167])).
fof(f426,plain,(
  ( ! [X26,X24,X25] : (s(bool,X24) = s(bool,X26) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X24),s(bool,X25)))) )),
  inference(superposition,[],[f167,f117])).
fof(f425,plain,(
  ( ! [X23,X22] : (s(bool,X22) = s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,'const.marker.unint_1'(s(bool,X22)))) )),
  inference(superposition,[],[f167,f114])).
fof(f406,plain,(
  ( ! [X2,X3] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',X3))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))),s('type.realax.real',X3)))) )),
  inference(forward_demodulation,[],[f397,f125])).
fof(f397,plain,(
  ( ! [X2,X3] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))),s('type.realax.real',X3))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3)))))))) )),
  inference(superposition,[],[f126,f125])).
fof(f390,plain,(
  ( ! [X0] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))))) )),
  inference(superposition,[],[f126,f94])).
fof(f245,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.bool.COND_3'(s(bool,X0),s(X1,X3),s(X1,X2))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f118,f101])).
fof(f199,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f194,f144])).
fof(f194,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f128,f101])).
fof(f167,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f101,f101])).
fof(f147,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f144,f104])).
fof(f132,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,X1),X2) = s(fun(X0,X1),X3) | s(X1,f(s(X0,sK2(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK2(X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK2(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK2(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f73])).
fof(f73,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f63])).
fof(f63,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.real.neg_rat|split|1',aHL_EXT)).
fof(f128,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X9,X10] : ((p(s(bool,X9)) => p(s(bool,X10))) => ((p(s(bool,X10)) => p(s(bool,X9))) => s(bool,X9) = s(bool,X10)))),
  file('thm.real.neg_rat|split|1','thm.bool.IMP_ANTISYM_AX')).
fof(f131,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f72])).
fof(f127,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))))) )),
  inference(cnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1] : s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X7,X6] : s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X7),s('type.realax.real',X6))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X7),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X6)))))),
  file('thm.real.neg_rat|split|1','thm.real.REAL_MUL_RNEG')).
fof(f126,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X1)))) )),
  inference(cnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0,X1] : s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X1)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X7,X6] : s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X7),s('type.realax.real',X6))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X7))),s('type.realax.real',X6)))),
  file('thm.real.neg_rat|split|1','thm.real.REAL_MUL_LNEG')).
fof(f125,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1)))))) )),
  inference(cnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1] : s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1)))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X7,X6] : s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X7),s('type.realax.real',X6))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X7),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X6)))))),
  file('thm.real.neg_rat|split|1','thm.real.real_div')).
fof(f122,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))) | ~p(s(bool,X2)) | s(X0,X3) != s(X0,X4) | s(bool,X1) != s(bool,X2)) )),
  inference(cnf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : (s(bool,X1) != s(bool,X2) | (p(s(bool,X2)) & s(X0,X3) != s(X0,X4)) | (~p(s(bool,X2)) & s(X0,X5) != s(X0,X6)) | s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(flattening,[],[f69])).
fof(f69,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s(bool,X1) != s(bool,X2) | (p(s(bool,X2)) & s(X0,X3) != s(X0,X4)) | (~p(s(bool,X2)) & s(X0,X5) != s(X0,X6))) | s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(ennf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s(bool,X1) = s(bool,X2) & (p(s(bool,X2)) => s(X0,X3) = s(X0,X4)) & (~p(s(bool,X2)) => s(X0,X5) = s(X0,X6))) => s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(flattening,[],[f57])).
fof(f57,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6] : ((s(bool,X1) = s(bool,X2) & (p(s(bool,X2)) => s(X0,X3) = s(X0,X4)) & (~p(s(bool,X2)) => s(X0,X5) = s(X0,X6))) => s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X12] : ! [X14,X15,X7,X16,X6,X17] : ((s(bool,X14) = s(bool,X15) & (p(s(bool,X15)) => s(X12,X7) = s(X12,X16)) & (~p(s(bool,X15)) => s(X12,X6) = s(X12,X17))) => s(X12,'const.bool.COND_3'(s(bool,X14),s(X12,X7),s(X12,X6))) = s(X12,'const.bool.COND_3'(s(bool,X15),s(X12,X16),s(X12,X17))))),
  file('thm.real.neg_rat|split|1','thm.bool.COND_CONG')).
fof(f123,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))) | s(X0,X5) != s(X0,X6) | p(s(bool,X2)) | s(bool,X1) != s(bool,X2)) )),
  inference(cnf_transformation,[],[f70])).
fof(f124,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))) | s(X0,X5) != s(X0,X6) | s(X0,X3) != s(X0,X4) | s(bool,X1) != s(bool,X2)) )),
  inference(cnf_transformation,[],[f70])).
fof(f117,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f51])).
fof(f51,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X12] : ! [X9,X10] : s(X12,X9) = s(X12,'const.bool.COND_3'(s(bool,'T_0'),s(X12,X9),s(X12,X10)))),
  file('thm.real.neg_rat|split|1','thm.bool.COND_CLAUSES|split|0')).
fof(f114,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,'const.marker.unint_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,'const.marker.unint_1'(s(X0,X1)))),
  inference(flattening,[],[f47])).
fof(f47,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,'const.marker.unint_1'(s(X0,X1)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X12] : ! [X7] : s(X12,X7) = s(X12,'const.marker.unint_1'(s(X12,X7)))),
  file('thm.real.neg_rat|split|1','thm.marker.unint_def')).
fof(f104,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.real.neg_rat|split|1','thm.bool.EQ_CLAUSES|split|0')).
fof(f105,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f80])).
fof(f103,plain,(
  ( ! [X0] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))))),
  inference(ennf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0] : (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ! [X0] : (~s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X7] : (~s('type.realax.real',X7) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X7))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X7))))))),
  file('thm.real.neg_rat|split|1','thm.real.REAL_NEG_INV')).
fof(f101,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.real.neg_rat|split|1',aHL_BOOL_CASES)).
fof(f94,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(cnf_transformation,[],[f26])).
fof(f26,axiom,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.real.neg_rat|split|1','thm.real.REAL_NEG_|48|')).
fof(f92,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.real.neg_rat|split|1',aHL_TRUTH)).
fof(f1982,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f1898,f92])).
fof(f1898,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl2)),
  inference(trivial_inequality_removal,[],[f1895])).
fof(f1895,plain,(
  s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))) != s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))) | ~p(s(bool,'T_0')) | ($spl0 | $spl2)),
  inference(superposition,[],[f1761,f117])).
fof(f1981,plain,(
  $false | ($spl0 | $spl2)),
  inference(global_subsumption,[],[f91,f90,f92,f94,f101,f103,f105,f104,f114,f117,f124,f123,f122,f125,f126,f127,f131,f128,f132,f147,f167,f199,f245,f390,f406,f425,f426,f440,f450,f451,f522,f770,f771,f778,f821,f962,f988,f990,f1029,f1031,f1203,f1205,f1207,f1209,f1054,f1048,f1046,f1043,f1042,f1206,f1195,f1190,f1129,f1128,f1127,f1105,f1092,f1091,f1060,f1058,f1055,f1045,f1044,f1033,f989,f945,f940,f872,f871,f854,f836,f835,f834,f766,f452,f427,f380,f341,f340,f307,f305,f304,f303,f259,f224,f169,f144,f118,f109,f108,f140,f1369,f1370,f1371,f1372,f1383,f1374,f1375,f1376,f1377,f1466,f1468,f1469,f1508,f1525,f1526,f1527,f1660,f1661,f1662,f1761,f1793,f1794,f1795,f1850,f1852,f1853,f1854,f1903,f1957,f1906,f1961,f1966,f1968,f1944,f1972,f1980])).
fof(f1973,plain,(
  $false | ($spl0 | $spl2)),
  inference(global_subsumption,[],[f91,f90,f92,f94,f101,f103,f105,f104,f114,f117,f124,f123,f122,f125,f126,f127,f131,f128,f132,f147,f167,f199,f245,f390,f406,f425,f426,f440,f450,f451,f522,f770,f771,f778,f821,f962,f988,f990,f1029,f1031,f1203,f1205,f1207,f1209,f1054,f1048,f1046,f1043,f1042,f1206,f1195,f1190,f1129,f1128,f1127,f1105,f1092,f1091,f1060,f1058,f1055,f1045,f1044,f1033,f989,f945,f940,f872,f871,f854,f836,f835,f834,f766,f452,f427,f380,f341,f340,f307,f305,f304,f303,f259,f224,f169,f144,f118,f109,f108,f140,f1369,f1370,f1371,f1372,f1383,f1374,f1375,f1376,f1377,f1466,f1468,f1469,f1508,f1525,f1526,f1527,f1660,f1661,f1662,f1761,f1793,f1794,f1795,f1850,f1852,f1853,f1854,f1903,f1957,f1906,f1961,f1966,f1968,f1944,f1972])).
fof(f1971,plain,(
  $false | ($spl0 | $spl2)),
  inference(global_subsumption,[],[f91,f90,f92,f94,f101,f103,f105,f104,f114,f117,f124,f123,f122,f125,f126,f127,f131,f128,f132,f147,f167,f199,f245,f390,f406,f425,f426,f440,f450,f451,f522,f770,f771,f778,f821,f962,f988,f990,f1029,f1031,f1203,f1205,f1207,f1209,f1054,f1048,f1046,f1043,f1042,f1206,f1195,f1190,f1129,f1128,f1127,f1105,f1092,f1091,f1060,f1058,f1055,f1045,f1044,f1033,f989,f945,f940,f872,f871,f854,f836,f835,f834,f766,f452,f427,f380,f341,f340,f307,f305,f304,f303,f259,f224,f169,f144,f118,f109,f108,f140,f1369,f1370,f1371,f1372,f1383,f1374,f1375,f1376,f1377,f1466,f1468,f1469,f1508,f1525,f1526,f1527,f1660,f1661,f1662,f1761,f1793,f1794,f1795,f1850,f1852,f1853,f1854,f1903,f1957,f1906,f1961,f1966,f1968,f1944])).
fof(f1969,plain,(
  $false | ($spl0 | $spl2)),
  inference(global_subsumption,[],[f91,f90,f92,f94,f101,f103,f105,f104,f114,f117,f124,f123,f122,f125,f126,f127,f131,f128,f132,f147,f167,f199,f245,f390,f406,f425,f426,f440,f450,f451,f522,f770,f771,f778,f821,f962,f988,f990,f1029,f1031,f1203,f1205,f1207,f1209,f1054,f1048,f1046,f1043,f1042,f1206,f1195,f1190,f1129,f1128,f1127,f1105,f1092,f1091,f1060,f1058,f1055,f1045,f1044,f1033,f989,f945,f940,f872,f871,f854,f836,f835,f834,f766,f452,f427,f380,f341,f340,f307,f305,f304,f303,f259,f224,f169,f144,f118,f109,f108,f140,f1369,f1370,f1371,f1372,f1383,f1374,f1375,f1376,f1377,f1466,f1468,f1469,f1508,f1525,f1526,f1527,f1660,f1661,f1662,f1761,f1793,f1794,f1795,f1850,f1852,f1853,f1854,f1903,f1957,f1906,f1961,f1966,f1968])).
fof(f1967,plain,(
  $false | ($spl0 | $spl2)),
  inference(global_subsumption,[],[f91,f90,f92,f94,f101,f103,f105,f104,f114,f117,f124,f123,f122,f125,f126,f127,f131,f128,f132,f147,f167,f199,f245,f390,f406,f425,f426,f440,f450,f451,f522,f770,f771,f778,f821,f962,f988,f990,f1029,f1031,f1203,f1205,f1207,f1209,f1054,f1048,f1046,f1043,f1042,f1206,f1195,f1190,f1129,f1128,f1127,f1105,f1092,f1091,f1060,f1058,f1055,f1045,f1044,f1033,f989,f945,f940,f872,f871,f854,f836,f835,f834,f766,f452,f427,f380,f341,f340,f307,f305,f304,f303,f259,f224,f169,f144,f118,f109,f108,f140,f1369,f1370,f1371,f1372,f1383,f1374,f1375,f1376,f1377,f1466,f1468,f1469,f1508,f1525,f1526,f1527,f1660,f1661,f1662,f1761,f1793,f1794,f1795,f1850,f1852,f1853,f1854,f1903,f1957,f1906,f1961,f1966])).
fof(f1962,plain,(
  $false | ($spl0 | $spl2)),
  inference(global_subsumption,[],[f91,f90,f92,f94,f101,f103,f105,f104,f114,f117,f124,f123,f122,f125,f126,f127,f131,f128,f132,f147,f167,f199,f245,f390,f406,f425,f426,f440,f450,f451,f522,f770,f771,f778,f821,f962,f988,f990,f1029,f1031,f1203,f1205,f1207,f1209,f1054,f1048,f1046,f1043,f1042,f1206,f1195,f1190,f1129,f1128,f1127,f1105,f1092,f1091,f1060,f1058,f1055,f1045,f1044,f1033,f989,f945,f940,f872,f871,f854,f836,f835,f834,f766,f452,f427,f380,f341,f340,f307,f305,f304,f303,f259,f224,f169,f144,f118,f109,f108,f140,f1369,f1370,f1371,f1372,f1383,f1374,f1375,f1376,f1377,f1466,f1468,f1469,f1508,f1525,f1526,f1527,f1660,f1661,f1662,f1761,f1793,f1794,f1795,f1850,f1852,f1853,f1854,f1903,f1957,f1906,f1961])).
fof(f1960,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f1907,f92])).
fof(f1907,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl2)),
  inference(backward_demodulation,[],[f1903,f144])).
fof(f1959,plain,(
  $false | ($spl0 | $spl2)),
  inference(global_subsumption,[],[f91,f90,f92,f94,f101,f103,f105,f104,f114,f117,f124,f123,f122,f125,f126,f127,f131,f128,f132,f147,f167,f199,f245,f390,f406,f425,f426,f440,f450,f451,f522,f770,f771,f778,f821,f962,f988,f990,f1029,f1031,f1203,f1205,f1207,f1209,f1054,f1048,f1046,f1043,f1042,f1206,f1195,f1190,f1129,f1128,f1127,f1105,f1092,f1091,f1060,f1058,f1055,f1045,f1044,f1033,f989,f945,f940,f872,f871,f854,f836,f835,f834,f766,f452,f427,f380,f341,f340,f307,f305,f304,f303,f259,f224,f169,f144,f118,f109,f108,f140,f1369,f1370,f1371,f1372,f1383,f1374,f1375,f1376,f1377,f1466,f1468,f1469,f1508,f1525,f1526,f1527,f1660,f1661,f1662,f1761,f1793,f1794,f1795,f1850,f1852,f1853,f1854,f1903,f1957,f1906])).
fof(f1958,plain,(
  $false | ($spl0 | $spl2)),
  inference(global_subsumption,[],[f91,f90,f92,f94,f101,f103,f105,f104,f114,f117,f124,f123,f122,f125,f126,f127,f131,f128,f132,f147,f167,f199,f245,f390,f406,f425,f426,f440,f450,f451,f522,f770,f771,f778,f821,f962,f988,f990,f1029,f1031,f1203,f1205,f1207,f1209,f1054,f1048,f1046,f1043,f1042,f1206,f1195,f1190,f1129,f1128,f1127,f1105,f1092,f1091,f1060,f1058,f1055,f1045,f1044,f1033,f989,f945,f940,f872,f871,f854,f836,f835,f834,f766,f452,f427,f380,f341,f340,f307,f305,f304,f303,f259,f224,f169,f144,f118,f109,f108,f140,f1369,f1370,f1371,f1372,f1383,f1374,f1375,f1376,f1377,f1466,f1468,f1469,f1508,f1525,f1526,f1527,f1660,f1661,f1662,f1761,f1793,f1794,f1795,f1850,f1852,f1853,f1854,f1903,f1957])).
fof(f1900,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f1879,f117])).
fof(f1879,plain,(
  s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))) != s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))))) | ($spl0 | $spl2)),
  inference(unit_resulting_resolution,[],[f92,f1761])).
fof(f1899,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f1880,f92])).
fof(f1880,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl2)),
  inference(unit_resulting_resolution,[],[f117,f1761])).
fof(f1881,plain,(
  $false | ($spl0 | $spl2)),
  inference(unit_resulting_resolution,[],[f92,f117,f1761])).
fof(f1739,plain,(
  $false | ($spl0 | $spl2 | $spl4)),
  inference(unit_resulting_resolution,[],[f117,f1378])).
fof(f1378,plain,(
  ( ! [X2] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))) != s('type.realax.real','const.bool.COND_3'(s(bool,X2),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0)))))))) ) | ($spl0 | $spl2 | $spl4)),
  inference(backward_demodulation,[],[f1369,f1368])).
fof(f1368,plain,(
  ( ! [X2] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK0))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X2),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0)))))))) ) | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f1367,f410])).
fof(f410,plain,(
  ( ! [X2] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK0))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X2),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))))) | p(s(bool,X2))) ) | $spl4),
  inference(backward_demodulation,[],[f406,f145])).
fof(f145,plain,(
  ( ! [X2] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK0))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X2),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1))),s('type.realax.real',sK0))))) | p(s(bool,X2))) ) | $spl4),
  inference(forward_demodulation,[],[f143,f114])).
fof(f1756,plain,(
  $false | ($spl0 | $spl2 | $spl4)),
  inference(trivial_inequality_removal,[],[f1753])).
fof(f1753,plain,(
  s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))) != s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))) | ($spl0 | $spl2 | $spl4)),
  inference(superposition,[],[f1378,f117])).
fof(f1365,plain,(
  $false | ($spl1 | $spl4)),
  inference(global_subsumption,[],[f91,f90,f92,f94,f144,f101,f103,f105,f104,f109,f108,f114,f117,f118,f124,f123,f122,f125,f126,f127,f131,f128,f132,f145,f141,f147,f167,f169,f199,f224,f245,f259,f303,f304,f305,f307,f308,f317,f323,f340,f341,f380,f390,f406,f407,f408,f410,f425,f426,f427,f440,f450,f451,f452,f522,f766,f770,f771,f778,f821,f834,f835,f836,f854,f871,f872,f940,f945,f962,f988,f989,f990,f1029,f1031,f1033,f1042,f1043,f1044,f1045,f1046,f1048,f1091,f1054,f1055,f1092,f1058,f1059,f1060,f1105,f1127,f1128,f1129,f1190,f1195,f1198,f1203,f1205,f1206,f1207,f1209,f1227,f1272,f1273,f1274,f1277,f1332,f1280,f1336,f1341,f1343,f1318,f1347,f1276,f1351,f1355,f1361,f1364])).
fof(f1364,plain,(
  ( ! [X23,X22] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X22),s('type.realax.real',sK0))))) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',X23),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X22),s('type.realax.real',sK0)))))))) ) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f1363,f1277])).
fof(f1363,plain,(
  ( ! [X23,X22] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X22),s('type.realax.real',sK0))))) = s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real',X23),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X22),s('type.realax.real',sK0)))))))) ) | $spl1),
  inference(forward_demodulation,[],[f1269,f1272])).
fof(f1269,plain,(
  ( ! [X23,X22] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X22),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK0))))) = s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real',X23),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X22),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK0)))))))) ) | $spl1),
  inference(superposition,[],[f118,f1227])).
fof(f1361,plain,(
  ( ! [X21,X20] : (s('type.realax.real',X21) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X20),s('type.realax.real',sK0))))),s('type.realax.real',X21)))) ) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f1360,f1277])).
fof(f1360,plain,(
  ( ! [X21,X20] : (s('type.realax.real',X21) = s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X20),s('type.realax.real',sK0))))),s('type.realax.real',X21)))) ) | $spl1),
  inference(forward_demodulation,[],[f1268,f1272])).
fof(f1268,plain,(
  ( ! [X21,X20] : (s('type.realax.real',X21) = s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X20),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK0))))),s('type.realax.real',X21)))) ) | $spl1),
  inference(superposition,[],[f118,f1227])).
fof(f1355,plain,(
  ( ! [X1] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',sK0))))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK0)))))) ) | $spl1),
  inference(forward_demodulation,[],[f1354,f125])).
fof(f1354,plain,(
  ( ! [X1] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK0))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))))))) ) | $spl1),
  inference(forward_demodulation,[],[f1257,f127])).
fof(f1257,plain,(
  ( ! [X1] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK0))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))))))) ) | $spl1),
  inference(superposition,[],[f126,f1227])).
fof(f1351,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X2),s('type.realax.real',sK0))))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK0)))))) ) | $spl1),
  inference(forward_demodulation,[],[f1350,f125])).
fof(f1350,plain,(
  ( ! [X2] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK0))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))))))) ) | $spl1),
  inference(forward_demodulation,[],[f1255,f127])).
fof(f1255,plain,(
  ( ! [X2] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK0))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))))))) ) | $spl1),
  inference(superposition,[],[f1227,f126])).
fof(f1276,plain,(
  ( ! [X0] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',sK0))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))))) ) | $spl1),
  inference(backward_demodulation,[],[f1272,f1227])).
fof(f1347,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f1329,f104])).
fof(f1329,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl4)),
  inference(condensation,[],[f1328])).
fof(f1328,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X0))) ) | ($spl1 | $spl4)),
  inference(condensation,[],[f1320])).
fof(f1320,plain,(
  ( ! [X47,X50,X49] : (s(bool,'T_0') = s(bool,X50) | p(s(bool,X49)) | p(s(bool,X50)) | s(bool,'T_0') = s(bool,X47)) ) | ($spl1 | $spl4)),
  inference(backward_demodulation,[],[f1277,f1046])).
fof(f1318,plain,(
  ( ! [X10,X11] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',X10),s('type.realax.real',X11))) | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X11))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X11)))))) ) | ($spl1 | $spl4)),
  inference(backward_demodulation,[],[f1277,f1206])).
fof(f1343,plain,(
  ( ! [X7] : (s(bool,'F_0') = s(bool,X7)) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f1316,f1277])).
fof(f1316,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X6),s(bool,X7))) | s(bool,'F_0') = s(bool,X7)) ) | ($spl1 | $spl4)),
  inference(backward_demodulation,[],[f1277,f1190])).
fof(f1341,plain,(
  ( ! [X6] : (s(bool,'F_0') = s(bool,X6)) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f1301,f1277])).
fof(f1301,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X6),s(bool,X7))) | s(bool,'F_0') = s(bool,X6)) ) | ($spl1 | $spl4)),
  inference(backward_demodulation,[],[f1277,f940])).
fof(f1336,plain,(
  ( ! [X3] : (s(bool,'F_0') = s(bool,X3)) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f1287,f1277])).
fof(f1287,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'const.marker.unint_1'(s(bool,X3))) | s(bool,'F_0') = s(bool,X3)) ) | ($spl1 | $spl4)),
  inference(backward_demodulation,[],[f1277,f305])).
fof(f1280,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) ) | ($spl1 | $spl4)),
  inference(backward_demodulation,[],[f1277,f118])).
fof(f1332,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f1278,f1277])).
fof(f1278,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl1 | $spl4)),
  inference(backward_demodulation,[],[f1277,f108])).
fof(f1277,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f1275,f245])).
fof(f1275,plain,(
  ( ! [X0] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl4)),
  inference(backward_demodulation,[],[f1272,f408])).
fof(f1274,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X2),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))))) | p(s(bool,X2))) ) | ($spl1 | $spl4)),
  inference(backward_demodulation,[],[f1272,f410])).
fof(f1273,plain,(
  ( ! [X99,X97,X98] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X99),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))))) | p(s(bool,'const.bool.COND_3'(s(bool,X97),s(bool,X98),s(bool,X99)))) | s(bool,'T_0') = s(bool,X97)) ) | ($spl1 | $spl4)),
  inference(backward_demodulation,[],[f1272,f1059])).
fof(f1272,plain,(
  ( ! [X1] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',sK0))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK0)))))) ) | $spl1),
  inference(forward_demodulation,[],[f1254,f125])).
fof(f1254,plain,(
  ( ! [X1] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK0))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))))) ) | $spl1),
  inference(superposition,[],[f1227,f127])).
fof(f1227,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK0))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0)))))))) ) | $spl1),
  inference(superposition,[],[f125,f1198])).
fof(f1198,plain,(
  s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK0))))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK0))))) | $spl1),
  inference(unit_resulting_resolution,[],[f141,f103])).
fof(f1059,plain,(
  ( ! [X99,X97,X98] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK0))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X99),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))))) | p(s(bool,'const.bool.COND_3'(s(bool,X97),s(bool,X98),s(bool,X99)))) | s(bool,'T_0') = s(bool,X97)) ) | $spl4),
  inference(superposition,[],[f410,f245])).
fof(f408,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK0))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl4),
  inference(backward_demodulation,[],[f406,f317])).
fof(f407,plain,(
  s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))) != s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK0))))) | $spl4),
  inference(backward_demodulation,[],[f406,f323])).
fof(f323,plain,(
  s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1))),s('type.realax.real',sK0))) != s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK0))))) | $spl4),
  inference(subsumption_resolution,[],[f316,f144])).
fof(f316,plain,(
  s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1))),s('type.realax.real',sK0))) != s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK0))))) | p(s(bool,'F_0')) | $spl4),
  inference(superposition,[],[f145,f118])).
fof(f317,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK0))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1))),s('type.realax.real',sK0))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl4),
  inference(subsumption_resolution,[],[f309,f144])).
fof(f309,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK0))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X0),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1))),s('type.realax.real',sK0))))) | p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,X0)) ) | $spl4),
  inference(superposition,[],[f145,f101])).
fof(f308,plain,(
  s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK0))))) != s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1))),s('type.realax.real',sK0))))) | $spl4),
  inference(unit_resulting_resolution,[],[f144,f145])).
fof(f1362,plain,(
  $false | ($spl1 | $spl4)),
  inference(global_subsumption,[],[f91,f90,f92,f94,f144,f101,f103,f105,f104,f109,f108,f114,f117,f118,f124,f123,f122,f125,f126,f127,f131,f128,f132,f145,f141,f147,f167,f169,f199,f224,f245,f259,f303,f304,f305,f307,f308,f317,f323,f340,f341,f380,f390,f406,f407,f408,f410,f425,f426,f427,f440,f450,f451,f452,f522,f766,f770,f771,f778,f821,f834,f835,f836,f854,f871,f872,f940,f945,f962,f988,f989,f990,f1029,f1031,f1033,f1042,f1043,f1044,f1045,f1046,f1048,f1091,f1054,f1055,f1092,f1058,f1059,f1060,f1105,f1127,f1128,f1129,f1190,f1195,f1198,f1203,f1205,f1206,f1207,f1209,f1227,f1272,f1273,f1274,f1277,f1332,f1280,f1336,f1341,f1343,f1318,f1347,f1276,f1351,f1355,f1361])).
fof(f1356,plain,(
  $false | ($spl1 | $spl4)),
  inference(global_subsumption,[],[f91,f90,f92,f94,f144,f101,f103,f105,f104,f109,f108,f114,f117,f118,f124,f123,f122,f125,f126,f127,f131,f128,f132,f145,f141,f147,f167,f169,f199,f224,f245,f259,f303,f304,f305,f307,f308,f317,f323,f340,f341,f380,f390,f406,f407,f408,f410,f425,f426,f427,f440,f450,f451,f452,f522,f766,f770,f771,f778,f821,f834,f835,f836,f854,f871,f872,f940,f945,f962,f988,f989,f990,f1029,f1031,f1033,f1042,f1043,f1044,f1045,f1046,f1048,f1091,f1054,f1055,f1092,f1058,f1059,f1060,f1105,f1127,f1128,f1129,f1190,f1195,f1198,f1203,f1205,f1206,f1207,f1209,f1227,f1272,f1273,f1274,f1277,f1332,f1280,f1336,f1341,f1343,f1318,f1347,f1276,f1351,f1355])).
fof(f1352,plain,(
  $false | ($spl1 | $spl4)),
  inference(global_subsumption,[],[f91,f90,f92,f94,f144,f101,f103,f105,f104,f109,f108,f114,f117,f118,f124,f123,f122,f125,f126,f127,f131,f128,f132,f145,f141,f147,f167,f169,f199,f224,f245,f259,f303,f304,f305,f307,f308,f317,f323,f340,f341,f380,f390,f406,f407,f408,f410,f425,f426,f427,f440,f450,f451,f452,f522,f766,f770,f771,f778,f821,f834,f835,f836,f854,f871,f872,f940,f945,f962,f988,f989,f990,f1029,f1031,f1033,f1042,f1043,f1044,f1045,f1046,f1048,f1091,f1054,f1055,f1092,f1058,f1059,f1060,f1105,f1127,f1128,f1129,f1190,f1195,f1198,f1203,f1205,f1206,f1207,f1209,f1227,f1272,f1273,f1274,f1277,f1332,f1280,f1336,f1341,f1343,f1318,f1347,f1276,f1351])).
fof(f1349,plain,(
  $false | ($spl1 | $spl4)),
  inference(global_subsumption,[],[f91,f90,f92,f94,f144,f101,f103,f105,f104,f109,f108,f114,f117,f118,f124,f123,f122,f125,f126,f127,f131,f128,f132,f145,f141,f147,f167,f169,f199,f224,f245,f259,f303,f304,f305,f307,f308,f317,f323,f340,f341,f380,f390,f406,f407,f408,f410,f425,f426,f427,f440,f450,f451,f452,f522,f766,f770,f771,f778,f821,f834,f835,f836,f854,f871,f872,f940,f945,f962,f988,f989,f990,f1029,f1031,f1033,f1042,f1043,f1044,f1045,f1046,f1048,f1091,f1054,f1055,f1092,f1058,f1059,f1060,f1105,f1127,f1128,f1129,f1190,f1195,f1198,f1203,f1205,f1206,f1207,f1209,f1227,f1272,f1273,f1274,f1277,f1332,f1280,f1336,f1341,f1343,f1318,f1347,f1276])).
fof(f1348,plain,(
  $false | ($spl1 | $spl4)),
  inference(global_subsumption,[],[f91,f90,f92,f94,f144,f101,f103,f105,f104,f109,f108,f114,f117,f118,f124,f123,f122,f125,f126,f127,f131,f128,f132,f145,f141,f147,f167,f169,f199,f224,f245,f259,f303,f304,f305,f307,f308,f317,f323,f340,f341,f380,f390,f406,f407,f408,f410,f425,f426,f427,f440,f450,f451,f452,f522,f766,f770,f771,f778,f821,f834,f835,f836,f854,f871,f872,f940,f945,f962,f988,f989,f990,f1029,f1031,f1033,f1042,f1043,f1044,f1045,f1046,f1048,f1091,f1054,f1055,f1092,f1058,f1059,f1060,f1105,f1127,f1128,f1129,f1190,f1195,f1198,f1203,f1205,f1206,f1207,f1209,f1227,f1272,f1273,f1274,f1277,f1332,f1280,f1336,f1341,f1343,f1318,f1347])).
fof(f1346,plain,(
  $false | ($spl1 | $spl4)),
  inference(global_subsumption,[],[f91,f90,f92,f94,f144,f101,f103,f105,f104,f109,f108,f114,f117,f118,f124,f123,f122,f125,f126,f127,f131,f128,f132,f145,f141,f147,f167,f169,f199,f224,f245,f259,f303,f304,f305,f307,f308,f317,f323,f340,f341,f380,f390,f406,f407,f408,f410,f425,f426,f427,f440,f450,f451,f452,f522,f766,f770,f771,f778,f821,f834,f835,f836,f854,f871,f872,f940,f945,f962,f988,f989,f990,f1029,f1031,f1033,f1042,f1043,f1044,f1045,f1046,f1048,f1091,f1054,f1055,f1092,f1058,f1059,f1060,f1105,f1127,f1128,f1129,f1190,f1195,f1198,f1203,f1205,f1206,f1207,f1209,f1227,f1272,f1273,f1274,f1277,f1332,f1280,f1336,f1341,f1343,f1318])).
fof(f1344,plain,(
  $false | ($spl1 | $spl4)),
  inference(global_subsumption,[],[f91,f90,f92,f94,f144,f101,f103,f105,f104,f109,f108,f114,f117,f118,f124,f123,f122,f125,f126,f127,f131,f128,f132,f145,f141,f147,f167,f169,f199,f224,f245,f259,f303,f304,f305,f307,f308,f317,f323,f340,f341,f380,f390,f406,f407,f408,f410,f425,f426,f427,f440,f450,f451,f452,f522,f766,f770,f771,f778,f821,f834,f835,f836,f854,f871,f872,f940,f945,f962,f988,f989,f990,f1029,f1031,f1033,f1042,f1043,f1044,f1045,f1046,f1048,f1091,f1054,f1055,f1092,f1058,f1059,f1060,f1105,f1127,f1128,f1129,f1190,f1195,f1198,f1203,f1205,f1206,f1207,f1209,f1227,f1272,f1273,f1274,f1277,f1332,f1280,f1336,f1341,f1343])).
fof(f1342,plain,(
  $false | ($spl1 | $spl4)),
  inference(global_subsumption,[],[f91,f90,f92,f94,f144,f101,f103,f105,f104,f109,f108,f114,f117,f118,f124,f123,f122,f125,f126,f127,f131,f128,f132,f145,f141,f147,f167,f169,f199,f224,f245,f259,f303,f304,f305,f307,f308,f317,f323,f340,f341,f380,f390,f406,f407,f408,f410,f425,f426,f427,f440,f450,f451,f452,f522,f766,f770,f771,f778,f821,f834,f835,f836,f854,f871,f872,f940,f945,f962,f988,f989,f990,f1029,f1031,f1033,f1042,f1043,f1044,f1045,f1046,f1048,f1091,f1054,f1055,f1092,f1058,f1059,f1060,f1105,f1127,f1128,f1129,f1190,f1195,f1198,f1203,f1205,f1206,f1207,f1209,f1227,f1272,f1273,f1274,f1277,f1332,f1280,f1336,f1341])).
fof(f1337,plain,(
  $false | ($spl1 | $spl4)),
  inference(global_subsumption,[],[f91,f90,f92,f94,f144,f101,f103,f105,f104,f109,f108,f114,f117,f118,f124,f123,f122,f125,f126,f127,f131,f128,f132,f145,f141,f147,f167,f169,f199,f224,f245,f259,f303,f304,f305,f307,f308,f317,f323,f340,f341,f380,f390,f406,f407,f408,f410,f425,f426,f427,f440,f450,f451,f452,f522,f766,f770,f771,f778,f821,f834,f835,f836,f854,f871,f872,f940,f945,f962,f988,f989,f990,f1029,f1031,f1033,f1042,f1043,f1044,f1045,f1046,f1048,f1091,f1054,f1055,f1092,f1058,f1059,f1060,f1105,f1127,f1128,f1129,f1190,f1195,f1198,f1203,f1205,f1206,f1207,f1209,f1227,f1272,f1273,f1274,f1277,f1332,f1280,f1336])).
fof(f1335,plain,(
  $false | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f1281,f92])).
fof(f1281,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl4)),
  inference(backward_demodulation,[],[f1277,f144])).
fof(f1334,plain,(
  $false | ($spl1 | $spl4)),
  inference(global_subsumption,[],[f91,f90,f92,f94,f144,f101,f103,f105,f104,f109,f108,f114,f117,f118,f124,f123,f122,f125,f126,f127,f131,f128,f132,f145,f141,f147,f167,f169,f199,f224,f245,f259,f303,f304,f305,f307,f308,f317,f323,f340,f341,f380,f390,f406,f407,f408,f410,f425,f426,f427,f440,f450,f451,f452,f522,f766,f770,f771,f778,f821,f834,f835,f836,f854,f871,f872,f940,f945,f962,f988,f989,f990,f1029,f1031,f1033,f1042,f1043,f1044,f1045,f1046,f1048,f1091,f1054,f1055,f1092,f1058,f1059,f1060,f1105,f1127,f1128,f1129,f1190,f1195,f1198,f1203,f1205,f1206,f1207,f1209,f1227,f1272,f1273,f1274,f1277,f1332,f1280])).
fof(f1333,plain,(
  $false | ($spl1 | $spl4)),
  inference(global_subsumption,[],[f91,f90,f92,f94,f144,f101,f103,f105,f104,f109,f108,f114,f117,f118,f124,f123,f122,f125,f126,f127,f131,f128,f132,f145,f141,f147,f167,f169,f199,f224,f245,f259,f303,f304,f305,f307,f308,f317,f323,f340,f341,f380,f390,f406,f407,f408,f410,f425,f426,f427,f440,f450,f451,f452,f522,f766,f770,f771,f778,f821,f834,f835,f836,f854,f871,f872,f940,f945,f962,f988,f989,f990,f1029,f1031,f1033,f1042,f1043,f1044,f1045,f1046,f1048,f1091,f1054,f1055,f1092,f1058,f1059,f1060,f1105,f1127,f1128,f1129,f1190,f1195,f1198,f1203,f1205,f1206,f1207,f1209,f1227,f1272,f1273,f1274,f1277,f1332])).
fof(f140,plain,(
  s('type.realax.real',sK0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl0),
  inference(cnf_transformation,[],[f140_D])).
fof(f140_D,plain,(
  s('type.realax.real',sK0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f143,plain,(
  ( ! [X2] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK0))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X2),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1))),s('type.realax.real',sK0))))) | p(s(bool,X2))) ) | $spl4),
  inference(cnf_transformation,[],[f143_D])).
fof(f143_D,plain,(
  ( ! [X2] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK0))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X2),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1))),s('type.realax.real',sK0))))) | p(s(bool,X2))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f90,plain,(
  ( ! [X2] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK0))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X2),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1))),s('type.realax.real',sK0))))) | s('type.realax.real',sK0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X2] : (((p(s(bool,X2)) | s('type.realax.real',sK0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (~p(s(bool,X2)) | s('type.realax.real',sK0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK0))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X2),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1))),s('type.realax.real',sK0))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f76])).
fof(f76,plain,(
  ? [X0,X1] : ! [X2] : (((p(s(bool,X2)) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (~p(s(bool,X2)) | s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X2),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))),s('type.realax.real',X0))))))),
  inference(nnf_transformation,[],[f66])).
fof(f66,plain,(
  ? [X0,X1] : ! [X2] : ((p(s(bool,X2)) <~> s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) | s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X2),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))),s('type.realax.real',X0))))))),
  inference(ennf_transformation,[],[f30])).
fof(f30,plain,(
  ~! [X0,X1] : ? [X2] : ((p(s(bool,X2)) <=> s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))))) = s('type.realax.real','const.bool.COND_3'(s(bool,X2),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X1),s('type.realax.real',X0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))),s('type.realax.real',X0))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ? [X8] : ((p(s(bool,X8)) <=> s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X7),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X6))))) = s('type.realax.real','const.bool.COND_3'(s(bool,X8),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X7),s('type.realax.real',X6))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X7))),s('type.realax.real',X6))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ? [X8] : ((p(s(bool,X8)) <=> s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X7),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X6))))) = s('type.realax.real','const.bool.COND_3'(s(bool,X8),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X7),s('type.realax.real',X6))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X7))),s('type.realax.real',X6))))))),
  file('thm.real.neg_rat|split|1','thm.real.neg_rat|split|1')).
fof(f141,plain,(
  s('type.realax.real',sK0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl1),
  inference(cnf_transformation,[],[f141_D])).
fof(f141_D,plain,(
  s('type.realax.real',sK0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f142,plain,(
  ( ! [X2] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK0))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X2),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1))),s('type.realax.real',sK0))))) | ~p(s(bool,X2))) ) | $spl2),
  inference(cnf_transformation,[],[f142_D])).
fof(f142_D,plain,(
  ( ! [X2] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK0))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X2),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1))),s('type.realax.real',sK0))))) | ~p(s(bool,X2))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f91,plain,(
  ( ! [X2] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK0))))) != s('type.realax.real','const.bool.COND_3'(s(bool,X2),s('type.realax.real','const.marker.unint_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',sK1),s('type.realax.real',sK0))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',sK1))),s('type.realax.real',sK0))))) | s('type.realax.real',sK0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f77])).
% SZS output end Proof for thm.real
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 50
% Passive clauses: 184
% Generated clauses: 1915
% Final active clauses: 25
% Final passive clauses: 37
% Input formulas: 28
% Initial clauses: 49
% 
% Duplicate literals: 94
% Trivial inequalities: 14
% Fw subsumption resolutions: 82
% Fw demodulations: 145
% Bw demodulations: 113
% Condensations: 161
% Global subsumptions: 21
% 
% Simple tautologies: 35
% Equational tautologies: 34
% Forward subsumptions: 1112
% Fw demodulations to eq. taut.: 2
% Bw demodulations to eq. taut.: 2
% 
% Unit resulting resolution: 62
% Factoring: 1
% Forward superposition: 522
% Backward superposition: 516
% Self superposition: 54
% Equality factoring: 26
% Equality resolution: 1
% 
% Split clauses: 2
% Split components: 4
% SAT solver clauses: 240
% SAT solver unit clauses: 15
% SAT solver binary clauses: 100
% 
% Sat splits: 2
% Sat splitting refutations: 29
% 
% Memory used [KB]: 2046
% Time elapsed: 0.057 s
% ------------------------------
% ------------------------------
% Success in time 1.366 s
