Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.pred_set
% (112000)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 121
% Passive clauses: 2137
% Generated clauses: 23517
% Final active clauses: 104
% Final passive clauses: 1751
% Input formulas: 8
% Initial clauses: 16
% 
% Duplicate literals: 1165
% Trivial inequalities: 306
% Fw subsumption resolutions: 3491
% Fw demodulations: 2588
% Bw demodulations: 12
% 
% Simple tautologies: 351
% Equational tautologies: 75
% Forward subsumptions: 10802
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 245
% Factoring: 3
% Forward superposition: 7481
% Backward superposition: 7453
% Self superposition: 309
% Equality factoring: 222
% Equality resolution: 25
% 
% Split clauses: 1260
% Split components: 2605
% SAT solver clauses: 2102
% SAT solver unit clauses: 63
% SAT solver binary clauses: 257
% 
% Sat splits: 1260
% Sat splitting refutations: 7
% 
% TWLsolver clauses: 2088
% TWLsolver calls for satisfiability: 164
% 
% Memory used [KB]: 9978
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (112017)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 101
% Passive clauses: 482
% Generated clauses: 9501
% Final active clauses: 79
% Final passive clauses: 280
% Input formulas: 8
% Initial clauses: 16
% 
% Duplicate literals: 790
% Trivial inequalities: 25
% Fw subsumption resolutions: 760
% Fw demodulations: 850
% Bw demodulations: 121
% Condensations: 1275
% 
% Simple tautologies: 106
% Equational tautologies: 49
% Forward subsumptions: 3925
% 
% Binary resolution: 311
% Factoring: 4
% Forward superposition: 2095
% Backward superposition: 2480
% Self superposition: 176
% Equality factoring: 124
% Equality resolution: 10
% 
% Split clauses: 46
% Split components: 93
% SAT solver clauses: 950
% SAT solver unit clauses: 19
% SAT solver binary clauses: 39
% 
% Sat splits: 46
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 950
% TWLsolver calls for satisfiability: 152
% 
% Memory used [KB]: 7675
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f3478,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f57,f64,f62,f77,f81,f79,f88,f417,f418,f424,f422,f419,f420,f427,f421,f429,f426,f530,f531,f532,f533,f540,f534,f542,f539,f666,f667,f673,f671,f668,f669,f676,f670,f678,f675,f778,f779,f785,f783,f780,f781,f788,f782,f790,f787,f959,f960,f967,f964,f961,f962,f970,f963,f972,f965,f969,f1263,f1266,f1272,f1273,f1275,f1277,f1278,f1280,f1281,f1282,f1283,f1284,f1286,f1287,f1288,f1290,f1291,f1292,f1293,f1294,f1295,f1296,f1297,f1298,f1299,f1300,f1301,f1302,f1304,f1305,f1306,f1307,f1309,f1310,f1311,f1313,f1314,f1316,f1317,f1318,f1319,f1320,f1321,f1322,f1323,f1324,f1325,f1327,f1328,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1347,f1348,f1349,f1350,f1351,f1352,f1353,f1354,f1355,f1356,f1357,f1358,f1359,f1360,f1361,f1421,f1422,f1425,f1428,f1429,f1431,f1433,f1434,f1436,f1437,f1438,f1479,f1480,f1366,f1363,f537,f535,f1741,f1743,f1744,f1748,f1752,f1754,f1755,f1756,f1758,f1760,f1761,f1763,f1765,f1766,f1767,f1808,f1814,f2069,f2071,f2073,f2074,f2079,f2084,f2086,f2088,f2089,f2090,f2092,f2094,f2096,f2097,f2099,f2101,f2103,f2104,f2105,f2145,f2151,f2677,f2680,f2720,f2729,f2731,f2736,f2738,f2748,f2754,f2756,f2759,f2761,f2763,f2765,f2767,f2769,f2771,f2775,f2784,f2789,f2792,f2794,f2796,f2798,f2800,f2804,f2805,f2843,f2846,f2848,f2850,f2853,f2856,f2859,f2866,f2869,f2871,f2891,f3304,f3307,f3314,f3315,f428,f3256,f3254,f3316,f3219,f3228,f3233,f3318,f3238,f541,f78,f3429,f3469])).
fof(f3469,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f3468,f38])).
fof(f38,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f12])).
fof(f12,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pred_set.SUBSET_ANTISYM',aHL_FALSITY)).
fof(f3468,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(forward_demodulation,[],[f3467,f62])).
fof(f3467,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X1))))) ) | $spl0),
  inference(subsumption_resolution,[],[f3441,f38])).
fof(f3441,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X1))))) ) | $spl0),
  inference(backward_demodulation,[],[f62,f43])).
fof(f43,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X1)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))))) & ((p(s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f31])).
fof(f31,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f30])).
fof(f30,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(nnf_transformation,[],[f22])).
fof(f22,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(ennf_transformation,[],[f16])).
fof(f16,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f15])).
fof(f15,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : ! [X7,X8] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))) <=> ! [X11] : (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X7)))) => p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X8))))))),
  file('thm.pred_set.SUBSET_ANTISYM','thm.pred_set.SUBSET_DEF')).
fof(f3429,plain,(
  $false | ($spl2 | $spl22)),
  inference(subsumption_resolution,[],[f63,f3323])).
fof(f3323,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl22),
  inference(subsumption_resolution,[],[f2960,f2806])).
fof(f2806,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | $spl22),
  inference(forward_demodulation,[],[f406,f539])).
fof(f406,plain,(
  s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) != s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0))))),
  inference(unit_resulting_resolution,[],[f36,f49])).
fof(f49,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK5(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK5(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK5(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK5(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f25])).
fof(f25,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f19])).
fof(f19,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f18])).
fof(f18,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.pred_set.SUBSET_ANTISYM',aHL_EXT)).
fof(f36,plain,(
  s(fun(sK0,bool),sK1) != s(fun(sK0,bool),sK2)),
  inference(cnf_transformation,[],[f26])).
fof(f26,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))) & s(fun(sK0,bool),sK1) != s(fun(sK0,bool),sK2)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f21])).
fof(f21,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) & s(fun(X0,bool),X1) != s(fun(X0,bool),X2))),
  inference(flattening,[],[f20])).
fof(f20,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))))) & s(fun(X0,bool),X1) != s(fun(X0,bool),X2))),
  inference(ennf_transformation,[],[f11])).
fof(f11,plain,(
  ~! [X0,X1,X2] : ((p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))))) => s(fun(X0,bool),X1) = s(fun(X0,bool),X2))),
  inference(flattening,[],[f10])).
fof(f10,plain,(
  ~! [X0] : ! [X1,X2] : ((p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))))) => s(fun(X0,bool),X1) = s(fun(X0,bool),X2))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : ((p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X7))))) => s(fun(X6,bool),X7) = s(fun(X6,bool),X8))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : ((p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X7))))) => s(fun(X6,bool),X7) = s(fun(X6,bool),X8))),
  file('thm.pred_set.SUBSET_ANTISYM','thm.pred_set.SUBSET_ANTISYM')).
fof(f2960,plain,(
  s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | $spl22),
  inference(superposition,[],[f2806,f39])).
fof(f39,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.SUBSET_ANTISYM',aHL_BOOL_CASES)).
fof(f63,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl2),
  inference(cnf_transformation,[],[f63_D])).
fof(f63_D,plain,(
  s(bool,'T_0') = s(bool,'F_0') <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f78,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl4),
  inference(cnf_transformation,[],[f78_D])).
fof(f78_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f541,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | $spl26),
  inference(cnf_transformation,[],[f541_D])).
fof(f541_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f3238,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl22),
  inference(duplicate_literal_removal,[],[f3143])).
fof(f3143,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl22),
  inference(superposition,[],[f2966,f39])).
fof(f2966,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl22),
  inference(trivial_inequality_removal,[],[f2961])).
fof(f2961,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl22),
  inference(superposition,[],[f2806,f50])).
fof(f50,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f39,f39])).
fof(f3318,plain,(
  s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) = s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) | $spl62),
  inference(cnf_transformation,[],[f3318_D])).
fof(f3318_D,plain,(
  s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) = s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f3233,plain,(
  ( ! [X9] : (s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) = s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X9)) ) | ($spl12 | $spl22)),
  inference(duplicate_literal_removal,[],[f3148])).
fof(f3148,plain,(
  ( ! [X9] : (s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) = s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | ($spl12 | $spl22)),
  inference(superposition,[],[f2966,f466])).
fof(f466,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(trivial_inequality_removal,[],[f460])).
fof(f460,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(superposition,[],[f430,f50])).
fof(f430,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) | $spl12),
  inference(backward_demodulation,[],[f426,f405])).
fof(f405,plain,(
  s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) != s(bool,g(s(sK0,sK5(sK2,sK1,bool,sK0))))),
  inference(unit_resulting_resolution,[],[f36,f49])).
fof(f3228,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl22),
  inference(duplicate_literal_removal,[],[f3167])).
fof(f3167,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl22),
  inference(superposition,[],[f39,f2966])).
fof(f3219,plain,(
  ( ! [X21] : (s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) = s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X21)) ) | ($spl12 | $spl22)),
  inference(duplicate_literal_removal,[],[f3178])).
fof(f3178,plain,(
  ( ! [X21] : (s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) = s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) ) | ($spl12 | $spl22)),
  inference(superposition,[],[f466,f2966])).
fof(f3316,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,X6) = s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0))))) ) | $spl60),
  inference(cnf_transformation,[],[f3316_D])).
fof(f3316_D,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,X6) = s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0))))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f3254,plain,(
  ( ! [X4,X3] : (p(s(bool,X4)) | s(bool,X4) = s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,X3)) ) | $spl22),
  inference(subsumption_resolution,[],[f3203,f52])).
fof(f52,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f38,f39])).
fof(f3203,plain,(
  ( ! [X4,X3] : (p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X4) = s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,X3)) ) | $spl22),
  inference(superposition,[],[f45,f2966])).
fof(f45,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f24])).
fof(f24,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f23])).
fof(f23,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f17])).
fof(f17,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X9,X10] : ((p(s(bool,X9)) => p(s(bool,X10))) => ((p(s(bool,X10)) => p(s(bool,X9))) => s(bool,X9) = s(bool,X10)))),
  file('thm.pred_set.SUBSET_ANTISYM','thm.bool.IMP_ANTISYM_AX')).
fof(f3256,plain,(
  ( ! [X6,X5] : (p(s(bool,X6)) | s(bool,X6) = s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,X5)) ) | $spl22),
  inference(subsumption_resolution,[],[f3204,f52])).
fof(f3204,plain,(
  ( ! [X6,X5] : (p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X6) = s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,X5)) ) | $spl22),
  inference(superposition,[],[f45,f2966])).
fof(f428,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) | $spl16),
  inference(cnf_transformation,[],[f428_D])).
fof(f428_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f3315,plain,(
  ( ! [X16] : (s(bool,'F_0') = s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X16)) ) | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f3271,f52])).
fof(f3271,plain,(
  ( ! [X16] : (p(s(bool,X16)) | s(bool,'F_0') = s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X16)) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f79,f466])).
fof(f3314,plain,(
  ( ! [X17] : (s(bool,'F_0') = s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,X17)) ) | ($spl6 | $spl22)),
  inference(subsumption_resolution,[],[f3272,f52])).
fof(f3272,plain,(
  ( ! [X17] : (p(s(bool,X17)) | s(bool,'F_0') = s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,X17)) ) | ($spl6 | $spl22)),
  inference(superposition,[],[f79,f2966])).
fof(f3307,plain,(
  $false | ($spl3 | $spl33 | $spl40 | $spl43)),
  inference(subsumption_resolution,[],[f3306,f37])).
fof(f37,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.SUBSET_ANTISYM',aHL_TRUTH)).
fof(f3306,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl33 | $spl40 | $spl43)),
  inference(forward_demodulation,[],[f3305,f784])).
fof(f784,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | $spl40),
  inference(cnf_transformation,[],[f784_D])).
fof(f784_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f3305,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1)))) | ($spl3 | $spl33 | $spl43)),
  inference(subsumption_resolution,[],[f3294,f38])).
fof(f3294,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1)))) | ($spl3 | $spl33 | $spl43)),
  inference(superposition,[],[f527,f2821])).
fof(f2821,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl33 | $spl43)),
  inference(forward_demodulation,[],[f2318,f1957])).
fof(f1957,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl33)),
  inference(unit_resulting_resolution,[],[f37,f64,f674,f376])).
fof(f376,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,X9)) | s(bool,X8) = s(bool,X9) | s(bool,X8) = s(bool,X10) | s(bool,'T_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f230,f37])).
fof(f230,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X9)) | s(bool,X8) = s(bool,X9) | s(bool,X8) = s(bool,X10) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f48,f50])).
fof(f48,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f24])).
fof(f674,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl33),
  inference(cnf_transformation,[],[f674_D])).
fof(f674_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f2318,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK2))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl33 | $spl43)),
  inference(unit_resulting_resolution,[],[f37,f674,f786,f376])).
fof(f786,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK2))) | $spl43),
  inference(cnf_transformation,[],[f786_D])).
fof(f786_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f527,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK1))))) )),
  inference(subsumption_resolution,[],[f520,f37])).
fof(f520,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2))))) )),
  inference(superposition,[],[f42,f58])).
fof(f58,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))),
  inference(subsumption_resolution,[],[f54,f38])).
fof(f54,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))),
  inference(superposition,[],[f34,f39])).
fof(f34,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),
  inference(cnf_transformation,[],[f26])).
fof(f42,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f32])).
fof(f3304,plain,(
  $false | ($spl3 | $spl30 | $spl33)),
  inference(subsumption_resolution,[],[f3303,f37])).
fof(f3303,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl30 | $spl33)),
  inference(forward_demodulation,[],[f3302,f672])).
fof(f672,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl30),
  inference(cnf_transformation,[],[f672_D])).
fof(f672_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f3302,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) | ($spl3 | $spl33)),
  inference(subsumption_resolution,[],[f3293,f38])).
fof(f3293,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) | ($spl3 | $spl33)),
  inference(superposition,[],[f527,f1957])).
fof(f2891,plain,(
  $false | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f2890,f37])).
fof(f2890,plain,(
  ~p(s(bool,'T_0')) | ($spl18 | $spl22)),
  inference(forward_demodulation,[],[f38,f2861])).
fof(f2861,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f2860,f37])).
fof(f2860,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,X0)) ) | ($spl18 | $spl22)),
  inference(forward_demodulation,[],[f2517,f539])).
fof(f2517,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0)))))) ) | $spl18),
  inference(condensation,[],[f2516])).
fof(f2516,plain,(
  ( ! [X6,X5] : (~p(s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0))))) | s(bool,X5) = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) ) | $spl18),
  inference(subsumption_resolution,[],[f2444,f80])).
fof(f80,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f74,f38])).
fof(f74,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f45,f39])).
fof(f2444,plain,(
  ( ! [X6,X5] : (~p(s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0))))) | ~p(s(bool,X6)) | s(bool,X5) = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) ) | $spl18),
  inference(superposition,[],[f48,f535])).
fof(f2871,plain,(
  $false | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f2870,f2861])).
fof(f2870,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl18 | $spl22)),
  inference(forward_demodulation,[],[f2602,f2861])).
fof(f2602,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK2)))),
  inference(unit_resulting_resolution,[],[f36,f41])).
fof(f41,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) )),
  inference(cnf_transformation,[],[f29])).
fof(f29,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f28])).
fof(f28,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f27])).
fof(f27,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f14])).
fof(f14,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f13])).
fof(f13,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X7,X8] : (s(fun(X6,bool),X7) = s(fun(X6,bool),X8) <=> ! [X11] : s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X7))) = s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X8))))),
  file('thm.pred_set.SUBSET_ANTISYM','thm.pred_set.EXTENSION')).
fof(f2869,plain,(
  $false | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f2868,f2861])).
fof(f2868,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl18 | $spl22)),
  inference(forward_demodulation,[],[f2603,f2861])).
fof(f2603,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK2)))),
  inference(unit_resulting_resolution,[],[f36,f41])).
fof(f2866,plain,(
  $false | ($spl3 | $spl12 | $spl18 | $spl22 | $spl33)),
  inference(subsumption_resolution,[],[f2864,f37])).
fof(f2864,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl12 | $spl18 | $spl22 | $spl33)),
  inference(backward_demodulation,[],[f2861,f2809])).
fof(f2809,plain,(
  ~p(s(bool,'F_0')) | ($spl3 | $spl12 | $spl33)),
  inference(backward_demodulation,[],[f2808,f2807])).
fof(f2807,plain,(
  ~p(s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0))))) | ($spl12 | $spl33)),
  inference(backward_demodulation,[],[f1958,f1964])).
fof(f1964,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK2)))) | $spl33),
  inference(unit_resulting_resolution,[],[f37,f674,f48])).
fof(f1958,plain,(
  s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl12 | $spl33)),
  inference(unit_resulting_resolution,[],[f37,f430,f674,f376])).
fof(f2808,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) | ($spl3 | $spl12 | $spl33)),
  inference(backward_demodulation,[],[f1957,f1958])).
fof(f2859,plain,(
  $false | ($spl3 | $spl18 | $spl22)),
  inference(subsumption_resolution,[],[f2858,f64])).
fof(f2858,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl18 | $spl22)),
  inference(condensation,[],[f2857])).
fof(f2857,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X0)) ) | ($spl18 | $spl22)),
  inference(forward_demodulation,[],[f2501,f539])).
fof(f2501,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl18),
  inference(duplicate_literal_removal,[],[f2441])).
fof(f2441,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl18),
  inference(superposition,[],[f39,f535])).
fof(f2856,plain,(
  $false | ($spl12 | $spl18 | $spl22)),
  inference(subsumption_resolution,[],[f2855,f430])).
fof(f2855,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) | ($spl12 | $spl18 | $spl22)),
  inference(condensation,[],[f2854])).
fof(f2854,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X9)) ) | ($spl12 | $spl18 | $spl22)),
  inference(forward_demodulation,[],[f2506,f539])).
fof(f2506,plain,(
  ( ! [X9] : (s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) = s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,X9)) ) | ($spl12 | $spl18)),
  inference(duplicate_literal_removal,[],[f2424])).
fof(f2424,plain,(
  ( ! [X9] : (s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) = s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | ($spl12 | $spl18)),
  inference(superposition,[],[f535,f466])).
fof(f2853,plain,(
  $false | ($spl3 | $spl18 | $spl22)),
  inference(subsumption_resolution,[],[f2852,f64])).
fof(f2852,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl18 | $spl22)),
  inference(condensation,[],[f2851])).
fof(f2851,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X0)) ) | ($spl18 | $spl22)),
  inference(forward_demodulation,[],[f2511,f539])).
fof(f2511,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl18),
  inference(duplicate_literal_removal,[],[f2419])).
fof(f2419,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl18),
  inference(superposition,[],[f535,f39])).
fof(f2850,plain,(
  $false | ($spl3 | $spl18 | $spl22)),
  inference(subsumption_resolution,[],[f2849,f64])).
fof(f2849,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl18 | $spl22)),
  inference(forward_demodulation,[],[f2413,f539])).
fof(f2413,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0)))) | ($spl3 | $spl18)),
  inference(unit_resulting_resolution,[],[f64,f535])).
fof(f2848,plain,(
  $false | ($spl12 | $spl18 | $spl22)),
  inference(subsumption_resolution,[],[f2847,f430])).
fof(f2847,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) | ($spl12 | $spl18 | $spl22)),
  inference(forward_demodulation,[],[f2414,f539])).
fof(f2414,plain,(
  s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) = s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0)))) | ($spl12 | $spl18)),
  inference(unit_resulting_resolution,[],[f430,f535])).
fof(f2846,plain,(
  $false | ($spl3 | $spl18 | $spl22 | $spl33)),
  inference(subsumption_resolution,[],[f2845,f64])).
fof(f2845,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl18 | $spl22 | $spl33)),
  inference(forward_demodulation,[],[f2844,f539])).
fof(f2844,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0)))) | ($spl3 | $spl18 | $spl33)),
  inference(forward_demodulation,[],[f2416,f1957])).
fof(f2416,plain,(
  s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl18 | $spl33)),
  inference(unit_resulting_resolution,[],[f674,f535])).
fof(f2843,plain,(
  $false | ($spl3 | $spl18 | $spl22 | $spl33 | $spl43)),
  inference(subsumption_resolution,[],[f2842,f64])).
fof(f2842,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl18 | $spl22 | $spl33 | $spl43)),
  inference(forward_demodulation,[],[f2841,f539])).
fof(f2841,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0)))) | ($spl3 | $spl18 | $spl33 | $spl43)),
  inference(forward_demodulation,[],[f2417,f2821])).
fof(f2417,plain,(
  s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl18 | $spl43)),
  inference(unit_resulting_resolution,[],[f786,f535])).
fof(f2805,plain,(
  $false | ($spl6 | $spl18 | $spl22 | $spl24)),
  inference(subsumption_resolution,[],[f406,f2739])).
fof(f2739,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0))))) ) | ($spl6 | $spl18 | $spl22 | $spl24)),
  inference(subsumption_resolution,[],[f2486,f2686])).
fof(f2686,plain,(
  ( ! [X17] : (s(bool,'T_0') = s(bool,X17)) ) | ($spl6 | $spl18 | $spl22 | $spl24)),
  inference(global_subsumption,[],[f2685,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f166,f376,f403,f404,f405,f407,f409,f417,f418,f419,f420,f421,f530,f531,f532,f533,f534,f539,f543,f579,f586,f603,f605,f613,f615,f643,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f959,f960,f961,f962,f963,f540,f406,f1366,f1363,f601,f608,f610,f527,f528,f1247,f1267,f602,f40,f41,f241,f240,f202,f201,f910,f529,f377,f127,f44,f43,f42,f412,f414,f1233,f617,f612,f607,f416,f411,f141,f125,f38])).
fof(f125,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(bool,bool),X1),s(fun(bool,bool),X0)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))) | s(bool,'T_0') = s(bool,sK4(X0,X1,bool))) )),
  inference(superposition,[],[f43,f39])).
fof(f141,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(bool,bool),X1),s(fun(bool,bool),X0)))) | s(bool,'T_0') = s(bool,sK4(X0,X1,bool))) )),
  inference(superposition,[],[f44,f39])).
fof(f411,plain,(
  ( ! [X17,X18,X16] : (s(X18,f(s(bool,'F_0'))) != s(X18,g(s(bool,'F_0'))) | s(fun(bool,X18),X16) = s(fun(bool,X18),X17) | s(bool,'T_0') = s(bool,sK5(X16,X17,X18,bool))) )),
  inference(superposition,[],[f49,f39])).
fof(f416,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,f(s(X16,sK5(X17,X18,bool,X16)))) | s(fun(X16,bool),X17) = s(fun(X16,bool),X18) | s(bool,'T_0') = s(bool,g(s(X16,sK5(X17,X18,bool,X16))))) )),
  inference(superposition,[],[f49,f39])).
fof(f607,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X16,sK3(X17,X18,X16)),s(fun(X16,bool),X17))) | s(fun(X16,bool),X17) = s(fun(X16,bool),X18) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X16,sK3(X17,X18,X16)),s(fun(X16,bool),X18)))) )),
  inference(superposition,[],[f41,f39])).
fof(f612,plain,(
  ( ! [X12,X13] : (s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X12))) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X13))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK3(X12,X13,bool))) )),
  inference(superposition,[],[f41,f39])).
fof(f617,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X16,sK3(X17,X18,X16)),s(fun(X16,bool),X18))) | s(fun(X16,bool),X17) = s(fun(X16,bool),X18) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X16,sK3(X17,X18,X16)),s(fun(X16,bool),X17)))) )),
  inference(superposition,[],[f41,f39])).
fof(f1233,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))) | s(bool,'T_0') = s(bool,sK4(X1,X0,bool)) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(bool,bool),X0),s(fun(bool,bool),X1)))) )),
  inference(subsumption_resolution,[],[f1226,f38])).
fof(f1226,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))) | s(bool,'T_0') = s(bool,sK4(X1,X0,bool)) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(bool,bool),X0),s(fun(bool,bool),X1)))) )),
  inference(superposition,[],[f125,f39])).
fof(f414,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X11) != s(bool,f(s(X8,sK5(X9,X10,bool,X8)))) | s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK5(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f49,f50])).
fof(f412,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,f(s(X0,sK5(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,X3) = s(bool,g(s(X0,sK5(X1,X2,bool,X0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f49,f50])).
fof(f44,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X2)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f32])).
fof(f127,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,sK4(X1,X2,X0)),s(fun(X0,bool),X2)))) )),
  inference(subsumption_resolution,[],[f126,f38])).
fof(f126,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,sK4(X1,X2,X0)),s(fun(X0,bool),X2)))) )),
  inference(superposition,[],[f43,f39])).
fof(f377,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X27,bool),X29),s(fun(X27,bool),X28)))) | s(bool,X30) = s(bool,'const.bool.IN_2'(s(X27,sK4(X28,X29,X27)),s(fun(X27,bool),X28))) | s(bool,'T_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f239,f37])).
fof(f239,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X27,bool),X29),s(fun(X27,bool),X28)))) | s(bool,X30) = s(bool,'const.bool.IN_2'(s(X27,sK4(X28,X29,X27)),s(fun(X27,bool),X28))) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f44,f50])).
fof(f529,plain,(
  ( ! [X6,X4,X2,X5,X3] : (~p(s(bool,'const.bool.IN_2'(s(X2,X5),s(fun(X2,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X2,X5),s(fun(X2,bool),X4)))) | s(bool,X6) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),X4))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f522,f37])).
fof(f522,plain,(
  ( ! [X6,X4,X2,X5,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(X2,X5),s(fun(X2,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X2,X5),s(fun(X2,bool),X4)))) | s(bool,X6) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X2,bool),X3),s(fun(X2,bool),X4))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f42,f50])).
fof(f910,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,sK4(X1,X2,X0)),s(fun(X0,bool),X2))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) )),
  inference(condensation,[],[f909])).
fof(f909,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X8,sK4(X10,X9,X8)),s(fun(X8,bool),X9))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X8,bool),X9),s(fun(X8,bool),X10))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(subsumption_resolution,[],[f906,f52])).
fof(f906,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X11)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X8,sK4(X10,X9,X8)),s(fun(X8,bool),X9))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X8,bool),X9),s(fun(X8,bool),X10))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f127,f50])).
fof(f201,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,'const.bool.IN_2'(s(bool,X36),s(fun(bool,bool),X34)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(bool,bool),X35),s(fun(bool,bool),X34)))) | s(bool,'T_0') = s(bool,X36) | s(bool,'T_0') = s(bool,sK4(X34,X35,bool))) )),
  inference(superposition,[],[f44,f50])).
fof(f202,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,'const.bool.IN_2'(s(bool,X39),s(fun(bool,bool),X38)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(bool,bool),X38),s(fun(bool,bool),X37)))) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,sK4(X37,X38,bool))) )),
  inference(superposition,[],[f43,f50])).
fof(f240,plain,(
  ( ! [X33,X31,X32] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X31)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(bool,bool),X32),s(fun(bool,bool),X31)))) | s(bool,X33) = s(bool,sK4(X31,X32,bool)) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f44,f50])).
fof(f241,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X35)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(bool,bool),X35),s(fun(bool,bool),X34)))) | s(bool,X36) = s(bool,sK4(X34,X35,bool)) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f43,f50])).
fof(f40,plain,(
  ( ! [X4,X2,X0,X1] : (s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) )),
  inference(cnf_transformation,[],[f29])).
fof(f602,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK2)))),
  inference(unit_resulting_resolution,[],[f36,f41])).
fof(f1267,plain,(
  ( ! [X2,X1] : (p(s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1)))) | s(bool,X2) = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f1255,f37])).
fof(f1255,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1)))) | s(bool,X2) = s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f528,f50])).
fof(f1247,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2)))) )),
  inference(subsumption_resolution,[],[f1240,f38])).
fof(f1240,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2)))) )),
  inference(superposition,[],[f527,f39])).
fof(f528,plain,(
  ( ! [X1] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK2)))) | p(s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1))))) )),
  inference(subsumption_resolution,[],[f521,f37])).
fof(f521,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK2)))) | p(s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1))))) )),
  inference(superposition,[],[f42,f60])).
fof(f610,plain,(
  ( ! [X6,X8,X7] : (s(bool,'const.bool.IN_2'(s(bool,X8),s(fun(bool,bool),X6))) != s(bool,'const.bool.IN_2'(s(bool,X8),s(fun(bool,bool),X7))) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7) | s(bool,'T_0') = s(bool,sK3(X6,X7,bool)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f41,f50])).
fof(f608,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X0))) != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X1))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,X2) = s(bool,sK3(X0,X1,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f41,f50])).
fof(f601,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK2)))),
  inference(unit_resulting_resolution,[],[f36,f41])).
fof(f643,plain,(
  ( ! [X6,X7] : (s(bool,X6) = s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) ) | $spl22),
  inference(subsumption_resolution,[],[f631,f37])).
fof(f631,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,X6) = s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) ) | $spl22),
  inference(superposition,[],[f376,f539])).
fof(f615,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X11) != s(bool,'const.bool.IN_2'(s(X8,sK3(X9,X10,X8)),s(fun(X8,bool),X10))) | s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X8,sK3(X9,X10,X8)),s(fun(X8,bool),X9))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f41,f50])).
fof(f613,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK3(X1,X2,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(X0,sK3(X1,X2,X0)),s(fun(X0,bool),X1))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f41,f50])).
fof(f605,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X11) != s(bool,'const.bool.IN_2'(s(X8,sK3(X9,X10,X8)),s(fun(X8,bool),X9))) | s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X8,sK3(X9,X10,X8)),s(fun(X8,bool),X10))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f41,f50])).
fof(f603,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK3(X1,X2,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(X0,sK3(X1,X2,X0)),s(fun(X0,bool),X2))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f41,f50])).
fof(f586,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl22),
  inference(trivial_inequality_removal,[],[f580])).
fof(f580,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl22),
  inference(superposition,[],[f543,f50])).
fof(f579,plain,(
  ~p(s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0))))) | $spl22),
  inference(unit_resulting_resolution,[],[f37,f543,f48])).
fof(f543,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | $spl22),
  inference(backward_demodulation,[],[f539,f406])).
fof(f409,plain,(
  ( ! [X10,X8,X11,X9] : (s(X10,f(s(bool,X11))) != s(X10,g(s(bool,X11))) | s(fun(bool,X10),X8) = s(fun(bool,X10),X9) | s(bool,'T_0') = s(bool,sK5(X8,X9,X10,bool)) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f49,f50])).
fof(f407,plain,(
  ( ! [X2,X0,X3,X1] : (s(X2,f(s(bool,'T_0'))) != s(X2,g(s(bool,'T_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,X3) = s(bool,sK5(X0,X1,X2,bool)) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f49,f50])).
fof(f404,plain,(
  ( ! [X33,X32] : (s(bool,X32) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))) | s(bool,X32) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f391,f37])).
fof(f391,plain,(
  ( ! [X33,X32] : (~p(s(bool,'T_0')) | s(bool,X32) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))) | s(bool,X32) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f376,f60])).
fof(f403,plain,(
  ( ! [X30,X31] : (s(bool,X30) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | s(bool,X30) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(subsumption_resolution,[],[f390,f37])).
fof(f390,plain,(
  ( ! [X30,X31] : (~p(s(bool,'T_0')) | s(bool,X30) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | s(bool,X30) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f376,f58])).
fof(f166,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f50,f50])).
fof(f60,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))),
  inference(subsumption_resolution,[],[f55,f38])).
fof(f55,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))),
  inference(superposition,[],[f35,f39])).
fof(f35,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),
  inference(cnf_transformation,[],[f26])).
fof(f2685,plain,(
  ( ! [X17] : (s(bool,'F_0') = s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,X17)) ) | ($spl6 | $spl18)),
  inference(subsumption_resolution,[],[f2658,f52])).
fof(f2658,plain,(
  ( ! [X17] : (p(s(bool,X17)) | s(bool,'F_0') = s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,X17)) ) | ($spl6 | $spl18)),
  inference(superposition,[],[f79,f535])).
fof(f2486,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,X0) = s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0))))) ) | $spl18),
  inference(equality_factoring,[],[f535])).
fof(f2804,plain,(
  $false | ($spl6 | $spl18 | $spl22 | $spl24 | $spl33)),
  inference(subsumption_resolution,[],[f2803,f37])).
fof(f2803,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl18 | $spl22 | $spl24 | $spl33)),
  inference(forward_demodulation,[],[f1964,f2686])).
fof(f2800,plain,(
  $false | ($spl6 | $spl18 | $spl22 | $spl24 | $spl33)),
  inference(subsumption_resolution,[],[f2799,f37])).
fof(f2799,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl18 | $spl22 | $spl24 | $spl33)),
  inference(forward_demodulation,[],[f1955,f2686])).
fof(f1955,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK2)))) | $spl33),
  inference(unit_resulting_resolution,[],[f37,f674,f48])).
fof(f2798,plain,(
  $false | ($spl6 | $spl18 | $spl22 | $spl24 | $spl33)),
  inference(subsumption_resolution,[],[f2797,f37])).
fof(f2797,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl18 | $spl22 | $spl24 | $spl33)),
  inference(forward_demodulation,[],[f1953,f2686])).
fof(f1953,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK2)))) | $spl33),
  inference(unit_resulting_resolution,[],[f674,f674,f674,f376])).
fof(f2796,plain,(
  $false | ($spl6 | $spl12 | $spl18 | $spl22 | $spl24 | $spl33)),
  inference(subsumption_resolution,[],[f2795,f37])).
fof(f2795,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl12 | $spl18 | $spl22 | $spl24 | $spl33)),
  inference(forward_demodulation,[],[f1951,f2686])).
fof(f1951,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK2)))) | ($spl12 | $spl33)),
  inference(unit_resulting_resolution,[],[f430,f430,f674,f376])).
fof(f2794,plain,(
  $false | ($spl3 | $spl6 | $spl18 | $spl22 | $spl24 | $spl33)),
  inference(subsumption_resolution,[],[f2793,f37])).
fof(f2793,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl6 | $spl18 | $spl22 | $spl24 | $spl33)),
  inference(forward_demodulation,[],[f1950,f2686])).
fof(f1950,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK2)))) | ($spl3 | $spl33)),
  inference(unit_resulting_resolution,[],[f64,f64,f674,f376])).
fof(f2792,plain,(
  $false | ($spl6 | $spl18 | $spl22 | $spl24 | $spl33)),
  inference(subsumption_resolution,[],[f2791,f37])).
fof(f2791,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl18 | $spl22 | $spl24 | $spl33)),
  inference(forward_demodulation,[],[f1949,f2686])).
fof(f1949,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK2)))) | $spl33),
  inference(unit_resulting_resolution,[],[f674,f674,f674,f376])).
fof(f2789,plain,(
  $false | ($spl6 | $spl18 | $spl22 | $spl24 | $spl33)),
  inference(subsumption_resolution,[],[f2788,f37])).
fof(f2788,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl18 | $spl22 | $spl24 | $spl33)),
  inference(forward_demodulation,[],[f1943,f2686])).
fof(f1943,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK2)))) | $spl33),
  inference(unit_resulting_resolution,[],[f674,f674,f674,f376])).
fof(f2784,plain,(
  $false | ($spl6 | $spl18 | $spl22 | $spl24 | $spl33)),
  inference(subsumption_resolution,[],[f2783,f37])).
fof(f2783,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl18 | $spl22 | $spl24 | $spl33)),
  inference(forward_demodulation,[],[f1932,f2686])).
fof(f1932,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK2)))) | $spl33),
  inference(unit_resulting_resolution,[],[f674,f52])).
fof(f2775,plain,(
  $false | ($spl6 | $spl18 | $spl22 | $spl24 | $spl43)),
  inference(subsumption_resolution,[],[f2774,f37])).
fof(f2774,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl18 | $spl22 | $spl24 | $spl43)),
  inference(forward_demodulation,[],[f2323,f2686])).
fof(f2323,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK2)))) | $spl43),
  inference(unit_resulting_resolution,[],[f37,f786,f48])).
fof(f2771,plain,(
  $false | ($spl6 | $spl18 | $spl22 | $spl24 | $spl43)),
  inference(subsumption_resolution,[],[f2770,f37])).
fof(f2770,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl18 | $spl22 | $spl24 | $spl43)),
  inference(forward_demodulation,[],[f2313,f2686])).
fof(f2313,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK2)))) | $spl43),
  inference(unit_resulting_resolution,[],[f37,f786,f48])).
fof(f2769,plain,(
  $false | ($spl6 | $spl18 | $spl22 | $spl24 | $spl43)),
  inference(subsumption_resolution,[],[f2768,f37])).
fof(f2768,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl18 | $spl22 | $spl24 | $spl43)),
  inference(forward_demodulation,[],[f2311,f2686])).
fof(f2311,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK2)))) | $spl43),
  inference(unit_resulting_resolution,[],[f786,f786,f786,f376])).
fof(f2767,plain,(
  $false | ($spl6 | $spl18 | $spl22 | $spl24 | $spl33 | $spl43)),
  inference(subsumption_resolution,[],[f2766,f37])).
fof(f2766,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl18 | $spl22 | $spl24 | $spl33 | $spl43)),
  inference(forward_demodulation,[],[f2310,f2686])).
fof(f2310,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK2)))) | ($spl33 | $spl43)),
  inference(unit_resulting_resolution,[],[f674,f674,f786,f376])).
fof(f2765,plain,(
  $false | ($spl6 | $spl12 | $spl18 | $spl22 | $spl24 | $spl43)),
  inference(subsumption_resolution,[],[f2764,f37])).
fof(f2764,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl12 | $spl18 | $spl22 | $spl24 | $spl43)),
  inference(forward_demodulation,[],[f2308,f2686])).
fof(f2308,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK2)))) | ($spl12 | $spl43)),
  inference(unit_resulting_resolution,[],[f430,f430,f786,f376])).
fof(f2763,plain,(
  $false | ($spl3 | $spl6 | $spl18 | $spl22 | $spl24 | $spl43)),
  inference(subsumption_resolution,[],[f2762,f37])).
fof(f2762,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl6 | $spl18 | $spl22 | $spl24 | $spl43)),
  inference(forward_demodulation,[],[f2307,f2686])).
fof(f2307,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK2)))) | ($spl3 | $spl43)),
  inference(unit_resulting_resolution,[],[f64,f64,f786,f376])).
fof(f2761,plain,(
  $false | ($spl6 | $spl18 | $spl22 | $spl24 | $spl43)),
  inference(subsumption_resolution,[],[f2760,f37])).
fof(f2760,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl18 | $spl22 | $spl24 | $spl43)),
  inference(forward_demodulation,[],[f2306,f2686])).
fof(f2306,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK2)))) | $spl43),
  inference(unit_resulting_resolution,[],[f786,f786,f786,f376])).
fof(f2759,plain,(
  $false | ($spl6 | $spl18 | $spl22 | $spl24 | $spl33 | $spl43)),
  inference(subsumption_resolution,[],[f2758,f37])).
fof(f2758,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl18 | $spl22 | $spl24 | $spl33 | $spl43)),
  inference(forward_demodulation,[],[f2305,f2686])).
fof(f2305,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK2)))) | ($spl33 | $spl43)),
  inference(unit_resulting_resolution,[],[f674,f786,f786,f376])).
fof(f2756,plain,(
  $false | ($spl6 | $spl18 | $spl22 | $spl24 | $spl43)),
  inference(subsumption_resolution,[],[f2755,f37])).
fof(f2755,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl18 | $spl22 | $spl24 | $spl43)),
  inference(forward_demodulation,[],[f2299,f2686])).
fof(f2299,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK2)))) | $spl43),
  inference(unit_resulting_resolution,[],[f786,f786,f786,f376])).
fof(f2754,plain,(
  $false | ($spl6 | $spl18 | $spl22 | $spl24 | $spl33 | $spl43)),
  inference(subsumption_resolution,[],[f2753,f37])).
fof(f2753,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl18 | $spl22 | $spl24 | $spl33 | $spl43)),
  inference(forward_demodulation,[],[f2298,f2686])).
fof(f2298,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK2)))) | ($spl33 | $spl43)),
  inference(unit_resulting_resolution,[],[f674,f786,f786,f376])).
fof(f2748,plain,(
  $false | ($spl6 | $spl18 | $spl22 | $spl24 | $spl43)),
  inference(subsumption_resolution,[],[f2747,f37])).
fof(f2747,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl18 | $spl22 | $spl24 | $spl43)),
  inference(forward_demodulation,[],[f2286,f2686])).
fof(f2286,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK2)))) | $spl43),
  inference(unit_resulting_resolution,[],[f786,f52])).
fof(f2738,plain,(
  $false | ($spl6 | $spl18 | $spl22 | $spl24 | $spl43)),
  inference(subsumption_resolution,[],[f2737,f2686])).
fof(f2737,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | ($spl6 | $spl18 | $spl22 | $spl24 | $spl43)),
  inference(forward_demodulation,[],[f2536,f2686])).
fof(f2536,plain,(
  s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl24 | $spl43)),
  inference(unit_resulting_resolution,[],[f786,f540])).
fof(f2736,plain,(
  $false | ($spl6 | $spl18 | $spl22 | $spl24 | $spl33)),
  inference(subsumption_resolution,[],[f2735,f2686])).
fof(f2735,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | ($spl6 | $spl18 | $spl22 | $spl24 | $spl33)),
  inference(forward_demodulation,[],[f2535,f2686])).
fof(f2535,plain,(
  s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl24 | $spl33)),
  inference(unit_resulting_resolution,[],[f674,f540])).
fof(f2731,plain,(
  $false | ($spl6 | $spl18 | $spl22 | $spl24)),
  inference(subsumption_resolution,[],[f2730,f2686])).
fof(f2730,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl6 | $spl18 | $spl22 | $spl24)),
  inference(forward_demodulation,[],[f2603,f2686])).
fof(f2729,plain,(
  $false | ($spl6 | $spl18 | $spl22 | $spl24)),
  inference(subsumption_resolution,[],[f2728,f2686])).
fof(f2728,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl6 | $spl18 | $spl22 | $spl24)),
  inference(forward_demodulation,[],[f2602,f2686])).
fof(f2720,plain,(
  $false | ($spl6 | $spl18 | $spl22 | $spl24)),
  inference(subsumption_resolution,[],[f2687,f37])).
fof(f2687,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl18 | $spl22 | $spl24)),
  inference(backward_demodulation,[],[f2686,f38])).
fof(f2680,plain,(
  $false | ($spl3 | $spl12 | $spl27 | $spl33 | $spl40 | $spl43)),
  inference(subsumption_resolution,[],[f2679,f37])).
fof(f2679,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl12 | $spl27 | $spl33 | $spl40 | $spl43)),
  inference(forward_demodulation,[],[f2678,f784])).
fof(f2678,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1)))) | ($spl3 | $spl12 | $spl27 | $spl33 | $spl43)),
  inference(subsumption_resolution,[],[f2667,f38])).
fof(f2667,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1)))) | ($spl3 | $spl12 | $spl27 | $spl33 | $spl43)),
  inference(superposition,[],[f527,f2332])).
fof(f2332,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl12 | $spl27 | $spl33 | $spl43)),
  inference(forward_demodulation,[],[f2318,f1973])).
fof(f1973,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl12 | $spl27 | $spl33)),
  inference(forward_demodulation,[],[f1959,f1589])).
fof(f1589,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0)))) | ($spl3 | $spl12 | $spl27)),
  inference(forward_demodulation,[],[f1575,f454])).
fof(f454,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) | ($spl3 | $spl12)),
  inference(unit_resulting_resolution,[],[f37,f64,f430,f376])).
fof(f1575,plain,(
  s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) = s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0)))) | ($spl12 | $spl27)),
  inference(unit_resulting_resolution,[],[f37,f430,f1536,f376])).
fof(f1536,plain,(
  s(bool,'T_0') != s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0)))) | $spl27),
  inference(backward_demodulation,[],[f1527,f406])).
fof(f1527,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | $spl27),
  inference(unit_resulting_resolution,[],[f38,f542,f80])).
fof(f1959,plain,(
  s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl27 | $spl33)),
  inference(unit_resulting_resolution,[],[f37,f1536,f674,f376])).
fof(f2677,plain,(
  $false | ($spl3 | $spl12 | $spl27 | $spl30 | $spl33)),
  inference(subsumption_resolution,[],[f2676,f37])).
fof(f2676,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl12 | $spl27 | $spl30 | $spl33)),
  inference(forward_demodulation,[],[f2675,f672])).
fof(f2675,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) | ($spl3 | $spl12 | $spl27 | $spl33)),
  inference(subsumption_resolution,[],[f2666,f38])).
fof(f2666,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) | ($spl3 | $spl12 | $spl27 | $spl33)),
  inference(superposition,[],[f527,f1973])).
fof(f2151,plain,(
  $false | ($spl41 | $spl44)),
  inference(subsumption_resolution,[],[f2136,f37])).
fof(f2136,plain,(
  ~p(s(bool,'T_0')) | ($spl41 | $spl44)),
  inference(backward_demodulation,[],[f2106,f2059])).
fof(f2059,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1)))) | $spl41),
  inference(unit_resulting_resolution,[],[f37,f785,f48])).
fof(f2106,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl41 | $spl44)),
  inference(subsumption_resolution,[],[f2067,f788])).
fof(f2067,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl41),
  inference(trivial_inequality_removal,[],[f2061])).
fof(f2061,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl41),
  inference(superposition,[],[f785,f50])).
fof(f2145,plain,(
  $false | ($spl41 | $spl44)),
  inference(subsumption_resolution,[],[f2107,f37])).
fof(f2107,plain,(
  ~p(s(bool,'T_0')) | ($spl41 | $spl44)),
  inference(backward_demodulation,[],[f2106,f38])).
fof(f2105,plain,(
  $false | ($spl41 | $spl47)),
  inference(subsumption_resolution,[],[f2009,f790])).
fof(f2009,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | $spl41),
  inference(unit_resulting_resolution,[],[f785,f39])).
fof(f2104,plain,(
  $false | ($spl3 | $spl41 | $spl47)),
  inference(subsumption_resolution,[],[f2010,f790])).
fof(f2010,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl41)),
  inference(unit_resulting_resolution,[],[f64,f785,f50])).
fof(f2103,plain,(
  $false | ($spl3 | $spl12 | $spl41 | $spl47)),
  inference(subsumption_resolution,[],[f2102,f790])).
fof(f2102,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl12 | $spl41)),
  inference(forward_demodulation,[],[f2011,f454])).
fof(f2011,plain,(
  s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl12 | $spl41)),
  inference(unit_resulting_resolution,[],[f430,f785,f50])).
fof(f2101,plain,(
  $false | ($spl3 | $spl12 | $spl27 | $spl41 | $spl47)),
  inference(subsumption_resolution,[],[f2100,f790])).
fof(f2100,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl12 | $spl27 | $spl41)),
  inference(forward_demodulation,[],[f2012,f1589])).
fof(f2012,plain,(
  s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl27 | $spl41)),
  inference(unit_resulting_resolution,[],[f1536,f785,f50])).
fof(f2099,plain,(
  $false | ($spl3 | $spl12 | $spl27 | $spl33 | $spl41 | $spl47)),
  inference(subsumption_resolution,[],[f2098,f790])).
fof(f2098,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl12 | $spl27 | $spl33 | $spl41)),
  inference(forward_demodulation,[],[f2013,f1973])).
fof(f2013,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl33 | $spl41)),
  inference(unit_resulting_resolution,[],[f674,f785,f50])).
fof(f2097,plain,(
  $false | ($spl3 | $spl41 | $spl47)),
  inference(subsumption_resolution,[],[f2016,f790])).
fof(f2016,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl41)),
  inference(unit_resulting_resolution,[],[f64,f785,f50])).
fof(f2096,plain,(
  $false | ($spl3 | $spl12 | $spl41 | $spl47)),
  inference(subsumption_resolution,[],[f2095,f790])).
fof(f2095,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl12 | $spl41)),
  inference(forward_demodulation,[],[f2017,f454])).
fof(f2017,plain,(
  s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl12 | $spl41)),
  inference(unit_resulting_resolution,[],[f430,f785,f50])).
fof(f2094,plain,(
  $false | ($spl3 | $spl12 | $spl27 | $spl41 | $spl47)),
  inference(subsumption_resolution,[],[f2093,f790])).
fof(f2093,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl12 | $spl27 | $spl41)),
  inference(forward_demodulation,[],[f2018,f1589])).
fof(f2018,plain,(
  s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl27 | $spl41)),
  inference(unit_resulting_resolution,[],[f1536,f785,f50])).
fof(f2092,plain,(
  $false | ($spl3 | $spl12 | $spl27 | $spl33 | $spl41 | $spl47)),
  inference(subsumption_resolution,[],[f2091,f790])).
fof(f2091,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl12 | $spl27 | $spl33 | $spl41)),
  inference(forward_demodulation,[],[f2019,f1973])).
fof(f2019,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl33 | $spl41)),
  inference(unit_resulting_resolution,[],[f674,f785,f50])).
fof(f2090,plain,(
  $false | ($spl41 | $spl47)),
  inference(subsumption_resolution,[],[f2023,f790])).
fof(f2023,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | $spl41),
  inference(unit_resulting_resolution,[],[f38,f785,f80])).
fof(f2089,plain,(
  $false | ($spl3 | $spl41 | $spl47)),
  inference(subsumption_resolution,[],[f2025,f790])).
fof(f2025,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl41)),
  inference(unit_resulting_resolution,[],[f37,f64,f785,f376])).
fof(f2088,plain,(
  $false | ($spl3 | $spl12 | $spl41 | $spl47)),
  inference(subsumption_resolution,[],[f2087,f790])).
fof(f2087,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl12 | $spl41)),
  inference(forward_demodulation,[],[f2026,f454])).
fof(f2026,plain,(
  s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl12 | $spl41)),
  inference(unit_resulting_resolution,[],[f37,f430,f785,f376])).
fof(f2086,plain,(
  $false | ($spl3 | $spl12 | $spl27 | $spl41 | $spl47)),
  inference(subsumption_resolution,[],[f2085,f790])).
fof(f2085,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl12 | $spl27 | $spl41)),
  inference(forward_demodulation,[],[f2027,f1589])).
fof(f2027,plain,(
  s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl27 | $spl41)),
  inference(unit_resulting_resolution,[],[f37,f1536,f785,f376])).
fof(f2084,plain,(
  $false | ($spl3 | $spl12 | $spl27 | $spl33 | $spl41 | $spl47)),
  inference(subsumption_resolution,[],[f2083,f790])).
fof(f2083,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl12 | $spl27 | $spl33 | $spl41)),
  inference(forward_demodulation,[],[f2028,f1973])).
fof(f2028,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl33 | $spl41)),
  inference(unit_resulting_resolution,[],[f37,f674,f785,f376])).
fof(f2079,plain,(
  $false | ($spl3 | $spl12 | $spl41 | $spl47)),
  inference(subsumption_resolution,[],[f2078,f790])).
fof(f2078,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl12 | $spl41)),
  inference(forward_demodulation,[],[f2036,f454])).
fof(f2036,plain,(
  s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl12 | $spl41)),
  inference(unit_resulting_resolution,[],[f785,f466])).
fof(f2074,plain,(
  $false | ($spl3 | $spl41 | $spl47)),
  inference(subsumption_resolution,[],[f2051,f790])).
fof(f2051,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl41)),
  inference(unit_resulting_resolution,[],[f37,f64,f785,f376])).
fof(f2073,plain,(
  $false | ($spl3 | $spl12 | $spl41 | $spl47)),
  inference(subsumption_resolution,[],[f2072,f790])).
fof(f2072,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl12 | $spl41)),
  inference(forward_demodulation,[],[f2052,f454])).
fof(f2052,plain,(
  s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl12 | $spl41)),
  inference(unit_resulting_resolution,[],[f37,f430,f785,f376])).
fof(f2071,plain,(
  $false | ($spl3 | $spl12 | $spl27 | $spl41 | $spl47)),
  inference(subsumption_resolution,[],[f2070,f790])).
fof(f2070,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl12 | $spl27 | $spl41)),
  inference(forward_demodulation,[],[f2053,f1589])).
fof(f2053,plain,(
  s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl27 | $spl41)),
  inference(unit_resulting_resolution,[],[f37,f1536,f785,f376])).
fof(f2069,plain,(
  $false | ($spl3 | $spl12 | $spl27 | $spl33 | $spl41 | $spl47)),
  inference(subsumption_resolution,[],[f2068,f790])).
fof(f2068,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl12 | $spl27 | $spl33 | $spl41)),
  inference(forward_demodulation,[],[f2054,f1973])).
fof(f2054,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | ($spl33 | $spl41)),
  inference(unit_resulting_resolution,[],[f37,f674,f785,f376])).
fof(f1814,plain,(
  $false | ($spl31 | $spl34)),
  inference(subsumption_resolution,[],[f1798,f37])).
fof(f1798,plain,(
  ~p(s(bool,'T_0')) | ($spl31 | $spl34)),
  inference(backward_demodulation,[],[f1768,f1731])).
fof(f1731,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl31),
  inference(unit_resulting_resolution,[],[f37,f673,f48])).
fof(f1768,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl31 | $spl34)),
  inference(subsumption_resolution,[],[f1739,f676])).
fof(f1739,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl31),
  inference(trivial_inequality_removal,[],[f1733])).
fof(f1733,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl31),
  inference(superposition,[],[f673,f50])).
fof(f1808,plain,(
  $false | ($spl31 | $spl34)),
  inference(subsumption_resolution,[],[f1769,f37])).
fof(f1769,plain,(
  ~p(s(bool,'T_0')) | ($spl31 | $spl34)),
  inference(backward_demodulation,[],[f1768,f38])).
fof(f1767,plain,(
  $false | ($spl31 | $spl37)),
  inference(subsumption_resolution,[],[f1688,f678])).
fof(f1688,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl31),
  inference(unit_resulting_resolution,[],[f673,f39])).
fof(f1766,plain,(
  $false | ($spl3 | $spl31 | $spl37)),
  inference(subsumption_resolution,[],[f1689,f678])).
fof(f1689,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl31)),
  inference(unit_resulting_resolution,[],[f64,f673,f50])).
fof(f1765,plain,(
  $false | ($spl3 | $spl12 | $spl31 | $spl37)),
  inference(subsumption_resolution,[],[f1764,f678])).
fof(f1764,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl12 | $spl31)),
  inference(forward_demodulation,[],[f1690,f454])).
fof(f1690,plain,(
  s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl12 | $spl31)),
  inference(unit_resulting_resolution,[],[f430,f673,f50])).
fof(f1763,plain,(
  $false | ($spl3 | $spl12 | $spl27 | $spl31 | $spl37)),
  inference(subsumption_resolution,[],[f1762,f678])).
fof(f1762,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl12 | $spl27 | $spl31)),
  inference(forward_demodulation,[],[f1691,f1589])).
fof(f1691,plain,(
  s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl27 | $spl31)),
  inference(unit_resulting_resolution,[],[f1536,f673,f50])).
fof(f1761,plain,(
  $false | ($spl3 | $spl31 | $spl37)),
  inference(subsumption_resolution,[],[f1694,f678])).
fof(f1694,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl31)),
  inference(unit_resulting_resolution,[],[f64,f673,f50])).
fof(f1760,plain,(
  $false | ($spl3 | $spl12 | $spl31 | $spl37)),
  inference(subsumption_resolution,[],[f1759,f678])).
fof(f1759,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl12 | $spl31)),
  inference(forward_demodulation,[],[f1695,f454])).
fof(f1695,plain,(
  s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl12 | $spl31)),
  inference(unit_resulting_resolution,[],[f430,f673,f50])).
fof(f1758,plain,(
  $false | ($spl3 | $spl12 | $spl27 | $spl31 | $spl37)),
  inference(subsumption_resolution,[],[f1757,f678])).
fof(f1757,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl12 | $spl27 | $spl31)),
  inference(forward_demodulation,[],[f1696,f1589])).
fof(f1696,plain,(
  s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl27 | $spl31)),
  inference(unit_resulting_resolution,[],[f1536,f673,f50])).
fof(f1756,plain,(
  $false | ($spl31 | $spl37)),
  inference(subsumption_resolution,[],[f1700,f678])).
fof(f1700,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl31),
  inference(unit_resulting_resolution,[],[f38,f673,f80])).
fof(f1755,plain,(
  $false | ($spl3 | $spl31 | $spl37)),
  inference(subsumption_resolution,[],[f1702,f678])).
fof(f1702,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl31)),
  inference(unit_resulting_resolution,[],[f37,f64,f673,f376])).
fof(f1754,plain,(
  $false | ($spl3 | $spl12 | $spl31 | $spl37)),
  inference(subsumption_resolution,[],[f1753,f678])).
fof(f1753,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl12 | $spl31)),
  inference(forward_demodulation,[],[f1703,f454])).
fof(f1703,plain,(
  s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl12 | $spl31)),
  inference(unit_resulting_resolution,[],[f37,f430,f673,f376])).
fof(f1752,plain,(
  $false | ($spl3 | $spl12 | $spl27 | $spl31 | $spl37)),
  inference(subsumption_resolution,[],[f1751,f678])).
fof(f1751,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl12 | $spl27 | $spl31)),
  inference(forward_demodulation,[],[f1704,f1589])).
fof(f1704,plain,(
  s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl27 | $spl31)),
  inference(unit_resulting_resolution,[],[f37,f1536,f673,f376])).
fof(f1748,plain,(
  $false | ($spl3 | $spl12 | $spl31 | $spl37)),
  inference(subsumption_resolution,[],[f1747,f678])).
fof(f1747,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl12 | $spl31)),
  inference(forward_demodulation,[],[f1711,f454])).
fof(f1711,plain,(
  s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl12 | $spl31)),
  inference(unit_resulting_resolution,[],[f673,f466])).
fof(f1744,plain,(
  $false | ($spl3 | $spl31 | $spl37)),
  inference(subsumption_resolution,[],[f1724,f678])).
fof(f1724,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl31)),
  inference(unit_resulting_resolution,[],[f37,f64,f673,f376])).
fof(f1743,plain,(
  $false | ($spl3 | $spl12 | $spl31 | $spl37)),
  inference(subsumption_resolution,[],[f1742,f678])).
fof(f1742,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl12 | $spl31)),
  inference(forward_demodulation,[],[f1725,f454])).
fof(f1725,plain,(
  s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl12 | $spl31)),
  inference(unit_resulting_resolution,[],[f37,f430,f673,f376])).
fof(f1741,plain,(
  $false | ($spl3 | $spl12 | $spl27 | $spl31 | $spl37)),
  inference(subsumption_resolution,[],[f1740,f678])).
fof(f1740,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl12 | $spl27 | $spl31)),
  inference(forward_demodulation,[],[f1726,f1589])).
fof(f1726,plain,(
  s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl27 | $spl31)),
  inference(unit_resulting_resolution,[],[f37,f1536,f673,f376])).
fof(f535,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl18),
  inference(cnf_transformation,[],[f535_D])).
fof(f535_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f537,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | $spl21),
  inference(cnf_transformation,[],[f537_D])).
fof(f537_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f1363,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,X0) = s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f406,f50])).
fof(f1366,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,X3) = s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f406,f50])).
fof(f1480,plain,(
  $false | ($spl21 | $spl24)),
  inference(subsumption_resolution,[],[f1454,f37])).
fof(f1454,plain,(
  ~p(s(bool,'T_0')) | ($spl21 | $spl24)),
  inference(backward_demodulation,[],[f1439,f1411])).
fof(f1411,plain,(
  ~p(s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0))))) | $spl21),
  inference(unit_resulting_resolution,[],[f37,f537,f48])).
fof(f1439,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl21 | $spl24)),
  inference(subsumption_resolution,[],[f1418,f540])).
fof(f1418,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl21),
  inference(trivial_inequality_removal,[],[f1413])).
fof(f1413,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl21),
  inference(superposition,[],[f537,f50])).
fof(f1479,plain,(
  $false | ($spl21 | $spl24)),
  inference(subsumption_resolution,[],[f1440,f37])).
fof(f1440,plain,(
  ~p(s(bool,'T_0')) | ($spl21 | $spl24)),
  inference(backward_demodulation,[],[f1439,f38])).
fof(f1438,plain,(
  $false | ($spl21 | $spl27)),
  inference(subsumption_resolution,[],[f1369,f542])).
fof(f1369,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | $spl21),
  inference(unit_resulting_resolution,[],[f537,f39])).
fof(f1437,plain,(
  $false | ($spl3 | $spl21 | $spl27)),
  inference(subsumption_resolution,[],[f1370,f542])).
fof(f1370,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | ($spl3 | $spl21)),
  inference(unit_resulting_resolution,[],[f64,f537,f50])).
fof(f1436,plain,(
  $false | ($spl3 | $spl12 | $spl21 | $spl27)),
  inference(subsumption_resolution,[],[f1435,f542])).
fof(f1435,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | ($spl3 | $spl12 | $spl21)),
  inference(forward_demodulation,[],[f1371,f454])).
fof(f1371,plain,(
  s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) = s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) | ($spl12 | $spl21)),
  inference(unit_resulting_resolution,[],[f430,f537,f50])).
fof(f1434,plain,(
  $false | ($spl3 | $spl21 | $spl27)),
  inference(subsumption_resolution,[],[f1375,f542])).
fof(f1375,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | ($spl3 | $spl21)),
  inference(unit_resulting_resolution,[],[f64,f537,f50])).
fof(f1433,plain,(
  $false | ($spl3 | $spl12 | $spl21 | $spl27)),
  inference(subsumption_resolution,[],[f1432,f542])).
fof(f1432,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | ($spl3 | $spl12 | $spl21)),
  inference(forward_demodulation,[],[f1376,f454])).
fof(f1376,plain,(
  s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) = s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) | ($spl12 | $spl21)),
  inference(unit_resulting_resolution,[],[f430,f537,f50])).
fof(f1431,plain,(
  $false | ($spl21 | $spl27)),
  inference(subsumption_resolution,[],[f1381,f542])).
fof(f1381,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | $spl21),
  inference(unit_resulting_resolution,[],[f38,f537,f80])).
fof(f1429,plain,(
  $false | ($spl3 | $spl21 | $spl27)),
  inference(subsumption_resolution,[],[f1384,f542])).
fof(f1384,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | ($spl3 | $spl21)),
  inference(unit_resulting_resolution,[],[f37,f64,f537,f376])).
fof(f1428,plain,(
  $false | ($spl3 | $spl12 | $spl21 | $spl27)),
  inference(subsumption_resolution,[],[f1427,f542])).
fof(f1427,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | ($spl3 | $spl12 | $spl21)),
  inference(forward_demodulation,[],[f1385,f454])).
fof(f1385,plain,(
  s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) = s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) | ($spl12 | $spl21)),
  inference(unit_resulting_resolution,[],[f37,f430,f537,f376])).
fof(f1425,plain,(
  $false | ($spl3 | $spl12 | $spl21 | $spl27)),
  inference(subsumption_resolution,[],[f1424,f542])).
fof(f1424,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | ($spl3 | $spl12 | $spl21)),
  inference(forward_demodulation,[],[f1393,f454])).
fof(f1393,plain,(
  s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) = s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) | ($spl12 | $spl21)),
  inference(unit_resulting_resolution,[],[f537,f466])).
fof(f1422,plain,(
  $false | ($spl3 | $spl21 | $spl27)),
  inference(subsumption_resolution,[],[f1404,f542])).
fof(f1404,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | ($spl3 | $spl21)),
  inference(unit_resulting_resolution,[],[f37,f64,f537,f376])).
fof(f1421,plain,(
  $false | ($spl3 | $spl12 | $spl21 | $spl27)),
  inference(subsumption_resolution,[],[f1420,f542])).
fof(f1420,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | ($spl3 | $spl12 | $spl21)),
  inference(forward_demodulation,[],[f1405,f454])).
fof(f1405,plain,(
  s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) = s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) | ($spl12 | $spl21)),
  inference(unit_resulting_resolution,[],[f37,f430,f537,f376])).
fof(f1361,plain,(
  $false | ($spl32 | $spl43)),
  inference(global_subsumption,[],[f786,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f675,f679,f729,f775,f778,f779,f780,f781,f782,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f406,f722,f713,f711,f707,f702,f701,f699,f688,f1261,f602])).
fof(f1261,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl32),
  inference(subsumption_resolution,[],[f1252,f37])).
fof(f1252,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl32),
  inference(superposition,[],[f528,f675])).
fof(f688,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl32),
  inference(unit_resulting_resolution,[],[f679,f679,f679,f376])).
fof(f699,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl32),
  inference(unit_resulting_resolution,[],[f38,f679,f80])).
fof(f701,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl32),
  inference(unit_resulting_resolution,[],[f679,f52])).
fof(f702,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl32),
  inference(unit_resulting_resolution,[],[f679,f39])).
fof(f707,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl32),
  inference(unit_resulting_resolution,[],[f679,f679,f679,f376])).
fof(f711,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl32),
  inference(unit_resulting_resolution,[],[f679,f679,f679,f376])).
fof(f713,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl32),
  inference(unit_resulting_resolution,[],[f37,f679,f48])).
fof(f722,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl32),
  inference(unit_resulting_resolution,[],[f37,f679,f48])).
fof(f775,plain,(
  ( ! [X6,X7] : (s(bool,X6) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) ) | $spl32),
  inference(subsumption_resolution,[],[f764,f37])).
fof(f764,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,X6) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) ) | $spl32),
  inference(superposition,[],[f376,f675])).
fof(f729,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl32),
  inference(trivial_inequality_removal,[],[f723])).
fof(f723,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl32),
  inference(superposition,[],[f679,f50])).
fof(f679,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl32),
  inference(backward_demodulation,[],[f675,f601])).
fof(f1360,plain,(
  $false | ($spl32 | $spl44)),
  inference(global_subsumption,[],[f788,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f675,f679,f729,f775,f778,f779,f780,f781,f782,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f406,f722,f713,f711,f707,f702,f701,f699,f688,f1261,f602])).
fof(f1359,plain,(
  $false | ($spl32 | $spl47)),
  inference(global_subsumption,[],[f790,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f675,f679,f729,f775,f778,f779,f780,f781,f782,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f406,f722,f713,f711,f707,f702,f701,f699,f688,f1261,f602])).
fof(f1358,plain,(
  $false | $spl32),
  inference(global_subsumption,[],[f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f675,f679,f729,f775,f778,f779,f780,f781,f782,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f406,f722,f713,f711,f707,f702,f701,f699,f688,f1261,f602])).
fof(f1357,plain,(
  $false | ($spl32 | $spl41)),
  inference(global_subsumption,[],[f785,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f675,f679,f729,f775,f778,f779,f780,f781,f782,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f406,f722,f713,f711,f707,f702,f701,f699,f688,f1261])).
fof(f1356,plain,(
  $false | $spl32),
  inference(global_subsumption,[],[f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f675,f679,f729,f775,f778,f779,f780,f781,f782,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f406,f722,f713,f711,f707,f702,f701,f699,f688,f1261])).
fof(f1355,plain,(
  $false | ($spl3 | $spl12 | $spl32)),
  inference(global_subsumption,[],[f1274,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f675,f679,f729,f775,f778,f779,f780,f781,f782,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1261,f406,f722,f713,f711,f707,f702,f701,f699,f688])).
fof(f1274,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl12 | $spl32)),
  inference(forward_demodulation,[],[f976,f454])).
fof(f976,plain,(
  s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl12 | $spl32)),
  inference(unit_resulting_resolution,[],[f679,f466])).
fof(f1354,plain,(
  $false | ($spl3 | $spl32)),
  inference(global_subsumption,[],[f680,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f675,f679,f729,f775,f778,f779,f780,f781,f782,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1261,f406,f722,f713,f711,f707,f702,f701,f699,f688])).
fof(f680,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl32)),
  inference(unit_resulting_resolution,[],[f37,f64,f679,f376])).
fof(f1353,plain,(
  $false | ($spl3 | $spl12 | $spl32)),
  inference(global_subsumption,[],[f1308,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f675,f679,f729,f775,f778,f779,f780,f781,f782,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1261,f406,f722,f713,f711,f707,f702,f701,f699,f688])).
fof(f1308,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl12 | $spl32)),
  inference(forward_demodulation,[],[f681,f454])).
fof(f681,plain,(
  s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl12 | $spl32)),
  inference(unit_resulting_resolution,[],[f37,f430,f679,f376])).
fof(f1352,plain,(
  $false | $spl32),
  inference(global_subsumption,[],[f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f675,f679,f729,f775,f778,f779,f780,f781,f782,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1261,f406,f722,f713,f711,f707,f702,f701,f699,f688])).
fof(f1351,plain,(
  $false | ($spl3 | $spl32)),
  inference(global_subsumption,[],[f689,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f675,f679,f729,f775,f778,f779,f780,f781,f782,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1261,f688,f406,f722,f713,f711,f707,f702,f701,f699])).
fof(f689,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl32)),
  inference(unit_resulting_resolution,[],[f64,f679,f50])).
fof(f1350,plain,(
  $false | ($spl3 | $spl12 | $spl32)),
  inference(global_subsumption,[],[f1312,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f675,f679,f729,f775,f778,f779,f780,f781,f782,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1261,f688,f406,f722,f713,f711,f707,f702,f701,f699])).
fof(f1312,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl12 | $spl32)),
  inference(forward_demodulation,[],[f690,f454])).
fof(f690,plain,(
  s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl12 | $spl32)),
  inference(unit_resulting_resolution,[],[f430,f679,f50])).
fof(f1349,plain,(
  $false | ($spl3 | $spl32)),
  inference(global_subsumption,[],[f694,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f675,f679,f729,f775,f778,f779,f780,f781,f782,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1261,f688,f406,f722,f713,f711,f707,f702,f701,f699])).
fof(f694,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl32)),
  inference(unit_resulting_resolution,[],[f64,f679,f50])).
fof(f1348,plain,(
  $false | ($spl3 | $spl12 | $spl32)),
  inference(global_subsumption,[],[f1315,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f675,f679,f729,f775,f778,f779,f780,f781,f782,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1261,f688,f406,f722,f713,f711,f707,f702,f701,f699])).
fof(f1315,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl12 | $spl32)),
  inference(forward_demodulation,[],[f695,f454])).
fof(f695,plain,(
  s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl12 | $spl32)),
  inference(unit_resulting_resolution,[],[f430,f679,f50])).
fof(f1347,plain,(
  $false | $spl32),
  inference(global_subsumption,[],[f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f675,f679,f729,f775,f778,f779,f780,f781,f782,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1261,f688,f406,f722,f713,f711,f707,f702,f701,f699])).
fof(f1346,plain,(
  $false | $spl32),
  inference(global_subsumption,[],[f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f675,f679,f729,f775,f778,f779,f780,f781,f782,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1261,f688,f699,f406,f722,f713,f711,f707,f702,f701])).
fof(f1345,plain,(
  $false | $spl32),
  inference(global_subsumption,[],[f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f675,f679,f729,f775,f778,f779,f780,f781,f782,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1261,f688,f699,f701,f406,f722,f713,f711,f707,f702])).
fof(f1344,plain,(
  $false | $spl32),
  inference(global_subsumption,[],[f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f675,f679,f729,f775,f778,f779,f780,f781,f782,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1261,f688,f699,f701,f702,f406,f722,f713,f711,f707])).
fof(f1343,plain,(
  $false | ($spl3 | $spl32)),
  inference(global_subsumption,[],[f708,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f675,f679,f729,f775,f778,f779,f780,f781,f782,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1261,f688,f699,f701,f702,f707,f406,f722,f713,f711])).
fof(f708,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) | ($spl3 | $spl32)),
  inference(unit_resulting_resolution,[],[f64,f64,f679,f376])).
fof(f1342,plain,(
  $false | ($spl12 | $spl32)),
  inference(global_subsumption,[],[f709,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f675,f679,f729,f775,f778,f779,f780,f781,f782,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1261,f688,f699,f701,f702,f707,f406,f722,f713,f711])).
fof(f709,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) | ($spl12 | $spl32)),
  inference(unit_resulting_resolution,[],[f430,f430,f679,f376])).
fof(f1341,plain,(
  $false | $spl32),
  inference(global_subsumption,[],[f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f675,f679,f729,f775,f778,f779,f780,f781,f782,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1261,f688,f699,f701,f702,f707,f406,f722,f713,f711])).
fof(f1340,plain,(
  $false | $spl32),
  inference(global_subsumption,[],[f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f675,f679,f729,f775,f778,f779,f780,f781,f782,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1261,f688,f699,f701,f702,f707,f711,f406,f722,f713])).
fof(f1339,plain,(
  $false | ($spl3 | $spl32)),
  inference(global_subsumption,[],[f715,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f675,f679,f729,f775,f778,f779,f780,f781,f782,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1261,f688,f699,f701,f702,f707,f711,f713,f406,f722])).
fof(f715,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl32)),
  inference(unit_resulting_resolution,[],[f37,f64,f679,f376])).
fof(f1338,plain,(
  $false | ($spl3 | $spl12 | $spl32)),
  inference(global_subsumption,[],[f1326,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f675,f679,f729,f775,f778,f779,f780,f781,f782,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1261,f688,f699,f701,f702,f707,f711,f713,f406,f722])).
fof(f1326,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl12 | $spl32)),
  inference(forward_demodulation,[],[f716,f454])).
fof(f716,plain,(
  s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl12 | $spl32)),
  inference(unit_resulting_resolution,[],[f37,f430,f679,f376])).
fof(f1337,plain,(
  $false | $spl32),
  inference(global_subsumption,[],[f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f675,f679,f729,f775,f778,f779,f780,f781,f782,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1261,f688,f699,f701,f702,f707,f711,f713,f406,f722])).
fof(f1336,plain,(
  $false | ($spl21 | $spl32)),
  inference(global_subsumption,[],[f537,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f675,f679,f729,f775,f778,f779,f780,f781,f782,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1261,f688,f699,f701,f702,f707,f711,f713,f722,f406])).
fof(f1335,plain,(
  $false | $spl32),
  inference(global_subsumption,[],[f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f675,f679,f729,f775,f778,f779,f780,f781,f782,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1261,f688,f699,f701,f702,f707,f711,f713,f722,f406])).
fof(f1334,plain,(
  $false | ($spl27 | $spl32)),
  inference(global_subsumption,[],[f542,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f675,f679,f729,f775,f778,f779,f780,f781,f782,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1261,f688,f699,f701,f702,f707,f711,f713,f722,f406])).
fof(f1333,plain,(
  $false | ($spl24 | $spl32)),
  inference(global_subsumption,[],[f540,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f675,f679,f729,f775,f778,f779,f780,f781,f782,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1261,f688,f699,f701,f702,f707,f711,f713,f722,f406])).
fof(f1332,plain,(
  $false | ($spl24 | $spl42)),
  inference(global_subsumption,[],[f540,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802,f815,f817,f818,f824,f829,f831,f841,f406])).
fof(f841,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1)))) | $spl42),
  inference(unit_resulting_resolution,[],[f37,f791,f48])).
fof(f831,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1)))) | $spl42),
  inference(unit_resulting_resolution,[],[f37,f791,f48])).
fof(f829,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1)))) | $spl42),
  inference(unit_resulting_resolution,[],[f791,f791,f791,f376])).
fof(f824,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1)))) | $spl42),
  inference(unit_resulting_resolution,[],[f791,f791,f791,f376])).
fof(f818,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | $spl42),
  inference(unit_resulting_resolution,[],[f791,f39])).
fof(f817,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1)))) | $spl42),
  inference(unit_resulting_resolution,[],[f791,f52])).
fof(f815,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | $spl42),
  inference(unit_resulting_resolution,[],[f38,f791,f80])).
fof(f802,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1)))) | $spl42),
  inference(unit_resulting_resolution,[],[f791,f791,f791,f376])).
fof(f1264,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1)))) | $spl42),
  inference(subsumption_resolution,[],[f1253,f37])).
fof(f1253,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1)))) | $spl42),
  inference(superposition,[],[f528,f787])).
fof(f901,plain,(
  ( ! [X6,X7] : (s(bool,X6) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK2))) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) ) | $spl42),
  inference(subsumption_resolution,[],[f890,f37])).
fof(f890,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,X6) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK2))) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) ) | $spl42),
  inference(superposition,[],[f376,f787])).
fof(f848,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl42),
  inference(trivial_inequality_removal,[],[f842])).
fof(f842,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl42),
  inference(superposition,[],[f791,f50])).
fof(f791,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | $spl42),
  inference(backward_demodulation,[],[f787,f602])).
fof(f1331,plain,(
  $false | ($spl27 | $spl42)),
  inference(global_subsumption,[],[f542,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802,f815,f817,f818,f824,f829,f831,f841,f406])).
fof(f1330,plain,(
  $false | $spl42),
  inference(global_subsumption,[],[f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802,f815,f817,f818,f824,f829,f831,f841,f406])).
fof(f1329,plain,(
  $false | ($spl21 | $spl42)),
  inference(global_subsumption,[],[f537,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802,f815,f817,f818,f824,f829,f831,f841])).
fof(f1328,plain,(
  $false | ($spl32 | $spl42)),
  inference(global_subsumption,[],[f722,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802,f815,f817,f818,f824,f829,f831,f841])).
fof(f1327,plain,(
  $false | ($spl3 | $spl12 | $spl32 | $spl42)),
  inference(global_subsumption,[],[f1326,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802,f815,f817,f818,f824,f829,f831,f841])).
fof(f1325,plain,(
  $false | ($spl3 | $spl32 | $spl42)),
  inference(global_subsumption,[],[f715,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802,f815,f817,f818,f824,f829,f831,f841])).
fof(f1324,plain,(
  $false | ($spl32 | $spl42)),
  inference(global_subsumption,[],[f713,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802,f815,f817,f818,f824,f829,f831,f841])).
fof(f1323,plain,(
  $false | ($spl32 | $spl42)),
  inference(global_subsumption,[],[f711,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802,f815,f817,f818,f824,f829,f831,f841])).
fof(f1322,plain,(
  $false | ($spl12 | $spl32 | $spl42)),
  inference(global_subsumption,[],[f709,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802,f815,f817,f818,f824,f829,f831,f841])).
fof(f1321,plain,(
  $false | ($spl3 | $spl32 | $spl42)),
  inference(global_subsumption,[],[f708,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802,f815,f817,f818,f824,f829,f831,f841])).
fof(f1320,plain,(
  $false | ($spl32 | $spl42)),
  inference(global_subsumption,[],[f707,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802,f815,f817,f818,f824,f829,f831,f841])).
fof(f1319,plain,(
  $false | ($spl32 | $spl42)),
  inference(global_subsumption,[],[f702,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802,f815,f817,f818,f824,f829,f831,f841])).
fof(f1318,plain,(
  $false | ($spl32 | $spl42)),
  inference(global_subsumption,[],[f701,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802,f815,f817,f818,f824,f829,f831,f841])).
fof(f1317,plain,(
  $false | ($spl32 | $spl42)),
  inference(global_subsumption,[],[f699,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802,f815,f817,f818,f824,f829,f831,f841])).
fof(f1316,plain,(
  $false | ($spl3 | $spl12 | $spl32 | $spl42)),
  inference(global_subsumption,[],[f1315,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802,f815,f817,f818,f824,f829,f831,f841])).
fof(f1314,plain,(
  $false | ($spl3 | $spl32 | $spl42)),
  inference(global_subsumption,[],[f694,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802,f815,f817,f818,f824,f829,f831,f841])).
fof(f1313,plain,(
  $false | ($spl3 | $spl12 | $spl32 | $spl42)),
  inference(global_subsumption,[],[f1312,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802,f815,f817,f818,f824,f829,f831,f841])).
fof(f1311,plain,(
  $false | ($spl3 | $spl32 | $spl42)),
  inference(global_subsumption,[],[f689,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802,f815,f817,f818,f824,f829,f831,f841])).
fof(f1310,plain,(
  $false | ($spl32 | $spl42)),
  inference(global_subsumption,[],[f688,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802,f815,f817,f818,f824,f829,f831,f841])).
fof(f1309,plain,(
  $false | ($spl3 | $spl12 | $spl32 | $spl42)),
  inference(global_subsumption,[],[f1308,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802,f815,f817,f818,f824,f829,f831,f841])).
fof(f1307,plain,(
  $false | ($spl3 | $spl32 | $spl42)),
  inference(global_subsumption,[],[f680,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802,f815,f817,f818,f824,f829,f831,f841])).
fof(f1306,plain,(
  $false | ($spl32 | $spl42)),
  inference(global_subsumption,[],[f836,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802,f815,f817,f818,f824,f829,f831,f841])).
fof(f836,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl32 | $spl42)),
  inference(unit_resulting_resolution,[],[f37,f679,f791,f376])).
fof(f1305,plain,(
  $false | $spl42),
  inference(global_subsumption,[],[f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802,f815,f817,f818,f824,f829,f831,f841])).
fof(f1304,plain,(
  $false | ($spl3 | $spl12 | $spl42)),
  inference(global_subsumption,[],[f1303,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f841,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802,f815,f817,f818,f824,f829,f831])).
fof(f1303,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl12 | $spl42)),
  inference(forward_demodulation,[],[f834,f454])).
fof(f834,plain,(
  s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl12 | $spl42)),
  inference(unit_resulting_resolution,[],[f37,f430,f791,f376])).
fof(f1302,plain,(
  $false | ($spl3 | $spl42)),
  inference(global_subsumption,[],[f833,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f841,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802,f815,f817,f818,f824,f829,f831])).
fof(f833,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl42)),
  inference(unit_resulting_resolution,[],[f37,f64,f791,f376])).
fof(f1301,plain,(
  $false | $spl42),
  inference(global_subsumption,[],[f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f841,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802,f815,f817,f818,f824,f829,f831])).
fof(f1300,plain,(
  $false | $spl42),
  inference(global_subsumption,[],[f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f841,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802,f815,f817,f818,f824,f829])).
fof(f1299,plain,(
  $false | ($spl32 | $spl42)),
  inference(global_subsumption,[],[f828,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f841,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802,f815,f817,f818,f824])).
fof(f828,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1)))) | ($spl32 | $spl42)),
  inference(unit_resulting_resolution,[],[f679,f679,f791,f376])).
fof(f1298,plain,(
  $false | ($spl12 | $spl42)),
  inference(global_subsumption,[],[f826,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f841,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802,f815,f817,f818,f824])).
fof(f826,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1)))) | ($spl12 | $spl42)),
  inference(unit_resulting_resolution,[],[f430,f430,f791,f376])).
fof(f1297,plain,(
  $false | ($spl3 | $spl42)),
  inference(global_subsumption,[],[f825,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f841,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802,f815,f817,f818,f824])).
fof(f825,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1)))) | ($spl3 | $spl42)),
  inference(unit_resulting_resolution,[],[f64,f64,f791,f376])).
fof(f1296,plain,(
  $false | $spl42),
  inference(global_subsumption,[],[f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f841,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802,f815,f817,f818,f824])).
fof(f1295,plain,(
  $false | ($spl32 | $spl42)),
  inference(global_subsumption,[],[f823,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f841,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802,f815,f817,f818])).
fof(f823,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) | ($spl32 | $spl42)),
  inference(unit_resulting_resolution,[],[f679,f791,f791,f376])).
fof(f1294,plain,(
  $false | $spl42),
  inference(global_subsumption,[],[f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f841,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802,f815,f817,f818])).
fof(f1293,plain,(
  $false | $spl42),
  inference(global_subsumption,[],[f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f841,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802,f815,f817])).
fof(f1292,plain,(
  $false | $spl42),
  inference(global_subsumption,[],[f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f841,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802,f815])).
fof(f1291,plain,(
  $false | ($spl32 | $spl42)),
  inference(global_subsumption,[],[f812,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f841,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802])).
fof(f812,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl32 | $spl42)),
  inference(unit_resulting_resolution,[],[f679,f791,f50])).
fof(f1290,plain,(
  $false | ($spl3 | $spl12 | $spl42)),
  inference(global_subsumption,[],[f1289,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f841,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802])).
fof(f1289,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl12 | $spl42)),
  inference(forward_demodulation,[],[f810,f454])).
fof(f810,plain,(
  s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl12 | $spl42)),
  inference(unit_resulting_resolution,[],[f430,f791,f50])).
fof(f1288,plain,(
  $false | ($spl3 | $spl42)),
  inference(global_subsumption,[],[f809,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f841,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802])).
fof(f809,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl42)),
  inference(unit_resulting_resolution,[],[f64,f791,f50])).
fof(f1287,plain,(
  $false | ($spl32 | $spl42)),
  inference(global_subsumption,[],[f806,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f841,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802])).
fof(f806,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl32 | $spl42)),
  inference(unit_resulting_resolution,[],[f679,f791,f50])).
fof(f1286,plain,(
  $false | ($spl3 | $spl12 | $spl42)),
  inference(global_subsumption,[],[f1285,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f841,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802])).
fof(f1285,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl12 | $spl42)),
  inference(forward_demodulation,[],[f804,f454])).
fof(f804,plain,(
  s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl12 | $spl42)),
  inference(unit_resulting_resolution,[],[f430,f791,f50])).
fof(f1284,plain,(
  $false | ($spl3 | $spl42)),
  inference(global_subsumption,[],[f803,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f841,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802])).
fof(f803,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl42)),
  inference(unit_resulting_resolution,[],[f64,f791,f50])).
fof(f1283,plain,(
  $false | $spl42),
  inference(global_subsumption,[],[f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f841,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264,f802])).
fof(f1282,plain,(
  $false | ($spl32 | $spl42)),
  inference(global_subsumption,[],[f801,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f841,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264])).
fof(f801,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1)))) | ($spl32 | $spl42)),
  inference(unit_resulting_resolution,[],[f679,f791,f791,f376])).
fof(f1281,plain,(
  $false | ($spl32 | $spl42)),
  inference(global_subsumption,[],[f795,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f841,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264])).
fof(f795,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl32 | $spl42)),
  inference(unit_resulting_resolution,[],[f37,f679,f791,f376])).
fof(f1280,plain,(
  $false | ($spl3 | $spl12 | $spl42)),
  inference(global_subsumption,[],[f1279,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f841,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264])).
fof(f1279,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl12 | $spl42)),
  inference(forward_demodulation,[],[f793,f454])).
fof(f793,plain,(
  s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl12 | $spl42)),
  inference(unit_resulting_resolution,[],[f37,f430,f791,f376])).
fof(f1278,plain,(
  $false | ($spl3 | $spl42)),
  inference(global_subsumption,[],[f792,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f841,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264])).
fof(f792,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl42)),
  inference(unit_resulting_resolution,[],[f37,f64,f791,f376])).
fof(f1277,plain,(
  $false | ($spl3 | $spl12 | $spl42)),
  inference(global_subsumption,[],[f1276,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f841,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264])).
fof(f1276,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl12 | $spl42)),
  inference(forward_demodulation,[],[f977,f454])).
fof(f977,plain,(
  s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl12 | $spl42)),
  inference(unit_resulting_resolution,[],[f791,f466])).
fof(f1275,plain,(
  $false | ($spl3 | $spl12 | $spl32 | $spl42)),
  inference(global_subsumption,[],[f1274,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f841,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264])).
fof(f1273,plain,(
  $false | ($spl32 | $spl42)),
  inference(global_subsumption,[],[f1261,f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f841,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264])).
fof(f1272,plain,(
  $false | $spl42),
  inference(global_subsumption,[],[f36,f35,f34,f37,f39,f48,f45,f49,f50,f52,f58,f60,f57,f77,f80,f81,f40,f41,f44,f43,f42,f38,f125,f127,f141,f166,f201,f202,f376,f377,f240,f241,f403,f404,f406,f405,f407,f409,f411,f412,f414,f416,f417,f418,f419,f420,f421,f527,f528,f529,f530,f531,f532,f533,f534,f602,f601,f603,f605,f607,f608,f610,f612,f613,f615,f617,f666,f667,f668,f669,f670,f778,f779,f780,f781,f782,f787,f791,f841,f848,f901,f910,f959,f960,f961,f962,f963,f1233,f1247,f1267,f1264])).
fof(f1266,plain,(
  $false | ($spl3 | $spl12 | $spl22 | $spl32 | $spl42)),
  inference(subsumption_resolution,[],[f1265,f38])).
fof(f1265,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl12 | $spl22 | $spl32 | $spl42)),
  inference(forward_demodulation,[],[f1264,f849])).
fof(f849,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl12 | $spl22 | $spl32 | $spl42)),
  inference(forward_demodulation,[],[f836,f730])).
fof(f730,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl12 | $spl22 | $spl32)),
  inference(forward_demodulation,[],[f717,f587])).
fof(f587,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | ($spl3 | $spl12 | $spl22)),
  inference(forward_demodulation,[],[f574,f454])).
fof(f574,plain,(
  s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) = s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) | ($spl12 | $spl22)),
  inference(unit_resulting_resolution,[],[f37,f430,f543,f376])).
fof(f717,plain,(
  s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl22 | $spl32)),
  inference(unit_resulting_resolution,[],[f37,f543,f679,f376])).
fof(f1263,plain,(
  $false | ($spl3 | $spl12 | $spl22 | $spl32)),
  inference(subsumption_resolution,[],[f1262,f38])).
fof(f1262,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl12 | $spl22 | $spl32)),
  inference(forward_demodulation,[],[f1261,f730])).
fof(f969,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl54),
  inference(cnf_transformation,[],[f969_D])).
fof(f969_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f965,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK5(X0,X1,bool,bool))) ) | $spl50),
  inference(cnf_transformation,[],[f965_D])).
fof(f965_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK5(X0,X1,bool,bool))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f972,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl59),
  inference(cnf_transformation,[],[f972_D])).
fof(f972_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f963,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK5(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f411,f39])).
fof(f970,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl56),
  inference(cnf_transformation,[],[f970_D])).
fof(f970_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f962,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK5(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f411,f50])).
fof(f961,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK5(X7,X8,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f411,f50])).
fof(f964,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl48),
  inference(cnf_transformation,[],[f964_D])).
fof(f964_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f967,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl53),
  inference(cnf_transformation,[],[f967_D])).
fof(f967_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f960,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK5(X3,X4,bool,bool)) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f411,f50])).
fof(f959,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK5(X0,X1,bool,bool)) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f411,f50])).
fof(f787,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK2))) | $spl42),
  inference(cnf_transformation,[],[f787_D])).
fof(f787_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f790,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | $spl47),
  inference(cnf_transformation,[],[f790_D])).
fof(f790_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f782,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK2)))),
  inference(superposition,[],[f602,f39])).
fof(f788,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X2)) ) | $spl44),
  inference(cnf_transformation,[],[f788_D])).
fof(f788_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f781,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK2)))) )),
  inference(superposition,[],[f602,f50])).
fof(f780,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f602,f50])).
fof(f783,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X0)) ) | $spl38),
  inference(cnf_transformation,[],[f783_D])).
fof(f783_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f785,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | $spl41),
  inference(cnf_transformation,[],[f785_D])).
fof(f785_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f779,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | s(bool,X1) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f602,f50])).
fof(f778,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK1))) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK1,sK2,sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f602,f50])).
fof(f675,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | $spl32),
  inference(cnf_transformation,[],[f675_D])).
fof(f675_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f678,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl37),
  inference(cnf_transformation,[],[f678_D])).
fof(f678_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f670,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK2)))),
  inference(superposition,[],[f601,f39])).
fof(f676,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X2)) ) | $spl34),
  inference(cnf_transformation,[],[f676_D])).
fof(f676_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f669,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK2)))) )),
  inference(superposition,[],[f601,f50])).
fof(f668,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f601,f50])).
fof(f671,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X0)) ) | $spl28),
  inference(cnf_transformation,[],[f671_D])).
fof(f671_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f673,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl31),
  inference(cnf_transformation,[],[f673_D])).
fof(f673_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f667,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,X1) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f601,f50])).
fof(f666,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,sK1,sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f601,f50])).
fof(f539,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0)))) | $spl22),
  inference(cnf_transformation,[],[f539_D])).
fof(f539_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0)))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f542,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | $spl27),
  inference(cnf_transformation,[],[f542_D])).
fof(f542_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f534,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0))))),
  inference(superposition,[],[f406,f39])).
fof(f540,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl24),
  inference(cnf_transformation,[],[f540_D])).
fof(f540_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f533,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0))))) )),
  inference(superposition,[],[f406,f50])).
fof(f532,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f406,f50])).
fof(f531,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,X1) = s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f406,f50])).
fof(f530,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,X0) = s(bool,g(s(sK0,sK5(sK1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f406,f50])).
fof(f426,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK5(sK2,sK1,bool,sK0)))) | $spl12),
  inference(cnf_transformation,[],[f426_D])).
fof(f426_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK5(sK2,sK1,bool,sK0)))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f429,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) | $spl17),
  inference(cnf_transformation,[],[f429_D])).
fof(f429_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f421,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK5(sK2,sK1,bool,sK0))))),
  inference(superposition,[],[f405,f39])).
fof(f427,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl14),
  inference(cnf_transformation,[],[f427_D])).
fof(f427_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f420,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,g(s(sK0,sK5(sK2,sK1,bool,sK0))))) )),
  inference(superposition,[],[f405,f50])).
fof(f419,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK5(sK2,sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f405,f50])).
fof(f422,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(sK0,sK5(sK2,sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(cnf_transformation,[],[f422_D])).
fof(f422_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(sK0,sK5(sK2,sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f424,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) | $spl11),
  inference(cnf_transformation,[],[f424_D])).
fof(f424_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f418,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) | s(bool,X1) = s(bool,g(s(sK0,sK5(sK2,sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f405,f50])).
fof(f417,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK0,sK5(sK2,sK1,bool,sK0)))) | s(bool,X0) = s(bool,g(s(sK0,sK5(sK2,sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f405,f50])).
fof(f88,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f82,f37])).
fof(f82,plain,(
  ~p(s(bool,'T_0')) | $spl4),
  inference(backward_demodulation,[],[f78,f38])).
fof(f79,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1)) ) | $spl6),
  inference(cnf_transformation,[],[f79_D])).
fof(f79_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1)) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f81,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f75,f52])).
fof(f75,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f45,f39])).
fof(f77,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f73,f52])).
fof(f73,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f45,f39])).
fof(f62,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl0),
  inference(cnf_transformation,[],[f62_D])).
fof(f62_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f64,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl3),
  inference(cnf_transformation,[],[f64_D])).
fof(f64_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f57,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f39])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 96
% Passive clauses: 319
% Generated clauses: 4003
% Final active clauses: 24
% Final passive clauses: 34
% Input formulas: 8
% Initial clauses: 16
% 
% Duplicate literals: 148
% Trivial inequalities: 31
% Fw subsumption resolutions: 261
% Fw demodulations: 607
% Bw demodulations: 208
% Condensations: 155
% Global subsumptions: 83
% 
% Simple tautologies: 55
% Equational tautologies: 199
% Forward subsumptions: 1931
% Bw demodulations to eq. taut.: 7
% 
% Unit resulting resolution: 593
% Factoring: 1
% Forward superposition: 650
% Backward superposition: 556
% Self superposition: 30
% Equality factoring: 34
% Equality resolution: 7
% 
% Split clauses: 38
% Split components: 81
% SAT solver clauses: 656
% SAT solver unit clauses: 51
% SAT solver binary clauses: 223
% 
% Sat splits: 38
% Sat splitting refutations: 176
% 
% Memory used [KB]: 1918
% Time elapsed: 0.101 s
% ------------------------------
% ------------------------------
% Success in time 1.426 s
