Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.pair
% (56119)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 101
% Passive clauses: 1585
% Generated clauses: 22040
% Final active clauses: 94
% Final passive clauses: 1350
% Input formulas: 10
% Initial clauses: 15
% 
% Duplicate literals: 1717
% Trivial inequalities: 250
% Fw subsumption resolutions: 2665
% Fw demodulations: 2239
% Bw demodulations: 6
% 
% Simple tautologies: 391
% Equational tautologies: 67
% Forward subsumptions: 11159
% 
% Binary resolution: 47
% Forward superposition: 7156
% Backward superposition: 7635
% Self superposition: 291
% Equality factoring: 199
% Equality resolution: 14
% 
% Split clauses: 1424
% Split components: 2881
% SAT solver clauses: 2118
% SAT solver unit clauses: 72
% SAT solver binary clauses: 398
% 
% Sat splits: 1424
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 2097
% TWLsolver calls for satisfiability: 100
% 
% Memory used [KB]: 12025
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pair
% (56131)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 80
% Passive clauses: 359
% Generated clauses: 12966
% Final active clauses: 66
% Final passive clauses: 205
% Input formulas: 10
% Initial clauses: 15
% 
% Duplicate literals: 1302
% Trivial inequalities: 181
% Fw subsumption resolutions: 689
% Fw demodulations: 865
% Bw demodulations: 50
% Condensations: 1822
% 
% Simple tautologies: 173
% Equational tautologies: 49
% Forward subsumptions: 5460
% Bw demodulations to eq. taut.: 28
% 
% Binary resolution: 22
% Forward superposition: 3678
% Backward superposition: 4082
% Self superposition: 245
% Equality factoring: 166
% Equality resolution: 9
% 
% Split clauses: 90
% Split components: 194
% SAT solver clauses: 704
% SAT solver unit clauses: 18
% SAT solver binary clauses: 46
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 3
% 
% Sat splits: 90
% Sat splitting refutations: 24
% 
% TWLsolver clauses: 702
% TWLsolver calls for satisfiability: 123
% 
% Memory used [KB]: 9083
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pair
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pair
% SZS output start Proof for thm.pair
fof(f92178,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f36,f55,f53,f51,f35,f54,f34,f52,f50,f64,f71,f412,f439,f437,f959,f962,f966,f838,f983,f984,f985,f961,f1002,f1022,f1029,f1030,f1031,f1064,f1066,f1069,f1070,f1072,f1073,f1228,f1229,f2036,f2294,f2299,f2300,f2303,f2305,f2105,f2311,f2293,f2442,f2451,f2443,f2444,f2445,f2455,f2446,f2548,f2551,f2972,f3044,f2292,f3170,f2481,f2454,f3301,f3304,f3302,f4869,f4263,f4882,f4853,f4890,f4892,f4895,f4896,f4833,f4832,f4883,f4825,f4824,f4811,f4878,f4809,f4900,f4395,f4398,f4907,f4401,f4909,f4911,f4437,f4922,f4925,f4758,f4897,f4752,f4928,f4749,f4929,f4523,f4893,f4912,f4745,f4930,f4965,f4980,f2550,f2549,f3652,f4988,f4989,f3732,f4992,f4990,f4993,f4995,f4996,f4997,f5065,f5066,f6722,f6738,f6739,f6755,f6754,f7380,f7381,f7404,f7750,f7753,f7755,f7758,f7769,f7778,f7782,f7786,f7790,f7794,f7875,f7876,f7688,f2447,f7887,f7893,f7894,f7898,f7902,f7905,f7617,f7900,f7907,f7616,f7908,f7615,f7911,f7614,f7912,f7613,f7910,f7913,f9589,f9593,f9594,f9612,f960,f9590,f10103,f10183,f11395,f11670,f16423,f16614,f24931,f24930,f24854,f24971,f24822,f24800,f24970,f25217,f24550,f25230,f24537,f25232,f25235,f25231,f25236,f25238,f24529,f24528,f24515,f25218,f25313,f25316,f25377,f25378,f25237,f25768,f25771,f25772,f25784,f24220,f25219,f25788,f23788,f25789,f25790,f25791,f20161,f25794,f19927,f19749,f19571,f25793,f25798,f25795,f19337,f25796,f19159,f25797,f18981,f18976,f18803,f25799,f18798,f25800,f2452,f25801,f25802,f13351,f25803,f25804,f25805,f13175,f25807,f4983,f13170,f25808,f2449,f2448,f13039,f25809,f9980,f9978,f9977,f9974,f6734,f25814,f6732,f4984,f4987,f6731,f6728,f2562,f2560,f25850,f25854,f25851,f26448,f26480,f29277,f29278,f30056,f30111,f30749,f30213,f30268,f30715,f30802,f30750,f30854,f31441,f33651,f33698,f33625,f33720,f33624,f33721,f33623,f33724,f34070,f34073,f34071,f34330,f34331,f34482,f34481,f34802,f34914,f35559,f25243,f36886,f36851,f36833,f36832,f36831,f36821,f36937,f36820,f37016,f36819,f37017,f36264,f37102,f36272,f37104,f37127,f37128,f37129,f37130,f37131,f37132,f37133,f37134,f36746,f37181,f36742,f37182,f37240,f37534,f37535,f37584,f37583,f37595,f37596,f38715,f38717,f36263,f38763,f36265,f38764,f36273,f38765,f37101,f40278,f40384,f42114,f41955,f42116,f42119,f41950,f42127,f42132,f42133,f37100,f42950,f42953,f42951,f43680,f44329,f47883,f47887,f47984,f47888,f47983,f48352,f48393,f48270,f48274,f48443,f48275,f48442,f48860,f48862,f48916,f48922,f48923,f48924,f48932,f48937,f48942,f48982,f48983,f48984,f48986,f49034,f49036,f49039,f49041,f49043,f49044,f49045,f49046,f49047,f49050,f37194,f36723,f49052,f37633,f49060,f49062,f49064,f49065,f49066,f49067,f49068,f49071,f39839,f49058,f49056,f49072,f49073,f49074,f49143,f49151,f49234,f49236,f49237,f49238,f49246,f49247,f49249,f7901,f43667,f43893,f43894,f49252,f43907,f49251,f45693,f45866,f45867,f49254,f45883,f49250,f49253,f49255,f49285,f49287,f49286,f49426,f25852,f49444,f49450,f49445,f49569,f49570,f49248,f49696,f53057,f53059,f56311,f56335,f57511,f57515,f57455,f57575,f60406,f60407,f63839,f63843,f63844,f63845,f63840,f49427,f64813,f64814,f64816,f64817,f67090,f67091,f67773,f67774,f67866,f67865,f67890,f67891,f68662,f68663,f69297,f69298,f69538,f69539,f69714,f69718,f69852,f69719,f69851,f70051,f70246,f70247,f70249,f70250,f70256,f70263,f70257,f70275,f70276,f70322,f70323,f70503,f70506,f70504,f72532,f72428,f72402,f72593,f73074,f73075,f73206,f72131,f73214,f72127,f73223,f73228,f73238,f73239,f73242,f73207,f74463,f74466,f74438,f74479,f74556,f74562,f74382,f74627,f74378,f74643,f74642,f74377,f74646,f74645,f74376,f74648,f74373,f74649,f74667,f74729,f74360,f74731,f74739,f74741,f74153,f74744,f74148,f74746,f74379,f74749,f74644,f74751,f74748,f74387,f74388,f74760,f74647,f74762,f74745,f74764,f74766,f74464,f74395,f74440,f74730,f69361,f7405,f74807,f75355,f75357,f75419,f76297,f76298,f76299,f76300,f80450,f80453,f69,f80451,f81751,f81791,f87116,f87159,f87115,f87160,f87025,f87029,f967,f87223,f87030,f4908,f87222,f88964,f89008,f4877,f91437,f92101,f92103])).
fof(f92103,plain,(
  $false | ($spl1 | $spl2 | $spl4 | $spl13 | $spl22)),
  inference(subsumption_resolution,[],[f91468,f91440])).
fof(f91440,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl2 | $spl4)),
  inference(subsumption_resolution,[],[f91439,f37])).
fof(f37,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pair.PROD_ALL_THM',aHL_TRUTH)).
fof(f91439,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl2 | $spl4)),
  inference(forward_demodulation,[],[f91438,f1074])).
fof(f1074,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3))) | $spl2),
  inference(unit_resulting_resolution,[],[f52,f58])).
fof(f58,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f39,f40])).
fof(f40,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pair.PROD_ALL_THM',aHL_BOOL_CASES)).
fof(f39,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f14])).
fof(f14,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pair.PROD_ALL_THM',aHL_FALSITY)).
fof(f91438,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f91328,f58])).
fof(f91328,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl4)),
  inference(superposition,[],[f2762,f1101])).
fof(f1101,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(subsumption_resolution,[],[f1097,f37])).
fof(f1097,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(superposition,[],[f51,f56])).
fof(f56,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f40,f40])).
fof(f2762,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X3,bool),X4),s(fun(sK1,bool),sK4),s('type.pair.prod'(X3,sK1),'const.pair.,_2'(s(X3,X5),s(sK1,sK2)))))) | ~p(s(bool,happ(s(fun(X3,bool),X4),s(X3,X5))))) ) | $spl4),
  inference(subsumption_resolution,[],[f2727,f37])).
fof(f2727,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X3,bool),X4),s(X3,X5)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X3,bool),X4),s(fun(sK1,bool),sK4),s('type.pair.prod'(X3,sK1),'const.pair.,_2'(s(X3,X5),s(sK1,sK2))))))) ) | $spl4),
  inference(superposition,[],[f785,f1085])).
fof(f1085,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | $spl4),
  inference(unit_resulting_resolution,[],[f54,f58])).
fof(f785,plain,(
  ( ! [X39,X37,X35,X38,X36,X34] : (~p(s(bool,happ(s(fun(X34,bool),X38),s(X34,X37)))) | ~p(s(bool,happ(s(fun(X35,bool),X39),s(X35,X36)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X35,bool),X39),s(fun(X34,bool),X38),s('type.pair.prod'(X35,X34),'const.pair.,_2'(s(X35,X36),s(X34,X37))))))) )),
  inference(forward_demodulation,[],[f768,f43])).
fof(f43,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))) )),
  inference(cnf_transformation,[],[f18])).
fof(f18,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(flattening,[],[f17])).
fof(f17,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X7,X6] : ! [X9,X8] : s(X6,X9) = s(X6,'const.pair.FST_1'(s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,X9),s(X7,X8)))))),
  file('thm.pair.PROD_ALL_THM','thm.pair.FST')).
fof(f768,plain,(
  ( ! [X39,X37,X35,X38,X36,X34] : (~p(s(bool,happ(s(fun(X34,bool),X38),s(X34,X37)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X35,bool),X39),s(fun(X34,bool),X38),s('type.pair.prod'(X35,X34),'const.pair.,_2'(s(X35,X36),s(X34,X37)))))) | ~p(s(bool,happ(s(fun(X35,bool),X39),s(X35,'const.pair.FST_1'(s('type.pair.prod'(X35,X34),'const.pair.,_2'(s(X35,X36),s(X34,X37))))))))) )),
  inference(superposition,[],[f48,f44])).
fof(f44,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))) )),
  inference(cnf_transformation,[],[f20])).
fof(f20,plain,(
  ! [X0,X1,X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(flattening,[],[f19])).
fof(f19,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6,X7] : ! [X9,X8] : s(X7,X8) = s(X7,'const.pair.SND_1'(s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,X9),s(X7,X8)))))),
  file('thm.pair.PROD_ALL_THM','thm.pair.SND')).
fof(f48,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),X3),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X4)))))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X0,bool),X2),s(fun(X1,bool),X3),s('type.pair.prod'(X0,X1),X4)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X4))))))) )),
  inference(cnf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X0,bool),X2),s(fun(X1,bool),X3),s('type.pair.prod'(X0,X1),X4)))) | (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X4)))))) & p(s(bool,happ(s(fun(X1,bool),X3),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X4)))))))) & (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X4)))))) | ~p(s(bool,happ(s(fun(X1,bool),X3),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X4)))))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X0,bool),X2),s(fun(X1,bool),X3),s('type.pair.prod'(X0,X1),X4))))))),
  inference(flattening,[],[f32])).
fof(f32,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X0,bool),X2),s(fun(X1,bool),X3),s('type.pair.prod'(X0,X1),X4)))) | (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X4)))))) & p(s(bool,happ(s(fun(X1,bool),X3),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X4)))))))) & ((~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X4)))))) | ~p(s(bool,happ(s(fun(X1,bool),X3),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X4))))))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X0,bool),X2),s(fun(X1,bool),X3),s('type.pair.prod'(X0,X1),X4))))))),
  inference(nnf_transformation,[],[f24])).
fof(f24,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X0,bool),X2),s(fun(X1,bool),X3),s('type.pair.prod'(X0,X1),X4)))) <=> (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X4)))))) & p(s(bool,happ(s(fun(X1,bool),X3),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X4))))))))),
  inference(flattening,[],[f23])).
fof(f23,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X0,bool),X2),s(fun(X1,bool),X3),s('type.pair.prod'(X0,X1),X4)))) <=> (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X4)))))) & p(s(bool,happ(s(fun(X1,bool),X3),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X4))))))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6,X7] : ! [X11,X10,X12] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X6,bool),X11),s(fun(X7,bool),X10),s('type.pair.prod'(X6,X7),X12)))) <=> (p(s(bool,happ(s(fun(X6,bool),X11),s(X6,'const.pair.FST_1'(s('type.pair.prod'(X6,X7),X12)))))) & p(s(bool,happ(s(fun(X7,bool),X10),s(X7,'const.pair.SND_1'(s('type.pair.prod'(X6,X7),X12))))))))),
  file('thm.pair.PROD_ALL_THM','thm.pair.PROD_ALL_def')).
fof(f91468,plain,(
  s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'T_0'),s(bool,'T_0'))))) | ($spl1 | $spl2 | $spl4 | $spl13 | $spl22)),
  inference(backward_demodulation,[],[f91440,f6406])).
fof(f6406,plain,(
  s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,'F_0'))))) | ($spl13 | $spl22)),
  inference(subsumption_resolution,[],[f6272,f439])).
fof(f6272,plain,(
  s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2))))) | ($spl13 | $spl22)),
  inference(superposition,[],[f439,f5991])).
fof(f5991,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl22),
  inference(condensation,[],[f5990])).
fof(f5990,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1)) ) | $spl22),
  inference(duplicate_literal_removal,[],[f5784])).
fof(f5784,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl22),
  inference(superposition,[],[f5520,f40])).
fof(f5520,plain,(
  ( ! [X8,X9] : (s(bool,X9) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,X8))))) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X9)) ) | $spl22),
  inference(duplicate_literal_removal,[],[f5475])).
fof(f5475,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X8) | s(bool,X9) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,X8))))) | s(bool,'T_0') = s(bool,X9)) ) | $spl22),
  inference(superposition,[],[f2293,f56])).
fof(f92101,plain,(
  $false | ($spl1 | $spl2 | $spl4)),
  inference(subsumption_resolution,[],[f91441,f37])).
fof(f91441,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl2 | $spl4)),
  inference(backward_demodulation,[],[f91440,f39])).
fof(f91437,plain,(
  $false | ($spl1 | $spl2 | $spl4 | $spl9)),
  inference(subsumption_resolution,[],[f91436,f37])).
fof(f91436,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl2 | $spl4 | $spl9)),
  inference(forward_demodulation,[],[f91435,f1074])).
fof(f91435,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) | ($spl1 | $spl4 | $spl9)),
  inference(subsumption_resolution,[],[f91327,f39])).
fof(f91327,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) | ($spl1 | $spl4 | $spl9)),
  inference(superposition,[],[f2762,f1103])).
fof(f1103,plain,(
  s(bool,'F_0') = s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2))))) | ($spl1 | $spl9)),
  inference(unit_resulting_resolution,[],[f51,f71,f200])).
fof(f200,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f37,f56])).
fof(f4877,plain,(
  ( ! [X41] : (s(bool,'T_0') = s(bool,X41)) ) | $spl48),
  inference(cnf_transformation,[],[f4877_D])).
fof(f4877_D,plain,(
  ( ! [X41] : (s(bool,'T_0') = s(bool,X41)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f89008,plain,(
  ( ! [X87] : (s(bool,'F_0') = s(bool,g(s(bool,f(s(bool,X87))))) | s(bool,X87) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) ) | $spl370),
  inference(cnf_transformation,[],[f89008_D])).
fof(f89008_D,plain,(
  ( ! [X87] : (s(bool,'F_0') = s(bool,g(s(bool,f(s(bool,X87))))) | s(bool,X87) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) ) <=> ~$spl370),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl370])])).
fof(f88964,plain,(
  ( ! [X87,X86] : (s(bool,'F_0') = s(bool,g(s(bool,f(s(bool,X87))))) | s(bool,'T_0') = s(bool,X86) | s(bool,X87) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) ) | ($spl2 | $spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl70 | $spl176 | $spl250)),
  inference(duplicate_literal_removal,[],[f88700])).
fof(f88700,plain,(
  ( ! [X87,X86] : (s(bool,'F_0') = s(bool,g(s(bool,f(s(bool,X87))))) | s(bool,'T_0') = s(bool,X86) | s(bool,'T_0') = s(bool,X86) | s(bool,X87) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) ) | ($spl2 | $spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl70 | $spl176 | $spl250)),
  inference(superposition,[],[f86458,f22297])).
fof(f22297,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3))) | s(bool,X0) = s(bool,X1)) ) | $spl2),
  inference(condensation,[],[f22296])).
fof(f22296,plain,(
  ( ! [X218,X220,X219] : (s(bool,'T_0') = s(bool,X218) | s(bool,X219) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3))) | s(bool,X219) = s(bool,X220) | s(bool,'T_0') = s(bool,X220)) ) | $spl2),
  inference(duplicate_literal_removal,[],[f18262])).
fof(f18262,plain,(
  ( ! [X218,X220,X219] : (s(bool,'T_0') = s(bool,X218) | s(bool,X219) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X218) | s(bool,X219) = s(bool,X220) | s(bool,'T_0') = s(bool,X220)) ) | $spl2),
  inference(superposition,[],[f146,f1074])).
fof(f146,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f56,f56])).
fof(f86458,plain,(
  ( ! [X131] : (s(bool,'F_0') = s(bool,g(s(bool,f(s(bool,X131))))) | s(bool,'T_0') = s(bool,X131)) ) | ($spl2 | $spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl70 | $spl176 | $spl250)),
  inference(forward_demodulation,[],[f86457,f75535])).
fof(f75535,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl2 | $spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl176 | $spl250)),
  inference(duplicate_literal_removal,[],[f75534])).
fof(f75534,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl2 | $spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl176 | $spl250)),
  inference(condensation,[],[f75533])).
fof(f75533,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl2 | $spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl176 | $spl250)),
  inference(subsumption_resolution,[],[f31569,f74775])).
fof(f74775,plain,(
  ( ! [X285,X284] : (s(bool,'F_0') != s(bool,X284) | s(bool,'T_0') = s(bool,X285) | s(bool,X284) = s(bool,X285)) ) | ($spl2 | $spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl250)),
  inference(subsumption_resolution,[],[f74774,f49260])).
fof(f49260,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(subsumption_resolution,[],[f41816,f439])).
fof(f41816,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2))))) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(superposition,[],[f439,f41609])).
fof(f41609,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X1)) ) | ($spl4 | $spl9 | $spl36 | $spl40)),
  inference(global_subsumption,[],[f41608,f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f64,f132,f133,f146,f156,f157,f200,f361,f363,f71,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f4170,f4171,f3832,f3890,f3878,f2550,f2549,f3652,f3732,f3828,f3727,f2700,f763,f2698,f761,f1935,f1480,f82,f104,f2706,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13852,f13225,f13350,f13427,f13415,f22573,f18154,f18155,f22291,f22279,f22275,f22219,f22217,f18360,f18362,f18363,f18802,f19158,f21203,f22901,f19338,f19402,f19403,f19570,f22984,f20200,f20161,f20156,f19927,f19749,f19571,f19566,f19337,f19159,f18981,f18976,f18803,f18798,f25801,f25802,f13351,f13346,f13175,f13170,f13039,f9980,f9978,f9977,f9974,f6734,f6732,f6731,f6728,f25850,f25854,f30902,f31517,f31518,f31523,f32190,f32954,f32955,f32960,f33745,f33747,f33748,f33749,f33888,f33897,f35168,f35169,f35170,f13038,f3651,f17264,f17266,f15352,f22930,f22929,f22928,f786,f550,f475,f48,f47,f46,f19925,f19747,f19569,f18801,f762,f2704,f19406,f19914,f19736,f19558,f18790,f2699,f785,f2763,f22931,f37580,f37578,f37577,f37576,f34476,f34474,f34472,f34425,f22898,f19565,f19331,f19153,f18797,f16976,f15069,f13826,f13345,f13169,f13033,f4154,f3827,f3726,f3646,f2559,f368,f366,f32792,f32194,f32193,f32192,f31357,f30906,f30905,f30904,f22293,f22277,f19401,f19400,f17276,f17226,f15341,f13224,f13223,f3759,f3758,f2874,f1972,f1473,f370,f365,f39,f2561])).
fof(f2561,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | $spl40),
  inference(cnf_transformation,[],[f2561_D])).
fof(f2561_D,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f365,plain,(
  ( ! [X17,X18,X16] : (s(X18,f(s(bool,'F_0'))) != s(X18,g(s(bool,'F_0'))) | s(fun(bool,X18),X16) = s(fun(bool,X18),X17) | s(bool,'T_0') = s(bool,sK6(X16,X17,X18,bool))) )),
  inference(superposition,[],[f45,f40])).
fof(f370,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,f(s(X16,sK6(X17,X18,bool,X16)))) | s(fun(X16,bool),X17) = s(fun(X16,bool),X18) | s(bool,'T_0') = s(bool,g(s(X16,sK6(X17,X18,bool,X16))))) )),
  inference(superposition,[],[f45,f40])).
fof(f1473,plain,(
  ( ! [X61,X59,X57,X60,X58] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X60),'const.pair.,_2'(s(bool,X59),s(X60,X61))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X57,bool),'const.pair.,_2'(s(X57,X58),s(bool,X59)))))) )),
  inference(superposition,[],[f82,f44])).
fof(f1972,plain,(
  ( ! [X94,X92,X90,X93,X91] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X93,bool),'const.pair.,_2'(s(X93,X94),s(bool,X92))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X90,bool),'const.pair.,_2'(s(X90,X91),s(bool,X92)))))) )),
  inference(superposition,[],[f104,f44])).
fof(f2874,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~p(s(bool,happ(s(fun(X5,bool),X6),s(X5,'const.pair.FST_1'(s('type.pair.prod'(X5,bool),X7)))))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X5,bool),X6),s(fun(bool,bool),X4),s('type.pair.prod'(X5,bool),X7)))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X5,bool),X7))) | s(bool,X8) = s(bool,happ(s(fun(bool,bool),X4),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f2867,f37])).
fof(f2867,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X5,bool),X6),s(fun(bool,bool),X4),s('type.pair.prod'(X5,bool),X7)))) | ~p(s(bool,happ(s(fun(X5,bool),X6),s(X5,'const.pair.FST_1'(s('type.pair.prod'(X5,bool),X7)))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X5,bool),X7))) | s(bool,X8) = s(bool,happ(s(fun(bool,bool),X4),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f761,f56])).
fof(f3758,plain,(
  ( ! [X134,X138,X136,X135,X137] : (s(bool,X134) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X135),s(bool,X134))))) | s(bool,'F_0') = s(bool,X135) | s(bool,X134) = s(bool,X136) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X137),'const.pair.,_2'(s(bool,X136),s(X137,X138)))))) )),
  inference(superposition,[],[f82,f132])).
fof(f3759,plain,(
  ( ! [X142,X140,X143,X141,X139] : (s(bool,X139) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X139),s(bool,X140))))) | s(bool,'F_0') = s(bool,X140) | s(bool,X139) = s(bool,X141) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X142),'const.pair.,_2'(s(bool,X141),s(X142,X143)))))) )),
  inference(superposition,[],[f104,f132])).
fof(f13223,plain,(
  ( ! [X218,X216,X215,X219,X217] : (s(bool,X215) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X216),s(bool,X215))))) | s(bool,'F_0') = s(bool,X216) | s(bool,X215) = s(bool,X217) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X218,bool),'const.pair.,_2'(s(X218,X219),s(bool,X217)))))) )),
  inference(superposition,[],[f82,f133])).
fof(f13224,plain,(
  ( ! [X222,X220,X223,X221,X224] : (s(bool,X220) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X220),s(bool,X221))))) | s(bool,'F_0') = s(bool,X221) | s(bool,X220) = s(bool,X222) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X223,bool),'const.pair.,_2'(s(X223,X224),s(bool,X222)))))) )),
  inference(superposition,[],[f104,f133])).
fof(f15341,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'F_0'),s(X1,X2)))))) )),
  inference(condensation,[],[f15340])).
fof(f15340,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),'const.pair.,_2'(s(bool,'F_0'),s(X2,X3))))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f14817])).
fof(f14817,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),'const.pair.,_2'(s(bool,'F_0'),s(X2,X3))))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f156,f40])).
fof(f17226,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,'F_0')))))) )),
  inference(condensation,[],[f17225])).
fof(f17225,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X2,bool),'const.pair.,_2'(s(X2,X3),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f16766])).
fof(f16766,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X2,bool),'const.pair.,_2'(s(X2,X3),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f157,f40])).
fof(f17276,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,X0)))))) )),
  inference(condensation,[],[f17275])).
fof(f17275,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,X3))))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f16657])).
fof(f16657,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,X3))))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f157,f40])).
fof(f19400,plain,(
  ( ! [X315,X316,X318,X314,X317] : (s(bool,X314) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X315),s(bool,X314))))) | s(bool,'F_0') = s(bool,X315) | s(bool,X316) = s(bool,X317) | s(bool,X314) = s(bool,X317) | s(bool,X316) = s(bool,X318) | s(bool,'T_0') = s(bool,X318)) )),
  inference(superposition,[],[f82,f146])).
fof(f19401,plain,(
  ( ! [X321,X323,X320,X322,X319] : (s(bool,X319) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X319),s(bool,X320))))) | s(bool,'F_0') = s(bool,X320) | s(bool,X321) = s(bool,X322) | s(bool,X319) = s(bool,X322) | s(bool,X321) = s(bool,X323) | s(bool,'T_0') = s(bool,X323)) )),
  inference(superposition,[],[f104,f146])).
fof(f22277,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X2,bool),'const.pair.,_2'(s(X2,X3),s(bool,X4))))) | s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X4)) )),
  inference(condensation,[],[f22276])).
fof(f22276,plain,(
  ( ! [X383,X379,X380,X382,X384,X381] : (s(bool,'T_0') = s(bool,X382) | s(bool,X383) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X379,bool),'const.pair.,_2'(s(X379,X380),s(bool,X381))))) | s(bool,X383) = s(bool,X384) | s(bool,'T_0') = s(bool,X384) | s(bool,'F_0') = s(bool,X381)) )),
  inference(duplicate_literal_removal,[],[f18289])).
fof(f18289,plain,(
  ( ! [X383,X379,X380,X382,X384,X381] : (s(bool,'T_0') = s(bool,X382) | s(bool,X383) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X379,bool),'const.pair.,_2'(s(X379,X380),s(bool,X381))))) | s(bool,'T_0') = s(bool,X382) | s(bool,X383) = s(bool,X384) | s(bool,'T_0') = s(bool,X384) | s(bool,'F_0') = s(bool,X381)) )),
  inference(superposition,[],[f146,f104])).
fof(f22293,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),'const.pair.,_2'(s(bool,X3),s(X2,X4))))) | s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X3)) )),
  inference(condensation,[],[f22292])).
fof(f22292,plain,(
  ( ! [X288,X285,X287,X289,X284,X286] : (s(bool,'T_0') = s(bool,X287) | s(bool,X288) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X284),'const.pair.,_2'(s(bool,X285),s(X284,X286))))) | s(bool,X288) = s(bool,X289) | s(bool,'T_0') = s(bool,X289) | s(bool,'F_0') = s(bool,X285)) )),
  inference(duplicate_literal_removal,[],[f18274])).
fof(f18274,plain,(
  ( ! [X288,X285,X287,X289,X284,X286] : (s(bool,'T_0') = s(bool,X287) | s(bool,X288) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X284),'const.pair.,_2'(s(bool,X285),s(X284,X286))))) | s(bool,'T_0') = s(bool,X287) | s(bool,X288) = s(bool,X289) | s(bool,'T_0') = s(bool,X289) | s(bool,'F_0') = s(bool,X285)) )),
  inference(superposition,[],[f146,f82])).
fof(f30904,plain,(
  ( ! [X198,X204,X202,X200,X199,X203,X201] : (s(bool,X198) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X199),s(bool,X198))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X200,bool),'const.pair.,_2'(s(X200,X201),s(bool,X199))))) | s(bool,X198) = s(bool,X202) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X203),'const.pair.,_2'(s(bool,X202),s(X203,X204)))))) )),
  inference(superposition,[],[f1473,f132])).
fof(f30905,plain,(
  ( ! [X210,X208,X206,X211,X209,X207,X205] : (s(bool,X205) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X206),s(bool,X205))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X207,bool),'const.pair.,_2'(s(X207,X208),s(bool,X206))))) | s(bool,X205) = s(bool,X209) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X210,bool),'const.pair.,_2'(s(X210,X211),s(bool,X209)))))) )),
  inference(superposition,[],[f1473,f133])).
fof(f30906,plain,(
  ( ! [X212,X218,X216,X215,X213,X217,X214] : (s(bool,X212) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X213),s(bool,X212))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X214,bool),'const.pair.,_2'(s(X214,X215),s(bool,X213))))) | s(bool,X216) = s(bool,X217) | s(bool,X212) = s(bool,X217) | s(bool,X216) = s(bool,X218) | s(bool,'T_0') = s(bool,X218)) )),
  inference(superposition,[],[f1473,f146])).
fof(f31357,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,X3))))) | s(bool,X4) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X5),'const.pair.,_2'(s(bool,X3),s(X5,X6))))) | s(bool,X0) = s(bool,X4)) )),
  inference(condensation,[],[f31356])).
fof(f31356,plain,(
  ( ! [X111,X109,X107,X105,X110,X108,X106,X104] : (s(bool,'T_0') = s(bool,X107) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X108,bool),'const.pair.,_2'(s(X108,X109),s(bool,X105))))) | s(bool,X110) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X104),'const.pair.,_2'(s(bool,X105),s(X104,X106))))) | s(bool,X110) = s(bool,X111) | s(bool,'T_0') = s(bool,X111)) )),
  inference(duplicate_literal_removal,[],[f30945])).
fof(f30945,plain,(
  ( ! [X111,X109,X107,X105,X110,X108,X106,X104] : (s(bool,'T_0') = s(bool,X107) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X108,bool),'const.pair.,_2'(s(X108,X109),s(bool,X105))))) | s(bool,X110) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X104),'const.pair.,_2'(s(bool,X105),s(X104,X106))))) | s(bool,'T_0') = s(bool,X107) | s(bool,X110) = s(bool,X111) | s(bool,'T_0') = s(bool,X111)) )),
  inference(superposition,[],[f1473,f146])).
fof(f32192,plain,(
  ( ! [X208,X206,X204,X202,X207,X205,X203] : (s(bool,X202) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X202),s(bool,X203))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X204,bool),'const.pair.,_2'(s(X204,X205),s(bool,X203))))) | s(bool,X202) = s(bool,X206) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X207),'const.pair.,_2'(s(bool,X206),s(X207,X208)))))) )),
  inference(superposition,[],[f1972,f132])).
fof(f32193,plain,(
  ( ! [X212,X210,X215,X213,X211,X209,X214] : (s(bool,X209) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X209),s(bool,X210))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X211,bool),'const.pair.,_2'(s(X211,X212),s(bool,X210))))) | s(bool,X209) = s(bool,X213) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X214,bool),'const.pair.,_2'(s(X214,X215),s(bool,X213)))))) )),
  inference(superposition,[],[f1972,f133])).
fof(f32194,plain,(
  ( ! [X218,X216,X222,X220,X219,X217,X221] : (s(bool,X216) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X216),s(bool,X217))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X218,bool),'const.pair.,_2'(s(X218,X219),s(bool,X217))))) | s(bool,X220) = s(bool,X221) | s(bool,X216) = s(bool,X221) | s(bool,X220) = s(bool,X222) | s(bool,'T_0') = s(bool,X222)) )),
  inference(superposition,[],[f1972,f146])).
fof(f32792,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,X3))))) | s(bool,X4) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X5,bool),'const.pair.,_2'(s(X5,X6),s(bool,X3))))) | s(bool,X0) = s(bool,X4)) )),
  inference(condensation,[],[f32791])).
fof(f32791,plain,(
  ( ! [X111,X109,X116,X114,X112,X110,X115,X113] : (s(bool,'T_0') = s(bool,X112) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X113,bool),'const.pair.,_2'(s(X113,X114),s(bool,X111))))) | s(bool,X115) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X109,bool),'const.pair.,_2'(s(X109,X110),s(bool,X111))))) | s(bool,X115) = s(bool,X116) | s(bool,'T_0') = s(bool,X116)) )),
  inference(duplicate_literal_removal,[],[f32303])).
fof(f32303,plain,(
  ( ! [X111,X109,X116,X114,X112,X110,X115,X113] : (s(bool,'T_0') = s(bool,X112) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X113,bool),'const.pair.,_2'(s(X113,X114),s(bool,X111))))) | s(bool,X115) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X109,bool),'const.pair.,_2'(s(X109,X110),s(bool,X111))))) | s(bool,'T_0') = s(bool,X112) | s(bool,X115) = s(bool,X116) | s(bool,'T_0') = s(bool,X116)) )),
  inference(superposition,[],[f1972,f146])).
fof(f366,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,f(s(X0,sK6(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,X3) = s(bool,g(s(X0,sK6(X1,X2,bool,X0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f45,f56])).
fof(f368,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X11) != s(bool,f(s(X8,sK6(X9,X10,bool,X8)))) | s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK6(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f45,f56])).
fof(f2559,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK6(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,f(s(X0,sK6(X1,X2,bool,X0))))) )),
  inference(trivial_inequality_removal,[],[f2553])).
fof(f2553,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK6(X1,X2,bool,X0)))) | s(bool,'T_0') = s(bool,f(s(X0,sK6(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f370,f40])).
fof(f3646,plain,(
  ( ! [X412,X414,X410,X413,X409,X411] : (s(bool,X412) != s(bool,f(s(X409,sK6(X410,X411,bool,X409)))) | s(fun(X409,bool),X410) = s(fun(X409,bool),X411) | s(bool,'T_0') = s(bool,X412) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X413),'const.pair.,_2'(s(bool,g(s(X409,sK6(X410,X411,bool,X409)))),s(X413,X414)))))) )),
  inference(superposition,[],[f45,f132])).
fof(f3726,plain,(
  ( ! [X397,X399,X401,X398,X400,X402] : (s(bool,X400) != s(bool,f(s(X397,sK6(X398,X399,bool,X397)))) | s(fun(X397,bool),X398) = s(fun(X397,bool),X399) | s(bool,'T_0') = s(bool,g(s(X397,sK6(X398,X399,bool,X397)))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X401),'const.pair.,_2'(s(bool,X400),s(X401,X402)))))) )),
  inference(superposition,[],[f45,f132])).
fof(f3827,plain,(
  ( ! [X358,X360,X362,X359,X361,X363] : (s(bool,'T_0') != s(bool,f(s(X358,sK6(X359,X360,bool,X358)))) | s(fun(X358,bool),X359) = s(fun(X358,bool),X360) | s(bool,X361) = s(bool,g(s(X358,sK6(X359,X360,bool,X358)))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X362),'const.pair.,_2'(s(bool,X361),s(X362,X363)))))) )),
  inference(superposition,[],[f45,f132])).
fof(f4154,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK6(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X3),'const.pair.,_2'(s(bool,f(s(X0,sK6(X1,X2,bool,X0)))),s(X3,X4)))))) )),
  inference(condensation,[],[f4153])).
fof(f4153,plain,(
  ( ! [X182,X180,X184,X183,X181,X179] : (s(fun(X179,bool),X180) = s(fun(X179,bool),X181) | s(bool,'T_0') = s(bool,g(s(X179,sK6(X180,X181,bool,X179)))) | s(bool,'T_0') = s(bool,X182) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X183),'const.pair.,_2'(s(bool,f(s(X179,sK6(X180,X181,bool,X179)))),s(X183,X184)))))) )),
  inference(subsumption_resolution,[],[f3608,f40])).
fof(f3608,plain,(
  ( ! [X182,X180,X184,X183,X181,X179] : (s(bool,'F_0') != s(bool,X182) | s(fun(X179,bool),X180) = s(fun(X179,bool),X181) | s(bool,'T_0') = s(bool,g(s(X179,sK6(X180,X181,bool,X179)))) | s(bool,'T_0') = s(bool,X182) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X183),'const.pair.,_2'(s(bool,f(s(X179,sK6(X180,X181,bool,X179)))),s(X183,X184)))))) )),
  inference(superposition,[],[f370,f132])).
fof(f13033,plain,(
  ( ! [X705,X700,X702,X704,X701,X703] : (s(bool,X703) != s(bool,f(s(X700,sK6(X701,X702,bool,X700)))) | s(fun(X700,bool),X701) = s(fun(X700,bool),X702) | s(bool,'T_0') = s(bool,X703) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X704,bool),'const.pair.,_2'(s(X704,X705),s(bool,g(s(X700,sK6(X701,X702,bool,X700))))))))) )),
  inference(superposition,[],[f45,f133])).
fof(f13169,plain,(
  ( ! [X689,X691,X687,X688,X690,X686] : (s(bool,X689) != s(bool,f(s(X686,sK6(X687,X688,bool,X686)))) | s(fun(X686,bool),X687) = s(fun(X686,bool),X688) | s(bool,'T_0') = s(bool,g(s(X686,sK6(X687,X688,bool,X686)))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X690,bool),'const.pair.,_2'(s(X690,X691),s(bool,X689)))))) )),
  inference(superposition,[],[f45,f133])).
fof(f13345,plain,(
  ( ! [X639,X635,X640,X636,X638,X637] : (s(bool,'T_0') != s(bool,f(s(X635,sK6(X636,X637,bool,X635)))) | s(fun(X635,bool),X636) = s(fun(X635,bool),X637) | s(bool,X638) = s(bool,g(s(X635,sK6(X636,X637,bool,X635)))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X639,bool),'const.pair.,_2'(s(X639,X640),s(bool,X638)))))) )),
  inference(superposition,[],[f45,f133])).
fof(f13826,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK6(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X3,bool),'const.pair.,_2'(s(X3,X4),s(bool,f(s(X0,sK6(X1,X2,bool,X0))))))))) )),
  inference(condensation,[],[f13825])).
fof(f13825,plain,(
  ( ! [X389,X391,X393,X390,X392,X394] : (s(fun(X389,bool),X390) = s(fun(X389,bool),X391) | s(bool,'T_0') = s(bool,g(s(X389,sK6(X390,X391,bool,X389)))) | s(bool,'T_0') = s(bool,X392) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X393,bool),'const.pair.,_2'(s(X393,X394),s(bool,f(s(X389,sK6(X390,X391,bool,X389))))))))) )),
  inference(subsumption_resolution,[],[f12977,f40])).
fof(f12977,plain,(
  ( ! [X389,X391,X393,X390,X392,X394] : (s(bool,'F_0') != s(bool,X392) | s(fun(X389,bool),X390) = s(fun(X389,bool),X391) | s(bool,'T_0') = s(bool,g(s(X389,sK6(X390,X391,bool,X389)))) | s(bool,'T_0') = s(bool,X392) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X393,bool),'const.pair.,_2'(s(X393,X394),s(bool,f(s(X389,sK6(X390,X391,bool,X389))))))))) )),
  inference(superposition,[],[f370,f133])).
fof(f15069,plain,(
  ( ! [X790,X795,X793,X791,X794,X792] : (s(bool,f(s(X790,sK6(X791,X792,bool,X790)))) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X793),'const.pair.,_2'(s(bool,X794),s(X793,X795))))) | s(fun(X790,bool),X791) = s(fun(X790,bool),X792) | s(bool,'T_0') = s(bool,X794) | s(bool,'T_0') = s(bool,g(s(X790,sK6(X791,X792,bool,X790))))) )),
  inference(superposition,[],[f45,f156])).
fof(f16976,plain,(
  ( ! [X887,X885,X888,X886,X884,X889] : (s(bool,f(s(X884,sK6(X885,X886,bool,X884)))) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X887,bool),'const.pair.,_2'(s(X887,X888),s(bool,X889))))) | s(fun(X884,bool),X885) = s(fun(X884,bool),X886) | s(bool,'T_0') = s(bool,X889) | s(bool,'T_0') = s(bool,g(s(X884,sK6(X885,X886,bool,X884))))) )),
  inference(superposition,[],[f45,f157])).
fof(f18797,plain,(
  ( ! [X978,X976,X975,X979,X977,X974] : (s(bool,X977) != s(bool,f(s(X974,sK6(X975,X976,bool,X974)))) | s(fun(X974,bool),X975) = s(fun(X974,bool),X976) | s(bool,X977) = s(bool,X978) | s(bool,'T_0') = s(bool,X978) | s(bool,X979) = s(bool,g(s(X974,sK6(X975,X976,bool,X974)))) | s(bool,'T_0') = s(bool,X979)) )),
  inference(superposition,[],[f45,f146])).
fof(f19153,plain,(
  ( ! [X978,X976,X975,X979,X977,X974] : (s(bool,X977) != s(bool,f(s(X974,sK6(X975,X976,bool,X974)))) | s(fun(X974,bool),X975) = s(fun(X974,bool),X976) | s(bool,X978) = s(bool,g(s(X974,sK6(X975,X976,bool,X974)))) | s(bool,'T_0') = s(bool,X977) | s(bool,X978) = s(bool,X979) | s(bool,'T_0') = s(bool,X979)) )),
  inference(superposition,[],[f45,f146])).
fof(f19331,plain,(
  ( ! [X971,X969,X967,X970,X968,X966] : (s(bool,X969) != s(bool,f(s(X966,sK6(X967,X968,bool,X966)))) | s(fun(X966,bool),X967) = s(fun(X966,bool),X968) | s(bool,X969) = s(bool,X970) | s(bool,'T_0') = s(bool,g(s(X966,sK6(X967,X968,bool,X966)))) | s(bool,X970) = s(bool,X971) | s(bool,'T_0') = s(bool,X971)) )),
  inference(superposition,[],[f45,f146])).
fof(f19565,plain,(
  ( ! [X887,X885,X890,X888,X886,X889] : (s(bool,'T_0') != s(bool,f(s(X885,sK6(X886,X887,bool,X885)))) | s(fun(X885,bool),X886) = s(fun(X885,bool),X887) | s(bool,X888) = s(bool,X889) | s(bool,X889) = s(bool,g(s(X885,sK6(X886,X887,bool,X885)))) | s(bool,X888) = s(bool,X890) | s(bool,'T_0') = s(bool,X890)) )),
  inference(superposition,[],[f45,f146])).
fof(f22898,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK6(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,X3) = s(bool,f(s(X0,sK6(X1,X2,bool,X0)))) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(condensation,[],[f22897])).
fof(f22897,plain,(
  ( ! [X540,X542,X538,X541,X543,X539] : (s(fun(X538,bool),X539) = s(fun(X538,bool),X540) | s(bool,'T_0') = s(bool,g(s(X538,sK6(X539,X540,bool,X538)))) | s(bool,X542) = s(bool,f(s(X538,sK6(X539,X540,bool,X538)))) | s(bool,'T_0') = s(bool,X541) | s(bool,X542) = s(bool,X543) | s(bool,'T_0') = s(bool,X543)) )),
  inference(subsumption_resolution,[],[f19077,f40])).
fof(f19077,plain,(
  ( ! [X540,X542,X538,X541,X543,X539] : (s(bool,'F_0') != s(bool,X541) | s(fun(X538,bool),X539) = s(fun(X538,bool),X540) | s(bool,'T_0') = s(bool,g(s(X538,sK6(X539,X540,bool,X538)))) | s(bool,X542) = s(bool,f(s(X538,sK6(X539,X540,bool,X538)))) | s(bool,'T_0') = s(bool,X541) | s(bool,X542) = s(bool,X543) | s(bool,'T_0') = s(bool,X543)) )),
  inference(superposition,[],[f370,f146])).
fof(f34425,plain,(
  ( ! [X327,X333,X329,X331,X332,X328,X330] : (~p(s(bool,happ(s(fun(bool,bool),X330),s(bool,'T_0')))) | ~p(s(bool,happ(s(fun(X331,bool),X332),s(X331,X333)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X331,bool),X332),s(fun(bool,bool),X330),s('type.pair.prod'(X331,bool),'const.pair.,_2'(s(X331,X333),s(bool,g(s(X327,sK6(X328,X329,bool,X327))))))))) | s(fun(X327,bool),X328) = s(fun(X327,bool),X329) | s(bool,'T_0') = s(bool,f(s(X327,sK6(X328,X329,bool,X327))))) )),
  inference(superposition,[],[f785,f2559])).
fof(f34472,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,f(s(X0,sK6(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,X3) = s(bool,g(s(X0,sK6(X1,X2,bool,X0)))) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(condensation,[],[f34471])).
fof(f34471,plain,(
  ( ! [X70,X74,X72,X71,X75,X73] : (s(bool,'T_0') = s(bool,X73) | s(fun(X70,bool),X71) = s(fun(X70,bool),X72) | s(bool,'T_0') = s(bool,f(s(X70,sK6(X71,X72,bool,X70)))) | s(bool,X74) = s(bool,g(s(X70,sK6(X71,X72,bool,X70)))) | s(bool,X74) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(duplicate_literal_removal,[],[f34342])).
fof(f34342,plain,(
  ( ! [X70,X74,X72,X71,X75,X73] : (s(bool,'T_0') = s(bool,X73) | s(fun(X70,bool),X71) = s(fun(X70,bool),X72) | s(bool,'T_0') = s(bool,f(s(X70,sK6(X71,X72,bool,X70)))) | s(bool,X74) = s(bool,g(s(X70,sK6(X71,X72,bool,X70)))) | s(bool,'T_0') = s(bool,X73) | s(bool,X74) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f2559,f146])).
fof(f34474,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,f(s(X0,sK6(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X3,bool),'const.pair.,_2'(s(X3,X4),s(bool,g(s(X0,sK6(X1,X2,bool,X0))))))))) )),
  inference(condensation,[],[f34473])).
fof(f34473,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (s(bool,'T_0') = s(bool,X40) | s(fun(X37,bool),X38) = s(fun(X37,bool),X39) | s(bool,'T_0') = s(bool,f(s(X37,sK6(X38,X39,bool,X37)))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X41,bool),'const.pair.,_2'(s(X41,X42),s(bool,g(s(X37,sK6(X38,X39,bool,X37))))))))) )),
  inference(duplicate_literal_removal,[],[f34338])).
fof(f34338,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (s(bool,'T_0') = s(bool,X40) | s(fun(X37,bool),X38) = s(fun(X37,bool),X39) | s(bool,'T_0') = s(bool,f(s(X37,sK6(X38,X39,bool,X37)))) | s(bool,'T_0') = s(bool,X40) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X41,bool),'const.pair.,_2'(s(X41,X42),s(bool,g(s(X37,sK6(X38,X39,bool,X37))))))))) )),
  inference(superposition,[],[f2559,f133])).
fof(f34476,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,f(s(X0,sK6(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X3),'const.pair.,_2'(s(bool,g(s(X0,sK6(X1,X2,bool,X0)))),s(X3,X4)))))) )),
  inference(condensation,[],[f34475])).
fof(f34475,plain,(
  ( ! [X26,X24,X23,X27,X25,X22] : (s(bool,'T_0') = s(bool,X25) | s(fun(X22,bool),X23) = s(fun(X22,bool),X24) | s(bool,'T_0') = s(bool,f(s(X22,sK6(X23,X24,bool,X22)))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X26),'const.pair.,_2'(s(bool,g(s(X22,sK6(X23,X24,bool,X22)))),s(X26,X27)))))) )),
  inference(duplicate_literal_removal,[],[f34336])).
fof(f34336,plain,(
  ( ! [X26,X24,X23,X27,X25,X22] : (s(bool,'T_0') = s(bool,X25) | s(fun(X22,bool),X23) = s(fun(X22,bool),X24) | s(bool,'T_0') = s(bool,f(s(X22,sK6(X23,X24,bool,X22)))) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X26),'const.pair.,_2'(s(bool,g(s(X22,sK6(X23,X24,bool,X22)))),s(X26,X27)))))) )),
  inference(superposition,[],[f2559,f132])).
fof(f37576,plain,(
  ( ! [X90,X88,X87,X85,X89,X86,X84] : (s(fun(X84,bool),X85) = s(fun(X84,bool),X86) | s(bool,X87) = s(bool,g(s(X84,sK6(X85,X86,bool,X84)))) | s(bool,'T_0') = s(bool,X87) | s(bool,X88) = s(bool,X89) | s(bool,X89) = s(bool,f(s(X84,sK6(X85,X86,bool,X84)))) | s(bool,X88) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(trivial_inequality_removal,[],[f37550])).
fof(f37550,plain,(
  ( ! [X90,X88,X87,X85,X89,X86,X84] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(X84,bool),X85) = s(fun(X84,bool),X86) | s(bool,X87) = s(bool,g(s(X84,sK6(X85,X86,bool,X84)))) | s(bool,'T_0') = s(bool,X87) | s(bool,X88) = s(bool,X89) | s(bool,X89) = s(bool,f(s(X84,sK6(X85,X86,bool,X84)))) | s(bool,X88) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f366,f146])).
fof(f37577,plain,(
  ( ! [X68,X66,X64,X69,X67,X65,X63] : (s(fun(X63,bool),X64) = s(fun(X63,bool),X65) | s(bool,X66) = s(bool,g(s(X63,sK6(X64,X65,bool,X63)))) | s(bool,'T_0') = s(bool,X66) | s(bool,X67) = s(bool,f(s(X63,sK6(X64,X65,bool,X63)))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X68,bool),'const.pair.,_2'(s(X68,X69),s(bool,X67)))))) )),
  inference(trivial_inequality_removal,[],[f37547])).
fof(f37547,plain,(
  ( ! [X68,X66,X64,X69,X67,X65,X63] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(X63,bool),X64) = s(fun(X63,bool),X65) | s(bool,X66) = s(bool,g(s(X63,sK6(X64,X65,bool,X63)))) | s(bool,'T_0') = s(bool,X66) | s(bool,X67) = s(bool,f(s(X63,sK6(X64,X65,bool,X63)))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X68,bool),'const.pair.,_2'(s(X68,X69),s(bool,X67)))))) )),
  inference(superposition,[],[f366,f133])).
fof(f37578,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (s(fun(X42,bool),X43) = s(fun(X42,bool),X44) | s(bool,X45) = s(bool,g(s(X42,sK6(X43,X44,bool,X42)))) | s(bool,'T_0') = s(bool,X45) | s(bool,X46) = s(bool,f(s(X42,sK6(X43,X44,bool,X42)))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X47),'const.pair.,_2'(s(bool,X46),s(X47,X48)))))) )),
  inference(trivial_inequality_removal,[],[f37544])).
fof(f37544,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(X42,bool),X43) = s(fun(X42,bool),X44) | s(bool,X45) = s(bool,g(s(X42,sK6(X43,X44,bool,X42)))) | s(bool,'T_0') = s(bool,X45) | s(bool,X46) = s(bool,f(s(X42,sK6(X43,X44,bool,X42)))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X47),'const.pair.,_2'(s(bool,X46),s(X47,X48)))))) )),
  inference(superposition,[],[f366,f132])).
fof(f37580,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,X11) = s(bool,g(s(X8,sK6(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11) | s(bool,X12) = s(bool,f(s(X8,sK6(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(trivial_inequality_removal,[],[f37538])).
fof(f37538,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,X11) = s(bool,g(s(X8,sK6(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11) | s(bool,X12) = s(bool,f(s(X8,sK6(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f366,f56])).
fof(f22931,plain,(
  ( ! [X540,X542,X536,X538,X534,X541,X537,X539,X535] : (~p(s(bool,happ(s(fun(X537,bool),X538),s(X537,X539)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X537,bool),X538),s(fun(X534,bool),X535),s('type.pair.prod'(X537,X534),'const.pair.,_2'(s(X537,X539),s(X534,X536)))))) | s(bool,X540) = s(bool,X541) | s(bool,X541) = s(bool,happ(s(fun(X534,bool),X535),s(X534,X536))) | s(bool,X540) = s(bool,X542) | s(bool,'T_0') = s(bool,X542)) )),
  inference(subsumption_resolution,[],[f19508,f37])).
fof(f19508,plain,(
  ( ! [X540,X542,X536,X538,X534,X541,X537,X539,X535] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X537,bool),X538),s(X537,X539)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X537,bool),X538),s(fun(X534,bool),X535),s('type.pair.prod'(X537,X534),'const.pair.,_2'(s(X537,X539),s(X534,X536)))))) | s(bool,X540) = s(bool,X541) | s(bool,X541) = s(bool,happ(s(fun(X534,bool),X535),s(X534,X536))) | s(bool,X540) = s(bool,X542) | s(bool,'T_0') = s(bool,X542)) )),
  inference(superposition,[],[f785,f146])).
fof(f2763,plain,(
  ( ! [X14,X12,X17,X15,X13,X18,X16] : (~p(s(bool,happ(s(fun(X15,bool),X16),s(X15,X17)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X15,bool),X16),s(fun(X12,bool),X13),s('type.pair.prod'(X15,X12),'const.pair.,_2'(s(X15,X17),s(X12,X14)))))) | s(bool,X18) = s(bool,happ(s(fun(X12,bool),X13),s(X12,X14))) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f2729,f37])).
fof(f2729,plain,(
  ( ! [X14,X12,X17,X15,X13,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X15,bool),X16),s(X15,X17)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X15,bool),X16),s(fun(X12,bool),X13),s('type.pair.prod'(X15,X12),'const.pair.,_2'(s(X15,X17),s(X12,X14)))))) | s(bool,X18) = s(bool,happ(s(fun(X12,bool),X13),s(X12,X14))) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f785,f56])).
fof(f2699,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (~p(s(bool,happ(s(fun(bool,bool),X6),s(bool,'T_0')))) | ~p(s(bool,happ(s(fun(X7,bool),X8),s(X7,X9)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X7,bool),X8),s(fun(bool,bool),X6),s('type.pair.prod'(X7,bool),'const.pair.,_2'(s(X7,X9),s(bool,X5)))))) | s(bool,X5) = s(bool,X10) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f785,f56])).
fof(f18790,plain,(
  ( ! [X926,X930,X928,X932,X927,X931,X929,X933] : (~p(s(bool,happ(s(fun(bool,bool),X928),s(bool,X927)))) | ~p(s(bool,happ(s(fun(X929,bool),X930),s(X929,X931)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X929,bool),X930),s(fun(bool,bool),X928),s('type.pair.prod'(X929,bool),'const.pair.,_2'(s(X929,X931),s(bool,X926)))))) | s(bool,X927) = s(bool,X932) | s(bool,'T_0') = s(bool,X932) | s(bool,X926) = s(bool,X933) | s(bool,'T_0') = s(bool,X933)) )),
  inference(superposition,[],[f785,f146])).
fof(f19558,plain,(
  ( ! [X843,X841,X845,X839,X842,X840,X846,X844] : (~p(s(bool,happ(s(fun(bool,bool),X840),s(bool,'T_0')))) | ~p(s(bool,happ(s(fun(X841,bool),X842),s(X841,X843)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X841,bool),X842),s(fun(bool,bool),X840),s('type.pair.prod'(X841,bool),'const.pair.,_2'(s(X841,X843),s(bool,X839)))))) | s(bool,X844) = s(bool,X845) | s(bool,X839) = s(bool,X845) | s(bool,X844) = s(bool,X846) | s(bool,'T_0') = s(bool,X846)) )),
  inference(superposition,[],[f785,f146])).
fof(f19736,plain,(
  ( ! [X926,X930,X928,X932,X927,X931,X929,X933] : (~p(s(bool,happ(s(fun(bool,bool),X928),s(bool,X927)))) | ~p(s(bool,happ(s(fun(X929,bool),X930),s(X929,X931)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X929,bool),X930),s(fun(bool,bool),X928),s('type.pair.prod'(X929,bool),'const.pair.,_2'(s(X929,X931),s(bool,X926)))))) | s(bool,X926) = s(bool,X932) | s(bool,X932) = s(bool,X933) | s(bool,'T_0') = s(bool,X933) | s(bool,'T_0') = s(bool,X927)) )),
  inference(superposition,[],[f785,f146])).
fof(f19914,plain,(
  ( ! [X926,X930,X928,X932,X927,X931,X929,X933] : (~p(s(bool,happ(s(fun(bool,bool),X928),s(bool,X927)))) | ~p(s(bool,happ(s(fun(X929,bool),X930),s(X929,X931)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X929,bool),X930),s(fun(bool,bool),X928),s('type.pair.prod'(X929,bool),'const.pair.,_2'(s(X929,X931),s(bool,X926)))))) | s(bool,X927) = s(bool,X932) | s(bool,X932) = s(bool,X933) | s(bool,'T_0') = s(bool,X933) | s(bool,'T_0') = s(bool,X926)) )),
  inference(superposition,[],[f785,f146])).
fof(f19406,plain,(
  ( ! [X350,X352,X354,X349,X351,X353,X355,X348] : (~p(s(bool,happ(s(fun(bool,bool),X349),s(bool,X348)))) | ~p(s(bool,happ(s(fun(X350,bool),X351),s(X350,X352)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X350,bool),X351),s(fun(bool,bool),X349),s('type.pair.prod'(X350,bool),'const.pair.,_2'(s(X350,X352),s(bool,'T_0')))))) | s(bool,X353) = s(bool,X354) | s(bool,X348) = s(bool,X354) | s(bool,X353) = s(bool,X355) | s(bool,'T_0') = s(bool,X355)) )),
  inference(superposition,[],[f785,f146])).
fof(f2704,plain,(
  ( ! [X39,X37,X35,X38,X36,X34] : (~p(s(bool,happ(s(fun(bool,bool),X35),s(bool,X34)))) | ~p(s(bool,happ(s(fun(X36,bool),X37),s(X36,X38)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X36,bool),X37),s(fun(bool,bool),X35),s('type.pair.prod'(X36,bool),'const.pair.,_2'(s(X36,X38),s(bool,'T_0')))))) | s(bool,X34) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f785,f56])).
fof(f762,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~p(s(bool,happ(s(fun(bool,bool),X6),s(bool,'T_0')))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X4,bool),X7),s(fun(bool,bool),X6),s('type.pair.prod'(X4,bool),X5)))) | ~p(s(bool,happ(s(fun(X4,bool),X7),s(X4,'const.pair.FST_1'(s('type.pair.prod'(X4,bool),X5)))))) | s(bool,X8) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X4,bool),X5))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f48,f56])).
fof(f18801,plain,(
  ( ! [X1002,X1000,X998,X1003,X1001,X999,X997] : (~p(s(bool,happ(s(fun(bool,bool),X1000),s(bool,X999)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X997,bool),X1001),s(fun(bool,bool),X1000),s('type.pair.prod'(X997,bool),X998)))) | ~p(s(bool,happ(s(fun(X997,bool),X1001),s(X997,'const.pair.FST_1'(s('type.pair.prod'(X997,bool),X998)))))) | s(bool,X999) = s(bool,X1002) | s(bool,'T_0') = s(bool,X1002) | s(bool,X1003) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X997,bool),X998))) | s(bool,'T_0') = s(bool,X1003)) )),
  inference(superposition,[],[f48,f146])).
fof(f19569,plain,(
  ( ! [X912,X914,X911,X909,X913,X910,X908] : (~p(s(bool,happ(s(fun(bool,bool),X910),s(bool,'T_0')))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X908,bool),X911),s(fun(bool,bool),X910),s('type.pair.prod'(X908,bool),X909)))) | ~p(s(bool,happ(s(fun(X908,bool),X911),s(X908,'const.pair.FST_1'(s('type.pair.prod'(X908,bool),X909)))))) | s(bool,X912) = s(bool,X913) | s(bool,X913) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X908,bool),X909))) | s(bool,X912) = s(bool,X914) | s(bool,'T_0') = s(bool,X914)) )),
  inference(superposition,[],[f48,f146])).
fof(f19747,plain,(
  ( ! [X1002,X1000,X998,X1003,X1001,X999,X997] : (~p(s(bool,happ(s(fun(bool,bool),X1000),s(bool,X999)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X997,bool),X1001),s(fun(bool,bool),X1000),s('type.pair.prod'(X997,bool),X998)))) | ~p(s(bool,happ(s(fun(X997,bool),X1001),s(X997,'const.pair.FST_1'(s('type.pair.prod'(X997,bool),X998)))))) | s(bool,X1002) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X997,bool),X998))) | s(bool,X1002) = s(bool,X1003) | s(bool,'T_0') = s(bool,X1003) | s(bool,'T_0') = s(bool,X999)) )),
  inference(superposition,[],[f48,f146])).
fof(f19925,plain,(
  ( ! [X1002,X1000,X998,X1003,X1001,X999,X997] : (~p(s(bool,happ(s(fun(bool,bool),X1000),s(bool,X999)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X997,bool),X1001),s(fun(bool,bool),X1000),s('type.pair.prod'(X997,bool),X998)))) | ~p(s(bool,happ(s(fun(X997,bool),X1001),s(X997,'const.pair.FST_1'(s('type.pair.prod'(X997,bool),X998)))))) | s(bool,X999) = s(bool,X1002) | s(bool,X1002) = s(bool,X1003) | s(bool,'T_0') = s(bool,X1003) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X997,bool),X998)))) )),
  inference(superposition,[],[f48,f146])).
fof(f46,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X0,bool),X2),s(fun(X1,bool),X3),s('type.pair.prod'(X0,X1),X4)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X4))))))) )),
  inference(cnf_transformation,[],[f33])).
fof(f47,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X0,bool),X2),s(fun(X1,bool),X3),s('type.pair.prod'(X0,X1),X4)))) | p(s(bool,happ(s(fun(X1,bool),X3),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X4))))))) )),
  inference(cnf_transformation,[],[f33])).
fof(f475,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (p(s(bool,happ(s(fun(X6,bool),X7),s(X6,'const.pair.FST_1'(s('type.pair.prod'(X6,X8),X10)))))) | s(bool,X11) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X6,bool),X7),s(fun(X8,bool),X9),s('type.pair.prod'(X6,X8),X10))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(subsumption_resolution,[],[f465,f37])).
fof(f465,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X6,bool),X7),s(X6,'const.pair.FST_1'(s('type.pair.prod'(X6,X8),X10)))))) | s(bool,X11) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X6,bool),X7),s(fun(X8,bool),X9),s('type.pair.prod'(X6,X8),X10))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f46,f56])).
fof(f550,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (p(s(bool,happ(s(fun(X8,bool),X9),s(X8,'const.pair.SND_1'(s('type.pair.prod'(X6,X8),X10)))))) | s(bool,X11) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X6,bool),X7),s(fun(X8,bool),X9),s('type.pair.prod'(X6,X8),X10))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(subsumption_resolution,[],[f540,f37])).
fof(f540,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X8,bool),X9),s(X8,'const.pair.SND_1'(s('type.pair.prod'(X6,X8),X10)))))) | s(bool,X11) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X6,bool),X7),s(fun(X8,bool),X9),s('type.pair.prod'(X6,X8),X10))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f47,f56])).
fof(f786,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (~p(s(bool,happ(s(fun(X7,bool),X9),s(X7,'const.pair.FST_1'(s('type.pair.prod'(X7,X5),X8)))))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X7,bool),X9),s(fun(X5,bool),X6),s('type.pair.prod'(X7,X5),X8)))) | s(bool,X10) = s(bool,happ(s(fun(X5,bool),X6),s(X5,'const.pair.SND_1'(s('type.pair.prod'(X7,X5),X8))))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f770,f37])).
fof(f770,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (~p(s(bool,'T_0')) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X7,bool),X9),s(fun(X5,bool),X6),s('type.pair.prod'(X7,X5),X8)))) | ~p(s(bool,happ(s(fun(X7,bool),X9),s(X7,'const.pair.FST_1'(s('type.pair.prod'(X7,X5),X8)))))) | s(bool,X10) = s(bool,happ(s(fun(X5,bool),X6),s(X5,'const.pair.SND_1'(s('type.pair.prod'(X7,X5),X8))))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f48,f56])).
fof(f22928,plain,(
  ( ! [X511,X516,X512,X514,X510,X517,X513,X515] : (p(s(bool,happ(s(fun(X512,bool),X513),s(X512,'const.pair.SND_1'(s('type.pair.prod'(X510,X512),X514)))))) | s(bool,X515) = s(bool,X516) | s(bool,X516) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X510,bool),X511),s(fun(X512,bool),X513),s('type.pair.prod'(X510,X512),X514))) | s(bool,X515) = s(bool,X517) | s(bool,'T_0') = s(bool,X517)) )),
  inference(subsumption_resolution,[],[f19505,f37])).
fof(f19505,plain,(
  ( ! [X511,X516,X512,X514,X510,X517,X513,X515] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X512,bool),X513),s(X512,'const.pair.SND_1'(s('type.pair.prod'(X510,X512),X514)))))) | s(bool,X515) = s(bool,X516) | s(bool,X516) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X510,bool),X511),s(fun(X512,bool),X513),s('type.pair.prod'(X510,X512),X514))) | s(bool,X515) = s(bool,X517) | s(bool,'T_0') = s(bool,X517)) )),
  inference(superposition,[],[f47,f146])).
fof(f22929,plain,(
  ( ! [X524,X520,X522,X518,X525,X521,X523,X519] : (p(s(bool,happ(s(fun(X518,bool),X519),s(X518,'const.pair.FST_1'(s('type.pair.prod'(X518,X520),X522)))))) | s(bool,X523) = s(bool,X524) | s(bool,X524) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X518,bool),X519),s(fun(X520,bool),X521),s('type.pair.prod'(X518,X520),X522))) | s(bool,X523) = s(bool,X525) | s(bool,'T_0') = s(bool,X525)) )),
  inference(subsumption_resolution,[],[f19506,f37])).
fof(f19506,plain,(
  ( ! [X524,X520,X522,X518,X525,X521,X523,X519] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X518,bool),X519),s(X518,'const.pair.FST_1'(s('type.pair.prod'(X518,X520),X522)))))) | s(bool,X523) = s(bool,X524) | s(bool,X524) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X518,bool),X519),s(fun(X520,bool),X521),s('type.pair.prod'(X518,X520),X522))) | s(bool,X523) = s(bool,X525) | s(bool,'T_0') = s(bool,X525)) )),
  inference(superposition,[],[f46,f146])).
fof(f22930,plain,(
  ( ! [X532,X528,X530,X526,X533,X529,X531,X527] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X528,bool),X530),s(fun(X526,bool),X527),s('type.pair.prod'(X528,X526),X529)))) | ~p(s(bool,happ(s(fun(X528,bool),X530),s(X528,'const.pair.FST_1'(s('type.pair.prod'(X528,X526),X529)))))) | s(bool,X531) = s(bool,X532) | s(bool,X532) = s(bool,happ(s(fun(X526,bool),X527),s(X526,'const.pair.SND_1'(s('type.pair.prod'(X528,X526),X529))))) | s(bool,X531) = s(bool,X533) | s(bool,'T_0') = s(bool,X533)) )),
  inference(subsumption_resolution,[],[f19507,f37])).
fof(f19507,plain,(
  ( ! [X532,X528,X530,X526,X533,X529,X531,X527] : (~p(s(bool,'T_0')) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X528,bool),X530),s(fun(X526,bool),X527),s('type.pair.prod'(X528,X526),X529)))) | ~p(s(bool,happ(s(fun(X528,bool),X530),s(X528,'const.pair.FST_1'(s('type.pair.prod'(X528,X526),X529)))))) | s(bool,X531) = s(bool,X532) | s(bool,X532) = s(bool,happ(s(fun(X526,bool),X527),s(X526,'const.pair.SND_1'(s('type.pair.prod'(X528,X526),X529))))) | s(bool,X531) = s(bool,X533) | s(bool,'T_0') = s(bool,X533)) )),
  inference(superposition,[],[f48,f146])).
fof(f15352,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,X2),s(X1,X3))))) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X4),'const.pair.,_2'(s(bool,X0),s(X4,X5))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(condensation,[],[f15351])).
fof(f15351,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67] : (s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X67),'const.pair.,_2'(s(bool,X68),s(X67,X69))))) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X70),'const.pair.,_2'(s(bool,X71),s(X70,X72))))) | s(bool,'T_0') = s(bool,X71) | s(bool,'T_0') = s(bool,X66) | s(bool,'T_0') = s(bool,X68)) )),
  inference(duplicate_literal_removal,[],[f14786])).
fof(f14786,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67] : (s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X67),'const.pair.,_2'(s(bool,X68),s(X67,X69))))) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X70),'const.pair.,_2'(s(bool,X71),s(X70,X72))))) | s(bool,'T_0') = s(bool,X71) | s(bool,'T_0') = s(bool,X66) | s(bool,'T_0') = s(bool,X68) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f156,f156])).
fof(f17266,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,X2),s(X1,X3))))) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X4,bool),'const.pair.,_2'(s(X4,X5),s(bool,X0))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(condensation,[],[f17265])).
fof(f17265,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67] : (s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X67),'const.pair.,_2'(s(bool,X68),s(X67,X69))))) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X70,bool),'const.pair.,_2'(s(X70,X71),s(bool,X72))))) | s(bool,'T_0') = s(bool,X72) | s(bool,'T_0') = s(bool,X66) | s(bool,'T_0') = s(bool,X68)) )),
  inference(duplicate_literal_removal,[],[f16668])).
fof(f16668,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67] : (s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X67),'const.pair.,_2'(s(bool,X68),s(X67,X69))))) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X70,bool),'const.pair.,_2'(s(X70,X71),s(bool,X72))))) | s(bool,'T_0') = s(bool,X72) | s(bool,'T_0') = s(bool,X66) | s(bool,'T_0') = s(bool,X68) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f157,f156])).
fof(f17264,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'const.pair.SND_1'(s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,X3))))) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X4,bool),'const.pair.,_2'(s(X4,X5),s(bool,X0))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(condensation,[],[f17263])).
fof(f17263,plain,(
  ( ! [X78,X76,X74,X79,X77,X75,X73] : (s(bool,'const.pair.SND_1'(s('type.pair.prod'(X74,bool),'const.pair.,_2'(s(X74,X75),s(bool,X76))))) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X77,bool),'const.pair.,_2'(s(X77,X78),s(bool,X79))))) | s(bool,'T_0') = s(bool,X79) | s(bool,'T_0') = s(bool,X73) | s(bool,'T_0') = s(bool,X76)) )),
  inference(duplicate_literal_removal,[],[f16669])).
fof(f16669,plain,(
  ( ! [X78,X76,X74,X79,X77,X75,X73] : (s(bool,'const.pair.SND_1'(s('type.pair.prod'(X74,bool),'const.pair.,_2'(s(X74,X75),s(bool,X76))))) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X77,bool),'const.pair.,_2'(s(X77,X78),s(bool,X79))))) | s(bool,'T_0') = s(bool,X79) | s(bool,'T_0') = s(bool,X73) | s(bool,'T_0') = s(bool,X76) | s(bool,'T_0') = s(bool,X73)) )),
  inference(superposition,[],[f157,f157])).
fof(f3651,plain,(
  ( ! [X441,X443,X444,X440,X442,X439] : (s(X441,f(s(bool,X442))) != s(X441,g(s(bool,X442))) | s(fun(bool,X441),X439) = s(fun(bool,X441),X440) | s(bool,'T_0') = s(bool,X442) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X443),'const.pair.,_2'(s(bool,sK6(X439,X440,X441,bool)),s(X443,X444)))))) )),
  inference(superposition,[],[f45,f132])).
fof(f13038,plain,(
  ( ! [X732,X734,X730,X733,X735,X731] : (s(X732,f(s(bool,X733))) != s(X732,g(s(bool,X733))) | s(fun(bool,X732),X730) = s(fun(bool,X732),X731) | s(bool,'T_0') = s(bool,X733) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X734,bool),'const.pair.,_2'(s(X734,X735),s(bool,sK6(X730,X731,X732,bool))))))) )),
  inference(superposition,[],[f45,f133])).
fof(f35170,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X34,X32] : (s(X29,f(s(bool,X28))) != s(X29,g(s(bool,X28))) | s(fun(bool,X29),X30) = s(fun(bool,X29),X31) | s(bool,X32) = s(bool,sK6(X30,X31,X29,bool)) | s(bool,'T_0') = s(bool,X32) | s(bool,X33) = s(bool,X34) | s(bool,X28) = s(bool,X34) | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f361,f146])).
fof(f35169,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22,X20] : (s(X21,f(s(bool,X20))) != s(X21,g(s(bool,X20))) | s(fun(bool,X21),X22) = s(fun(bool,X21),X23) | s(bool,X24) = s(bool,sK6(X22,X23,X21,bool)) | s(bool,'T_0') = s(bool,X24) | s(bool,X20) = s(bool,X25) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X26,bool),'const.pair.,_2'(s(X26,X27),s(bool,X25)))))) )),
  inference(superposition,[],[f361,f133])).
fof(f35168,plain,(
  ( ! [X14,X12,X19,X17,X15,X13,X18,X16] : (s(X13,f(s(bool,X12))) != s(X13,g(s(bool,X12))) | s(fun(bool,X13),X14) = s(fun(bool,X13),X15) | s(bool,X16) = s(bool,sK6(X14,X15,X13,bool)) | s(bool,'T_0') = s(bool,X16) | s(bool,X12) = s(bool,X17) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X18),'const.pair.,_2'(s(bool,X17),s(X18,X19)))))) )),
  inference(superposition,[],[f361,f132])).
fof(f33897,plain,(
  ( ! [X30,X28,X29] : (s(bool,X28) = s(bool,X29) | s(bool,X29) = s(bool,f(s(bool,'T_0'))) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) ) | ($spl9 | $spl40)),
  inference(subsumption_resolution,[],[f33787,f71])).
fof(f33787,plain,(
  ( ! [X30,X28,X29] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X28) = s(bool,X29) | s(bool,X29) = s(bool,f(s(bool,'T_0'))) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) ) | $spl40),
  inference(superposition,[],[f2561,f146])).
fof(f33888,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl9 | $spl40)),
  inference(subsumption_resolution,[],[f33775,f71])).
fof(f33775,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X0) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl40),
  inference(superposition,[],[f2561,f56])).
fof(f33749,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') = s(bool,f(s(bool,X12))) | s(bool,X13) = s(bool,X14) | s(bool,X12) = s(bool,X14) | s(bool,X13) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) ) | $spl40),
  inference(superposition,[],[f2561,f146])).
fof(f33748,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') = s(bool,f(s(bool,X8))) | s(bool,X8) = s(bool,X9) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X10,bool),'const.pair.,_2'(s(X10,X11),s(bool,X9)))))) ) | $spl40),
  inference(superposition,[],[f2561,f133])).
fof(f33747,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') = s(bool,f(s(bool,X4))) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X6),'const.pair.,_2'(s(bool,X5),s(X6,X7)))))) ) | $spl40),
  inference(superposition,[],[f2561,f132])).
fof(f33745,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,f(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl40),
  inference(superposition,[],[f2561,f56])).
fof(f32960,plain,(
  ( ! [X123,X121,X118,X116,X122,X120,X119,X117] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X119,bool),'const.pair.,_2'(s(X119,X120),s(bool,X118))))) | s(bool,X121) = s(bool,X122) | s(bool,X122) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X116,bool),'const.pair.,_2'(s(X116,X117),s(bool,X118))))) | s(bool,X121) = s(bool,X123) | s(bool,'T_0') = s(bool,X123)) ) | $spl9),
  inference(subsumption_resolution,[],[f32470,f71])).
fof(f32470,plain,(
  ( ! [X123,X121,X118,X116,X122,X120,X119,X117] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X119,bool),'const.pair.,_2'(s(X119,X120),s(bool,X118))))) | s(bool,X121) = s(bool,X122) | s(bool,X122) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X116,bool),'const.pair.,_2'(s(X116,X117),s(bool,X118))))) | s(bool,X121) = s(bool,X123) | s(bool,'T_0') = s(bool,X123)) )),
  inference(superposition,[],[f1972,f146])).
fof(f32955,plain,(
  ( ! [X94,X92,X99,X97,X95,X93,X98,X96] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X95,bool),'const.pair.,_2'(s(X95,X96),s(bool,X94))))) | s(bool,X97) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X92,bool),'const.pair.,_2'(s(X92,X93),s(bool,X94))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X98,bool),'const.pair.,_2'(s(X98,X99),s(bool,X97)))))) ) | $spl9),
  inference(subsumption_resolution,[],[f32467,f71])).
fof(f32467,plain,(
  ( ! [X94,X92,X99,X97,X95,X93,X98,X96] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X95,bool),'const.pair.,_2'(s(X95,X96),s(bool,X94))))) | s(bool,X97) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X92,bool),'const.pair.,_2'(s(X92,X93),s(bool,X94))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X98,bool),'const.pair.,_2'(s(X98,X99),s(bool,X97)))))) )),
  inference(superposition,[],[f1972,f133])).
fof(f32954,plain,(
  ( ! [X70,X68,X74,X72,X71,X69,X75,X73] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X71,bool),'const.pair.,_2'(s(X71,X72),s(bool,X70))))) | s(bool,X73) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X68,bool),'const.pair.,_2'(s(X68,X69),s(bool,X70))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X74),'const.pair.,_2'(s(bool,X73),s(X74,X75)))))) ) | $spl9),
  inference(subsumption_resolution,[],[f32464,f71])).
fof(f32464,plain,(
  ( ! [X70,X68,X74,X72,X71,X69,X75,X73] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X71,bool),'const.pair.,_2'(s(X71,X72),s(bool,X70))))) | s(bool,X73) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X68,bool),'const.pair.,_2'(s(X68,X69),s(bool,X70))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X74),'const.pair.,_2'(s(bool,X73),s(X74,X75)))))) )),
  inference(superposition,[],[f1972,f132])).
fof(f32190,plain,(
  ( ! [X196,X194,X192,X195,X193] : (s(bool,X192) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X192),s(bool,X193))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X194,bool),'const.pair.,_2'(s(X194,X195),s(bool,X193))))) | s(bool,X192) = s(bool,X196) | s(bool,'T_0') = s(bool,X196)) )),
  inference(superposition,[],[f1972,f56])).
fof(f31523,plain,(
  ( ! [X111,X109,X107,X112,X110,X108,X106,X113] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X109),'const.pair.,_2'(s(bool,X108),s(X109,X110))))) | s(bool,X111) = s(bool,X112) | s(bool,X112) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X106,bool),'const.pair.,_2'(s(X106,X107),s(bool,X108))))) | s(bool,X111) = s(bool,X113) | s(bool,'T_0') = s(bool,X113)) ) | $spl9),
  inference(subsumption_resolution,[],[f31102,f71])).
fof(f31102,plain,(
  ( ! [X111,X109,X107,X112,X110,X108,X106,X113] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X109),'const.pair.,_2'(s(bool,X108),s(X109,X110))))) | s(bool,X111) = s(bool,X112) | s(bool,X112) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X106,bool),'const.pair.,_2'(s(X106,X107),s(bool,X108))))) | s(bool,X111) = s(bool,X113) | s(bool,'T_0') = s(bool,X113)) )),
  inference(superposition,[],[f1473,f146])).
fof(f31518,plain,(
  ( ! [X88,X87,X85,X83,X89,X86,X84,X82] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X85),'const.pair.,_2'(s(bool,X84),s(X85,X86))))) | s(bool,X87) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X82,bool),'const.pair.,_2'(s(X82,X83),s(bool,X84))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X88,bool),'const.pair.,_2'(s(X88,X89),s(bool,X87)))))) ) | $spl9),
  inference(subsumption_resolution,[],[f31099,f71])).
fof(f31099,plain,(
  ( ! [X88,X87,X85,X83,X89,X86,X84,X82] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X85),'const.pair.,_2'(s(bool,X84),s(X85,X86))))) | s(bool,X87) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X82,bool),'const.pair.,_2'(s(X82,X83),s(bool,X84))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X88,bool),'const.pair.,_2'(s(X88,X89),s(bool,X87)))))) )),
  inference(superposition,[],[f1473,f133])).
fof(f31517,plain,(
  ( ! [X61,X59,X64,X62,X60,X58,X65,X63] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X61),'const.pair.,_2'(s(bool,X60),s(X61,X62))))) | s(bool,X63) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X58,bool),'const.pair.,_2'(s(X58,X59),s(bool,X60))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X64),'const.pair.,_2'(s(bool,X63),s(X64,X65)))))) ) | $spl9),
  inference(subsumption_resolution,[],[f31096,f71])).
fof(f31096,plain,(
  ( ! [X61,X59,X64,X62,X60,X58,X65,X63] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X61),'const.pair.,_2'(s(bool,X60),s(X61,X62))))) | s(bool,X63) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X58,bool),'const.pair.,_2'(s(X58,X59),s(bool,X60))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X64),'const.pair.,_2'(s(bool,X63),s(X64,X65)))))) )),
  inference(superposition,[],[f1473,f132])).
fof(f30902,plain,(
  ( ! [X191,X189,X192,X190,X188] : (s(bool,X188) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X189),s(bool,X188))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X190,bool),'const.pair.,_2'(s(X190,X191),s(bool,X189))))) | s(bool,X188) = s(bool,X192) | s(bool,'T_0') = s(bool,X192)) )),
  inference(superposition,[],[f1473,f56])).
fof(f13346,plain,(
  ( ! [X644,X642,X645,X641,X643] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X641) = s(fun(bool,bool),X642) | s(bool,'T_0') = s(bool,sK6(X641,X642,bool,bool)) | s(bool,X643) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X644,bool),'const.pair.,_2'(s(X644,X645),s(bool,X643)))))) )),
  inference(superposition,[],[f365,f133])).
fof(f19566,plain,(
  ( ! [X895,X893,X894,X892,X891] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X891) = s(fun(bool,bool),X892) | s(bool,'T_0') = s(bool,sK6(X891,X892,bool,bool)) | s(bool,X893) = s(bool,X894) | s(bool,X894) = s(bool,g(s(bool,'F_0'))) | s(bool,X893) = s(bool,X895) | s(bool,'T_0') = s(bool,X895)) )),
  inference(superposition,[],[f365,f146])).
fof(f20156,plain,(
  ( ! [X895,X893,X894,X892,X891] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X891) = s(fun(bool,bool),X892) | s(bool,'T_0') = s(bool,sK6(X891,X892,bool,bool)) | s(bool,X893) = s(bool,X894) | s(bool,X894) = s(bool,X895) | s(bool,'T_0') = s(bool,X895) | s(bool,X893) = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f365,f146])).
fof(f20200,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X2) = s(bool,X3) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2)) )),
  inference(equality_factoring,[],[f146])).
fof(f22984,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f20193,f56])).
fof(f20193,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,X2)) )),
  inference(equality_factoring,[],[f146])).
fof(f19570,plain,(
  ( ! [X920,X918,X916,X915,X919,X917] : (s(X917,f(s(bool,'T_0'))) != s(X917,g(s(bool,'T_0'))) | s(fun(bool,X917),X915) = s(fun(bool,X917),X916) | s(bool,X918) = s(bool,X919) | s(bool,X919) = s(bool,sK6(X915,X916,X917,bool)) | s(bool,X918) = s(bool,X920) | s(bool,'T_0') = s(bool,X920)) )),
  inference(superposition,[],[f45,f146])).
fof(f19403,plain,(
  ( ! [X333,X335,X331,X332,X334,X330] : (s(bool,X330) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X330),s(bool,X331))))) | s(bool,X331) = s(bool,X332) | s(bool,'T_0') = s(bool,X332) | s(bool,X333) = s(bool,X334) | s(bool,X330) = s(bool,X334) | s(bool,X333) = s(bool,X335) | s(bool,'T_0') = s(bool,X335)) )),
  inference(superposition,[],[f133,f146])).
fof(f19402,plain,(
  ( ! [X325,X327,X329,X324,X326,X328] : (s(bool,X324) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X325),s(bool,X324))))) | s(bool,X325) = s(bool,X326) | s(bool,'T_0') = s(bool,X326) | s(bool,X327) = s(bool,X328) | s(bool,X324) = s(bool,X328) | s(bool,X327) = s(bool,X329) | s(bool,'T_0') = s(bool,X329)) )),
  inference(superposition,[],[f132,f146])).
fof(f19338,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X0)) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X2) | s(bool,X1) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f37,f146])).
fof(f22901,plain,(
  ( ! [X436,X432,X434,X431,X433,X435] : (s(X432,f(s(bool,X431))) != s(X432,g(s(bool,X431))) | s(fun(bool,X432),X433) = s(fun(bool,X432),X434) | s(bool,'T_0') = s(bool,sK6(X433,X434,X432,bool)) | s(bool,X431) = s(bool,X435) | s(bool,X435) = s(bool,X436) | s(bool,'T_0') = s(bool,X436)) ) | $spl9),
  inference(subsumption_resolution,[],[f19234,f71])).
fof(f19234,plain,(
  ( ! [X436,X432,X434,X431,X433,X435] : (s(X432,f(s(bool,X431))) != s(X432,g(s(bool,X431))) | s(fun(bool,X432),X433) = s(fun(bool,X432),X434) | s(bool,'T_0') = s(bool,sK6(X433,X434,X432,bool)) | s(bool,X431) = s(bool,X435) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X435) = s(bool,X436) | s(bool,'T_0') = s(bool,X436)) )),
  inference(superposition,[],[f365,f146])).
fof(f21203,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X2)) )),
  inference(condensation,[],[f21202])).
fof(f21202,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,X25)) | s(bool,'T_0') = s(bool,X24) | s(bool,X25) = s(bool,X26) | s(bool,X26) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(duplicate_literal_removal,[],[f19165])).
fof(f19165,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,X25)) | s(bool,'T_0') = s(bool,X24) | s(bool,X25) = s(bool,X26) | s(bool,'T_0') = s(bool,X24) | s(bool,X26) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f58,f146])).
fof(f19158,plain,(
  ( ! [X1009,X1006,X1004,X1008,X1007,X1005] : (s(X1006,f(s(bool,X1007))) != s(X1006,g(s(bool,X1007))) | s(fun(bool,X1006),X1004) = s(fun(bool,X1006),X1005) | s(bool,X1008) = s(bool,sK6(X1004,X1005,X1006,bool)) | s(bool,'T_0') = s(bool,X1007) | s(bool,X1008) = s(bool,X1009) | s(bool,'T_0') = s(bool,X1009)) )),
  inference(superposition,[],[f45,f146])).
fof(f18802,plain,(
  ( ! [X1009,X1006,X1004,X1008,X1007,X1005] : (s(X1006,f(s(bool,X1007))) != s(X1006,g(s(bool,X1007))) | s(fun(bool,X1006),X1004) = s(fun(bool,X1006),X1005) | s(bool,X1007) = s(bool,X1008) | s(bool,'T_0') = s(bool,X1008) | s(bool,X1009) = s(bool,sK6(X1004,X1005,X1006,bool)) | s(bool,'T_0') = s(bool,X1009)) )),
  inference(superposition,[],[f45,f146])).
fof(f18363,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (s(bool,X24) = s(bool,X25) | s(bool,X26) = s(bool,X27) | s(bool,X25) = s(bool,X27) | s(bool,X26) = s(bool,X28) | s(bool,'T_0') = s(bool,X28) | s(bool,X24) = s(bool,X29) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X30),'const.pair.,_2'(s(bool,X29),s(X30,X31)))))) )),
  inference(superposition,[],[f146,f132])).
fof(f18362,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (s(bool,X16) = s(bool,X17) | s(bool,X18) = s(bool,X19) | s(bool,X17) = s(bool,X19) | s(bool,X18) = s(bool,X20) | s(bool,'T_0') = s(bool,X20) | s(bool,X16) = s(bool,X21) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X22,bool),'const.pair.,_2'(s(X22,X23),s(bool,X21)))))) )),
  inference(superposition,[],[f146,f133])).
fof(f18360,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X2) = s(bool,X3) | s(bool,X1) = s(bool,X3) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4) | s(bool,X5) = s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7) | s(bool,X0) = s(bool,X5)) )),
  inference(superposition,[],[f146,f146])).
fof(f22217,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X2,bool),'const.pair.,_2'(s(X2,X3),s(bool,X4))))) | s(bool,X1) = s(bool,X5) | s(bool,X0) = s(bool,X5) | s(bool,'T_0') = s(bool,X4)) )),
  inference(condensation,[],[f22216])).
fof(f22216,plain,(
  ( ! [X134,X132,X130,X135,X133,X131,X129] : (s(bool,X133) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X130,bool),'const.pair.,_2'(s(X130,X131),s(bool,X132))))) | s(bool,X133) = s(bool,X134) | s(bool,'T_0') = s(bool,X129) | s(bool,X134) = s(bool,X135) | s(bool,'T_0') = s(bool,X135) | s(bool,'T_0') = s(bool,X132)) )),
  inference(duplicate_literal_removal,[],[f18312])).
fof(f18312,plain,(
  ( ! [X134,X132,X130,X135,X133,X131,X129] : (s(bool,X133) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X130,bool),'const.pair.,_2'(s(X130,X131),s(bool,X132))))) | s(bool,X133) = s(bool,X134) | s(bool,'T_0') = s(bool,X129) | s(bool,X134) = s(bool,X135) | s(bool,'T_0') = s(bool,X135) | s(bool,'T_0') = s(bool,X132) | s(bool,'T_0') = s(bool,X129)) )),
  inference(superposition,[],[f146,f157])).
fof(f22219,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),'const.pair.,_2'(s(bool,X3),s(X2,X4))))) | s(bool,X1) = s(bool,X5) | s(bool,X0) = s(bool,X5) | s(bool,'T_0') = s(bool,X3)) )),
  inference(condensation,[],[f22218])).
fof(f22218,plain,(
  ( ! [X127,X125,X123,X128,X126,X124,X122] : (s(bool,X126) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X123),'const.pair.,_2'(s(bool,X124),s(X123,X125))))) | s(bool,X126) = s(bool,X127) | s(bool,'T_0') = s(bool,X122) | s(bool,X127) = s(bool,X128) | s(bool,'T_0') = s(bool,X128) | s(bool,'T_0') = s(bool,X124)) )),
  inference(duplicate_literal_removal,[],[f18311])).
fof(f18311,plain,(
  ( ! [X127,X125,X123,X128,X126,X124,X122] : (s(bool,X126) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X123),'const.pair.,_2'(s(bool,X124),s(X123,X125))))) | s(bool,X126) = s(bool,X127) | s(bool,'T_0') = s(bool,X122) | s(bool,X127) = s(bool,X128) | s(bool,'T_0') = s(bool,X128) | s(bool,'T_0') = s(bool,X124) | s(bool,'T_0') = s(bool,X122)) )),
  inference(superposition,[],[f146,f156])).
fof(f22275,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X2)) )),
  inference(condensation,[],[f22274])).
fof(f22274,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X0) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(duplicate_literal_removal,[],[f18292])).
fof(f18292,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X0) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f146,f40])).
fof(f22279,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X2,bool),'const.pair.,_2'(s(X2,X3),s(bool,X4))))) | s(bool,X0) = s(bool,X1) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(condensation,[],[f22278])).
fof(f22278,plain,(
  ( ! [X377,X372,X374,X376,X378,X373,X375] : (s(bool,'T_0') = s(bool,X375) | s(bool,X376) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X372,bool),'const.pair.,_2'(s(X372,X373),s(bool,X374))))) | s(bool,X376) = s(bool,X377) | s(bool,'T_0') = s(bool,X377) | s(bool,X374) = s(bool,X378) | s(bool,'T_0') = s(bool,X378)) )),
  inference(duplicate_literal_removal,[],[f18288])).
fof(f18288,plain,(
  ( ! [X377,X372,X374,X376,X378,X373,X375] : (s(bool,'T_0') = s(bool,X375) | s(bool,X376) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X372,bool),'const.pair.,_2'(s(X372,X373),s(bool,X374))))) | s(bool,'T_0') = s(bool,X375) | s(bool,X376) = s(bool,X377) | s(bool,'T_0') = s(bool,X377) | s(bool,X374) = s(bool,X378) | s(bool,'T_0') = s(bool,X378)) )),
  inference(superposition,[],[f146,f133])).
fof(f22291,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),'const.pair.,_2'(s(bool,X3),s(X2,X4))))) | s(bool,X0) = s(bool,X1) | s(bool,X3) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(condensation,[],[f22290])).
fof(f22290,plain,(
  ( ! [X292,X294,X290,X296,X293,X295,X291] : (s(bool,'T_0') = s(bool,X293) | s(bool,X294) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X290),'const.pair.,_2'(s(bool,X291),s(X290,X292))))) | s(bool,X294) = s(bool,X295) | s(bool,'T_0') = s(bool,X295) | s(bool,X291) = s(bool,X296) | s(bool,'T_0') = s(bool,X296)) )),
  inference(duplicate_literal_removal,[],[f18275])).
fof(f18275,plain,(
  ( ! [X292,X294,X290,X296,X293,X295,X291] : (s(bool,'T_0') = s(bool,X293) | s(bool,X294) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X290),'const.pair.,_2'(s(bool,X291),s(X290,X292))))) | s(bool,'T_0') = s(bool,X293) | s(bool,X294) = s(bool,X295) | s(bool,'T_0') = s(bool,X295) | s(bool,X291) = s(bool,X296) | s(bool,'T_0') = s(bool,X296)) )),
  inference(superposition,[],[f146,f132])).
fof(f18155,plain,(
  ( ! [X391,X393,X395,X396,X392,X394] : (s(bool,X393) = s(bool,X394) | s(bool,X394) = s(bool,X395) | s(bool,'T_0') = s(bool,X395) | s(bool,X396) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X391,bool),'const.pair.,_2'(s(X391,X392),s(bool,X393))))) | s(bool,'T_0') = s(bool,X396)) )),
  inference(superposition,[],[f146,f44])).
fof(f18154,plain,(
  ( ! [X389,X385,X387,X388,X390,X386] : (s(bool,X386) = s(bool,X388) | s(bool,X388) = s(bool,X389) | s(bool,'T_0') = s(bool,X389) | s(bool,X390) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X385),'const.pair.,_2'(s(bool,X386),s(X385,X387))))) | s(bool,'T_0') = s(bool,X390)) )),
  inference(superposition,[],[f146,f43])).
fof(f22573,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X2) | s(bool,X3) = s(bool,X4) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X5) | s(bool,X1) = s(bool,X3)) )),
  inference(condensation,[],[f22572])).
fof(f22572,plain,(
  ( ! [X158,X156,X161,X159,X157,X155,X160] : (s(bool,X155) = s(bool,X156) | s(bool,X156) = s(bool,X157) | s(bool,'T_0') = s(bool,X157) | s(bool,'T_0') = s(bool,X158) | s(bool,X159) = s(bool,X160) | s(bool,X160) = s(bool,X161) | s(bool,'T_0') = s(bool,X161) | s(bool,X155) = s(bool,X159)) )),
  inference(duplicate_literal_removal,[],[f18113])).
fof(f18113,plain,(
  ( ! [X158,X156,X161,X159,X157,X155,X160] : (s(bool,X155) = s(bool,X156) | s(bool,X156) = s(bool,X157) | s(bool,'T_0') = s(bool,X157) | s(bool,'T_0') = s(bool,X158) | s(bool,'T_0') = s(bool,X158) | s(bool,X159) = s(bool,X160) | s(bool,X160) = s(bool,X161) | s(bool,'T_0') = s(bool,X161) | s(bool,X155) = s(bool,X159)) )),
  inference(superposition,[],[f146,f146])).
fof(f13415,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,X3))))) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f133])).
fof(f13427,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f13409])).
fof(f13409,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f133])).
fof(f13350,plain,(
  ( ! [X666,X668,X670,X667,X669,X665] : (s(X667,f(s(bool,'T_0'))) != s(X667,g(s(bool,'T_0'))) | s(fun(bool,X667),X665) = s(fun(bool,X667),X666) | s(bool,X668) = s(bool,sK6(X665,X666,X667,bool)) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X669,bool),'const.pair.,_2'(s(X669,X670),s(bool,X668)))))) )),
  inference(superposition,[],[f45,f133])).
fof(f13225,plain,(
  ( ! [X229,X227,X225,X230,X228,X226] : (s(bool,X225) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X226),s(bool,X225))))) | s(bool,X226) = s(bool,X227) | s(bool,'T_0') = s(bool,X227) | s(bool,X225) = s(bool,X228) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X229,bool),'const.pair.,_2'(s(X229,X230),s(bool,X228)))))) )),
  inference(superposition,[],[f132,f133])).
fof(f13852,plain,(
  ( ! [X292,X294,X296,X293,X295,X297] : (s(X293,f(s(bool,X292))) != s(X293,g(s(bool,X292))) | s(fun(bool,X293),X294) = s(fun(bool,X293),X295) | s(bool,'T_0') = s(bool,sK6(X294,X295,X293,bool)) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X296,bool),'const.pair.,_2'(s(X296,X297),s(bool,X292)))))) ) | $spl9),
  inference(subsumption_resolution,[],[f13095,f71])).
fof(f13095,plain,(
  ( ! [X292,X294,X296,X293,X295,X297] : (s(X293,f(s(bool,X292))) != s(X293,g(s(bool,X292))) | s(fun(bool,X293),X294) = s(fun(bool,X293),X295) | s(bool,'T_0') = s(bool,sK6(X294,X295,X293,bool)) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X296,bool),'const.pair.,_2'(s(X296,X297),s(bool,X292)))))) )),
  inference(superposition,[],[f365,f133])).
fof(f12891,plain,(
  ( ! [X255,X253,X251,X254,X256,X252] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X254,bool),'const.pair.,_2'(s(X254,X255),s(bool,X253))))) | s(bool,X256) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X251,bool),'const.pair.,_2'(s(X251,X252),s(bool,X253))))) | s(bool,'T_0') = s(bool,X256)) )),
  inference(superposition,[],[f133,f44])).
fof(f12809,plain,(
  ( ! [X85,X83,X84,X82] : (s(bool,X82) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X82),s(bool,X83))))) | s(bool,X83) = s(bool,X84) | s(bool,'T_0') = s(bool,X84) | s(bool,X82) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) )),
  inference(superposition,[],[f133,f56])).
fof(f12808,plain,(
  ( ! [X80,X78,X76,X81,X79,X77] : (s(bool,X76) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X76),s(bool,X77))))) | s(bool,X77) = s(bool,X78) | s(bool,'T_0') = s(bool,X78) | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X80),'const.pair.,_2'(s(bool,X79),s(X80,X81)))))) )),
  inference(superposition,[],[f133,f132])).
fof(f12807,plain,(
  ( ! [X70,X74,X72,X71,X75,X73] : (s(bool,X70) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X70),s(bool,X71))))) | s(bool,X71) = s(bool,X72) | s(bool,'T_0') = s(bool,X72) | s(bool,X70) = s(bool,X73) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X74,bool),'const.pair.,_2'(s(X74,X75),s(bool,X73)))))) )),
  inference(superposition,[],[f133,f133])).
fof(f12791,plain,(
  ( ! [X175,X173,X171,X176,X174,X172] : (s(bool,'T_0') = s(bool,X173) | s(bool,X174) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X171,bool),'const.pair.,_2'(s(X171,X172),s(bool,X173))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X175,bool),'const.pair.,_2'(s(X175,X176),s(bool,X174)))))) )),
  inference(superposition,[],[f133,f44])).
fof(f12752,plain,(
  ( ! [X21,X19,X17,X20,X18,X16] : (s(bool,X16) = s(bool,X17) | s(bool,X17) = s(bool,X18) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X19,bool),'const.pair.,_2'(s(X19,X20),s(bool,X18))))) | s(bool,X16) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f133,f56])).
fof(f12751,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (s(bool,X8) = s(bool,X9) | s(bool,X9) = s(bool,X10) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X11,bool),'const.pair.,_2'(s(X11,X12),s(bool,X10))))) | s(bool,X8) = s(bool,X13) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X14),'const.pair.,_2'(s(bool,X13),s(X14,X15)))))) )),
  inference(superposition,[],[f133,f132])).
fof(f12750,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X3,bool),'const.pair.,_2'(s(X3,X4),s(bool,X2))))) | s(bool,X0) = s(bool,X5) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X6,bool),'const.pair.,_2'(s(X6,X7),s(bool,X5)))))) )),
  inference(superposition,[],[f133,f133])).
fof(f13741,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,X2))))) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X4,bool),'const.pair.,_2'(s(X4,X5),s(bool,X3)))))) )),
  inference(condensation,[],[f13740])).
fof(f13740,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (s(bool,'T_0') = s(bool,X60) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X61,bool),'const.pair.,_2'(s(X61,X62),s(bool,X59))))) | s(bool,X59) = s(bool,X63) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X64,bool),'const.pair.,_2'(s(X64,X65),s(bool,X63)))))) )),
  inference(duplicate_literal_removal,[],[f12664])).
fof(f12664,plain,(
  ( ! [X61,X59,X64,X62,X60,X65,X63] : (s(bool,'T_0') = s(bool,X60) | s(bool,'T_0') = s(bool,X60) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X61,bool),'const.pair.,_2'(s(X61,X62),s(bool,X59))))) | s(bool,X59) = s(bool,X63) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X64,bool),'const.pair.,_2'(s(X64,X65),s(bool,X63)))))) )),
  inference(superposition,[],[f133,f133])).
fof(f2706,plain,(
  ( ! [X47,X50,X48,X46,X49] : (~p(s(bool,happ(s(fun(bool,bool),X47),s(bool,X46)))) | ~p(s(bool,happ(s(fun(X48,bool),X49),s(X48,X50)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X48,bool),X49),s(fun(bool,bool),X47),s('type.pair.prod'(X48,bool),'const.pair.,_2'(s(X48,X50),s(bool,'F_0')))))) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f785,f40])).
fof(f104,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,X2))))) | s(bool,'F_0') = s(bool,X2)) )),
  inference(superposition,[],[f44,f40])).
fof(f82,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,X1),s(X0,X2))))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f43,f40])).
fof(f1480,plain,(
  ( ! [X17,X18,X16] : (s(bool,X16) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X17),s(bool,X16))))) | s(bool,'F_0') = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f82,f56])).
fof(f1935,plain,(
  ( ! [X17,X18,X16] : (s(bool,X16) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X16),s(bool,X17))))) | s(bool,'F_0') = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f104,f56])).
fof(f761,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(bool,bool),X2),s(bool,'F_0')))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X0,bool),X3),s(fun(bool,bool),X2),s('type.pair.prod'(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,bool),X1)))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X0,bool),X1)))) )),
  inference(superposition,[],[f48,f40])).
fof(f2698,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | ~p(s(bool,happ(s(fun(X2,bool),X3),s(X2,X4)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X2,bool),X3),s(fun(bool,bool),X1),s('type.pair.prod'(X2,bool),'const.pair.,_2'(s(X2,X4),s(bool,X0)))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f785,f40])).
fof(f763,plain,(
  ( ! [X12,X10,X13,X11,X9] : (~p(s(bool,happ(s(fun(bool,bool),X12),s(bool,X11)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X9,bool),X13),s(fun(bool,bool),X12),s('type.pair.prod'(X9,bool),X10)))) | ~p(s(bool,happ(s(fun(X9,bool),X13),s(X9,'const.pair.FST_1'(s('type.pair.prod'(X9,bool),X10)))))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X9,bool),X10)))) )),
  inference(superposition,[],[f48,f56])).
fof(f2700,plain,(
  ( ! [X14,X12,X15,X13,X11,X16] : (~p(s(bool,happ(s(fun(bool,bool),X13),s(bool,X12)))) | ~p(s(bool,happ(s(fun(X14,bool),X15),s(X14,X16)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X14,bool),X15),s(fun(bool,bool),X13),s('type.pair.prod'(X14,bool),'const.pair.,_2'(s(X14,X16),s(bool,X11)))))) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f785,f56])).
fof(f3727,plain,(
  ( ! [X405,X407,X403,X404,X406] : (s(bool,X403) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X404) = s(fun(bool,bool),X405) | s(bool,'T_0') = s(bool,sK6(X404,X405,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X406),'const.pair.,_2'(s(bool,X403),s(X406,X407)))))) )),
  inference(superposition,[],[f365,f132])).
fof(f3828,plain,(
  ( ! [X364,X366,X368,X365,X367] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X364) = s(fun(bool,bool),X365) | s(bool,'T_0') = s(bool,sK6(X364,X365,bool,bool)) | s(bool,X366) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X367),'const.pair.,_2'(s(bool,X366),s(X367,X368)))))) )),
  inference(superposition,[],[f365,f132])).
fof(f3878,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,X2),s(X1,X3))))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f132])).
fof(f3890,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'T_0'),s(X1,X2))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f3872])).
fof(f3872,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'T_0'),s(X1,X2))))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f132])).
fof(f3832,plain,(
  ( ! [X389,X391,X393,X388,X390,X392] : (s(X390,f(s(bool,'T_0'))) != s(X390,g(s(bool,'T_0'))) | s(fun(bool,X390),X388) = s(fun(bool,X390),X389) | s(bool,X391) = s(bool,sK6(X388,X389,X390,bool)) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X392),'const.pair.,_2'(s(bool,X391),s(X392,X393)))))) )),
  inference(superposition,[],[f45,f132])).
fof(f4171,plain,(
  ( ! [X134,X138,X136,X135,X139,X137] : (s(X135,f(s(bool,X134))) != s(X135,g(s(bool,X134))) | s(fun(bool,X135),X136) = s(fun(bool,X135),X137) | s(bool,'T_0') = s(bool,sK6(X136,X137,X135,bool)) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X138),'const.pair.,_2'(s(bool,X134),s(X138,X139)))))) ) | $spl9),
  inference(subsumption_resolution,[],[f3681,f71])).
fof(f3681,plain,(
  ( ! [X134,X138,X136,X135,X139,X137] : (s(X135,f(s(bool,X134))) != s(X135,g(s(bool,X134))) | s(fun(bool,X135),X136) = s(fun(bool,X135),X137) | s(bool,'T_0') = s(bool,sK6(X136,X137,X135,bool)) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X138),'const.pair.,_2'(s(bool,X134),s(X138,X139)))))) )),
  inference(superposition,[],[f365,f132])).
fof(f4170,plain,(
  ( ! [X132,X133,X131] : (s(bool,'T_0') != s(bool,X131) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X132),'const.pair.,_2'(s(bool,X131),s(X132,X133)))))) ) | $spl9),
  inference(subsumption_resolution,[],[f3680,f71])).
fof(f3680,plain,(
  ( ! [X132,X133,X131] : (s(bool,'T_0') != s(bool,X131) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X132),'const.pair.,_2'(s(bool,X131),s(X132,X133)))))) ) | $spl9),
  inference(superposition,[],[f71,f132])).
fof(f3536,plain,(
  ( ! [X47,X48,X46,X49] : (s(bool,X46) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X47),s(bool,X46))))) | s(bool,X47) = s(bool,X48) | s(bool,'T_0') = s(bool,X48) | s(bool,X46) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(superposition,[],[f132,f56])).
fof(f3535,plain,(
  ( ! [X45,X43,X41,X44,X42,X40] : (s(bool,X40) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X41),s(bool,X40))))) | s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,X42) | s(bool,X40) = s(bool,X43) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X44),'const.pair.,_2'(s(bool,X43),s(X44,X45)))))) )),
  inference(superposition,[],[f132,f132])).
fof(f3525,plain,(
  ( ! [X185,X182,X186,X184,X183,X181] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X184),'const.pair.,_2'(s(bool,X183),s(X184,X185))))) | s(bool,X186) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X181,bool),'const.pair.,_2'(s(X181,X182),s(bool,X183))))) | s(bool,'T_0') = s(bool,X186)) )),
  inference(superposition,[],[f132,f44])).
fof(f3490,plain,(
  ( ! [X111,X116,X114,X112,X115,X113] : (s(bool,'T_0') = s(bool,X113) | s(bool,X114) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X111,bool),'const.pair.,_2'(s(X111,X112),s(bool,X113))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X115),'const.pair.,_2'(s(bool,X114),s(X115,X116)))))) )),
  inference(superposition,[],[f132,f44])).
fof(f3465,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (s(bool,X8) = s(bool,X9) | s(bool,X9) = s(bool,X10) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X11),'const.pair.,_2'(s(bool,X10),s(X11,X12))))) | s(bool,X8) = s(bool,X13) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f132,f56])).
fof(f3464,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X3),'const.pair.,_2'(s(bool,X2),s(X3,X4))))) | s(bool,X0) = s(bool,X5) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X6),'const.pair.,_2'(s(bool,X5),s(X6,X7)))))) )),
  inference(superposition,[],[f132,f132])).
fof(f3463,plain,(
  ( ! [X175,X173,X171,X174,X172,X170] : (s(bool,X172) = s(bool,X173) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X170,bool),'const.pair.,_2'(s(X170,X171),s(bool,X172))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X174),'const.pair.,_2'(s(bool,X173),s(X174,X175)))))) )),
  inference(superposition,[],[f132,f44])).
fof(f4122,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,X1),s(X0,X2))))) | s(bool,X1) = s(bool,X3) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X4),'const.pair.,_2'(s(bool,X3),s(X4,X5)))))) )),
  inference(condensation,[],[f4121])).
fof(f4121,plain,(
  ( ! [X39,X43,X41,X38,X44,X42,X40] : (s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X40),'const.pair.,_2'(s(bool,X38),s(X40,X41))))) | s(bool,X38) = s(bool,X42) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X43),'const.pair.,_2'(s(bool,X42),s(X43,X44)))))) )),
  inference(duplicate_literal_removal,[],[f3401])).
fof(f3401,plain,(
  ( ! [X39,X43,X41,X38,X44,X42,X40] : (s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X40),'const.pair.,_2'(s(bool,X38),s(X40,X41))))) | s(bool,X38) = s(bool,X42) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X43),'const.pair.,_2'(s(bool,X42),s(X43,X44)))))) )),
  inference(superposition,[],[f132,f132])).
fof(f363,plain,(
  ( ! [X10,X8,X11,X9] : (s(X10,f(s(bool,X11))) != s(X10,g(s(bool,X11))) | s(fun(bool,X10),X8) = s(fun(bool,X10),X9) | s(bool,'T_0') = s(bool,sK6(X8,X9,X10,bool)) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f45,f56])).
fof(f361,plain,(
  ( ! [X2,X0,X3,X1] : (s(X2,f(s(bool,'T_0'))) != s(X2,g(s(bool,'T_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,X3) = s(bool,sK6(X0,X1,X2,bool)) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f45,f56])).
fof(f157,plain,(
  ( ! [X21,X19,X20,X18] : (s(bool,X21) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X18,bool),'const.pair.,_2'(s(X18,X19),s(bool,X20))))) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f56,f44])).
fof(f156,plain,(
  ( ! [X14,X17,X15,X16] : (s(bool,X17) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X14),'const.pair.,_2'(s(bool,X15),s(X14,X16))))) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f56,f43])).
fof(f133,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X26,bool),'const.pair.,_2'(s(X26,X27),s(bool,X28)))))) )),
  inference(superposition,[],[f56,f44])).
fof(f132,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,X23) = s(bool,X25) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X22),'const.pair.,_2'(s(bool,X23),s(X22,X24)))))) )),
  inference(superposition,[],[f56,f43])).
fof(f45,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK6(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK6(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK6(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK6(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6])],[f26])).
fof(f26,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f22])).
fof(f22,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f21])).
fof(f21,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.pair.PROD_ALL_THM',aHL_EXT)).
fof(f41608,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,X1) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X1)) ) | ($spl4 | $spl9 | $spl36 | $spl40)),
  inference(condensation,[],[f41607])).
fof(f41607,plain,(
  ( ! [X24,X23,X22] : (s(bool,'F_0') = s(bool,X22) | s(bool,X22) = s(bool,X23) | s(bool,X24) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X24)) ) | ($spl4 | $spl9 | $spl36 | $spl40)),
  inference(subsumption_resolution,[],[f41506,f30026])).
fof(f30026,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') != s(bool,X5) | s(bool,X5) = s(bool,X6) | s(bool,'F_0') = s(bool,X6)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f71,f28992])).
fof(f28992,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f25792,f25792])).
fof(f25792,plain,(
  ( ! [X349] : (s(bool,X349) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,'F_0') = s(bool,X349)) ) | ($spl4 | $spl9 | $spl36)),
  inference(global_subsumption,[],[f23712,f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f67,f64,f132,f133,f146,f156,f157,f200,f361,f363,f71,f54,f1085,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f4170,f4171,f3832,f3890,f3878,f2550,f2549,f3652,f3732,f3828,f3727,f2700,f763,f2874,f2698,f761,f3651,f786,f550,f475,f48,f47,f46,f762,f2762,f2699,f2704,f785,f2763,f4154,f3827,f3726,f3646,f368,f366,f3759,f3758,f1935,f1480,f1972,f1473,f39,f82,f104,f365,f370,f2706,f2559,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13826,f13033,f13038,f13852,f13169,f13223,f13224,f13225,f13345,f13350,f13427,f13415,f15352,f15341,f15069,f17276,f17266,f17264,f17226,f16976,f22573,f18154,f18155,f22295,f22293,f22291,f22279,f22277,f22275,f22219,f22217,f18360,f18362,f18363,f18790,f18797,f18801,f18802,f22898,f19153,f19158,f21203,f22901,f19331,f19338,f19400,f19401,f19402,f19403,f19406,f22928,f22929,f22930,f22931,f19558,f19565,f19569,f19570,f22984,f20200,f23229,f23233,f23235,f24931,f24930,f24854,f24822,f24800,f23782,f23787,f25261,f25262,f24214,f25784,f25785,f24220,f24215,f25788,f23788,f23783])).
fof(f23783,plain,(
  ( ! [X635,X636,X637] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X635) = s(fun(bool,bool),X636) | s(bool,'T_0') = s(bool,sK6(X635,X636,bool,bool)) | s(bool,X637) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X637) = s(bool,g(s(bool,'F_0')))) ) | $spl4),
  inference(superposition,[],[f365,f22295])).
fof(f24215,plain,(
  ( ! [X695,X696,X694] : (s(bool,X694) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X695) = s(fun(bool,bool),X696) | s(bool,'T_0') = s(bool,sK6(X695,X696,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,X694) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | $spl4),
  inference(superposition,[],[f365,f22295])).
fof(f25785,plain,(
  ( ! [X695,X696,X694] : (s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,X694) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X695) = s(fun(bool,bool),X696) | s(bool,'T_0') = s(bool,sK6(X695,X696,bool,bool)) | s(bool,'T_0') = s(bool,X694)) ) | $spl4),
  inference(forward_demodulation,[],[f24405,f1085])).
fof(f24405,plain,(
  ( ! [X695,X696,X694] : (s(bool,X694) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X695) = s(fun(bool,bool),X696) | s(bool,'T_0') = s(bool,sK6(X695,X696,bool,bool)) | s(bool,'T_0') = s(bool,X694) | s(bool,g(s(bool,'F_0'))) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | $spl4),
  inference(superposition,[],[f365,f22295])).
fof(f24214,plain,(
  ( ! [X693,X691,X692,X690] : (s(bool,X693) != s(bool,f(s(X690,sK6(X691,X692,bool,X690)))) | s(fun(X690,bool),X691) = s(fun(X690,bool),X692) | s(bool,'T_0') = s(bool,g(s(X690,sK6(X691,X692,bool,X690)))) | s(bool,X693) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | $spl4),
  inference(superposition,[],[f45,f22295])).
fof(f25262,plain,(
  ( ! [X350,X352,X349,X351] : (s(X350,f(s(bool,X349))) != s(X350,g(s(bool,X349))) | s(fun(bool,X350),X351) = s(fun(bool,X350),X352) | s(bool,'T_0') = s(bool,sK6(X351,X352,X350,bool)) | s(bool,X349) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl9)),
  inference(subsumption_resolution,[],[f24117,f71])).
fof(f24117,plain,(
  ( ! [X350,X352,X349,X351] : (s(X350,f(s(bool,X349))) != s(X350,g(s(bool,X349))) | s(fun(bool,X350),X351) = s(fun(bool,X350),X352) | s(bool,'T_0') = s(bool,sK6(X351,X352,X350,bool)) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X349) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | $spl4),
  inference(superposition,[],[f365,f22295])).
fof(f25261,plain,(
  ( ! [X348] : (s(bool,'T_0') != s(bool,X348) | s(bool,X348) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl9)),
  inference(subsumption_resolution,[],[f24116,f71])).
fof(f24116,plain,(
  ( ! [X348] : (s(bool,'T_0') != s(bool,X348) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X348) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl9)),
  inference(superposition,[],[f71,f22295])).
fof(f23787,plain,(
  ( ! [X652,X654,X653,X651] : (s(X653,f(s(bool,'T_0'))) != s(X653,g(s(bool,'T_0'))) | s(fun(bool,X653),X651) = s(fun(bool,X653),X652) | s(bool,X654) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X654) = s(bool,sK6(X651,X652,X653,bool))) ) | $spl4),
  inference(superposition,[],[f45,f22295])).
fof(f23782,plain,(
  ( ! [X633,X631,X632,X634] : (s(bool,'T_0') != s(bool,f(s(X631,sK6(X632,X633,bool,X631)))) | s(fun(X631,bool),X632) = s(fun(X631,bool),X633) | s(bool,X634) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X634) = s(bool,g(s(X631,sK6(X632,X633,bool,X631))))) ) | $spl4),
  inference(superposition,[],[f45,f22295])).
fof(f23235,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,X32) = s(bool,X33) | s(bool,X34) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X33) = s(bool,X34) | s(bool,X35) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X32) = s(bool,X35)) ) | $spl4),
  inference(superposition,[],[f22295,f22295])).
fof(f23233,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (s(bool,X20) = s(bool,X21) | s(bool,X22) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X21) = s(bool,X22) | s(bool,X23) = s(bool,X24) | s(bool,X20) = s(bool,X24) | s(bool,X23) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) ) | $spl4),
  inference(superposition,[],[f22295,f146])).
fof(f23229,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X2) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | $spl4),
  inference(superposition,[],[f22295,f56])).
fof(f22295,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X0) = s(bool,X1)) ) | $spl4),
  inference(condensation,[],[f22294])).
fof(f22294,plain,(
  ( ! [X222,X223,X221] : (s(bool,'T_0') = s(bool,X221) | s(bool,X222) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X222) = s(bool,X223) | s(bool,'T_0') = s(bool,X223)) ) | $spl4),
  inference(duplicate_literal_removal,[],[f18263])).
fof(f18263,plain,(
  ( ! [X222,X223,X221] : (s(bool,'T_0') = s(bool,X221) | s(bool,X222) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,'T_0') = s(bool,X221) | s(bool,X222) = s(bool,X223) | s(bool,'T_0') = s(bool,X223)) ) | $spl4),
  inference(superposition,[],[f146,f1085])).
fof(f67,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | $spl4),
  inference(subsumption_resolution,[],[f62,f39])).
fof(f62,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | $spl4),
  inference(superposition,[],[f54,f40])).
fof(f23712,plain,(
  ( ! [X349] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,X349) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,'F_0') = s(bool,X349)) ) | ($spl4 | $spl36)),
  inference(superposition,[],[f11203,f22295])).
fof(f11203,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl36),
  inference(condensation,[],[f11202])).
fof(f11202,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,f(s(bool,'F_0')))) ) | $spl36),
  inference(subsumption_resolution,[],[f11171,f40])).
fof(f11171,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,f(s(bool,'F_0')))) ) | $spl36),
  inference(superposition,[],[f2455,f40])).
fof(f41506,plain,(
  ( ! [X24,X23,X22] : (s(bool,'F_0') = s(bool,X22) | s(bool,X22) = s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,X24) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X24)) ) | $spl40),
  inference(superposition,[],[f2561,f146])).
fof(f74774,plain,(
  ( ! [X285,X284] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') != s(bool,X284) | s(bool,'T_0') = s(bool,X285) | s(bool,X284) = s(bool,X285)) ) | ($spl2 | $spl4 | $spl9 | $spl36 | $spl250)),
  inference(forward_demodulation,[],[f71517,f1074])).
fof(f71517,plain,(
  ( ! [X285,X284] : (s(bool,'F_0') != s(bool,X284) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3))))) | s(bool,'T_0') = s(bool,X285) | s(bool,X284) = s(bool,X285)) ) | ($spl2 | $spl4 | $spl9 | $spl36 | $spl250)),
  inference(superposition,[],[f49698,f22297])).
fof(f49698,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,'T_0') = s(bool,f(s(bool,X1)))) ) | ($spl4 | $spl9 | $spl36 | $spl250)),
  inference(global_subsumption,[],[f36718,f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f64,f132,f133,f146,f156,f157,f200,f361,f363,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f3832,f3890,f3878,f2550,f2549,f3652,f3732,f3828,f3727,f2700,f763,f2698,f761,f1935,f1480,f82,f104,f2706,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13225,f13350,f13427,f13415,f22573,f18154,f18155,f22291,f22279,f22275,f22219,f22217,f18360,f18362,f18363,f18802,f19158,f21203,f19338,f19402,f19403,f19570,f22984,f20200,f20161,f20156,f19927,f19749,f19571,f19566,f19337,f19159,f18981,f18976,f18803,f18798,f25801,f25802,f13351,f13346,f13175,f13170,f13039,f9980,f9978,f9977,f9974,f6734,f6732,f6731,f6728,f25850,f25854,f30902,f32190,f35168,f35169,f35170,f17226,f15341,f365,f39,f49052,f16976,f17264,f17266,f13038,f3651,f15352,f19925,f19747,f19569,f18801,f762,f22930,f22929,f22928,f786,f550,f475,f48,f47,f46,f2704,f19406,f19914,f19736,f19558,f18790,f2699,f785,f2763,f22931,f37580,f37578,f37577,f37576,f34476,f34474,f34472,f34425,f22898,f19565,f19331,f19153,f18797,f15069,f13826,f13345,f13169,f13033,f4154,f3827,f3726,f3646,f2559,f368,f366,f32792,f32194,f32193,f32192,f31357,f30906,f30905,f30904,f22293,f22277,f19401,f19400,f17276,f13224,f13223,f3759,f3758,f2874,f1972,f1473,f370])).
fof(f36718,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,'T_0') = s(bool,f(s(bool,X1))) | s(bool,X0) = s(bool,X1)) ) | ($spl4 | $spl9 | $spl36)),
  inference(equality_factoring,[],[f30029])).
fof(f30029,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') = s(bool,f(s(bool,X20))) | s(bool,X20) = s(bool,X21) | s(bool,'F_0') = s(bool,X21)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f11203,f28992])).
fof(f31569,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl176),
  inference(superposition,[],[f30750,f56])).
fof(f86457,plain,(
  ( ! [X131] : (s(bool,'F_0') = s(bool,g(s(bool,f(s(bool,X131))))) | s(bool,X131) = s(bool,f(s(bool,'F_0')))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl70)),
  inference(subsumption_resolution,[],[f85858,f37])).
fof(f85858,plain,(
  ( ! [X131] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,g(s(bool,f(s(bool,X131))))) | s(bool,X131) = s(bool,f(s(bool,'F_0')))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl70)),
  inference(superposition,[],[f4928,f49702])).
fof(f49702,plain,(
  ( ! [X123] : (s(bool,'T_0') = s(bool,f(s(bool,X123))) | s(bool,X123) = s(bool,f(s(bool,'F_0')))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(backward_demodulation,[],[f49268,f49269])).
fof(f49269,plain,(
  ( ! [X123] : (s(bool,f(s(bool,'F_0'))) = s(bool,f(s(bool,X123))) | s(bool,X123) = s(bool,f(s(bool,'F_0')))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(subsumption_resolution,[],[f44751,f49260])).
fof(f44751,plain,(
  ( ! [X123] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,f(s(bool,'F_0'))) = s(bool,f(s(bool,X123))) | s(bool,X123) = s(bool,f(s(bool,'F_0')))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f37032,f37032])).
fof(f37032,plain,(
  ( ! [X143] : (s(bool,'T_0') = s(bool,f(s(bool,X143))) | s(bool,X143) = s(bool,f(s(bool,'F_0')))) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f36091,f71])).
fof(f36091,plain,(
  ( ! [X143] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(bool,X143))) | s(bool,X143) = s(bool,f(s(bool,'F_0')))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30029,f11203])).
fof(f49268,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl4 | $spl9 | $spl36)),
  inference(global_subsumption,[],[f44738,f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f64,f132,f133,f146,f156,f157,f200,f361,f363,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f3832,f3890,f3878,f2550,f2549,f3652,f3732,f3828,f3727,f2700,f763,f2698,f761,f2455,f1935,f1480,f82,f104,f2706,f11126,f11125,f11188,f11161,f11162,f11222,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13225,f13350,f13427,f13415,f15429,f17356,f22573,f18154,f18155,f22291,f22279,f22275,f22219,f22217,f18360,f18362,f18363,f22711,f18802,f19158,f21203,f19338,f19402,f19403,f19570,f22984,f20200,f20161,f20156,f25794,f19927,f19749,f19571,f19566,f25798,f19337,f19159,f18981,f18976,f18803,f18798,f25801,f25802,f13351,f13346,f25804,f13292,f13175,f13170,f13039,f11230,f9980,f9978,f9977,f9974,f6734,f6732,f6731,f6728,f25850,f25854,f30902,f32190,f35168,f35169,f35170,f13038,f3651,f17264,f17266,f15352,f22930,f22929,f22928,f786,f550,f475,f48,f47,f46,f19925,f19747,f19569,f18801,f762,f11302,f2704,f19406,f19914,f19736,f19558,f18790,f2699,f785,f2763,f22931,f37580,f37578,f37577,f37576,f34476,f34474,f34472,f34425,f22898,f19565,f19331,f19153,f18797,f16976,f15069,f13826,f13345,f13169,f13033,f4154,f3827,f3726,f3646,f2559,f368,f366,f32792,f32194,f32193,f32192,f31431,f31357,f30906,f30905,f30904,f25810,f25806,f22303,f22293,f22277,f21173,f19401,f19400,f17276,f17226,f15341,f13224,f13223,f11203,f3759,f3758,f2874,f1972,f1473,f370,f365,f39])).
fof(f21173,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) != s(bool,f(s(bool,'F_0'))) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X2)) ) | $spl36),
  inference(condensation,[],[f21172])).
fof(f21172,plain,(
  ( ! [X167,X165,X166,X164] : (s(bool,X165) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X164) | s(bool,X165) = s(bool,X166) | s(bool,X166) = s(bool,X167) | s(bool,'T_0') = s(bool,X167)) ) | $spl36),
  inference(duplicate_literal_removal,[],[f19187])).
fof(f19187,plain,(
  ( ! [X167,X165,X166,X164] : (s(bool,X165) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X164) | s(bool,X165) = s(bool,X166) | s(bool,'T_0') = s(bool,X164) | s(bool,X166) = s(bool,X167) | s(bool,'T_0') = s(bool,X167)) ) | $spl36),
  inference(superposition,[],[f2455,f146])).
fof(f22303,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,f(s(bool,'F_0'))) | s(bool,X0) = s(bool,X1)) ) | $spl36),
  inference(condensation,[],[f22302])).
fof(f22302,plain,(
  ( ! [X202,X200,X201] : (s(bool,'T_0') = s(bool,X200) | s(bool,X201) = s(bool,f(s(bool,'F_0'))) | s(bool,X201) = s(bool,X202) | s(bool,'T_0') = s(bool,X202)) ) | $spl36),
  inference(duplicate_literal_removal,[],[f18257])).
fof(f18257,plain,(
  ( ! [X202,X200,X201] : (s(bool,'T_0') = s(bool,X200) | s(bool,X201) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X200) | s(bool,X201) = s(bool,X202) | s(bool,'T_0') = s(bool,X202)) ) | $spl36),
  inference(superposition,[],[f146,f11203])).
fof(f25806,plain,(
  ( ! [X329,X328,X330] : (s(bool,'F_0') = s(bool,X328) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X329,bool),'const.pair.,_2'(s(X329,X330),s(bool,X328)))))) ) | $spl36),
  inference(global_subsumption,[],[f13292,f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f64,f132,f133,f146,f156,f157,f200,f361,f363,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f3832,f3890,f3878,f2550,f2549,f3652,f3732,f3828,f3727,f2700,f763,f2874,f2698,f761,f3651,f786,f550,f475,f48,f47,f46,f762,f2699,f2704,f785,f2763,f4154,f3827,f3726,f3646,f368,f366,f3759,f3758,f1935,f1480,f1972,f1473,f39,f82,f104,f365,f370,f2706,f2559,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13826,f13033,f13038,f13169,f13223,f13224,f13225,f13345,f13350,f13427,f13415,f15352,f15341,f15069,f17276,f17266,f17264,f17226,f16976,f22573,f18154,f18155,f22293,f22291,f22279,f22277,f22275,f22219,f22217,f18360,f18362,f18363,f18790,f18797,f18801,f18802,f22898,f19153,f19158,f21203,f19331,f19338,f19400,f19401,f19402,f19403,f19406,f22928,f22929,f22930,f22931,f19558,f19565,f19569,f19570,f22984,f20200,f20161,f20156,f19927,f19749,f19571,f19566,f19337,f19159,f18981,f18976,f18803,f18798,f25801,f25802,f13351,f13346])).
fof(f25810,plain,(
  ( ! [X12,X13,X11] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X12),'const.pair.,_2'(s(bool,X11),s(X12,X13))))) | s(bool,'F_0') = s(bool,X11)) ) | $spl36),
  inference(global_subsumption,[],[f11230,f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f64,f132,f133,f146,f156,f157,f200,f361,f363,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f3832,f3890,f3878,f2550,f2549,f3652,f3732,f3828,f3727,f2700,f763,f2874,f2698,f761,f3651,f786,f550,f475,f48,f47,f46,f762,f2699,f2704,f785,f2763,f4154,f3827,f3726,f3646,f368,f366,f3759,f3758,f1935,f1480,f1972,f1473,f39,f82,f104,f365,f370,f2706,f2559,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13826,f13033,f13038,f13169,f13223,f13224,f13225,f13345,f13350,f13427,f13415,f15352,f15341,f15069,f17276,f17266,f17264,f17226,f16976,f22573,f18154,f18155,f22293,f22291,f22279,f22277,f22275,f22219,f22217,f18360,f18362,f18363,f18790,f18797,f18801,f18802,f22898,f19153,f19158,f21203,f19331,f19338,f19400,f19401,f19402,f19403,f19406,f22928,f22929,f22930,f22931,f19558,f19565,f19569,f19570,f22984,f20200,f20161,f20156,f19927,f19749,f19571,f19566,f19337,f19159,f18981,f18976,f18803,f18798,f25801,f25802,f13351,f13346,f13175,f13170,f13039])).
fof(f31431,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,X3))))) | s(bool,X3) = s(bool,X4) | s(bool,X0) = s(bool,X4)) ) | $spl36),
  inference(condensation,[],[f31430])).
fof(f31430,plain,(
  ( ! [X109,X116,X114,X110,X115,X113] : (s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X113,bool),'const.pair.,_2'(s(X113,X114),s(bool,X109))))) | s(bool,X109) = s(bool,X115) | s(bool,'T_0') = s(bool,X110) | s(bool,X115) = s(bool,X116) | s(bool,'T_0') = s(bool,X116)) ) | $spl36),
  inference(subsumption_resolution,[],[f30823,f15429])).
fof(f30823,plain,(
  ( ! [X111,X109,X116,X114,X112,X110,X115,X113] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X111),'const.pair.,_2'(s(bool,X110),s(X111,X112))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X113,bool),'const.pair.,_2'(s(X113,X114),s(bool,X109))))) | s(bool,X109) = s(bool,X115) | s(bool,'T_0') = s(bool,X110) | s(bool,X115) = s(bool,X116) | s(bool,'T_0') = s(bool,X116)) )),
  inference(superposition,[],[f1473,f146])).
fof(f11302,plain,(
  ( ! [X61,X59,X62,X60] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X60,bool),X61),s(fun(bool,bool),X59),s('type.pair.prod'(X60,bool),'const.pair.,_2'(s(X60,X62),s(bool,'T_0')))))) | ~p(s(bool,happ(s(fun(bool,bool),X59),s(bool,'T_0')))) | ~p(s(bool,happ(s(fun(X60,bool),X61),s(X60,X62))))) ) | $spl36),
  inference(forward_demodulation,[],[f11274,f11203])).
fof(f11274,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,happ(s(fun(bool,bool),X59),s(bool,'T_0')))) | ~p(s(bool,happ(s(fun(X60,bool),X61),s(X60,X62)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X60,bool),X61),s(fun(bool,bool),X59),s('type.pair.prod'(X60,bool),'const.pair.,_2'(s(X60,X62),s(bool,f(s(bool,'F_0'))))))))) ) | $spl36),
  inference(superposition,[],[f785,f11203])).
fof(f11230,plain,(
  ( ! [X12,X13,X11] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X11) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X12),'const.pair.,_2'(s(bool,X11),s(X12,X13)))))) ) | $spl36),
  inference(superposition,[],[f11203,f132])).
fof(f13292,plain,(
  ( ! [X329,X328,X330] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X328) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X329,bool),'const.pair.,_2'(s(X329,X330),s(bool,X328)))))) ) | $spl36),
  inference(superposition,[],[f11203,f133])).
fof(f22711,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl36),
  inference(condensation,[],[f22710])).
fof(f22710,plain,(
  ( ! [X516,X514,X513,X515] : (s(bool,X513) != s(bool,X514) | s(bool,'T_0') = s(bool,X514) | s(bool,X513) = s(bool,X515) | s(bool,'T_0') = s(bool,X515) | s(bool,'T_0') = s(bool,X516)) ) | $spl36),
  inference(subsumption_resolution,[],[f18716,f2455])).
fof(f18716,plain,(
  ( ! [X516,X514,X513,X515] : (s(bool,X513) != s(bool,X514) | s(bool,'T_0') = s(bool,X514) | s(bool,X513) = s(bool,X515) | s(bool,'T_0') = s(bool,X515) | s(bool,X516) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X516)) ) | $spl36),
  inference(superposition,[],[f2455,f146])).
fof(f17356,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,X0))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl36),
  inference(forward_demodulation,[],[f17161,f11203])).
fof(f17161,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,f(s(bool,'F_0'))) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,X0)))))) ) | $spl36),
  inference(condensation,[],[f17160])).
fof(f17160,plain,(
  ( ! [X152,X151,X153,X150] : (s(bool,f(s(bool,'F_0'))) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X151,bool),'const.pair.,_2'(s(X151,X152),s(bool,X153))))) | s(bool,'T_0') = s(bool,X150) | s(bool,'T_0') = s(bool,X153)) ) | $spl36),
  inference(duplicate_literal_removal,[],[f16872])).
fof(f16872,plain,(
  ( ! [X152,X151,X153,X150] : (s(bool,f(s(bool,'F_0'))) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X151,bool),'const.pair.,_2'(s(X151,X152),s(bool,X153))))) | s(bool,'T_0') = s(bool,X150) | s(bool,'T_0') = s(bool,X153) | s(bool,'T_0') = s(bool,X150)) ) | $spl36),
  inference(superposition,[],[f2455,f157])).
fof(f15429,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,X0),s(X1,X2))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl36),
  inference(forward_demodulation,[],[f15254,f11203])).
fof(f15254,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,f(s(bool,'F_0'))) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,X0),s(X1,X2)))))) ) | $spl36),
  inference(condensation,[],[f15253])).
fof(f15253,plain,(
  ( ! [X138,X136,X139,X137] : (s(bool,f(s(bool,'F_0'))) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X137),'const.pair.,_2'(s(bool,X138),s(X137,X139))))) | s(bool,'T_0') = s(bool,X136) | s(bool,'T_0') = s(bool,X138)) ) | $spl36),
  inference(duplicate_literal_removal,[],[f14975])).
fof(f14975,plain,(
  ( ! [X138,X136,X139,X137] : (s(bool,f(s(bool,'F_0'))) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X137),'const.pair.,_2'(s(bool,X138),s(X137,X139))))) | s(bool,'T_0') = s(bool,X136) | s(bool,'T_0') = s(bool,X138) | s(bool,'T_0') = s(bool,X136)) ) | $spl36),
  inference(superposition,[],[f2455,f156])).
fof(f11222,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl36),
  inference(superposition,[],[f11203,f40])).
fof(f11162,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl36),
  inference(superposition,[],[f2455,f40])).
fof(f11161,plain,(
  ( ! [X76,X77,X75] : (s(bool,X77) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X75,bool),'const.pair.,_2'(s(X75,X76),s(bool,X77)))))) ) | $spl36),
  inference(superposition,[],[f2455,f44])).
fof(f11188,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,X1),s(X0,X2))))) | s(bool,X1) != s(bool,f(s(bool,'F_0')))) ) | $spl36),
  inference(condensation,[],[f11187])).
fof(f11187,plain,(
  ( ! [X14,X15,X13,X16] : (s(bool,X14) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X15),'const.pair.,_2'(s(bool,X14),s(X15,X16)))))) ) | $spl36),
  inference(duplicate_literal_removal,[],[f11133])).
fof(f11133,plain,(
  ( ! [X14,X15,X13,X16] : (s(bool,X14) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X15),'const.pair.,_2'(s(bool,X14),s(X15,X16)))))) ) | $spl36),
  inference(superposition,[],[f2455,f132])).
fof(f11125,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,f(s(bool,'F_0'))),s(X0,X1)))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f43,f2455])).
fof(f11126,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,f(s(bool,'F_0')))))))) ) | $spl36),
  inference(unit_resulting_resolution,[],[f44,f2455])).
fof(f44738,plain,(
  ( ! [X94,X95] : (s(bool,X94) = s(bool,f(s(bool,X94))) | s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | s(bool,X94) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f37032,f56])).
fof(f87222,plain,(
  ( ! [X109,X105,X108] : (s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X108,bool),'const.pair.,_2'(s(X108,X109),s(bool,f(s(bool,f(s(bool,X105))))))))) | s(bool,'T_0') = s(bool,X105)) ) | $spl366),
  inference(cnf_transformation,[],[f87222_D])).
fof(f87222_D,plain,(
  ( ! [X109,X105,X108] : (s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X108,bool),'const.pair.,_2'(s(X108,X109),s(bool,f(s(bool,f(s(bool,X105))))))))) | s(bool,'T_0') = s(bool,X105)) ) <=> ~$spl366),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl366])])).
fof(f4908,plain,(
  ( ! [X74,X73] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X73,bool),'const.pair.,_2'(s(X73,X74),s(bool,'F_0')))))) ) | $spl64),
  inference(cnf_transformation,[],[f4908_D])).
fof(f4908_D,plain,(
  ( ! [X74,X73] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X73,bool),'const.pair.,_2'(s(X73,X74),s(bool,'F_0')))))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f87030,plain,(
  ( ! [X127,X130,X128,X131,X129] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X128,bool),'const.pair.,_2'(s(X128,X129),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X130,bool),'const.pair.,_2'(s(X130,X131),s(bool,f(s(bool,f(s(bool,X127))))))))) | s(bool,'T_0') = s(bool,X127)) ) | ($spl176 | $spl204)),
  inference(superposition,[],[f1972,f74384])).
fof(f74384,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X0))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl176 | $spl204)),
  inference(condensation,[],[f74383])).
fof(f74383,plain,(
  ( ! [X273,X272] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X273))))) | s(bool,'F_0') = s(bool,X272) | s(bool,'T_0') = s(bool,X273)) ) | ($spl176 | $spl204)),
  inference(duplicate_literal_removal,[],[f74083])).
fof(f74083,plain,(
  ( ! [X273,X272] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X273))))) | s(bool,'F_0') = s(bool,X272) | s(bool,'F_0') = s(bool,X272) | s(bool,'T_0') = s(bool,X273)) ) | ($spl176 | $spl204)),
  inference(superposition,[],[f37102,f31623])).
fof(f31623,plain,(
  ( ! [X6,X7] : (s(bool,X7) = s(bool,f(s(bool,X6))) | s(bool,'F_0') = s(bool,X7) | s(bool,'T_0') = s(bool,X6)) ) | $spl176),
  inference(superposition,[],[f30750,f40])).
fof(f87223,plain,(
  ( ! [X123,X126] : (s(bool,X126) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,f(s(bool,f(s(bool,X123))))),s(bool,X126))))) | s(bool,'T_0') = s(bool,X123) | s(bool,'T_0') = s(bool,X126)) ) | $spl368),
  inference(cnf_transformation,[],[f87223_D])).
fof(f87223_D,plain,(
  ( ! [X123,X126] : (s(bool,X126) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,f(s(bool,f(s(bool,X123))))),s(bool,X126))))) | s(bool,'T_0') = s(bool,X123) | s(bool,'T_0') = s(bool,X126)) ) <=> ~$spl368),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl368])])).
fof(f967,plain,(
  ( ! [X72,X73] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X72),'const.pair.,_2'(s(bool,'F_0'),s(X72,X73)))))) ) | $spl18),
  inference(cnf_transformation,[],[f967_D])).
fof(f967_D,plain,(
  ( ! [X72,X73] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X72),'const.pair.,_2'(s(bool,'F_0'),s(X72,X73)))))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f87029,plain,(
  ( ! [X125,X123,X126,X124] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X124),'const.pair.,_2'(s(bool,'F_0'),s(X124,X125))))) | s(bool,X126) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,f(s(bool,f(s(bool,X123))))),s(bool,X126))))) | s(bool,'T_0') = s(bool,X126) | s(bool,'T_0') = s(bool,X123)) ) | ($spl16 | $spl176 | $spl204)),
  inference(superposition,[],[f1557,f74384])).
fof(f1557,plain,(
  ( ! [X39,X41,X38,X40] : (s(bool,X38) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X39),s(bool,X38))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X40),'const.pair.,_2'(s(bool,X39),s(X40,X41))))) | s(bool,'T_0') = s(bool,X38)) ) | $spl16),
  inference(superposition,[],[f1257,f40])).
fof(f1257,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X40),'const.pair.,_2'(s(bool,X38),s(X40,X41))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X37),'const.pair.,_2'(s(bool,X38),s(X37,X39)))))) ) | $spl16),
  inference(superposition,[],[f961,f43])).
fof(f87025,plain,(
  ( ! [X109,X107,X105,X108,X106] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X106),'const.pair.,_2'(s(bool,'F_0'),s(X106,X107))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X108,bool),'const.pair.,_2'(s(X108,X109),s(bool,f(s(bool,f(s(bool,X105))))))))) | s(bool,'T_0') = s(bool,X105)) ) | ($spl176 | $spl204)),
  inference(superposition,[],[f1473,f74384])).
fof(f87160,plain,(
  ( ! [X87] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X87))))) | s(bool,X87) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) ) | $spl364),
  inference(cnf_transformation,[],[f87160_D])).
fof(f87160_D,plain,(
  ( ! [X87] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X87))))) | s(bool,X87) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) ) <=> ~$spl364),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl364])])).
fof(f87115,plain,(
  ( ! [X87,X86] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X87))))) | s(bool,'T_0') = s(bool,X86) | s(bool,X87) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) ) | ($spl2 | $spl176 | $spl204)),
  inference(duplicate_literal_removal,[],[f86857])).
fof(f86857,plain,(
  ( ! [X87,X86] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X87))))) | s(bool,'T_0') = s(bool,X86) | s(bool,'T_0') = s(bool,X86) | s(bool,X87) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) ) | ($spl2 | $spl176 | $spl204)),
  inference(superposition,[],[f74384,f22297])).
fof(f87159,plain,(
  ( ! [X79] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X79))))) | s(bool,X79) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | $spl362),
  inference(cnf_transformation,[],[f87159_D])).
fof(f87159_D,plain,(
  ( ! [X79] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X79))))) | s(bool,X79) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) <=> ~$spl362),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl362])])).
fof(f87116,plain,(
  ( ! [X78,X79] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X79))))) | s(bool,'T_0') = s(bool,X78) | s(bool,X79) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl176 | $spl204)),
  inference(duplicate_literal_removal,[],[f86853])).
fof(f86853,plain,(
  ( ! [X78,X79] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X79))))) | s(bool,'T_0') = s(bool,X78) | s(bool,'T_0') = s(bool,X78) | s(bool,X79) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl176 | $spl204)),
  inference(superposition,[],[f74384,f22295])).
fof(f81791,plain,(
  ( ! [X87] : (s(bool,'F_0') != s(bool,f(s(bool,X87))) | s(bool,X87) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) ) | $spl360),
  inference(cnf_transformation,[],[f81791_D])).
fof(f81791_D,plain,(
  ( ! [X87] : (s(bool,'F_0') != s(bool,f(s(bool,X87))) | s(bool,X87) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) ) <=> ~$spl360),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl360])])).
fof(f81751,plain,(
  ( ! [X87,X86] : (s(bool,'F_0') != s(bool,f(s(bool,X87))) | s(bool,'T_0') = s(bool,X86) | s(bool,X87) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) ) | ($spl2 | $spl4 | $spl40 | $spl130 | $spl176 | $spl234)),
  inference(duplicate_literal_removal,[],[f81658])).
fof(f81658,plain,(
  ( ! [X87,X86] : (s(bool,'F_0') != s(bool,f(s(bool,X87))) | s(bool,'T_0') = s(bool,X86) | s(bool,'T_0') = s(bool,X86) | s(bool,X87) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) ) | ($spl2 | $spl4 | $spl40 | $spl130 | $spl176 | $spl234)),
  inference(superposition,[],[f74629,f22297])).
fof(f74629,plain,(
  ( ! [X281] : (s(bool,'F_0') != s(bool,f(s(bool,X281))) | s(bool,'T_0') = s(bool,X281)) ) | ($spl4 | $spl40 | $spl130 | $spl176 | $spl234)),
  inference(global_subsumption,[],[f74628,f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f67,f64,f132,f133,f146,f156,f157,f200,f361,f363,f54,f1085,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f3832,f3890,f3878,f2550,f2549,f3652,f3732,f3828,f3727,f2700,f763,f2698,f761,f1935,f1480,f82,f104,f2706,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13225,f13350,f13427,f13415,f22573,f18154,f18155,f22295,f22291,f22279,f22275,f22219,f22217,f18360,f18362,f18363,f18802,f19158,f21203,f19338,f19402,f19403,f19570,f22984,f20200,f23229,f23233,f23235,f24931,f24930,f24854,f24822,f24800,f23787,f25784,f25785,f24220,f24215,f25788,f23788,f23783,f20161,f20156,f19927,f19749,f19571,f19566,f19337,f19159,f18981,f18976,f18803,f18798,f25801,f25802,f13351,f13346,f13175,f13170,f13039,f9980,f9978,f9977,f9974,f6734,f6732,f6731,f6728,f24431,f25850,f25854,f28579,f30902,f32190,f35168,f35169,f35170,f35172,f13038,f3651,f15352,f22930,f22929,f22928,f786,f550,f475,f48,f47,f46,f25855,f19925,f19747,f19569,f18801,f762,f2762,f2704,f19406,f25856,f19914,f19736,f19558,f18790,f2699,f785,f2763,f22931,f37580,f37578,f37577,f37576,f37574,f34476,f34474,f34472,f34425,f24214,f23782,f22898,f19565,f19331,f19153,f18797,f15069,f13826,f13345,f13169,f13033,f4154,f3827,f3726,f3646,f2559,f368,f366,f32792,f32194,f32193,f32192,f31357,f30906,f30905,f30904,f22293,f22277,f19401,f19400,f17276,f17226,f15341,f13224,f13223,f3759,f3758,f2874,f1972,f1473,f370,f365,f39,f25219,f17266,f17264,f16976,f63512,f63513,f63598,f63517,f63518,f63519,f63528,f63530,f63532,f63533,f63554,f63621,f63623,f63625,f63627,f64145,f64146,f64062,f64063,f64064,f64065,f64066,f64075,f64077,f64079,f64080,f64101,f64168,f64170,f64172,f64174,f64629,f64630,f64631,f64640,f64642,f64644,f64645,f64666,f64670,f64675,f64689,f64820,f64821,f64828,f66956,f66957,f66958,f66967,f66969,f66971,f66972,f67074,f67081,f67781,f67874,f67881,f68646,f68653,f69246,f69154,f69156,f69169,f69170,f69280,f69281,f69288])).
fof(f69288,plain,(
  ( ! [X94,X99,X97,X95,X93,X100,X98,X96] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X93,bool),X94),s(fun(bool,bool),X96),s('type.pair.prod'(X93,bool),X95)))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X93,bool),X95))) | s(bool,X97) = s(bool,happ(s(fun(bool,bool),X96),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X97) | s(bool,X98) = s(bool,X99) | s(bool,X99) = s(bool,happ(s(fun(X93,bool),X94),s(X93,'const.pair.FST_1'(s('type.pair.prod'(X93,bool),X95))))) | s(bool,X98) = s(bool,X100) | s(bool,'T_0') = s(bool,X100)) )),
  inference(subsumption_resolution,[],[f69210,f37])).
fof(f69210,plain,(
  ( ! [X94,X99,X97,X95,X93,X100,X98,X96] : (~p(s(bool,'T_0')) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X93,bool),X94),s(fun(bool,bool),X96),s('type.pair.prod'(X93,bool),X95)))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X93,bool),X95))) | s(bool,X97) = s(bool,happ(s(fun(bool,bool),X96),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X97) | s(bool,X98) = s(bool,X99) | s(bool,X99) = s(bool,happ(s(fun(X93,bool),X94),s(X93,'const.pair.FST_1'(s('type.pair.prod'(X93,bool),X95))))) | s(bool,X98) = s(bool,X100) | s(bool,'T_0') = s(bool,X100)) )),
  inference(superposition,[],[f2874,f146])).
fof(f69281,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X5,bool),X6),s(fun(bool,bool),X8),s('type.pair.prod'(X5,bool),X7)))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X5,bool),X7))) | s(bool,X9) = s(bool,happ(s(fun(bool,bool),X8),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X9) | s(bool,X10) = s(bool,happ(s(fun(X5,bool),X6),s(X5,'const.pair.FST_1'(s('type.pair.prod'(X5,bool),X7))))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f69198,f37])).
fof(f69198,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (~p(s(bool,'T_0')) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X5,bool),X6),s(fun(bool,bool),X8),s('type.pair.prod'(X5,bool),X7)))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X5,bool),X7))) | s(bool,X9) = s(bool,happ(s(fun(bool,bool),X8),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X9) | s(bool,X10) = s(bool,happ(s(fun(X5,bool),X6),s(X5,'const.pair.FST_1'(s('type.pair.prod'(X5,bool),X7))))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f2874,f56])).
fof(f69280,plain,(
  ( ! [X255,X257,X259,X260,X256,X258] : (s(bool,'T_0') = s(bool,X257) | ~p(s(bool,happ(s(fun(X255,bool),X258),s(X255,X256)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X255,bool),X258),s(fun(bool,bool),X259),s('type.pair.prod'(X255,bool),'const.pair.,_2'(s(X255,X256),s(bool,X257)))))) | s(bool,X260) = s(bool,happ(s(fun(bool,bool),X259),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X260)) )),
  inference(forward_demodulation,[],[f69194,f44])).
fof(f69194,plain,(
  ( ! [X255,X257,X259,X260,X256,X258] : (~p(s(bool,happ(s(fun(X255,bool),X258),s(X255,X256)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X255,bool),X258),s(fun(bool,bool),X259),s('type.pair.prod'(X255,bool),'const.pair.,_2'(s(X255,X256),s(bool,X257)))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X255,bool),'const.pair.,_2'(s(X255,X256),s(bool,X257))))) | s(bool,X260) = s(bool,happ(s(fun(bool,bool),X259),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X260)) )),
  inference(superposition,[],[f2874,f43])).
fof(f69170,plain,(
  ( ! [X134,X138,X136,X135,X133,X139,X137] : (~p(s(bool,happ(s(fun(bool,bool),X135),s(bool,X134)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X135),s(fun(bool,bool),X136),s('type.pair.prod'(bool,bool),X133)))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),X133))) | s(bool,X137) = s(bool,happ(s(fun(bool,bool),X136),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X137) | s(bool,X134) = s(bool,X138) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),X133))) | s(bool,X138) = s(bool,X139) | s(bool,'T_0') = s(bool,X139)) )),
  inference(superposition,[],[f2874,f146])).
fof(f69169,plain,(
  ( ! [X127,X132,X130,X128,X126,X131,X129] : (~p(s(bool,happ(s(fun(bool,bool),X128),s(bool,X127)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X128),s(fun(bool,bool),X129),s('type.pair.prod'(bool,bool),X126)))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),X126))) | s(bool,X130) = s(bool,happ(s(fun(bool,bool),X129),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X130) | s(bool,X131) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),X126))) | s(bool,'T_0') = s(bool,X127) | s(bool,X131) = s(bool,X132) | s(bool,'T_0') = s(bool,X132)) )),
  inference(superposition,[],[f2874,f146])).
fof(f69156,plain,(
  ( ! [X45,X43,X41,X44,X42] : (~p(s(bool,happ(s(fun(bool,bool),X43),s(bool,X42)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X43),s(fun(bool,bool),X44),s('type.pair.prod'(bool,bool),X41)))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),X41))) | s(bool,X45) = s(bool,happ(s(fun(bool,bool),X44),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X45) | s(bool,'T_0') = s(bool,X42) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),X41)))) )),
  inference(superposition,[],[f2874,f56])).
fof(f69154,plain,(
  ( ! [X35,X33,X34,X32] : (~p(s(bool,happ(s(fun(bool,bool),X33),s(bool,'F_0')))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X33),s(fun(bool,bool),X34),s('type.pair.prod'(bool,bool),X32)))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),X32))) | s(bool,X35) = s(bool,happ(s(fun(bool,bool),X34),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),X32)))) )),
  inference(superposition,[],[f2874,f40])).
fof(f69246,plain,(
  ( ! [X14,X17,X15,X13,X18,X16] : (s(bool,'T_0') = s(bool,X14) | ~p(s(bool,happ(s(fun(bool,bool),X16),s(bool,X15)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X16),s(fun(bool,bool),X17),s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X13),s(bool,X14)))))) | s(bool,X18) = s(bool,happ(s(fun(bool,bool),X17),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X15)) )),
  inference(forward_demodulation,[],[f69151,f44])).
fof(f69151,plain,(
  ( ! [X14,X17,X15,X13,X18,X16] : (~p(s(bool,happ(s(fun(bool,bool),X16),s(bool,X15)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X16),s(fun(bool,bool),X17),s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X13),s(bool,X14)))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X13),s(bool,X14))))) | s(bool,X18) = s(bool,happ(s(fun(bool,bool),X17),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f2874,f156])).
fof(f68653,plain,(
  ( ! [X111,X109,X107,X114,X112,X110,X108,X106,X113] : (~p(s(bool,happ(s(fun(X108,bool),X109),s(X108,X110)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X108,bool),X109),s(fun(bool,bool),X106),s('type.pair.prod'(X108,bool),'const.pair.,_2'(s(X108,X110),s(bool,'T_0')))))) | s(bool,X107) = s(bool,X111) | s(bool,'T_0') = s(bool,X111) | s(bool,X112) = s(bool,X113) | s(bool,X113) = s(bool,happ(s(fun(bool,bool),X106),s(bool,X107))) | s(bool,X112) = s(bool,X114) | s(bool,'T_0') = s(bool,X114)) )),
  inference(subsumption_resolution,[],[f68556,f37])).
fof(f68556,plain,(
  ( ! [X111,X109,X107,X114,X112,X110,X108,X106,X113] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X108,bool),X109),s(X108,X110)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X108,bool),X109),s(fun(bool,bool),X106),s('type.pair.prod'(X108,bool),'const.pair.,_2'(s(X108,X110),s(bool,'T_0')))))) | s(bool,X107) = s(bool,X111) | s(bool,'T_0') = s(bool,X111) | s(bool,X112) = s(bool,X113) | s(bool,X113) = s(bool,happ(s(fun(bool,bool),X106),s(bool,X107))) | s(bool,X112) = s(bool,X114) | s(bool,'T_0') = s(bool,X114)) )),
  inference(superposition,[],[f2704,f146])).
fof(f68646,plain,(
  ( ! [X6,X12,X10,X8,X7,X11,X9] : (~p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X8,bool),X9),s(fun(bool,bool),X6),s('type.pair.prod'(X8,bool),'const.pair.,_2'(s(X8,X10),s(bool,'T_0')))))) | s(bool,X7) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,X12) = s(bool,happ(s(fun(bool,bool),X6),s(bool,X7))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f68544,f37])).
fof(f68544,plain,(
  ( ! [X6,X12,X10,X8,X7,X11,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X8,bool),X9),s(fun(bool,bool),X6),s('type.pair.prod'(X8,bool),'const.pair.,_2'(s(X8,X10),s(bool,'T_0')))))) | s(bool,X7) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,X12) = s(bool,happ(s(fun(bool,bool),X6),s(bool,X7))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f2704,f56])).
fof(f67881,plain,(
  ( ! [X111,X109,X107,X114,X112,X110,X108,X106,X113] : (~p(s(bool,happ(s(fun(X107,bool),X108),s(X107,X109)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X107,bool),X108),s(fun(bool,bool),X106),s('type.pair.prod'(X107,bool),'const.pair.,_2'(s(X107,X109),s(bool,X110)))))) | s(bool,X110) = s(bool,X111) | s(bool,'T_0') = s(bool,X111) | s(bool,X112) = s(bool,X113) | s(bool,X113) = s(bool,happ(s(fun(bool,bool),X106),s(bool,'T_0'))) | s(bool,X112) = s(bool,X114) | s(bool,'T_0') = s(bool,X114)) )),
  inference(subsumption_resolution,[],[f67822,f37])).
fof(f67822,plain,(
  ( ! [X111,X109,X107,X114,X112,X110,X108,X106,X113] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X107,bool),X108),s(X107,X109)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X107,bool),X108),s(fun(bool,bool),X106),s('type.pair.prod'(X107,bool),'const.pair.,_2'(s(X107,X109),s(bool,X110)))))) | s(bool,X110) = s(bool,X111) | s(bool,'T_0') = s(bool,X111) | s(bool,X112) = s(bool,X113) | s(bool,X113) = s(bool,happ(s(fun(bool,bool),X106),s(bool,'T_0'))) | s(bool,X112) = s(bool,X114) | s(bool,'T_0') = s(bool,X114)) )),
  inference(superposition,[],[f2699,f146])).
fof(f67874,plain,(
  ( ! [X6,X12,X10,X8,X7,X11,X9] : (~p(s(bool,happ(s(fun(X7,bool),X8),s(X7,X9)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X7,bool),X8),s(fun(bool,bool),X6),s('type.pair.prod'(X7,bool),'const.pair.,_2'(s(X7,X9),s(bool,X10)))))) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,X12) = s(bool,happ(s(fun(bool,bool),X6),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f67810,f37])).
fof(f67810,plain,(
  ( ! [X6,X12,X10,X8,X7,X11,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X7,bool),X8),s(X7,X9)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X7,bool),X8),s(fun(bool,bool),X6),s('type.pair.prod'(X7,bool),'const.pair.,_2'(s(X7,X9),s(bool,X10)))))) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,X12) = s(bool,happ(s(fun(bool,bool),X6),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f2699,f56])).
fof(f67781,plain,(
  ( ! [X39,X37,X45,X43,X41,X38,X36,X44,X42,X40] : (~p(s(bool,happ(s(fun(bool,bool),X37),s(bool,X36)))) | ~p(s(bool,happ(s(fun(X38,bool),X39),s(X38,X40)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X38,bool),X39),s(fun(bool,bool),X37),s('type.pair.prod'(X38,bool),'const.pair.,_2'(s(X38,X40),s(bool,X41)))))) | s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,X42) | s(bool,X43) = s(bool,X44) | s(bool,X36) = s(bool,X44) | s(bool,X43) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f2699,f146])).
fof(f67081,plain,(
  ( ! [X111,X109,X107,X114,X112,X110,X108,X106,X113] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X106,bool),X107),s(fun(X108,bool),X110),s('type.pair.prod'(X106,X108),X109)))) | s(bool,X111) = s(bool,happ(s(fun(X108,bool),X110),s(X108,'const.pair.SND_1'(s('type.pair.prod'(X106,X108),X109))))) | s(bool,'T_0') = s(bool,X111) | s(bool,X112) = s(bool,X113) | s(bool,X113) = s(bool,happ(s(fun(X106,bool),X107),s(X106,'const.pair.FST_1'(s('type.pair.prod'(X106,X108),X109))))) | s(bool,X112) = s(bool,X114) | s(bool,'T_0') = s(bool,X114)) )),
  inference(subsumption_resolution,[],[f67012,f37])).
fof(f67012,plain,(
  ( ! [X111,X109,X107,X114,X112,X110,X108,X106,X113] : (~p(s(bool,'T_0')) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X106,bool),X107),s(fun(X108,bool),X110),s('type.pair.prod'(X106,X108),X109)))) | s(bool,X111) = s(bool,happ(s(fun(X108,bool),X110),s(X108,'const.pair.SND_1'(s('type.pair.prod'(X106,X108),X109))))) | s(bool,'T_0') = s(bool,X111) | s(bool,X112) = s(bool,X113) | s(bool,X113) = s(bool,happ(s(fun(X106,bool),X107),s(X106,'const.pair.FST_1'(s('type.pair.prod'(X106,X108),X109))))) | s(bool,X112) = s(bool,X114) | s(bool,'T_0') = s(bool,X114)) )),
  inference(superposition,[],[f786,f146])).
fof(f67074,plain,(
  ( ! [X6,X12,X10,X8,X7,X11,X9] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X6,bool),X7),s(fun(X8,bool),X10),s('type.pair.prod'(X6,X8),X9)))) | s(bool,X11) = s(bool,happ(s(fun(X8,bool),X10),s(X8,'const.pair.SND_1'(s('type.pair.prod'(X6,X8),X9))))) | s(bool,'T_0') = s(bool,X11) | s(bool,X12) = s(bool,happ(s(fun(X6,bool),X7),s(X6,'const.pair.FST_1'(s('type.pair.prod'(X6,X8),X9))))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f67000,f37])).
fof(f67000,plain,(
  ( ! [X6,X12,X10,X8,X7,X11,X9] : (~p(s(bool,'T_0')) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X6,bool),X7),s(fun(X8,bool),X10),s('type.pair.prod'(X6,X8),X9)))) | s(bool,X11) = s(bool,happ(s(fun(X8,bool),X10),s(X8,'const.pair.SND_1'(s('type.pair.prod'(X6,X8),X9))))) | s(bool,'T_0') = s(bool,X11) | s(bool,X12) = s(bool,happ(s(fun(X6,bool),X7),s(X6,'const.pair.FST_1'(s('type.pair.prod'(X6,X8),X9))))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f786,f56])).
fof(f66972,plain,(
  ( ! [X158,X156,X154,X159,X157,X155,X153,X160] : (~p(s(bool,happ(s(fun(bool,bool),X156),s(bool,X155)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X156),s(fun(X153,bool),X157),s('type.pair.prod'(bool,X153),X154)))) | s(bool,X158) = s(bool,happ(s(fun(X153,bool),X157),s(X153,'const.pair.SND_1'(s('type.pair.prod'(bool,X153),X154))))) | s(bool,'T_0') = s(bool,X158) | s(bool,X155) = s(bool,X159) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X153),X154))) | s(bool,X159) = s(bool,X160) | s(bool,'T_0') = s(bool,X160)) )),
  inference(superposition,[],[f786,f146])).
fof(f66971,plain,(
  ( ! [X146,X152,X151,X149,X147,X145,X150,X148] : (~p(s(bool,happ(s(fun(bool,bool),X148),s(bool,X147)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X148),s(fun(X145,bool),X149),s('type.pair.prod'(bool,X145),X146)))) | s(bool,X150) = s(bool,happ(s(fun(X145,bool),X149),s(X145,'const.pair.SND_1'(s('type.pair.prod'(bool,X145),X146))))) | s(bool,'T_0') = s(bool,X150) | s(bool,X151) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X145),X146))) | s(bool,'T_0') = s(bool,X147) | s(bool,X151) = s(bool,X152) | s(bool,'T_0') = s(bool,X152)) )),
  inference(superposition,[],[f786,f146])).
fof(f66969,plain,(
  ( ! [X134,X132,X130,X136,X135,X133,X131,X129] : (~p(s(bool,happ(s(fun(bool,bool),X131),s(bool,'T_0')))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X131),s(fun(X129,bool),X132),s('type.pair.prod'(bool,X129),X130)))) | s(bool,X133) = s(bool,happ(s(fun(X129,bool),X132),s(X129,'const.pair.SND_1'(s('type.pair.prod'(bool,X129),X130))))) | s(bool,'T_0') = s(bool,X133) | s(bool,X134) = s(bool,X135) | s(bool,X135) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X129),X130))) | s(bool,X134) = s(bool,X136) | s(bool,'T_0') = s(bool,X136)) )),
  inference(superposition,[],[f786,f146])).
fof(f66967,plain,(
  ( ! [X118,X116,X114,X120,X119,X117,X115,X113] : (~p(s(bool,happ(s(fun(bool,bool),X116),s(bool,X115)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X116),s(fun(X113,bool),X117),s('type.pair.prod'(bool,X113),X114)))) | s(bool,X118) = s(bool,happ(s(fun(X113,bool),X117),s(X113,'const.pair.SND_1'(s('type.pair.prod'(bool,X113),X114))))) | s(bool,'T_0') = s(bool,X118) | s(bool,X115) = s(bool,X119) | s(bool,'T_0') = s(bool,X119) | s(bool,X120) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X113),X114))) | s(bool,'T_0') = s(bool,X120)) )),
  inference(superposition,[],[f786,f146])).
fof(f66958,plain,(
  ( ! [X47,X52,X50,X48,X51,X49] : (~p(s(bool,happ(s(fun(bool,bool),X50),s(bool,X49)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X50),s(fun(X47,bool),X51),s('type.pair.prod'(bool,X47),X48)))) | s(bool,X52) = s(bool,happ(s(fun(X47,bool),X51),s(X47,'const.pair.SND_1'(s('type.pair.prod'(bool,X47),X48))))) | s(bool,'T_0') = s(bool,X52) | s(bool,'T_0') = s(bool,X49) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X47),X48)))) )),
  inference(superposition,[],[f786,f56])).
fof(f66957,plain,(
  ( ! [X45,X43,X41,X46,X44,X42] : (~p(s(bool,happ(s(fun(bool,bool),X43),s(bool,'T_0')))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X43),s(fun(X41,bool),X44),s('type.pair.prod'(bool,X41),X42)))) | s(bool,X45) = s(bool,happ(s(fun(X41,bool),X44),s(X41,'const.pair.SND_1'(s('type.pair.prod'(bool,X41),X42))))) | s(bool,'T_0') = s(bool,X45) | s(bool,X46) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X41),X42))) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f786,f56])).
fof(f66956,plain,(
  ( ! [X39,X37,X38,X36,X40] : (~p(s(bool,happ(s(fun(bool,bool),X38),s(bool,'F_0')))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X38),s(fun(X36,bool),X39),s('type.pair.prod'(bool,X36),X37)))) | s(bool,X40) = s(bool,happ(s(fun(X36,bool),X39),s(X36,'const.pair.SND_1'(s('type.pair.prod'(bool,X36),X37))))) | s(bool,'T_0') = s(bool,X40) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X36),X37)))) )),
  inference(superposition,[],[f786,f40])).
fof(f64828,plain,(
  ( ! [X146,X144,X152,X151,X149,X147,X145,X143,X150,X148] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X143,bool),X144),s(fun(X146,bool),X147),s('type.pair.prod'(X143,X146),'const.pair.,_2'(s(X143,X145),s(X146,X148)))))) | s(bool,X149) = s(bool,happ(s(fun(X146,bool),X147),s(X146,X148))) | s(bool,'T_0') = s(bool,X149) | s(bool,X150) = s(bool,X151) | s(bool,X151) = s(bool,happ(s(fun(X143,bool),X144),s(X143,X145))) | s(bool,X150) = s(bool,X152) | s(bool,'T_0') = s(bool,X152)) )),
  inference(subsumption_resolution,[],[f64729,f37])).
fof(f64729,plain,(
  ( ! [X146,X144,X152,X151,X149,X147,X145,X143,X150,X148] : (~p(s(bool,'T_0')) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X143,bool),X144),s(fun(X146,bool),X147),s('type.pair.prod'(X143,X146),'const.pair.,_2'(s(X143,X145),s(X146,X148)))))) | s(bool,X149) = s(bool,happ(s(fun(X146,bool),X147),s(X146,X148))) | s(bool,'T_0') = s(bool,X149) | s(bool,X150) = s(bool,X151) | s(bool,X151) = s(bool,happ(s(fun(X143,bool),X144),s(X143,X145))) | s(bool,X150) = s(bool,X152) | s(bool,'T_0') = s(bool,X152)) )),
  inference(superposition,[],[f2763,f146])).
fof(f64821,plain,(
  ( ! [X37,X35,X33,X31,X38,X36,X34,X32] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X31,bool),X32),s(fun(X34,bool),X35),s('type.pair.prod'(X31,X34),'const.pair.,_2'(s(X31,X33),s(X34,X36)))))) | s(bool,X37) = s(bool,happ(s(fun(X34,bool),X35),s(X34,X36))) | s(bool,'T_0') = s(bool,X37) | s(bool,X38) = s(bool,happ(s(fun(X31,bool),X32),s(X31,X33))) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f64717,f37])).
fof(f64717,plain,(
  ( ! [X37,X35,X33,X31,X38,X36,X34,X32] : (~p(s(bool,'T_0')) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X31,bool),X32),s(fun(X34,bool),X35),s('type.pair.prod'(X31,X34),'const.pair.,_2'(s(X31,X33),s(X34,X36)))))) | s(bool,X37) = s(bool,happ(s(fun(X34,bool),X35),s(X34,X36))) | s(bool,'T_0') = s(bool,X37) | s(bool,X38) = s(bool,happ(s(fun(X31,bool),X32),s(X31,X33))) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f2763,f56])).
fof(f64820,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X4) = s(bool,happ(s(fun(X5,bool),X6),s(X5,X7))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(sK1,bool),sK4),s(fun(X5,bool),X6),s('type.pair.prod'(sK1,X5),'const.pair.,_2'(s(sK1,sK2),s(X5,X7)))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl4),
  inference(duplicate_literal_removal,[],[f64819])).
fof(f64819,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X4) = s(bool,happ(s(fun(X5,bool),X6),s(X5,X7))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(sK1,bool),sK4),s(fun(X5,bool),X6),s('type.pair.prod'(sK1,X5),'const.pair.,_2'(s(sK1,sK2),s(X5,X7)))))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl4),
  inference(condensation,[],[f64818])).
fof(f64818,plain,(
  ( ! [X14,X19,X17,X15,X18,X16] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(sK1,bool),sK4),s(fun(X15,bool),X16),s('type.pair.prod'(sK1,X15),'const.pair.,_2'(s(sK1,sK2),s(X15,X17)))))) | s(bool,X18) = s(bool,happ(s(fun(X15,bool),X16),s(X15,X17))) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X19) | s(bool,X14) = s(bool,X19)) ) | $spl4),
  inference(subsumption_resolution,[],[f64714,f200])).
fof(f64714,plain,(
  ( ! [X14,X19,X17,X15,X18,X16] : (~p(s(bool,X14)) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(sK1,bool),sK4),s(fun(X15,bool),X16),s('type.pair.prod'(sK1,X15),'const.pair.,_2'(s(sK1,sK2),s(X15,X17)))))) | s(bool,X18) = s(bool,happ(s(fun(X15,bool),X16),s(X15,X17))) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X19) | s(bool,X14) = s(bool,X19)) ) | $spl4),
  inference(superposition,[],[f2763,f22295])).
fof(f64689,plain,(
  ( ! [X441,X436,X438,X434,X440,X437,X439,X435] : (~p(s(bool,happ(s(fun(bool,bool),X437),s(bool,'T_0')))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X437),s(fun(X438,bool),X439),s('type.pair.prod'(bool,X438),'const.pair.,_2'(s(bool,g(s(X434,sK6(X435,X436,bool,X434)))),s(X438,X440)))))) | s(bool,X441) = s(bool,happ(s(fun(X438,bool),X439),s(X438,X440))) | s(bool,'T_0') = s(bool,X441) | s(fun(X434,bool),X435) = s(fun(X434,bool),X436) | s(bool,'T_0') = s(bool,f(s(X434,sK6(X435,X436,bool,X434))))) )),
  inference(superposition,[],[f2763,f2559])).
fof(f64675,plain,(
  ( ! [X356,X358,X354,X357,X353,X355] : (~p(s(bool,happ(s(fun(bool,bool),X354),s(bool,X353)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X354),s(fun(X355,bool),X356),s('type.pair.prod'(bool,X355),'const.pair.,_2'(s(bool,'F_0'),s(X355,X357)))))) | s(bool,X358) = s(bool,happ(s(fun(X355,bool),X356),s(X355,X357))) | s(bool,'T_0') = s(bool,X358) | s(bool,'T_0') = s(bool,X353)) )),
  inference(superposition,[],[f2763,f40])).
fof(f64670,plain,(
  ( ! [X315,X321,X316,X318,X314,X320,X322,X317,X319] : (~p(s(bool,happ(s(fun(bool,bool),X315),s(bool,X314)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X315),s(fun(X316,bool),X317),s('type.pair.prod'(bool,X316),'const.pair.,_2'(s(bool,'T_0'),s(X316,X318)))))) | s(bool,X319) = s(bool,happ(s(fun(X316,bool),X317),s(X316,X318))) | s(bool,'T_0') = s(bool,X319) | s(bool,X320) = s(bool,X321) | s(bool,X314) = s(bool,X321) | s(bool,X320) = s(bool,X322) | s(bool,'T_0') = s(bool,X322)) )),
  inference(superposition,[],[f2763,f146])).
fof(f64666,plain,(
  ( ! [X282,X288,X285,X287,X283,X284,X286] : (~p(s(bool,happ(s(fun(bool,bool),X283),s(bool,X282)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X283),s(fun(X284,bool),X285),s('type.pair.prod'(bool,X284),'const.pair.,_2'(s(bool,'T_0'),s(X284,X286)))))) | s(bool,X287) = s(bool,happ(s(fun(X284,bool),X285),s(X284,X286))) | s(bool,'T_0') = s(bool,X287) | s(bool,X282) = s(bool,X288) | s(bool,'T_0') = s(bool,X288)) )),
  inference(superposition,[],[f2763,f56])).
fof(f64645,plain,(
  ( ! [X134,X140,X138,X136,X135,X133,X141,X139,X137] : (~p(s(bool,happ(s(fun(bool,bool),X135),s(bool,X134)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X135),s(fun(X136,bool),X137),s('type.pair.prod'(bool,X136),'const.pair.,_2'(s(bool,X133),s(X136,X138)))))) | s(bool,X139) = s(bool,happ(s(fun(X136,bool),X137),s(X136,X138))) | s(bool,'T_0') = s(bool,X139) | s(bool,X134) = s(bool,X140) | s(bool,'T_0') = s(bool,X133) | s(bool,X140) = s(bool,X141) | s(bool,'T_0') = s(bool,X141)) )),
  inference(superposition,[],[f2763,f146])).
fof(f64644,plain,(
  ( ! [X127,X125,X132,X130,X128,X126,X124,X131,X129] : (~p(s(bool,happ(s(fun(bool,bool),X126),s(bool,X125)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X126),s(fun(X127,bool),X128),s('type.pair.prod'(bool,X127),'const.pair.,_2'(s(bool,X124),s(X127,X129)))))) | s(bool,X130) = s(bool,happ(s(fun(X127,bool),X128),s(X127,X129))) | s(bool,'T_0') = s(bool,X130) | s(bool,X124) = s(bool,X131) | s(bool,'T_0') = s(bool,X125) | s(bool,X131) = s(bool,X132) | s(bool,'T_0') = s(bool,X132)) )),
  inference(superposition,[],[f2763,f146])).
fof(f64642,plain,(
  ( ! [X111,X109,X107,X114,X112,X110,X108,X106,X113] : (~p(s(bool,happ(s(fun(bool,bool),X107),s(bool,'T_0')))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X107),s(fun(X108,bool),X109),s('type.pair.prod'(bool,X108),'const.pair.,_2'(s(bool,X106),s(X108,X110)))))) | s(bool,X111) = s(bool,happ(s(fun(X108,bool),X109),s(X108,X110))) | s(bool,'T_0') = s(bool,X111) | s(bool,X112) = s(bool,X113) | s(bool,X106) = s(bool,X113) | s(bool,X112) = s(bool,X114) | s(bool,'T_0') = s(bool,X114)) )),
  inference(superposition,[],[f2763,f146])).
fof(f64640,plain,(
  ( ! [X94,X92,X90,X88,X95,X93,X91,X89,X96] : (~p(s(bool,happ(s(fun(bool,bool),X90),s(bool,X89)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X90),s(fun(X91,bool),X92),s('type.pair.prod'(bool,X91),'const.pair.,_2'(s(bool,X88),s(X91,X93)))))) | s(bool,X94) = s(bool,happ(s(fun(X91,bool),X92),s(X91,X93))) | s(bool,'T_0') = s(bool,X94) | s(bool,X89) = s(bool,X95) | s(bool,'T_0') = s(bool,X95) | s(bool,X88) = s(bool,X96) | s(bool,'T_0') = s(bool,X96)) )),
  inference(superposition,[],[f2763,f146])).
fof(f64631,plain,(
  ( ! [X14,X19,X17,X15,X13,X18,X16] : (~p(s(bool,happ(s(fun(bool,bool),X15),s(bool,X14)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X15),s(fun(X16,bool),X17),s('type.pair.prod'(bool,X16),'const.pair.,_2'(s(bool,X13),s(X16,X18)))))) | s(bool,X19) = s(bool,happ(s(fun(X16,bool),X17),s(X16,X18))) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f2763,f56])).
fof(f64630,plain,(
  ( ! [X6,X12,X10,X8,X7,X11,X9] : (~p(s(bool,happ(s(fun(bool,bool),X7),s(bool,'T_0')))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X7),s(fun(X8,bool),X9),s('type.pair.prod'(bool,X8),'const.pair.,_2'(s(bool,X6),s(X8,X10)))))) | s(bool,X11) = s(bool,happ(s(fun(X8,bool),X9),s(X8,X10))) | s(bool,'T_0') = s(bool,X11) | s(bool,X6) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f2763,f56])).
fof(f64629,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X1),s(fun(X2,bool),X3),s('type.pair.prod'(bool,X2),'const.pair.,_2'(s(bool,X0),s(X2,X4)))))) | s(bool,X5) = s(bool,happ(s(fun(X2,bool),X3),s(X2,X4))) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2763,f40])).
fof(f64174,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X1,bool),X2),s(fun(X3,bool),X4),s('type.pair.prod'(X1,X3),X5))) | s(bool,X6) = s(bool,happ(s(fun(X3,bool),X4),s(X3,'const.pair.SND_1'(s('type.pair.prod'(X1,X3),X5))))) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(condensation,[],[f64173])).
fof(f64173,plain,(
  ( ! [X127,X125,X132,X130,X128,X126,X124,X131,X129] : (s(bool,X129) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X126,bool),X130),s(fun(X124,bool),X125),s('type.pair.prod'(X126,X124),X127))) | s(bool,'T_0') = s(bool,X129) | s(bool,X131) = s(bool,happ(s(fun(X124,bool),X125),s(X124,'const.pair.SND_1'(s('type.pair.prod'(X126,X124),X127))))) | s(bool,'T_0') = s(bool,X128) | s(bool,X131) = s(bool,X132) | s(bool,'T_0') = s(bool,X132)) )),
  inference(subsumption_resolution,[],[f64119,f58])).
fof(f64119,plain,(
  ( ! [X127,X125,X132,X130,X128,X126,X124,X131,X129] : (p(s(bool,X128)) | s(bool,X129) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X126,bool),X130),s(fun(X124,bool),X125),s('type.pair.prod'(X126,X124),X127))) | s(bool,'T_0') = s(bool,X129) | s(bool,X131) = s(bool,happ(s(fun(X124,bool),X125),s(X124,'const.pair.SND_1'(s('type.pair.prod'(X126,X124),X127))))) | s(bool,'T_0') = s(bool,X128) | s(bool,X131) = s(bool,X132) | s(bool,'T_0') = s(bool,X132)) )),
  inference(superposition,[],[f550,f146])).
fof(f64172,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X1,bool),X2),s(fun(X3,bool),X4),s('type.pair.prod'(X1,X3),X5))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X6,bool),'const.pair.,_2'(s(X6,X7),s(bool,happ(s(fun(X3,bool),X4),s(X3,'const.pair.SND_1'(s('type.pair.prod'(X1,X3),X5)))))))))) )),
  inference(condensation,[],[f64171])).
fof(f64171,plain,(
  ( ! [X61,X68,X66,X64,X62,X69,X67,X65,X63] : (s(bool,X66) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X63,bool),X67),s(fun(X61,bool),X62),s('type.pair.prod'(X63,X61),X64))) | s(bool,'T_0') = s(bool,X66) | s(bool,'T_0') = s(bool,X65) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X68,bool),'const.pair.,_2'(s(X68,X69),s(bool,happ(s(fun(X61,bool),X62),s(X61,'const.pair.SND_1'(s('type.pair.prod'(X63,X61),X64)))))))))) )),
  inference(subsumption_resolution,[],[f64112,f58])).
fof(f64112,plain,(
  ( ! [X61,X68,X66,X64,X62,X69,X67,X65,X63] : (p(s(bool,X65)) | s(bool,X66) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X63,bool),X67),s(fun(X61,bool),X62),s('type.pair.prod'(X63,X61),X64))) | s(bool,'T_0') = s(bool,X66) | s(bool,'T_0') = s(bool,X65) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X68,bool),'const.pair.,_2'(s(X68,X69),s(bool,happ(s(fun(X61,bool),X62),s(X61,'const.pair.SND_1'(s('type.pair.prod'(X63,X61),X64)))))))))) )),
  inference(superposition,[],[f550,f133])).
fof(f64170,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X1,bool),X2),s(fun(X3,bool),X4),s('type.pair.prod'(X1,X3),X5))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X6),'const.pair.,_2'(s(bool,happ(s(fun(X3,bool),X4),s(X3,'const.pair.SND_1'(s('type.pair.prod'(X1,X3),X5))))),s(X6,X7)))))) )),
  inference(condensation,[],[f64169])).
fof(f64169,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X42,X40] : (s(bool,X39) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X36,bool),X40),s(fun(X34,bool),X35),s('type.pair.prod'(X36,X34),X37))) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,X38) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X41),'const.pair.,_2'(s(bool,happ(s(fun(X34,bool),X35),s(X34,'const.pair.SND_1'(s('type.pair.prod'(X36,X34),X37))))),s(X41,X42)))))) )),
  inference(subsumption_resolution,[],[f64109,f58])).
fof(f64109,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X42,X40] : (p(s(bool,X38)) | s(bool,X39) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X36,bool),X40),s(fun(X34,bool),X35),s('type.pair.prod'(X36,X34),X37))) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,X38) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X41),'const.pair.,_2'(s(bool,happ(s(fun(X34,bool),X35),s(X34,'const.pair.SND_1'(s('type.pair.prod'(X36,X34),X37))))),s(X41,X42)))))) )),
  inference(superposition,[],[f550,f132])).
fof(f64168,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X4) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X2,bool),X5),s(fun(X0,bool),X1),s('type.pair.prod'(X2,X0),X3))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.pair.SND_1'(s('type.pair.prod'(X2,X0),X3)))))) )),
  inference(subsumption_resolution,[],[f64104,f39])).
fof(f64104,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'F_0')) | s(bool,X4) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X2,bool),X5),s(fun(X0,bool),X1),s('type.pair.prod'(X2,X0),X3))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.pair.SND_1'(s('type.pair.prod'(X2,X0),X3)))))) )),
  inference(superposition,[],[f550,f40])).
fof(f64101,plain,(
  ( ! [X288,X285,X287,X283,X289,X284,X286] : (p(s(bool,happ(s(fun(X283,bool),X287),s(X283,X286)))) | s(bool,X288) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X284,bool),X289),s(fun(X283,bool),X287),s('type.pair.prod'(X284,X283),'const.pair.,_2'(s(X284,X285),s(X283,X286))))) | s(bool,'T_0') = s(bool,X288)) )),
  inference(superposition,[],[f550,f44])).
fof(f64080,plain,(
  ( ! [X158,X156,X161,X159,X157,X155,X162,X160] : (p(s(bool,happ(s(fun(bool,bool),X158),s(bool,X157)))) | s(bool,X159) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X155,bool),X160),s(fun(bool,bool),X158),s('type.pair.prod'(X155,bool),X156))) | s(bool,'T_0') = s(bool,X159) | s(bool,X157) = s(bool,X161) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X155,bool),X156))) | s(bool,X161) = s(bool,X162) | s(bool,'T_0') = s(bool,X162)) )),
  inference(superposition,[],[f550,f146])).
fof(f64079,plain,(
  ( ! [X154,X152,X151,X149,X147,X153,X150,X148] : (p(s(bool,happ(s(fun(bool,bool),X150),s(bool,X149)))) | s(bool,X151) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X147,bool),X152),s(fun(bool,bool),X150),s('type.pair.prod'(X147,bool),X148))) | s(bool,'T_0') = s(bool,X151) | s(bool,X153) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X147,bool),X148))) | s(bool,'T_0') = s(bool,X149) | s(bool,X153) = s(bool,X154) | s(bool,'T_0') = s(bool,X154)) )),
  inference(superposition,[],[f550,f146])).
fof(f64077,plain,(
  ( ! [X134,X132,X138,X136,X135,X133,X131,X137] : (p(s(bool,happ(s(fun(bool,bool),X133),s(bool,'T_0')))) | s(bool,X134) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X131,bool),X135),s(fun(bool,bool),X133),s('type.pair.prod'(X131,bool),X132))) | s(bool,'T_0') = s(bool,X134) | s(bool,X136) = s(bool,X137) | s(bool,X137) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X131,bool),X132))) | s(bool,X136) = s(bool,X138) | s(bool,'T_0') = s(bool,X138)) )),
  inference(superposition,[],[f550,f146])).
fof(f64075,plain,(
  ( ! [X121,X118,X116,X122,X120,X119,X117,X115] : (p(s(bool,happ(s(fun(bool,bool),X118),s(bool,X117)))) | s(bool,X119) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X115,bool),X120),s(fun(bool,bool),X118),s('type.pair.prod'(X115,bool),X116))) | s(bool,'T_0') = s(bool,X119) | s(bool,X117) = s(bool,X121) | s(bool,'T_0') = s(bool,X121) | s(bool,X122) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X115,bool),X116))) | s(bool,'T_0') = s(bool,X122)) )),
  inference(superposition,[],[f550,f146])).
fof(f64066,plain,(
  ( ! [X54,X52,X50,X53,X51,X49] : (p(s(bool,happ(s(fun(bool,bool),X52),s(bool,X51)))) | s(bool,X53) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X49,bool),X54),s(fun(bool,bool),X52),s('type.pair.prod'(X49,bool),X50))) | s(bool,'T_0') = s(bool,X53) | s(bool,'T_0') = s(bool,X51) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X49,bool),X50)))) )),
  inference(superposition,[],[f550,f56])).
fof(f64065,plain,(
  ( ! [X47,X45,X43,X48,X46,X44] : (p(s(bool,happ(s(fun(bool,bool),X45),s(bool,'T_0')))) | s(bool,X46) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X43,bool),X47),s(fun(bool,bool),X45),s('type.pair.prod'(X43,bool),X44))) | s(bool,'T_0') = s(bool,X46) | s(bool,X48) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X43,bool),X44))) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f550,f56])).
fof(f64064,plain,(
  ( ! [X39,X41,X38,X42,X40] : (p(s(bool,happ(s(fun(bool,bool),X40),s(bool,'F_0')))) | s(bool,X41) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X38,bool),X42),s(fun(bool,bool),X40),s('type.pair.prod'(X38,bool),X39))) | s(bool,'T_0') = s(bool,X41) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X38,bool),X39)))) )),
  inference(superposition,[],[f550,f40])).
fof(f64063,plain,(
  ( ! [X37,X35,X33,X31,X36,X34,X32] : (p(s(bool,happ(s(fun(bool,bool),X34),s(bool,'T_0')))) | s(bool,X35) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X31,bool),X36),s(fun(bool,bool),X34),s('type.pair.prod'(X31,bool),'const.pair.,_2'(s(X31,X32),s(bool,X33))))) | s(bool,'T_0') = s(bool,X35) | s(bool,X33) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f550,f133])).
fof(f64062,plain,(
  ( ! [X30,X28,X26,X24,X29,X27,X25] : (p(s(bool,happ(s(fun(bool,bool),X28),s(bool,X27)))) | s(bool,X29) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X24,bool),X30),s(fun(bool,bool),X28),s('type.pair.prod'(X24,bool),'const.pair.,_2'(s(X24,X25),s(bool,X26))))) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f550,f157])).
fof(f64146,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (s(bool,'F_0') = s(bool,X10) | p(s(bool,happ(s(fun(bool,bool),X11),s(bool,'T_0')))) | s(bool,X12) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X8,bool),X13),s(fun(bool,bool),X11),s('type.pair.prod'(X8,bool),'const.pair.,_2'(s(X8,X9),s(bool,X10))))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(forward_demodulation,[],[f64060,f44])).
fof(f64060,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (p(s(bool,happ(s(fun(bool,bool),X11),s(bool,'T_0')))) | s(bool,X12) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X8,bool),X13),s(fun(bool,bool),X11),s('type.pair.prod'(X8,bool),'const.pair.,_2'(s(X8,X9),s(bool,X10))))) | s(bool,'T_0') = s(bool,X12) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X14,bool),'const.pair.,_2'(s(X14,X15),s(bool,X10)))))) )),
  inference(superposition,[],[f550,f1972])).
fof(f64145,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X2) | p(s(bool,happ(s(fun(bool,bool),X3),s(bool,'F_0')))) | s(bool,X4) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X0,bool),X5),s(fun(bool,bool),X3),s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,X2))))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(forward_demodulation,[],[f64059,f44])).
fof(f64059,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,happ(s(fun(bool,bool),X3),s(bool,'F_0')))) | s(bool,X4) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X0,bool),X5),s(fun(bool,bool),X3),s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,X2))))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X6,bool),'const.pair.,_2'(s(X6,X7),s(bool,X2)))))) )),
  inference(superposition,[],[f550,f1972])).
fof(f63627,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X1,bool),X2),s(fun(X3,bool),X4),s('type.pair.prod'(X1,X3),X5))) | s(bool,X6) = s(bool,happ(s(fun(X1,bool),X2),s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X3),X5))))) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(condensation,[],[f63626])).
fof(f63626,plain,(
  ( ! [X127,X125,X132,X130,X128,X126,X124,X131,X129] : (s(bool,X129) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X124,bool),X125),s(fun(X126,bool),X130),s('type.pair.prod'(X124,X126),X127))) | s(bool,'T_0') = s(bool,X129) | s(bool,X131) = s(bool,happ(s(fun(X124,bool),X125),s(X124,'const.pair.FST_1'(s('type.pair.prod'(X124,X126),X127))))) | s(bool,'T_0') = s(bool,X128) | s(bool,X131) = s(bool,X132) | s(bool,'T_0') = s(bool,X132)) )),
  inference(subsumption_resolution,[],[f63572,f58])).
fof(f63572,plain,(
  ( ! [X127,X125,X132,X130,X128,X126,X124,X131,X129] : (p(s(bool,X128)) | s(bool,X129) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X124,bool),X125),s(fun(X126,bool),X130),s('type.pair.prod'(X124,X126),X127))) | s(bool,'T_0') = s(bool,X129) | s(bool,X131) = s(bool,happ(s(fun(X124,bool),X125),s(X124,'const.pair.FST_1'(s('type.pair.prod'(X124,X126),X127))))) | s(bool,'T_0') = s(bool,X128) | s(bool,X131) = s(bool,X132) | s(bool,'T_0') = s(bool,X132)) )),
  inference(superposition,[],[f475,f146])).
fof(f63625,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X1,bool),X2),s(fun(X3,bool),X4),s('type.pair.prod'(X1,X3),X5))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X6,bool),'const.pair.,_2'(s(X6,X7),s(bool,happ(s(fun(X1,bool),X2),s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X3),X5)))))))))) )),
  inference(condensation,[],[f63624])).
fof(f63624,plain,(
  ( ! [X61,X68,X66,X64,X62,X69,X67,X65,X63] : (s(bool,X66) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X61,bool),X62),s(fun(X63,bool),X67),s('type.pair.prod'(X61,X63),X64))) | s(bool,'T_0') = s(bool,X66) | s(bool,'T_0') = s(bool,X65) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X68,bool),'const.pair.,_2'(s(X68,X69),s(bool,happ(s(fun(X61,bool),X62),s(X61,'const.pair.FST_1'(s('type.pair.prod'(X61,X63),X64)))))))))) )),
  inference(subsumption_resolution,[],[f63565,f58])).
fof(f63565,plain,(
  ( ! [X61,X68,X66,X64,X62,X69,X67,X65,X63] : (p(s(bool,X65)) | s(bool,X66) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X61,bool),X62),s(fun(X63,bool),X67),s('type.pair.prod'(X61,X63),X64))) | s(bool,'T_0') = s(bool,X66) | s(bool,'T_0') = s(bool,X65) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X68,bool),'const.pair.,_2'(s(X68,X69),s(bool,happ(s(fun(X61,bool),X62),s(X61,'const.pair.FST_1'(s('type.pair.prod'(X61,X63),X64)))))))))) )),
  inference(superposition,[],[f475,f133])).
fof(f63623,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X1,bool),X2),s(fun(X3,bool),X4),s('type.pair.prod'(X1,X3),X5))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X6),'const.pair.,_2'(s(bool,happ(s(fun(X1,bool),X2),s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X3),X5))))),s(X6,X7)))))) )),
  inference(condensation,[],[f63622])).
fof(f63622,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X42,X40] : (s(bool,X39) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X34,bool),X35),s(fun(X36,bool),X40),s('type.pair.prod'(X34,X36),X37))) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,X38) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X41),'const.pair.,_2'(s(bool,happ(s(fun(X34,bool),X35),s(X34,'const.pair.FST_1'(s('type.pair.prod'(X34,X36),X37))))),s(X41,X42)))))) )),
  inference(subsumption_resolution,[],[f63562,f58])).
fof(f63562,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X42,X40] : (p(s(bool,X38)) | s(bool,X39) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X34,bool),X35),s(fun(X36,bool),X40),s('type.pair.prod'(X34,X36),X37))) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,X38) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X41),'const.pair.,_2'(s(bool,happ(s(fun(X34,bool),X35),s(X34,'const.pair.FST_1'(s('type.pair.prod'(X34,X36),X37))))),s(X41,X42)))))) )),
  inference(superposition,[],[f475,f132])).
fof(f63621,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,X4) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X0,bool),X1),s(fun(X2,bool),X5),s('type.pair.prod'(X0,X2),X3))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X2),X3)))))) )),
  inference(subsumption_resolution,[],[f63557,f39])).
fof(f63557,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'F_0')) | s(bool,X4) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X0,bool),X1),s(fun(X2,bool),X5),s('type.pair.prod'(X0,X2),X3))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X2),X3)))))) )),
  inference(superposition,[],[f475,f40])).
fof(f63554,plain,(
  ( ! [X282,X285,X287,X281,X283,X284,X286] : (p(s(bool,happ(s(fun(X281,bool),X285),s(X281,X283)))) | s(bool,X286) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X281,bool),X285),s(fun(X282,bool),X287),s('type.pair.prod'(X281,X282),'const.pair.,_2'(s(X281,X283),s(X282,X284))))) | s(bool,'T_0') = s(bool,X286)) )),
  inference(superposition,[],[f475,f43])).
fof(f63533,plain,(
  ( ! [X158,X156,X154,X159,X157,X155,X153,X160] : (p(s(bool,happ(s(fun(bool,bool),X156),s(bool,X155)))) | s(bool,X157) = s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X156),s(fun(X153,bool),X158),s('type.pair.prod'(bool,X153),X154))) | s(bool,'T_0') = s(bool,X157) | s(bool,X155) = s(bool,X159) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X153),X154))) | s(bool,X159) = s(bool,X160) | s(bool,'T_0') = s(bool,X160)) )),
  inference(superposition,[],[f475,f146])).
fof(f63532,plain,(
  ( ! [X146,X152,X151,X149,X147,X145,X150,X148] : (p(s(bool,happ(s(fun(bool,bool),X148),s(bool,X147)))) | s(bool,X149) = s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X148),s(fun(X145,bool),X150),s('type.pair.prod'(bool,X145),X146))) | s(bool,'T_0') = s(bool,X149) | s(bool,X151) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X145),X146))) | s(bool,'T_0') = s(bool,X147) | s(bool,X151) = s(bool,X152) | s(bool,'T_0') = s(bool,X152)) )),
  inference(superposition,[],[f475,f146])).
fof(f63530,plain,(
  ( ! [X134,X132,X130,X136,X135,X133,X131,X129] : (p(s(bool,happ(s(fun(bool,bool),X131),s(bool,'T_0')))) | s(bool,X132) = s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X131),s(fun(X129,bool),X133),s('type.pair.prod'(bool,X129),X130))) | s(bool,'T_0') = s(bool,X132) | s(bool,X134) = s(bool,X135) | s(bool,X135) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X129),X130))) | s(bool,X134) = s(bool,X136) | s(bool,'T_0') = s(bool,X136)) )),
  inference(superposition,[],[f475,f146])).
fof(f63528,plain,(
  ( ! [X118,X116,X114,X120,X119,X117,X115,X113] : (p(s(bool,happ(s(fun(bool,bool),X116),s(bool,X115)))) | s(bool,X117) = s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X116),s(fun(X113,bool),X118),s('type.pair.prod'(bool,X113),X114))) | s(bool,'T_0') = s(bool,X117) | s(bool,X115) = s(bool,X119) | s(bool,'T_0') = s(bool,X119) | s(bool,X120) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X113),X114))) | s(bool,'T_0') = s(bool,X120)) )),
  inference(superposition,[],[f475,f146])).
fof(f63519,plain,(
  ( ! [X47,X52,X50,X48,X51,X49] : (p(s(bool,happ(s(fun(bool,bool),X50),s(bool,X49)))) | s(bool,X51) = s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X50),s(fun(X47,bool),X52),s('type.pair.prod'(bool,X47),X48))) | s(bool,'T_0') = s(bool,X51) | s(bool,'T_0') = s(bool,X49) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X47),X48)))) )),
  inference(superposition,[],[f475,f56])).
fof(f63518,plain,(
  ( ! [X45,X43,X41,X46,X44,X42] : (p(s(bool,happ(s(fun(bool,bool),X43),s(bool,'T_0')))) | s(bool,X44) = s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X43),s(fun(X41,bool),X45),s('type.pair.prod'(bool,X41),X42))) | s(bool,'T_0') = s(bool,X44) | s(bool,X46) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X41),X42))) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f475,f56])).
fof(f63517,plain,(
  ( ! [X39,X37,X38,X36,X40] : (p(s(bool,happ(s(fun(bool,bool),X38),s(bool,'F_0')))) | s(bool,X39) = s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X38),s(fun(X36,bool),X40),s('type.pair.prod'(bool,X36),X37))) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X36),X37)))) )),
  inference(superposition,[],[f475,f40])).
fof(f63598,plain,(
  ( ! [X14,X19,X17,X15,X18,X16] : (s(bool,'F_0') = s(bool,X15) | p(s(bool,happ(s(fun(bool,bool),X17),s(bool,'T_0')))) | s(bool,X18) = s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X17),s(fun(X14,bool),X19),s('type.pair.prod'(bool,X14),'const.pair.,_2'(s(bool,X15),s(X14,X16))))) | s(bool,'T_0') = s(bool,X18)) )),
  inference(forward_demodulation,[],[f63514,f44])).
fof(f63514,plain,(
  ( ! [X14,X21,X19,X17,X15,X20,X18,X16] : (p(s(bool,happ(s(fun(bool,bool),X17),s(bool,'T_0')))) | s(bool,X18) = s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X17),s(fun(X14,bool),X19),s('type.pair.prod'(bool,X14),'const.pair.,_2'(s(bool,X15),s(X14,X16))))) | s(bool,'T_0') = s(bool,X18) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X20,bool),'const.pair.,_2'(s(X20,X21),s(bool,X15)))))) )),
  inference(superposition,[],[f475,f1473])).
fof(f63513,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,happ(s(fun(bool,bool),X11),s(bool,X10)))) | s(bool,X12) = s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X11),s(fun(X7,bool),X13),s('type.pair.prod'(bool,X7),'const.pair.,_2'(s(bool,X8),s(X7,X9))))) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f475,f156])).
fof(f63512,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(bool,bool),X3),s(bool,'T_0')))) | s(bool,X4) = s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X3),s(fun(X0,bool),X5),s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,X1),s(X0,X2))))) | s(bool,'T_0') = s(bool,X4) | s(bool,X1) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f475,f132])).
fof(f37574,plain,(
  ( ! [X152,X151,X149,X153,X150] : (s(fun(X149,bool),X150) = s(fun(X149,bool),X151) | s(bool,X152) = s(bool,g(s(X149,sK6(X150,X151,bool,X149)))) | s(bool,'T_0') = s(bool,X152) | s(bool,X153) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X153) = s(bool,f(s(X149,sK6(X150,X151,bool,X149))))) ) | $spl4),
  inference(trivial_inequality_removal,[],[f37560])).
fof(f37560,plain,(
  ( ! [X152,X151,X149,X153,X150] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(X149,bool),X150) = s(fun(X149,bool),X151) | s(bool,X152) = s(bool,g(s(X149,sK6(X150,X151,bool,X149)))) | s(bool,'T_0') = s(bool,X152) | s(bool,X153) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X153) = s(bool,f(s(X149,sK6(X150,X151,bool,X149))))) ) | $spl4),
  inference(superposition,[],[f366,f22295])).
fof(f25856,plain,(
  ( ! [X660,X656,X658,X661,X657,X659] : (s(bool,'T_0') = s(bool,X656) | ~p(s(bool,happ(s(fun(bool,bool),X658),s(bool,X657)))) | ~p(s(bool,happ(s(fun(X659,bool),X660),s(X659,X661)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X659,bool),X660),s(fun(bool,bool),X658),s('type.pair.prod'(X659,bool),'const.pair.,_2'(s(X659,X661),s(bool,X656)))))) | s(bool,'T_0') = s(bool,X657)) ) | $spl4),
  inference(forward_demodulation,[],[f24397,f1085])).
fof(f24397,plain,(
  ( ! [X660,X656,X658,X661,X657,X659] : (~p(s(bool,happ(s(fun(bool,bool),X658),s(bool,X657)))) | ~p(s(bool,happ(s(fun(X659,bool),X660),s(X659,X661)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X659,bool),X660),s(fun(bool,bool),X658),s('type.pair.prod'(X659,bool),'const.pair.,_2'(s(X659,X661),s(bool,X656)))))) | s(bool,'T_0') = s(bool,X657) | s(bool,X656) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | $spl4),
  inference(superposition,[],[f785,f22295])).
fof(f25855,plain,(
  ( ! [X707,X708,X706,X709,X705] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X705,bool),X706))) | ~p(s(bool,happ(s(fun(bool,bool),X708),s(bool,X707)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X705,bool),X709),s(fun(bool,bool),X708),s('type.pair.prod'(X705,bool),X706)))) | ~p(s(bool,happ(s(fun(X705,bool),X709),s(X705,'const.pair.FST_1'(s('type.pair.prod'(X705,bool),X706)))))) | s(bool,'T_0') = s(bool,X707)) ) | $spl4),
  inference(forward_demodulation,[],[f24408,f1085])).
fof(f24408,plain,(
  ( ! [X707,X708,X706,X709,X705] : (~p(s(bool,happ(s(fun(bool,bool),X708),s(bool,X707)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X705,bool),X709),s(fun(bool,bool),X708),s('type.pair.prod'(X705,bool),X706)))) | ~p(s(bool,happ(s(fun(X705,bool),X709),s(X705,'const.pair.FST_1'(s('type.pair.prod'(X705,bool),X706)))))) | s(bool,'T_0') = s(bool,X707) | s(bool,'const.pair.SND_1'(s('type.pair.prod'(X705,bool),X706))) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | $spl4),
  inference(superposition,[],[f48,f22295])).
fof(f35172,plain,(
  ( ! [X47,X45,X48,X46,X44,X49] : (s(X45,f(s(bool,X44))) != s(X45,g(s(bool,X44))) | s(fun(bool,X45),X46) = s(fun(bool,X45),X47) | s(bool,X48) = s(bool,sK6(X46,X47,X45,bool)) | s(bool,'T_0') = s(bool,X48) | s(bool,X49) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X44) = s(bool,X49)) ) | $spl4),
  inference(superposition,[],[f361,f22295])).
fof(f28579,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) != s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X2)) ) | $spl4),
  inference(condensation,[],[f28578])).
fof(f28578,plain,(
  ( ! [X54,X56,X55,X53] : (s(bool,X54) != s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,'T_0') = s(bool,X53) | s(bool,X54) = s(bool,X55) | s(bool,X55) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) ) | $spl4),
  inference(duplicate_literal_removal,[],[f28508])).
fof(f28508,plain,(
  ( ! [X54,X56,X55,X53] : (s(bool,X54) != s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,'T_0') = s(bool,X53) | s(bool,X54) = s(bool,X55) | s(bool,'T_0') = s(bool,X53) | s(bool,X55) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) ) | $spl4),
  inference(superposition,[],[f24431,f146])).
fof(f24431,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,'T_0') = s(bool,X0)) ) | $spl4),
  inference(condensation,[],[f24429])).
fof(f24429,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,'T_0') = s(bool,X1) | s(bool,X0) = s(bool,X1)) ) | $spl4),
  inference(equality_factoring,[],[f22295])).
fof(f74628,plain,(
  ( ! [X280,X281] : (s(bool,'F_0') != s(bool,f(s(bool,X281))) | s(bool,'F_0') = s(bool,X280) | s(bool,'T_0') = s(bool,X281)) ) | ($spl40 | $spl176 | $spl234)),
  inference(forward_demodulation,[],[f74381,f2561])).
fof(f74381,plain,(
  ( ! [X280,X281] : (s(bool,f(s(bool,'T_0'))) != s(bool,f(s(bool,X281))) | s(bool,'F_0') = s(bool,X280) | s(bool,'T_0') = s(bool,X281)) ) | ($spl176 | $spl234)),
  inference(duplicate_literal_removal,[],[f74087])).
fof(f74087,plain,(
  ( ! [X280,X281] : (s(bool,f(s(bool,'T_0'))) != s(bool,f(s(bool,X281))) | s(bool,'F_0') = s(bool,X280) | s(bool,'F_0') = s(bool,X280) | s(bool,'T_0') = s(bool,X281)) ) | ($spl176 | $spl234)),
  inference(superposition,[],[f42127,f31623])).
fof(f80451,plain,(
  ( ! [X2,X1] : (s(bool,X1) != s(bool,g(s(bool,X1))) | s(bool,'T_0') = s(bool,X2) | s(bool,X1) = s(bool,X2)) ) | $spl358),
  inference(cnf_transformation,[],[f80451_D])).
fof(f80451_D,plain,(
  ( ! [X2,X1] : (s(bool,X1) != s(bool,g(s(bool,X1))) | s(bool,'T_0') = s(bool,X2) | s(bool,X1) = s(bool,X2)) ) <=> ~$spl358),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl358])])).
fof(f69,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl6),
  inference(cnf_transformation,[],[f69_D])).
fof(f69_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f80453,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,X1) != s(bool,g(s(bool,X1))) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl4 | $spl9 | $spl36 | $spl228)),
  inference(condensation,[],[f80452])).
fof(f80452,plain,(
  ( ! [X39,X37,X38,X40] : (s(bool,X38) != s(bool,g(s(bool,X38))) | s(bool,'F_0') = s(bool,X37) | s(bool,X37) = s(bool,X39) | s(bool,X38) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) ) | ($spl4 | $spl9 | $spl36 | $spl228)),
  inference(subsumption_resolution,[],[f80348,f30026])).
fof(f80348,plain,(
  ( ! [X39,X37,X38,X40] : (s(bool,X38) != s(bool,g(s(bool,X38))) | s(bool,'F_0') = s(bool,X37) | s(bool,X37) = s(bool,X39) | s(bool,'T_0') = s(bool,X39) | s(bool,X38) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) ) | $spl228),
  inference(superposition,[],[f38764,f146])).
fof(f80450,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,X1) != s(bool,g(s(bool,X1))) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl4 | $spl9 | $spl36 | $spl228)),
  inference(condensation,[],[f80449])).
fof(f80449,plain,(
  ( ! [X35,X33,X36,X34] : (s(bool,X34) != s(bool,g(s(bool,X34))) | s(bool,'F_0') = s(bool,X33) | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35) | s(bool,X33) = s(bool,X36)) ) | ($spl4 | $spl9 | $spl36 | $spl228)),
  inference(subsumption_resolution,[],[f80347,f30026])).
fof(f80347,plain,(
  ( ! [X35,X33,X36,X34] : (s(bool,X34) != s(bool,g(s(bool,X34))) | s(bool,'F_0') = s(bool,X33) | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35) | s(bool,X33) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) ) | $spl228),
  inference(superposition,[],[f38764,f146])).
fof(f76300,plain,(
  ( ! [X666,X662,X667,X663,X665] : (s(bool,X666) = s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X665),s(fun(X662,bool),X667),s('type.pair.prod'(bool,X662),X663))) | s(bool,'T_0') = s(bool,f(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X662),X663))))) | s(bool,'T_0') = s(bool,X666)) ) | $spl356),
  inference(cnf_transformation,[],[f76300_D])).
fof(f76300_D,plain,(
  ( ! [X666,X662,X667,X663,X665] : (s(bool,X666) = s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X665),s(fun(X662,bool),X667),s('type.pair.prod'(bool,X662),X663))) | s(bool,'T_0') = s(bool,f(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X662),X663))))) | s(bool,'T_0') = s(bool,X666)) ) <=> ~$spl356),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl356])])).
fof(f76299,plain,(
  ( ! [X666,X662,X664,X667,X663,X665] : (s(bool,X666) = s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X665),s(fun(X662,bool),X667),s('type.pair.prod'(bool,X662),X663))) | s(bool,'T_0') = s(bool,X666) | s(bool,'T_0') = s(bool,f(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X662),X663))))) | s(bool,'F_0') = s(bool,X664)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl290)),
  inference(subsumption_resolution,[],[f65535,f68413])).
fof(f68413,plain,(
  ( ! [X105,X104] : (~p(s(bool,happ(s(fun(bool,bool),X105),s(bool,X104)))) | s(bool,'F_0') = s(bool,X104)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl290)),
  inference(subsumption_resolution,[],[f68282,f49260])).
fof(f68282,plain,(
  ( ! [X105,X104] : (~p(s(bool,happ(s(fun(bool,bool),X105),s(bool,X104)))) | s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X104)) ) | ($spl4 | $spl9 | $spl36 | $spl290)),
  inference(superposition,[],[f67866,f30029])).
fof(f65535,plain,(
  ( ! [X666,X662,X664,X667,X663,X665] : (p(s(bool,happ(s(fun(bool,bool),X665),s(bool,X664)))) | s(bool,X666) = s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X665),s(fun(X662,bool),X667),s('type.pair.prod'(bool,X662),X663))) | s(bool,'T_0') = s(bool,X666) | s(bool,'T_0') = s(bool,f(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X662),X663))))) | s(bool,'F_0') = s(bool,X664)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f475,f30029])).
fof(f76298,plain,(
  ( ! [X677,X673,X675,X676,X672] : (s(bool,X676) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X672,bool),X677),s(fun(bool,bool),X675),s('type.pair.prod'(X672,bool),X673))) | s(bool,'T_0') = s(bool,f(s(bool,'const.pair.SND_1'(s('type.pair.prod'(X672,bool),X673))))) | s(bool,'T_0') = s(bool,X676)) ) | $spl354),
  inference(cnf_transformation,[],[f76298_D])).
fof(f76298_D,plain,(
  ( ! [X677,X673,X675,X676,X672] : (s(bool,X676) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X672,bool),X677),s(fun(bool,bool),X675),s('type.pair.prod'(X672,bool),X673))) | s(bool,'T_0') = s(bool,f(s(bool,'const.pair.SND_1'(s('type.pair.prod'(X672,bool),X673))))) | s(bool,'T_0') = s(bool,X676)) ) <=> ~$spl354),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl354])])).
fof(f76297,plain,(
  ( ! [X677,X673,X675,X676,X672,X674] : (s(bool,X676) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X672,bool),X677),s(fun(bool,bool),X675),s('type.pair.prod'(X672,bool),X673))) | s(bool,'T_0') = s(bool,X676) | s(bool,'T_0') = s(bool,f(s(bool,'const.pair.SND_1'(s('type.pair.prod'(X672,bool),X673))))) | s(bool,'F_0') = s(bool,X674)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl290)),
  inference(subsumption_resolution,[],[f65537,f68413])).
fof(f65537,plain,(
  ( ! [X677,X673,X675,X676,X672,X674] : (p(s(bool,happ(s(fun(bool,bool),X675),s(bool,X674)))) | s(bool,X676) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X672,bool),X677),s(fun(bool,bool),X675),s('type.pair.prod'(X672,bool),X673))) | s(bool,'T_0') = s(bool,X676) | s(bool,'T_0') = s(bool,f(s(bool,'const.pair.SND_1'(s('type.pair.prod'(X672,bool),X673))))) | s(bool,'F_0') = s(bool,X674)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f550,f30029])).
fof(f75419,plain,(
  $false | ($spl4 | $spl9 | $spl36 | $spl42 | $spl47 | $spl218 | $spl278)),
  inference(subsumption_resolution,[],[f3304,f74886])).
fof(f74886,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl4 | $spl9 | $spl36 | $spl42 | $spl218 | $spl278)),
  inference(subsumption_resolution,[],[f74885,f3890])).
fof(f74885,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'T_0'),s(bool,X0)))))) ) | ($spl4 | $spl9 | $spl36 | $spl42 | $spl218 | $spl278)),
  inference(condensation,[],[f74821])).
fof(f74821,plain,(
  ( ! [X64,X63] : (s(bool,X64) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'T_0'),s(bool,X64))))) | s(bool,'T_0') = s(bool,X63) | s(bool,'T_0') = s(bool,X64)) ) | ($spl4 | $spl9 | $spl36 | $spl42 | $spl218 | $spl278)),
  inference(backward_demodulation,[],[f74820,f3170])).
fof(f74820,plain,(
  ( ! [X487] : (s(bool,'T_0') = s(bool,f(s(bool,X487)))) ) | ($spl4 | $spl9 | $spl36 | $spl218 | $spl278)),
  inference(global_subsumption,[],[f65692,f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f64,f132,f133,f146,f156,f157,f200,f361,f363,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f3832,f3890,f3878,f2550,f2549,f3652,f3732,f3828,f3727,f2700,f763,f2698,f761,f1935,f1480,f82,f104,f2706,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13225,f13350,f13427,f13415,f22573,f18154,f18155,f22291,f22279,f22275,f22219,f22217,f18360,f18362,f18363,f18802,f19158,f21203,f19338,f19402,f19403,f19570,f22984,f20200,f20161,f20156,f19927,f19749,f19571,f19566,f19337,f19159,f18981,f18976,f18803,f18798,f25801,f25802,f13351,f13346,f13175,f13170,f13039,f9980,f9978,f9977,f9974,f6734,f6732,f6731,f6728,f25850,f25854,f30902,f32190,f35168,f35169,f35170,f37182,f13038,f3651,f15352,f22930,f22929,f22928,f786,f550,f475,f48,f47,f46,f19925,f19747,f19569,f18801,f762,f2704,f19406,f19914,f19736,f19558,f18790,f2699,f785,f2763,f22931,f37580,f37578,f37577,f37576,f34476,f34474,f34472,f34425,f22898,f19565,f19331,f19153,f18797,f15069,f13826,f13345,f13169,f13033,f4154,f3827,f3726,f3646,f2559,f368,f366,f32792,f32194,f32193,f32192,f31357,f30906,f30905,f30904,f22293,f22277,f19401,f19400,f17276,f17226,f15341,f13224,f13223,f3759,f3758,f2874,f1972,f1473,f370,f365,f39,f17266,f17264,f16976,f63512,f63513,f63598,f63517,f63518,f63519,f63528,f63530,f63532,f63533,f63554,f63621,f63623,f63625,f63627,f64145,f64146,f64062,f64063,f64064,f64065,f64066,f64075,f64077,f64079,f64080,f64101,f64168,f64170,f64172,f64174,f64629,f64630,f64631,f64640,f64642,f64644,f64645,f64666,f64670,f64675,f64689,f64821,f64828,f66956,f66957,f66958,f66967,f66969,f66971,f66972,f67074,f67081,f67781,f67874,f67881,f68646,f68653,f69246,f69154,f69156,f69169,f69170,f69280,f69281,f69288,f60407])).
fof(f65692,plain,(
  ( ! [X484,X486,X488,X485,X487,X483,X489] : (p(s(bool,X487)) | s(bool,X488) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X483,bool),X484),s(fun(X485,bool),X489),s('type.pair.prod'(X483,X485),X486))) | s(bool,'T_0') = s(bool,X488) | s(bool,'T_0') = s(bool,f(s(bool,X487))) | s(bool,'F_0') = s(bool,happ(s(fun(X483,bool),X484),s(X483,'const.pair.FST_1'(s('type.pair.prod'(X483,X485),X486)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f475,f30029])).
fof(f75357,plain,(
  $false | ($spl4 | $spl9 | $spl13 | $spl22 | $spl36 | $spl42 | $spl218 | $spl278)),
  inference(subsumption_resolution,[],[f74912,f74886])).
fof(f74912,plain,(
  s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'T_0'),s(bool,'T_0'))))) | ($spl4 | $spl9 | $spl13 | $spl22 | $spl36 | $spl42 | $spl218 | $spl278)),
  inference(backward_demodulation,[],[f74886,f6406])).
fof(f75355,plain,(
  $false | ($spl4 | $spl9 | $spl36 | $spl42 | $spl218 | $spl278)),
  inference(subsumption_resolution,[],[f74887,f37])).
fof(f74887,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl9 | $spl36 | $spl42 | $spl218 | $spl278)),
  inference(backward_demodulation,[],[f74886,f39])).
fof(f74807,plain,(
  ( ! [X190] : (s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,X190))))) | s(bool,'F_0') = s(bool,X190)) ) | $spl352),
  inference(cnf_transformation,[],[f74807_D])).
fof(f74807_D,plain,(
  ( ! [X190] : (s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,X190))))) | s(bool,'F_0') = s(bool,X190)) ) <=> ~$spl352),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl352])])).
fof(f7405,plain,(
  ( ! [X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0'))))) ) | $spl92),
  inference(cnf_transformation,[],[f7405_D])).
fof(f7405_D,plain,(
  ( ! [X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0'))))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f69361,plain,(
  ( ! [X191,X190] : (~p(s(bool,happ(s(fun(bool,bool),X191),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,X190))))) | s(bool,'F_0') = s(bool,X190)) ) | ($spl4 | $spl9 | $spl36 | $spl204 | $spl290)),
  inference(superposition,[],[f68414,f37102])).
fof(f68414,plain,(
  ( ! [X107,X106] : (~p(s(bool,happ(s(fun(bool,bool),X107),s(bool,X106)))) | s(bool,'T_0') = s(bool,f(s(bool,X106)))) ) | ($spl4 | $spl9 | $spl36 | $spl290)),
  inference(subsumption_resolution,[],[f68283,f71])).
fof(f68283,plain,(
  ( ! [X107,X106] : (~p(s(bool,happ(s(fun(bool,bool),X107),s(bool,X106)))) | s(bool,'T_0') = s(bool,f(s(bool,X106))) | s(bool,'T_0') = s(bool,'F_0')) ) | ($spl4 | $spl9 | $spl36 | $spl290)),
  inference(superposition,[],[f67866,f30029])).
fof(f74730,plain,(
  ( ! [X320,X322] : (s(bool,'T_0') = s(bool,f(s(bool,X322))) | s(bool,'T_0') = s(bool,X320) | s(bool,X322) = s(bool,f(s(bool,X320)))) ) | $spl342),
  inference(cnf_transformation,[],[f74730_D])).
fof(f74730_D,plain,(
  ( ! [X320,X322] : (s(bool,'T_0') = s(bool,f(s(bool,X322))) | s(bool,'T_0') = s(bool,X320) | s(bool,X322) = s(bool,f(s(bool,X320)))) ) <=> ~$spl342),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl342])])).
fof(f74440,plain,(
  ( ! [X138,X139,X137] : (s(bool,X138) = s(bool,f(s(bool,X139))) | s(bool,'F_0') = s(bool,X137) | s(bool,'T_0') = s(bool,X139) | s(bool,'T_0') = s(bool,f(s(bool,X138)))) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(duplicate_literal_removal,[],[f73803])).
fof(f73803,plain,(
  ( ! [X138,X139,X137] : (s(bool,X138) = s(bool,f(s(bool,X139))) | s(bool,'F_0') = s(bool,X137) | s(bool,'T_0') = s(bool,X139) | s(bool,'T_0') = s(bool,f(s(bool,X138))) | s(bool,'F_0') = s(bool,X137)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(superposition,[],[f31623,f30029])).
fof(f74395,plain,(
  ( ! [X163,X164,X162] : (s(bool,'F_0') = s(bool,X163) | s(bool,'T_0') = s(bool,X162) | s(bool,'T_0') = s(bool,f(s(bool,X164))) | s(bool,X164) = s(bool,f(s(bool,X162)))) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(duplicate_literal_removal,[],[f74003])).
fof(f74003,plain,(
  ( ! [X163,X164,X162] : (s(bool,'F_0') = s(bool,X163) | s(bool,'F_0') = s(bool,X163) | s(bool,'T_0') = s(bool,X162) | s(bool,'T_0') = s(bool,f(s(bool,X164))) | s(bool,X164) = s(bool,f(s(bool,X162)))) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(superposition,[],[f31623,f30029])).
fof(f74464,plain,(
  ( ! [X2,X3,X1] : (s(bool,X1) = s(bool,f(s(bool,X2))) | s(bool,'T_0') = s(bool,X3) | s(bool,X1) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) ) | $spl326),
  inference(cnf_transformation,[],[f74464_D])).
fof(f74464_D,plain,(
  ( ! [X2,X3,X1] : (s(bool,X1) = s(bool,f(s(bool,X2))) | s(bool,'T_0') = s(bool,X3) | s(bool,X1) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl326),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl326])])).
fof(f74766,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,X1) = s(bool,f(s(bool,X2))) | s(bool,X1) = s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(condensation,[],[f74765])).
fof(f74765,plain,(
  ( ! [X59,X57,X60,X58,X56] : (s(bool,X58) = s(bool,f(s(bool,X57))) | s(bool,X58) = s(bool,X59) | s(bool,'T_0') = s(bool,X59) | s(bool,X56) = s(bool,X60) | s(bool,'F_0') = s(bool,X56) | s(bool,'T_0') = s(bool,X57)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(subsumption_resolution,[],[f74025,f30026])).
fof(f74025,plain,(
  ( ! [X59,X57,X60,X58,X56] : (s(bool,X58) = s(bool,f(s(bool,X57))) | s(bool,X58) = s(bool,X59) | s(bool,'T_0') = s(bool,X59) | s(bool,X56) = s(bool,X60) | s(bool,'T_0') = s(bool,X60) | s(bool,'F_0') = s(bool,X56) | s(bool,'T_0') = s(bool,X57)) ) | $spl176),
  inference(superposition,[],[f146,f31623])).
fof(f74764,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,X1) = s(bool,f(s(bool,X2))) | s(bool,X1) = s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(condensation,[],[f74763])).
fof(f74763,plain,(
  ( ! [X61,X64,X62,X65,X63] : (s(bool,X63) = s(bool,f(s(bool,X62))) | s(bool,X61) = s(bool,X64) | s(bool,X63) = s(bool,X65) | s(bool,'T_0') = s(bool,X65) | s(bool,'F_0') = s(bool,X61) | s(bool,'T_0') = s(bool,X62)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(subsumption_resolution,[],[f74026,f30026])).
fof(f74026,plain,(
  ( ! [X61,X64,X62,X65,X63] : (s(bool,X63) = s(bool,f(s(bool,X62))) | s(bool,X61) = s(bool,X64) | s(bool,'T_0') = s(bool,X64) | s(bool,X63) = s(bool,X65) | s(bool,'T_0') = s(bool,X65) | s(bool,'F_0') = s(bool,X61) | s(bool,'T_0') = s(bool,X62)) ) | $spl176),
  inference(superposition,[],[f146,f31623])).
fof(f74745,plain,(
  ( ! [X508] : (p(s(bool,f(s(bool,X508)))) | s(bool,'T_0') = s(bool,X508)) ) | $spl346),
  inference(cnf_transformation,[],[f74745_D])).
fof(f74745_D,plain,(
  ( ! [X508] : (p(s(bool,f(s(bool,X508)))) | s(bool,'T_0') = s(bool,X508)) ) <=> ~$spl346),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl346])])).
fof(f74762,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,f(s(bool,X1)))) | s(bool,'T_0') = s(bool,X1)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(condensation,[],[f74761])).
fof(f74761,plain,(
  ( ! [X118,X116,X117] : (p(s(bool,f(s(bool,X117)))) | s(bool,X116) = s(bool,X118) | s(bool,'F_0') = s(bool,X116) | s(bool,'T_0') = s(bool,X117)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(subsumption_resolution,[],[f74037,f30026])).
fof(f74037,plain,(
  ( ! [X118,X116,X117] : (p(s(bool,f(s(bool,X117)))) | s(bool,X116) = s(bool,X118) | s(bool,'T_0') = s(bool,X118) | s(bool,'F_0') = s(bool,X116) | s(bool,'T_0') = s(bool,X117)) ) | $spl176),
  inference(superposition,[],[f200,f31623])).
fof(f74647,plain,(
  ( ! [X318,X317] : (~p(s(bool,happ(s(fun(bool,bool),X318),s(bool,f(s(bool,X317)))))) | s(bool,'T_0') = s(bool,X317)) ) | $spl338),
  inference(cnf_transformation,[],[f74647_D])).
fof(f74647_D,plain,(
  ( ! [X318,X317] : (~p(s(bool,happ(s(fun(bool,bool),X318),s(bool,f(s(bool,X317)))))) | s(bool,'T_0') = s(bool,X317)) ) <=> ~$spl338),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl338])])).
fof(f74760,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X2) | ~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,f(s(bool,X2)))))) | s(bool,'F_0') = s(bool,X0)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl176 | $spl290)),
  inference(global_subsumption,[],[f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f67,f64,f132,f133,f146,f156,f157,f200,f361,f363,f502,f496,f71,f54,f439,f1085,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f4170,f4171,f3832,f3890,f3878,f2550,f2549,f3652,f3732,f3828,f3727,f2700,f763,f2698,f761,f2455,f1935,f1480,f82,f104,f2706,f11126,f11125,f11188,f11200,f11161,f11162,f11222,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13852,f13225,f13350,f13427,f13415,f15429,f17356,f22573,f18154,f18155,f22295,f22291,f22279,f22275,f22219,f22217,f18360,f18362,f18363,f22711,f18802,f19158,f21203,f22901,f22902,f22905,f19338,f19402,f19403,f19570,f22984,f20200,f23229,f23233,f23235,f24931,f24930,f24854,f24822,f24800,f23787,f25217,f25261,f25262,f25265,f25280,f25784,f25785,f24220,f24215,f25787,f25788,f23788,f23783,f25790,f23712,f20161,f20156,f25794,f19927,f19749,f19571,f19566,f25798,f19337,f19159,f18981,f18976,f18803,f18798,f25801,f25802,f13351,f13346,f25804,f13292,f13175,f13170,f13039,f11230,f9980,f9978,f9977,f9974,f6734,f6732,f6731,f6728,f24431,f25850,f25854,f28579,f28795,f28928,f28929,f28799,f28801,f28992,f29101,f29277,f29642,f30026,f30749,f30802,f30902,f31517,f31518,f31523,f32112,f32190,f32954,f32955,f32960,f33699,f33700,f33637,f33624,f33745,f33747,f33748,f33749,f33751,f33888,f33897,f34113,f34117,f35168,f35169,f35170,f35172,f36886,f36851,f36833,f36832,f37018,f37019,f37020,f37021,f36024,f36821,f36820,f37037,f37041,f36264,f36272,f36746,f37183,f37187,f37191,f36506,f36507,f36509,f36510,f37202,f37203,f36577,f36580,f37595,f11203,f11293,f36263,f41608,f41609,f42049,f41681,f41683,f41684,f41685,f41687,f42116,f42119,f41950,f42132,f45396,f45400,f49038,f49040,f49042,f31629,f31632,f31906,f31910,f49049,f49051,f37194,f49053,f36723,f39729,f49061,f49063,f39735,f39738,f39989,f39992,f49070,f49078,f45327,f45330,f44862,f36220,f49246,f36225,f43667,f43893,f43894,f43907,f45693,f45866,f45867,f45883,f49260,f44738,f37032,f49695,f49248,f49261,f49263,f49266,f49269,f49276,f49697,f49281,f36718,f50029,f50031,f50032,f50033,f50035,f45545,f13038,f3651,f15352,f22930,f22929,f22928,f786,f550,f475,f48,f47,f46,f25855,f19925,f19747,f19569,f18801,f762,f2762,f11302,f2704,f19406,f25856,f19914,f19736,f19558,f18790,f2699,f785,f2763,f22931,f45539,f37580,f37578,f37577,f37576,f37574,f37033,f36728,f34476,f34474,f34472,f34425,f24214,f23782,f22898,f19565,f19331,f19153,f18797,f15069,f13826,f13345,f13169,f13033,f4154,f3827,f3726,f3646,f2559,f368,f366,f49702,f49268,f49265,f49037,f45453,f45254,f37594,f37211,f37040,f37038,f36953,f36870,f36709,f36702,f36513,f36501,f36494,f36344,f36337,f36111,f36110,f36025,f36023,f36022,f34518,f34208,f34207,f34119,f34116,f34115,f33645,f33643,f33629,f33627,f33552,f33551,f32991,f32792,f32361,f32194,f32193,f32192,f32154,f32152,f32151,f32126,f32105,f32101,f31911,f31907,f31623,f31621,f31431,f31357,f31001,f30906,f30905,f30904,f30750,f30404,f30179,f30178,f30174,f30172,f30080,f30030,f30029,f30027,f30024,f29905,f29904,f29900,f29898,f29715,f29714,f29581,f29580,f29283,f29256,f29255,f29252,f29247,f29245,f29243,f29241,f25810,f25806,f25792,f25786,f22303,f22293,f22277,f21173,f19401,f19400,f17276,f17226,f15341,f13224,f13223,f3759,f3758,f2874,f1972,f1473,f370,f365,f39,f36343,f45544,f2561,f46456,f45588,f45538,f45454,f45392,f45391,f36874,f36872,f36708,f36696,f36575,f36504,f36500,f36488,f36331,f36184,f35841,f35840,f31903,f31902,f13854,f13853,f11201,f17266,f17264,f16976,f60404,f63512,f63513,f63598,f63517,f63518,f63519,f63528,f63530,f63532,f63533,f63614,f63544,f63546,f63618,f63554,f63621,f63623,f63625,f63627,f63641,f63838,f63839,f63845,f63730,f63731,f63853,f63855,f63862,f63864,f63865,f63869,f64145,f64146,f64062,f64063,f64064,f64065,f64066,f64075,f64077,f64079,f64080,f64161,f64091,f64093,f64165,f64101,f64168,f64170,f64172,f64174,f64188,f64325,f64324,f64321,f64457,f64459,f64460,f64461,f64463,f64629,f64630,f64631,f64640,f64642,f64644,f64645,f64768,f64656,f64658,f64772,f64666,f64670,f64673,f64675,f64782,f64783,f64689,f64820,f64821,f64828,f64836,f64839,f65344,f66452,f65529,f65535,f65537,f65728,f65734,f65736,f65754,f65830,f66588,f66589,f65985,f65991,f65993,f66956,f66957,f66958,f66967,f66969,f66971,f66972,f67066,f66983,f66985,f66987,f66988,f66989,f67070,f67074,f67081,f67089,f67094,f67781,f67872,f67873,f67874,f67881,f67889,f67894,f67866,f68252,f68256,f68413,f68414,f68415,f68419,f68423,f68430,f68434,f68646,f68653,f68661,f68666,f69129,f69131,f69133,f69134,f69138,f69246,f69154,f69156,f69169,f69170,f69181,f69187,f69280,f69281,f69288,f69296,f69301,f69521,f69523,f69525,f69532,f69536,f69542,f70246,f70249,f70256,f70263,f70275,f70160,f70161,f70322,f70198,f70206,f74463,f74466,f74438,f73853,f73854,f74556,f74562,f74376,f74112,f74190,f74667,f74196,f74197,f74202,f74204,f74207,f74208,f74360,f74739,f74741,f74153,f74148,f74751,f74387,f74759])).
fof(f74759,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'F_0') = s(bool,X0) | ~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,f(s(bool,X2)))))) | ~p(s(bool,happ(s(fun(X3,bool),X4),s(X3,X5)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X3,bool),X4),s(fun(bool,bool),X1),s('type.pair.prod'(X3,bool),'const.pair.,_2'(s(X3,X5),s(bool,'T_0')))))) | s(bool,'T_0') = s(bool,X2)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(condensation,[],[f74758])).
fof(f74758,plain,(
  ( ! [X175,X173,X171,X176,X174,X172,X170] : (~p(s(bool,happ(s(fun(bool,bool),X172),s(bool,f(s(bool,X171)))))) | ~p(s(bool,happ(s(fun(X173,bool),X174),s(X173,X175)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X173,bool),X174),s(fun(bool,bool),X172),s('type.pair.prod'(X173,bool),'const.pair.,_2'(s(X173,X175),s(bool,'T_0')))))) | s(bool,X170) = s(bool,X176) | s(bool,'F_0') = s(bool,X170) | s(bool,'T_0') = s(bool,X171)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(subsumption_resolution,[],[f74048,f30026])).
fof(f74048,plain,(
  ( ! [X175,X173,X171,X176,X174,X172,X170] : (~p(s(bool,happ(s(fun(bool,bool),X172),s(bool,f(s(bool,X171)))))) | ~p(s(bool,happ(s(fun(X173,bool),X174),s(X173,X175)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X173,bool),X174),s(fun(bool,bool),X172),s('type.pair.prod'(X173,bool),'const.pair.,_2'(s(X173,X175),s(bool,'T_0')))))) | s(bool,X170) = s(bool,X176) | s(bool,'T_0') = s(bool,X176) | s(bool,'F_0') = s(bool,X170) | s(bool,'T_0') = s(bool,X171)) ) | $spl176),
  inference(superposition,[],[f2704,f31623])).
fof(f74208,plain,(
  ( ! [X745,X743,X744,X742] : (s(X744,f(s(bool,f(s(bool,X745))))) != s(X744,g(s(bool,f(s(bool,X745))))) | s(fun(bool,X744),X742) = s(fun(bool,X744),X743) | s(bool,'F_0') = s(bool,sK6(X742,X743,X744,bool)) | s(bool,'T_0') = s(bool,X745)) ) | $spl176),
  inference(superposition,[],[f45,f31623])).
fof(f74207,plain,(
  ( ! [X738,X741,X737,X739,X740] : (~p(s(bool,happ(s(fun(bool,bool),X740),s(bool,f(s(bool,X739)))))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X737,bool),X741),s(fun(bool,bool),X740),s('type.pair.prod'(X737,bool),X738)))) | ~p(s(bool,happ(s(fun(X737,bool),X741),s(X737,'const.pair.FST_1'(s('type.pair.prod'(X737,bool),X738)))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X737,bool),X738))) | s(bool,'T_0') = s(bool,X739)) ) | $spl176),
  inference(superposition,[],[f48,f31623])).
fof(f74204,plain,(
  ( ! [X724,X726,X722,X725,X723] : (~p(s(bool,happ(s(fun(bool,bool),X724),s(bool,f(s(bool,X723)))))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X724),s(fun(bool,bool),X725),s('type.pair.prod'(bool,bool),X722)))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),X722))) | s(bool,X726) = s(bool,happ(s(fun(bool,bool),X725),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X726) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),X722))) | s(bool,'T_0') = s(bool,X723)) ) | $spl176),
  inference(superposition,[],[f2874,f31623])).
fof(f74202,plain,(
  ( ! [X712,X714,X710,X713,X715,X711] : (~p(s(bool,happ(s(fun(bool,bool),X713),s(bool,f(s(bool,X712)))))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X713),s(fun(X710,bool),X714),s('type.pair.prod'(bool,X710),X711)))) | s(bool,X715) = s(bool,happ(s(fun(X710,bool),X714),s(X710,'const.pair.SND_1'(s('type.pair.prod'(bool,X710),X711))))) | s(bool,'T_0') = s(bool,X715) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X710),X711))) | s(bool,'T_0') = s(bool,X712)) ) | $spl176),
  inference(superposition,[],[f786,f31623])).
fof(f74197,plain,(
  ( ! [X693,X691,X692,X694] : (s(bool,f(s(bool,X694))) != s(bool,f(s(X691,sK6(X692,X693,bool,X691)))) | s(fun(X691,bool),X692) = s(fun(X691,bool),X693) | s(bool,'F_0') = s(bool,g(s(X691,sK6(X692,X693,bool,X691)))) | s(bool,'T_0') = s(bool,X694)) ) | $spl176),
  inference(superposition,[],[f45,f31623])).
fof(f74196,plain,(
  ( ! [X681,X677,X679,X680,X682,X676,X678] : (~p(s(bool,happ(s(fun(bool,bool),X678),s(bool,f(s(bool,X677)))))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X678),s(fun(X679,bool),X680),s('type.pair.prod'(bool,X679),'const.pair.,_2'(s(bool,X676),s(X679,X681)))))) | s(bool,X682) = s(bool,happ(s(fun(X679,bool),X680),s(X679,X681))) | s(bool,'T_0') = s(bool,X682) | s(bool,'F_0') = s(bool,X676) | s(bool,'T_0') = s(bool,X677)) ) | $spl176),
  inference(superposition,[],[f2763,f31623])).
fof(f74190,plain,(
  ( ! [X648,X650,X646,X649,X651,X647] : (~p(s(bool,happ(s(fun(bool,bool),X648),s(bool,f(s(bool,X647)))))) | ~p(s(bool,happ(s(fun(X649,bool),X650),s(X649,X651)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X649,bool),X650),s(fun(bool,bool),X648),s('type.pair.prod'(X649,bool),'const.pair.,_2'(s(X649,X651),s(bool,X646)))))) | s(bool,'F_0') = s(bool,X646) | s(bool,'T_0') = s(bool,X647)) ) | $spl176),
  inference(superposition,[],[f785,f31623])).
fof(f74112,plain,(
  ( ! [X383,X385,X384] : (s(bool,X383) != s(bool,f(s(bool,X384))) | s(bool,X383) = s(bool,X385) | s(bool,'F_0') = s(bool,X385) | s(bool,'F_0') = s(bool,f(s(bool,X383))) | s(bool,'T_0') = s(bool,X384)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(superposition,[],[f49037,f31623])).
fof(f73854,plain,(
  ( ! [X304,X306,X305,X307] : (s(bool,X306) = s(bool,f(s(bool,X307))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X304,bool),'const.pair.,_2'(s(X304,X305),s(bool,X306))))) | s(bool,'T_0') = s(bool,X307)) ) | $spl176),
  inference(superposition,[],[f31623,f44])).
fof(f73853,plain,(
  ( ! [X300,X302,X301,X303] : (s(bool,X301) = s(bool,f(s(bool,X303))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X300),'const.pair.,_2'(s(bool,X301),s(X300,X302))))) | s(bool,'T_0') = s(bool,X303)) ) | $spl176),
  inference(superposition,[],[f31623,f43])).
fof(f70206,plain,(
  ( ! [X151,X150] : (s(bool,X150) != s(bool,f(s(bool,'F_0'))) | s(bool,X150) = s(bool,X151) | s(bool,'F_0') = s(bool,X151) | s(bool,'F_0') = s(bool,f(s(bool,X150)))) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(superposition,[],[f49037,f30750])).
fof(f70198,plain,(
  ( ! [X127,X128] : (s(bool,X127) != s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X127) = s(bool,X128) | s(bool,'F_0') = s(bool,X128) | s(bool,'F_0') = s(bool,f(s(bool,X127)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f49037,f25792])).
fof(f70161,plain,(
  ( ! [X300,X302,X301,X299] : (s(bool,X301) != s(bool,f(s(bool,X301))) | s(bool,X302) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X299,bool),'const.pair.,_2'(s(X299,X300),s(bool,X301))))) | s(bool,'F_0') = s(bool,X302)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f49037,f44])).
fof(f70160,plain,(
  ( ! [X296,X298,X295,X297] : (s(bool,X296) != s(bool,f(s(bool,X296))) | s(bool,X298) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X295),'const.pair.,_2'(s(bool,X296),s(X295,X297))))) | s(bool,'F_0') = s(bool,X298)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f49037,f43])).
fof(f69542,plain,(
  ( ! [X138,X139] : (s(bool,'T_0') = s(bool,f(s(bool,X139))) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(bool,bool),X138),s(bool,X139)))))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl290)),
  inference(subsumption_resolution,[],[f69541,f37])).
fof(f69541,plain,(
  ( ! [X138,X139] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,f(s(bool,X139))) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(bool,bool),X138),s(bool,X139)))))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl290)),
  inference(forward_demodulation,[],[f69435,f49268])).
fof(f69435,plain,(
  ( ! [X138,X139] : (~p(s(bool,f(s(bool,'F_0')))) | s(bool,'T_0') = s(bool,f(s(bool,X139))) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(bool,bool),X138),s(bool,X139)))))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl290)),
  inference(superposition,[],[f68414,f49702])).
fof(f69536,plain,(
  ( ! [X101,X102,X100] : (s(bool,'T_0') = s(bool,f(s(bool,X101))) | s(bool,X102) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X102) = s(bool,happ(s(fun(bool,bool),X100),s(bool,X101)))) ) | ($spl4 | $spl9 | $spl36 | $spl290)),
  inference(subsumption_resolution,[],[f69421,f37])).
fof(f69421,plain,(
  ( ! [X101,X102,X100] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,f(s(bool,X101))) | s(bool,X102) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X102) = s(bool,happ(s(fun(bool,bool),X100),s(bool,X101)))) ) | ($spl4 | $spl9 | $spl36 | $spl290)),
  inference(superposition,[],[f68414,f22295])).
fof(f69532,plain,(
  ( ! [X59,X57,X58,X56,X55] : (s(bool,'T_0') = s(bool,f(s(bool,X56))) | s(bool,X57) = s(bool,X58) | s(bool,X58) = s(bool,happ(s(fun(bool,bool),X55),s(bool,X56))) | s(bool,X57) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) ) | ($spl4 | $spl9 | $spl36 | $spl290)),
  inference(subsumption_resolution,[],[f69411,f37])).
fof(f69411,plain,(
  ( ! [X59,X57,X58,X56,X55] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,f(s(bool,X56))) | s(bool,X57) = s(bool,X58) | s(bool,X58) = s(bool,happ(s(fun(bool,bool),X55),s(bool,X56))) | s(bool,X57) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) ) | ($spl4 | $spl9 | $spl36 | $spl290)),
  inference(superposition,[],[f68414,f146])).
fof(f69525,plain,(
  ( ! [X43,X41,X44,X42,X40] : (s(bool,'T_0') = s(bool,f(s(bool,X41))) | s(bool,X42) = s(bool,happ(s(fun(bool,bool),X40),s(bool,X41))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X43,bool),'const.pair.,_2'(s(X43,X44),s(bool,X42)))))) ) | ($spl4 | $spl9 | $spl36 | $spl290)),
  inference(subsumption_resolution,[],[f69408,f37])).
fof(f69408,plain,(
  ( ! [X43,X41,X44,X42,X40] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,f(s(bool,X41))) | s(bool,X42) = s(bool,happ(s(fun(bool,bool),X40),s(bool,X41))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X43,bool),'const.pair.,_2'(s(X43,X44),s(bool,X42)))))) ) | ($spl4 | $spl9 | $spl36 | $spl290)),
  inference(superposition,[],[f68414,f133])).
fof(f69523,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,'T_0') = s(bool,f(s(bool,X26))) | s(bool,X27) = s(bool,happ(s(fun(bool,bool),X25),s(bool,X26))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X28),'const.pair.,_2'(s(bool,X27),s(X28,X29)))))) ) | ($spl4 | $spl9 | $spl36 | $spl290)),
  inference(subsumption_resolution,[],[f69405,f37])).
fof(f69405,plain,(
  ( ! [X28,X26,X29,X27,X25] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,f(s(bool,X26))) | s(bool,X27) = s(bool,happ(s(fun(bool,bool),X25),s(bool,X26))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X28),'const.pair.,_2'(s(bool,X27),s(X28,X29)))))) ) | ($spl4 | $spl9 | $spl36 | $spl290)),
  inference(superposition,[],[f68414,f132])).
fof(f69521,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,f(s(bool,X4))) | s(bool,X5) = s(bool,happ(s(fun(bool,bool),X3),s(bool,X4))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl4 | $spl9 | $spl36 | $spl290)),
  inference(subsumption_resolution,[],[f69399,f37])).
fof(f69399,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,f(s(bool,X4))) | s(bool,X5) = s(bool,happ(s(fun(bool,bool),X3),s(bool,X4))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl4 | $spl9 | $spl36 | $spl290)),
  inference(superposition,[],[f68414,f56])).
fof(f69301,plain,(
  ( ! [X251,X249,X252,X250,X248] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X248,bool),X249),s(fun(bool,bool),X251),s('type.pair.prod'(X248,bool),X250)))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X248,bool),X250))) | s(bool,X252) = s(bool,happ(s(fun(bool,bool),X251),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X252) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(X248,bool),X249),s(X248,'const.pair.FST_1'(s('type.pair.prod'(X248,bool),X250)))))))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(subsumption_resolution,[],[f69300,f37])).
fof(f69300,plain,(
  ( ! [X251,X249,X252,X250,X248] : (~p(s(bool,'T_0')) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X248,bool),X249),s(fun(bool,bool),X251),s('type.pair.prod'(X248,bool),X250)))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X248,bool),X250))) | s(bool,X252) = s(bool,happ(s(fun(bool,bool),X251),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X252) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(X248,bool),X249),s(X248,'const.pair.FST_1'(s('type.pair.prod'(X248,bool),X250)))))))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(forward_demodulation,[],[f69234,f49268])).
fof(f69234,plain,(
  ( ! [X251,X249,X252,X250,X248] : (~p(s(bool,f(s(bool,'F_0')))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X248,bool),X249),s(fun(bool,bool),X251),s('type.pair.prod'(X248,bool),X250)))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X248,bool),X250))) | s(bool,X252) = s(bool,happ(s(fun(bool,bool),X251),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X252) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(X248,bool),X249),s(X248,'const.pair.FST_1'(s('type.pair.prod'(X248,bool),X250)))))))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(superposition,[],[f2874,f49702])).
fof(f69296,plain,(
  ( ! [X189,X187,X190,X188,X186] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X186,bool),X187),s(fun(bool,bool),X189),s('type.pair.prod'(X186,bool),X188)))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X186,bool),X188))) | s(bool,X190) = s(bool,happ(s(fun(bool,bool),X189),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X190) | s(bool,'F_0') = s(bool,happ(s(fun(X186,bool),X187),s(X186,'const.pair.FST_1'(s('type.pair.prod'(X186,bool),X188)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f69295,f37])).
fof(f69295,plain,(
  ( ! [X189,X187,X190,X188,X186] : (~p(s(bool,'T_0')) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X186,bool),X187),s(fun(bool,bool),X189),s('type.pair.prod'(X186,bool),X188)))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X186,bool),X188))) | s(bool,X190) = s(bool,happ(s(fun(bool,bool),X189),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X190) | s(bool,'F_0') = s(bool,happ(s(fun(X186,bool),X187),s(X186,'const.pair.FST_1'(s('type.pair.prod'(X186,bool),X188)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(forward_demodulation,[],[f69223,f1085])).
fof(f69223,plain,(
  ( ! [X189,X187,X190,X188,X186] : (~p(s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X186,bool),X187),s(fun(bool,bool),X189),s('type.pair.prod'(X186,bool),X188)))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X186,bool),X188))) | s(bool,X190) = s(bool,happ(s(fun(bool,bool),X189),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X190) | s(bool,'F_0') = s(bool,happ(s(fun(X186,bool),X187),s(X186,'const.pair.FST_1'(s('type.pair.prod'(X186,bool),X188)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f2874,f25792])).
fof(f69187,plain,(
  ( ! [X229,X227,X230,X228,X226] : (~p(s(bool,happ(s(fun(bool,bool),X227),s(bool,'F_0')))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X227),s(fun(bool,bool),X228),s('type.pair.prod'(bool,bool),X226)))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),X226))) | s(bool,X229) = s(bool,happ(s(fun(bool,bool),X228),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X229) | s(bool,'T_0') = s(bool,f(s(bool,X230))) | s(bool,X230) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),X226)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f2874,f30029])).
fof(f69181,plain,(
  ( ! [X198,X196,X200,X199,X197] : (~p(s(bool,happ(s(fun(bool,bool),X197),s(bool,'F_0')))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X197),s(fun(bool,bool),X198),s('type.pair.prod'(bool,bool),X196)))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),X196))) | s(bool,X199) = s(bool,happ(s(fun(bool,bool),X198),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X199) | s(bool,X200) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),X196))) | s(bool,'F_0') = s(bool,X200)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f2874,f28992])).
fof(f69138,plain,(
  ( ! [X101,X102,X100] : (s(bool,'F_0') = s(bool,X101) | s(bool,X102) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X102) = s(bool,happ(s(fun(bool,bool),X100),s(bool,X101)))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl290)),
  inference(subsumption_resolution,[],[f69031,f37])).
fof(f69031,plain,(
  ( ! [X101,X102,X100] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,X101) | s(bool,X102) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X102) = s(bool,happ(s(fun(bool,bool),X100),s(bool,X101)))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl290)),
  inference(superposition,[],[f68413,f22295])).
fof(f69134,plain,(
  ( ! [X59,X57,X58,X56,X55] : (s(bool,'F_0') = s(bool,X56) | s(bool,X57) = s(bool,X58) | s(bool,X58) = s(bool,happ(s(fun(bool,bool),X55),s(bool,X56))) | s(bool,X57) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl290)),
  inference(subsumption_resolution,[],[f69021,f37])).
fof(f69021,plain,(
  ( ! [X59,X57,X58,X56,X55] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,X56) | s(bool,X57) = s(bool,X58) | s(bool,X58) = s(bool,happ(s(fun(bool,bool),X55),s(bool,X56))) | s(bool,X57) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl290)),
  inference(superposition,[],[f68413,f146])).
fof(f69133,plain,(
  ( ! [X43,X41,X44,X42,X40] : (s(bool,'F_0') = s(bool,X41) | s(bool,X42) = s(bool,happ(s(fun(bool,bool),X40),s(bool,X41))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X43,bool),'const.pair.,_2'(s(X43,X44),s(bool,X42)))))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl290)),
  inference(subsumption_resolution,[],[f69018,f37])).
fof(f69018,plain,(
  ( ! [X43,X41,X44,X42,X40] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,X41) | s(bool,X42) = s(bool,happ(s(fun(bool,bool),X40),s(bool,X41))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X43,bool),'const.pair.,_2'(s(X43,X44),s(bool,X42)))))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl290)),
  inference(superposition,[],[f68413,f133])).
fof(f69131,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,'F_0') = s(bool,X26) | s(bool,X27) = s(bool,happ(s(fun(bool,bool),X25),s(bool,X26))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X28),'const.pair.,_2'(s(bool,X27),s(X28,X29)))))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl290)),
  inference(subsumption_resolution,[],[f69015,f37])).
fof(f69015,plain,(
  ( ! [X28,X26,X29,X27,X25] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,X26) | s(bool,X27) = s(bool,happ(s(fun(bool,bool),X25),s(bool,X26))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X28),'const.pair.,_2'(s(bool,X27),s(X28,X29)))))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl290)),
  inference(superposition,[],[f68413,f132])).
fof(f69129,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') = s(bool,X4) | s(bool,X5) = s(bool,happ(s(fun(bool,bool),X3),s(bool,X4))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl290)),
  inference(subsumption_resolution,[],[f69009,f37])).
fof(f69009,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,X4) | s(bool,X5) = s(bool,happ(s(fun(bool,bool),X3),s(bool,X4))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl290)),
  inference(superposition,[],[f68413,f56])).
fof(f68666,plain,(
  ( ! [X288,X290,X285,X287,X289,X286] : (~p(s(bool,happ(s(fun(X287,bool),X288),s(X287,X289)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X287,bool),X288),s(fun(bool,bool),X285),s('type.pair.prod'(X287,bool),'const.pair.,_2'(s(X287,X289),s(bool,'T_0')))))) | s(bool,X286) = s(bool,X290) | s(bool,'T_0') = s(bool,X290) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(bool,bool),X285),s(bool,X286)))))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(subsumption_resolution,[],[f68665,f37])).
fof(f68665,plain,(
  ( ! [X288,X290,X285,X287,X289,X286] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X287,bool),X288),s(X287,X289)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X287,bool),X288),s(fun(bool,bool),X285),s('type.pair.prod'(X287,bool),'const.pair.,_2'(s(X287,X289),s(bool,'T_0')))))) | s(bool,X286) = s(bool,X290) | s(bool,'T_0') = s(bool,X290) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(bool,bool),X285),s(bool,X286)))))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(forward_demodulation,[],[f68580,f49268])).
fof(f68580,plain,(
  ( ! [X288,X290,X285,X287,X289,X286] : (~p(s(bool,f(s(bool,'F_0')))) | ~p(s(bool,happ(s(fun(X287,bool),X288),s(X287,X289)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X287,bool),X288),s(fun(bool,bool),X285),s('type.pair.prod'(X287,bool),'const.pair.,_2'(s(X287,X289),s(bool,'T_0')))))) | s(bool,X286) = s(bool,X290) | s(bool,'T_0') = s(bool,X290) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(bool,bool),X285),s(bool,X286)))))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(superposition,[],[f2704,f49702])).
fof(f68661,plain,(
  ( ! [X212,X216,X215,X213,X217,X214] : (~p(s(bool,happ(s(fun(X214,bool),X215),s(X214,X216)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X214,bool),X215),s(fun(bool,bool),X212),s('type.pair.prod'(X214,bool),'const.pair.,_2'(s(X214,X216),s(bool,'T_0')))))) | s(bool,X213) = s(bool,X217) | s(bool,'T_0') = s(bool,X217) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X212),s(bool,X213)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f68660,f37])).
fof(f68660,plain,(
  ( ! [X212,X216,X215,X213,X217,X214] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X214,bool),X215),s(X214,X216)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X214,bool),X215),s(fun(bool,bool),X212),s('type.pair.prod'(X214,bool),'const.pair.,_2'(s(X214,X216),s(bool,'T_0')))))) | s(bool,X213) = s(bool,X217) | s(bool,'T_0') = s(bool,X217) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X212),s(bool,X213)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(forward_demodulation,[],[f68569,f1085])).
fof(f68569,plain,(
  ( ! [X212,X216,X215,X213,X217,X214] : (~p(s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) | ~p(s(bool,happ(s(fun(X214,bool),X215),s(X214,X216)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X214,bool),X215),s(fun(bool,bool),X212),s('type.pair.prod'(X214,bool),'const.pair.,_2'(s(X214,X216),s(bool,'T_0')))))) | s(bool,X213) = s(bool,X217) | s(bool,'T_0') = s(bool,X217) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X212),s(bool,X213)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f2704,f25792])).
fof(f68434,plain,(
  ( ! [X76,X77] : (s(bool,X77) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X77) = s(bool,happ(s(fun(bool,bool),X76),s(bool,'T_0')))) ) | ($spl4 | $spl290)),
  inference(subsumption_resolution,[],[f68307,f37])).
fof(f68307,plain,(
  ( ! [X76,X77] : (~p(s(bool,'T_0')) | s(bool,X77) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X77) = s(bool,happ(s(fun(bool,bool),X76),s(bool,'T_0')))) ) | ($spl4 | $spl290)),
  inference(superposition,[],[f67866,f22295])).
fof(f68430,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,X42) = s(bool,X43) | s(bool,X43) = s(bool,happ(s(fun(bool,bool),X41),s(bool,'T_0'))) | s(bool,X42) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) ) | $spl290),
  inference(subsumption_resolution,[],[f68297,f37])).
fof(f68297,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,'T_0')) | s(bool,X42) = s(bool,X43) | s(bool,X43) = s(bool,happ(s(fun(bool,bool),X41),s(bool,'T_0'))) | s(bool,X42) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) ) | $spl290),
  inference(superposition,[],[f67866,f146])).
fof(f68423,plain,(
  ( ! [X30,X31,X29,X32] : (s(bool,X30) = s(bool,happ(s(fun(bool,bool),X29),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X31,bool),'const.pair.,_2'(s(X31,X32),s(bool,X30)))))) ) | $spl290),
  inference(subsumption_resolution,[],[f68294,f37])).
fof(f68294,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,'T_0')) | s(bool,X30) = s(bool,happ(s(fun(bool,bool),X29),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X31,bool),'const.pair.,_2'(s(X31,X32),s(bool,X30)))))) ) | $spl290),
  inference(superposition,[],[f67866,f133])).
fof(f68419,plain,(
  ( ! [X19,X17,X20,X18] : (s(bool,X18) = s(bool,happ(s(fun(bool,bool),X17),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X19),'const.pair.,_2'(s(bool,X18),s(X19,X20)))))) ) | $spl290),
  inference(subsumption_resolution,[],[f68291,f37])).
fof(f68291,plain,(
  ( ! [X19,X17,X20,X18] : (~p(s(bool,'T_0')) | s(bool,X18) = s(bool,happ(s(fun(bool,bool),X17),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X19),'const.pair.,_2'(s(bool,X18),s(X19,X20)))))) ) | $spl290),
  inference(superposition,[],[f67866,f132])).
fof(f68415,plain,(
  ( ! [X2,X1] : (s(bool,X2) = s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl290),
  inference(subsumption_resolution,[],[f68285,f37])).
fof(f68285,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | s(bool,X2) = s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl290),
  inference(superposition,[],[f67866,f56])).
fof(f68256,plain,(
  ( ! [X19,X17,X20,X18,X16] : (~p(s(bool,happ(s(fun(bool,bool),X17),s(bool,X16)))) | s(bool,X18) = s(bool,X19) | s(bool,X16) = s(bool,X19) | s(bool,X18) = s(bool,X20) | s(bool,'T_0') = s(bool,X20)) ) | $spl290),
  inference(superposition,[],[f67866,f146])).
fof(f68252,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X0)))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl290),
  inference(superposition,[],[f67866,f56])).
fof(f67894,plain,(
  ( ! [X288,X290,X285,X287,X289,X286] : (~p(s(bool,happ(s(fun(X286,bool),X287),s(X286,X288)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X286,bool),X287),s(fun(bool,bool),X285),s('type.pair.prod'(X286,bool),'const.pair.,_2'(s(X286,X288),s(bool,X289)))))) | s(bool,X289) = s(bool,X290) | s(bool,'T_0') = s(bool,X290) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(bool,bool),X285),s(bool,'T_0')))))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(subsumption_resolution,[],[f67893,f37])).
fof(f67893,plain,(
  ( ! [X288,X290,X285,X287,X289,X286] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X286,bool),X287),s(X286,X288)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X286,bool),X287),s(fun(bool,bool),X285),s('type.pair.prod'(X286,bool),'const.pair.,_2'(s(X286,X288),s(bool,X289)))))) | s(bool,X289) = s(bool,X290) | s(bool,'T_0') = s(bool,X290) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(bool,bool),X285),s(bool,'T_0')))))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(forward_demodulation,[],[f67846,f49268])).
fof(f67846,plain,(
  ( ! [X288,X290,X285,X287,X289,X286] : (~p(s(bool,f(s(bool,'F_0')))) | ~p(s(bool,happ(s(fun(X286,bool),X287),s(X286,X288)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X286,bool),X287),s(fun(bool,bool),X285),s('type.pair.prod'(X286,bool),'const.pair.,_2'(s(X286,X288),s(bool,X289)))))) | s(bool,X289) = s(bool,X290) | s(bool,'T_0') = s(bool,X290) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(bool,bool),X285),s(bool,'T_0')))))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(superposition,[],[f2699,f49702])).
fof(f67889,plain,(
  ( ! [X212,X216,X215,X213,X217,X214] : (~p(s(bool,happ(s(fun(X213,bool),X214),s(X213,X215)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X213,bool),X214),s(fun(bool,bool),X212),s('type.pair.prod'(X213,bool),'const.pair.,_2'(s(X213,X215),s(bool,X216)))))) | s(bool,X216) = s(bool,X217) | s(bool,'T_0') = s(bool,X217) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X212),s(bool,'T_0')))) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f67888,f37])).
fof(f67888,plain,(
  ( ! [X212,X216,X215,X213,X217,X214] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X213,bool),X214),s(X213,X215)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X213,bool),X214),s(fun(bool,bool),X212),s('type.pair.prod'(X213,bool),'const.pair.,_2'(s(X213,X215),s(bool,X216)))))) | s(bool,X216) = s(bool,X217) | s(bool,'T_0') = s(bool,X217) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X212),s(bool,'T_0')))) ) | ($spl4 | $spl9 | $spl36)),
  inference(forward_demodulation,[],[f67835,f1085])).
fof(f67835,plain,(
  ( ! [X212,X216,X215,X213,X217,X214] : (~p(s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) | ~p(s(bool,happ(s(fun(X213,bool),X214),s(X213,X215)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X213,bool),X214),s(fun(bool,bool),X212),s('type.pair.prod'(X213,bool),'const.pair.,_2'(s(X213,X215),s(bool,X216)))))) | s(bool,X216) = s(bool,X217) | s(bool,'T_0') = s(bool,X217) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X212),s(bool,'T_0')))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f2699,f25792])).
fof(f67873,plain,(
  ( ! [X261,X263,X265,X267,X262,X264,X266] : (~p(s(bool,happ(s(fun(bool,bool),X262),s(bool,X261)))) | ~p(s(bool,happ(s(fun(X263,bool),X264),s(X263,X265)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X263,bool),X264),s(fun(bool,bool),X262),s('type.pair.prod'(X263,bool),'const.pair.,_2'(s(X263,X265),s(bool,X266)))))) | s(bool,X266) = s(bool,X267) | s(bool,'T_0') = s(bool,X267) | s(bool,'T_0') = s(bool,f(s(bool,X261)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f67808,f71])).
fof(f67808,plain,(
  ( ! [X261,X263,X265,X267,X262,X264,X266] : (~p(s(bool,happ(s(fun(bool,bool),X262),s(bool,X261)))) | ~p(s(bool,happ(s(fun(X263,bool),X264),s(X263,X265)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X263,bool),X264),s(fun(bool,bool),X262),s('type.pair.prod'(X263,bool),'const.pair.,_2'(s(X263,X265),s(bool,X266)))))) | s(bool,X266) = s(bool,X267) | s(bool,'T_0') = s(bool,X267) | s(bool,'T_0') = s(bool,f(s(bool,X261))) | s(bool,'T_0') = s(bool,'F_0')) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f2699,f30029])).
fof(f67872,plain,(
  ( ! [X255,X257,X259,X260,X254,X256,X258] : (~p(s(bool,happ(s(fun(bool,bool),X255),s(bool,X254)))) | ~p(s(bool,happ(s(fun(X256,bool),X257),s(X256,X258)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X256,bool),X257),s(fun(bool,bool),X255),s('type.pair.prod'(X256,bool),'const.pair.,_2'(s(X256,X258),s(bool,X259)))))) | s(bool,X259) = s(bool,X260) | s(bool,'T_0') = s(bool,X260) | s(bool,'F_0') = s(bool,X254)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(subsumption_resolution,[],[f67807,f49260])).
fof(f67807,plain,(
  ( ! [X255,X257,X259,X260,X254,X256,X258] : (~p(s(bool,happ(s(fun(bool,bool),X255),s(bool,X254)))) | ~p(s(bool,happ(s(fun(X256,bool),X257),s(X256,X258)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X256,bool),X257),s(fun(bool,bool),X255),s('type.pair.prod'(X256,bool),'const.pair.,_2'(s(X256,X258),s(bool,X259)))))) | s(bool,X259) = s(bool,X260) | s(bool,'T_0') = s(bool,X260) | s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X254)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f2699,f30029])).
fof(f67094,plain,(
  ( ! [X288,X290,X285,X287,X289,X286] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X285,bool),X286),s(fun(X287,bool),X289),s('type.pair.prod'(X285,X287),X288)))) | s(bool,X290) = s(bool,happ(s(fun(X287,bool),X289),s(X287,'const.pair.SND_1'(s('type.pair.prod'(X285,X287),X288))))) | s(bool,'T_0') = s(bool,X290) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(X285,bool),X286),s(X285,'const.pair.FST_1'(s('type.pair.prod'(X285,X287),X288)))))))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(subsumption_resolution,[],[f67093,f37])).
fof(f67093,plain,(
  ( ! [X288,X290,X285,X287,X289,X286] : (~p(s(bool,'T_0')) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X285,bool),X286),s(fun(X287,bool),X289),s('type.pair.prod'(X285,X287),X288)))) | s(bool,X290) = s(bool,happ(s(fun(X287,bool),X289),s(X287,'const.pair.SND_1'(s('type.pair.prod'(X285,X287),X288))))) | s(bool,'T_0') = s(bool,X290) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(X285,bool),X286),s(X285,'const.pair.FST_1'(s('type.pair.prod'(X285,X287),X288)))))))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(forward_demodulation,[],[f67036,f49268])).
fof(f67036,plain,(
  ( ! [X288,X290,X285,X287,X289,X286] : (~p(s(bool,f(s(bool,'F_0')))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X285,bool),X286),s(fun(X287,bool),X289),s('type.pair.prod'(X285,X287),X288)))) | s(bool,X290) = s(bool,happ(s(fun(X287,bool),X289),s(X287,'const.pair.SND_1'(s('type.pair.prod'(X285,X287),X288))))) | s(bool,'T_0') = s(bool,X290) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(X285,bool),X286),s(X285,'const.pair.FST_1'(s('type.pair.prod'(X285,X287),X288)))))))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(superposition,[],[f786,f49702])).
fof(f67089,plain,(
  ( ! [X212,X216,X215,X213,X217,X214] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X212,bool),X213),s(fun(X214,bool),X216),s('type.pair.prod'(X212,X214),X215)))) | s(bool,X217) = s(bool,happ(s(fun(X214,bool),X216),s(X214,'const.pair.SND_1'(s('type.pair.prod'(X212,X214),X215))))) | s(bool,'T_0') = s(bool,X217) | s(bool,'F_0') = s(bool,happ(s(fun(X212,bool),X213),s(X212,'const.pair.FST_1'(s('type.pair.prod'(X212,X214),X215)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f67088,f37])).
fof(f67088,plain,(
  ( ! [X212,X216,X215,X213,X217,X214] : (~p(s(bool,'T_0')) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X212,bool),X213),s(fun(X214,bool),X216),s('type.pair.prod'(X212,X214),X215)))) | s(bool,X217) = s(bool,happ(s(fun(X214,bool),X216),s(X214,'const.pair.SND_1'(s('type.pair.prod'(X212,X214),X215))))) | s(bool,'T_0') = s(bool,X217) | s(bool,'F_0') = s(bool,happ(s(fun(X212,bool),X213),s(X212,'const.pair.FST_1'(s('type.pair.prod'(X212,X214),X215)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(forward_demodulation,[],[f67025,f1085])).
fof(f67025,plain,(
  ( ! [X212,X216,X215,X213,X217,X214] : (~p(s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X212,bool),X213),s(fun(X214,bool),X216),s('type.pair.prod'(X212,X214),X215)))) | s(bool,X217) = s(bool,happ(s(fun(X214,bool),X216),s(X214,'const.pair.SND_1'(s('type.pair.prod'(X212,X214),X215))))) | s(bool,'T_0') = s(bool,X217) | s(bool,'F_0') = s(bool,happ(s(fun(X212,bool),X213),s(X212,'const.pair.FST_1'(s('type.pair.prod'(X212,X214),X215)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f786,f25792])).
fof(f67070,plain,(
  ( ! [X288,X285,X287,X284,X286] : (~p(s(bool,happ(s(fun(bool,bool),X286),s(bool,'T_0')))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X286),s(fun(X284,bool),X287),s('type.pair.prod'(bool,X284),X285)))) | s(bool,X288) = s(bool,happ(s(fun(X284,bool),X287),s(X284,'const.pair.SND_1'(s('type.pair.prod'(bool,X284),X285))))) | s(bool,'T_0') = s(bool,X288) | s(bool,'T_0') = s(bool,f(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X284),X285)))))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(forward_demodulation,[],[f66993,f49268])).
fof(f66993,plain,(
  ( ! [X288,X285,X287,X284,X286] : (~p(s(bool,happ(s(fun(bool,bool),X286),s(bool,f(s(bool,'F_0')))))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X286),s(fun(X284,bool),X287),s('type.pair.prod'(bool,X284),X285)))) | s(bool,X288) = s(bool,happ(s(fun(X284,bool),X287),s(X284,'const.pair.SND_1'(s('type.pair.prod'(bool,X284),X285))))) | s(bool,'T_0') = s(bool,X288) | s(bool,'T_0') = s(bool,f(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X284),X285)))))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(superposition,[],[f786,f49702])).
fof(f66989,plain,(
  ( ! [X263,X265,X267,X268,X264,X266] : (~p(s(bool,happ(s(fun(bool,bool),X265),s(bool,'F_0')))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X265),s(fun(X263,bool),X266),s('type.pair.prod'(bool,X263),X264)))) | s(bool,X267) = s(bool,happ(s(fun(X263,bool),X266),s(X263,'const.pair.SND_1'(s('type.pair.prod'(bool,X263),X264))))) | s(bool,'T_0') = s(bool,X267) | s(bool,'T_0') = s(bool,f(s(bool,X268))) | s(bool,X268) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X263),X264)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f786,f30029])).
fof(f66988,plain,(
  ( ! [X261,X257,X259,X260,X262,X258] : (~p(s(bool,happ(s(fun(bool,bool),X260),s(bool,X259)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X260),s(fun(X257,bool),X261),s('type.pair.prod'(bool,X257),X258)))) | s(bool,X262) = s(bool,happ(s(fun(X257,bool),X261),s(X257,'const.pair.SND_1'(s('type.pair.prod'(bool,X257),X258))))) | s(bool,'T_0') = s(bool,X262) | s(bool,'T_0') = s(bool,f(s(bool,X259))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X257),X258)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f786,f30029])).
fof(f66987,plain,(
  ( ! [X255,X253,X251,X254,X252,X256] : (~p(s(bool,happ(s(fun(bool,bool),X254),s(bool,X253)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X254),s(fun(X251,bool),X255),s('type.pair.prod'(bool,X251),X252)))) | s(bool,X256) = s(bool,happ(s(fun(X251,bool),X255),s(X251,'const.pair.SND_1'(s('type.pair.prod'(bool,X251),X252))))) | s(bool,'T_0') = s(bool,X256) | s(bool,'T_0') = s(bool,f(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X251),X252))))) | s(bool,'F_0') = s(bool,X253)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f786,f30029])).
fof(f66985,plain,(
  ( ! [X243,X241,X239,X244,X242,X240] : (~p(s(bool,happ(s(fun(bool,bool),X242),s(bool,X241)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X242),s(fun(X239,bool),X243),s('type.pair.prod'(bool,X239),X240)))) | s(bool,X244) = s(bool,happ(s(fun(X239,bool),X243),s(X239,'const.pair.SND_1'(s('type.pair.prod'(bool,X239),X240))))) | s(bool,'T_0') = s(bool,X244) | s(bool,'F_0') = s(bool,X241) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X239),X240)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f786,f28992])).
fof(f66983,plain,(
  ( ! [X231,X229,X227,X230,X228,X232] : (~p(s(bool,happ(s(fun(bool,bool),X229),s(bool,'F_0')))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X229),s(fun(X227,bool),X230),s('type.pair.prod'(bool,X227),X228)))) | s(bool,X231) = s(bool,happ(s(fun(X227,bool),X230),s(X227,'const.pair.SND_1'(s('type.pair.prod'(bool,X227),X228))))) | s(bool,'T_0') = s(bool,X231) | s(bool,X232) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X227),X228))) | s(bool,'F_0') = s(bool,X232)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f786,f28992])).
fof(f67066,plain,(
  ( ! [X222,X225,X223,X226,X224] : (~p(s(bool,happ(s(fun(bool,bool),X224),s(bool,'T_0')))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X224),s(fun(X222,bool),X225),s('type.pair.prod'(bool,X222),X223)))) | s(bool,X226) = s(bool,happ(s(fun(X222,bool),X225),s(X222,'const.pair.SND_1'(s('type.pair.prod'(bool,X222),X223))))) | s(bool,'T_0') = s(bool,X226) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X222),X223)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(forward_demodulation,[],[f66982,f1085])).
fof(f66982,plain,(
  ( ! [X222,X225,X223,X226,X224] : (~p(s(bool,happ(s(fun(bool,bool),X224),s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X224),s(fun(X222,bool),X225),s('type.pair.prod'(bool,X222),X223)))) | s(bool,X226) = s(bool,happ(s(fun(X222,bool),X225),s(X222,'const.pair.SND_1'(s('type.pair.prod'(bool,X222),X223))))) | s(bool,'T_0') = s(bool,X226) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X222),X223)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f786,f25792])).
fof(f65993,plain,(
  ( ! [X580,X576,X578,X581,X577,X579] : (p(s(bool,happ(s(fun(bool,bool),X578),s(bool,'F_0')))) | s(bool,X579) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X576,bool),X580),s(fun(bool,bool),X578),s('type.pair.prod'(X576,bool),X577))) | s(bool,'T_0') = s(bool,X579) | s(bool,'T_0') = s(bool,f(s(bool,X581))) | s(bool,X581) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X576,bool),X577)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f550,f30029])).
fof(f65991,plain,(
  ( ! [X569,X571,X567,X568,X570,X566] : (p(s(bool,happ(s(fun(bool,bool),X568),s(bool,'F_0')))) | s(bool,X569) = s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X568),s(fun(X566,bool),X570),s('type.pair.prod'(bool,X566),X567))) | s(bool,'T_0') = s(bool,X569) | s(bool,'T_0') = s(bool,f(s(bool,X571))) | s(bool,X571) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X566),X567)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f475,f30029])).
fof(f65985,plain,(
  ( ! [X540,X542,X536,X538,X541,X537,X539] : (~p(s(bool,happ(s(fun(bool,bool),X537),s(bool,'F_0')))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X537),s(fun(X538,bool),X539),s('type.pair.prod'(bool,X538),'const.pair.,_2'(s(bool,X536),s(X538,X540)))))) | s(bool,X541) = s(bool,happ(s(fun(X538,bool),X539),s(X538,X540))) | s(bool,'T_0') = s(bool,X541) | s(bool,'T_0') = s(bool,f(s(bool,X542))) | s(bool,X536) = s(bool,X542)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f2763,f30029])).
fof(f66589,plain,(
  ( ! [X397,X399,X401,X403,X398,X400,X402] : (s(bool,X401) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X397,bool),X398),s(fun(X399,bool),X402),s('type.pair.prod'(X397,X399),X400))) | s(bool,'T_0') = s(bool,X401) | s(bool,'T_0') = s(bool,f(s(bool,X403))) | s(bool,X403) = s(bool,happ(s(fun(X397,bool),X398),s(X397,'const.pair.FST_1'(s('type.pair.prod'(X397,X399),X400)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f65951,f39])).
fof(f65951,plain,(
  ( ! [X397,X399,X401,X403,X398,X400,X402] : (p(s(bool,'F_0')) | s(bool,X401) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X397,bool),X398),s(fun(X399,bool),X402),s('type.pair.prod'(X397,X399),X400))) | s(bool,'T_0') = s(bool,X401) | s(bool,'T_0') = s(bool,f(s(bool,X403))) | s(bool,X403) = s(bool,happ(s(fun(X397,bool),X398),s(X397,'const.pair.FST_1'(s('type.pair.prod'(X397,X399),X400)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f475,f30029])).
fof(f66588,plain,(
  ( ! [X372,X374,X370,X373,X375,X369,X371] : (s(bool,X373) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X371,bool),X374),s(fun(X369,bool),X370),s('type.pair.prod'(X371,X369),X372))) | s(bool,'T_0') = s(bool,X373) | s(bool,'T_0') = s(bool,f(s(bool,X375))) | s(bool,X375) = s(bool,happ(s(fun(X369,bool),X370),s(X369,'const.pair.SND_1'(s('type.pair.prod'(X371,X369),X372)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f65947,f39])).
fof(f65947,plain,(
  ( ! [X372,X374,X370,X373,X375,X369,X371] : (p(s(bool,'F_0')) | s(bool,X373) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X371,bool),X374),s(fun(X369,bool),X370),s('type.pair.prod'(X371,X369),X372))) | s(bool,'T_0') = s(bool,X373) | s(bool,'T_0') = s(bool,f(s(bool,X375))) | s(bool,X375) = s(bool,happ(s(fun(X369,bool),X370),s(X369,'const.pair.SND_1'(s('type.pair.prod'(X371,X369),X372)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f550,f30029])).
fof(f65830,plain,(
  ( ! [X350,X352,X349,X351,X347,X353,X348] : (~p(s(bool,happ(s(fun(bool,bool),X348),s(bool,X347)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X348),s(fun(X349,bool),X350),s('type.pair.prod'(bool,X349),'const.pair.,_2'(s(bool,'F_0'),s(X349,X351)))))) | s(bool,X352) = s(bool,happ(s(fun(X349,bool),X350),s(X349,X351))) | s(bool,'T_0') = s(bool,X352) | s(bool,'T_0') = s(bool,f(s(bool,X353))) | s(bool,X347) = s(bool,X353)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f2763,f30029])).
fof(f65754,plain,(
  ( ! [X39,X37,X38] : (s(bool,X38) = s(bool,f(s(bool,X37))) | s(bool,'T_0') = s(bool,f(s(bool,X38))) | s(bool,'T_0') = s(bool,f(s(bool,X39))) | s(bool,X37) = s(bool,X39)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(superposition,[],[f49702,f30029])).
fof(f65736,plain,(
  ( ! [X677,X673,X675,X676,X672,X674] : (p(s(bool,happ(s(fun(bool,bool),X675),s(bool,X674)))) | s(bool,X676) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X672,bool),X677),s(fun(bool,bool),X675),s('type.pair.prod'(X672,bool),X673))) | s(bool,'T_0') = s(bool,X676) | s(bool,'T_0') = s(bool,f(s(bool,X674))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X672,bool),X673)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f550,f30029])).
fof(f65734,plain,(
  ( ! [X666,X662,X664,X667,X663,X665] : (p(s(bool,happ(s(fun(bool,bool),X665),s(bool,X664)))) | s(bool,X666) = s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X665),s(fun(X662,bool),X667),s('type.pair.prod'(bool,X662),X663))) | s(bool,'T_0') = s(bool,X666) | s(bool,'T_0') = s(bool,f(s(bool,X664))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X662),X663)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f475,f30029])).
fof(f65728,plain,(
  ( ! [X633,X629,X631,X632,X634,X628,X630] : (~p(s(bool,happ(s(fun(bool,bool),X630),s(bool,X629)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X630),s(fun(X631,bool),X632),s('type.pair.prod'(bool,X631),'const.pair.,_2'(s(bool,X628),s(X631,X633)))))) | s(bool,X634) = s(bool,happ(s(fun(X631,bool),X632),s(X631,X633))) | s(bool,'T_0') = s(bool,X634) | s(bool,'T_0') = s(bool,f(s(bool,X629))) | s(bool,'F_0') = s(bool,X628)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f2763,f30029])).
fof(f65529,plain,(
  ( ! [X633,X629,X631,X632,X634,X628,X630] : (~p(s(bool,happ(s(fun(bool,bool),X630),s(bool,X629)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X630),s(fun(X631,bool),X632),s('type.pair.prod'(bool,X631),'const.pair.,_2'(s(bool,X628),s(X631,X633)))))) | s(bool,X634) = s(bool,happ(s(fun(X631,bool),X632),s(X631,X633))) | s(bool,'T_0') = s(bool,X634) | s(bool,'T_0') = s(bool,f(s(bool,X628))) | s(bool,'F_0') = s(bool,X629)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f2763,f30029])).
fof(f66452,plain,(
  ( ! [X350,X349] : (s(bool,X350) != s(bool,f(s(bool,X349))) | s(bool,'T_0') = s(bool,f(s(bool,X350))) | s(bool,'F_0') = s(bool,X349)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(subsumption_resolution,[],[f65445,f49260])).
fof(f65445,plain,(
  ( ! [X350,X349] : (s(bool,X350) != s(bool,f(s(bool,X349))) | s(bool,'T_0') = s(bool,f(s(bool,X350))) | s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X349)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(superposition,[],[f64321,f30029])).
fof(f65344,plain,(
  ( ! [X356,X358,X360,X362,X357,X359,X361] : (~p(s(bool,happ(s(fun(bool,bool),X357),s(bool,'T_0')))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X357),s(fun(X358,bool),X359),s('type.pair.prod'(bool,X358),'const.pair.,_2'(s(bool,f(s(bool,X356))),s(X358,X360)))))) | s(bool,X361) = s(bool,happ(s(fun(X358,bool),X359),s(X358,X360))) | s(bool,'T_0') = s(bool,X361) | s(bool,X356) = s(bool,X362) | s(bool,'F_0') = s(bool,X362)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f2763,f30029])).
fof(f64839,plain,(
  ( ! [X325,X327,X323,X324,X326,X322,X328] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X322,bool),X323),s(fun(X325,bool),X326),s('type.pair.prod'(X322,X325),'const.pair.,_2'(s(X322,X324),s(X325,X327)))))) | s(bool,X328) = s(bool,happ(s(fun(X325,bool),X326),s(X325,X327))) | s(bool,'T_0') = s(bool,X328) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(X322,bool),X323),s(X322,X324)))))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(subsumption_resolution,[],[f64838,f37])).
fof(f64838,plain,(
  ( ! [X325,X327,X323,X324,X326,X322,X328] : (~p(s(bool,'T_0')) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X322,bool),X323),s(fun(X325,bool),X326),s('type.pair.prod'(X322,X325),'const.pair.,_2'(s(X322,X324),s(X325,X327)))))) | s(bool,X328) = s(bool,happ(s(fun(X325,bool),X326),s(X325,X327))) | s(bool,'T_0') = s(bool,X328) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(X322,bool),X323),s(X322,X324)))))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(forward_demodulation,[],[f64750,f49268])).
fof(f64750,plain,(
  ( ! [X325,X327,X323,X324,X326,X322,X328] : (~p(s(bool,f(s(bool,'F_0')))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X322,bool),X323),s(fun(X325,bool),X326),s('type.pair.prod'(X322,X325),'const.pair.,_2'(s(X322,X324),s(X325,X327)))))) | s(bool,X328) = s(bool,happ(s(fun(X325,bool),X326),s(X325,X327))) | s(bool,'T_0') = s(bool,X328) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(X322,bool),X323),s(X322,X324)))))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(superposition,[],[f2763,f49702])).
fof(f64836,plain,(
  ( ! [X263,X265,X267,X262,X268,X264,X266] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X262,bool),X263),s(fun(X265,bool),X266),s('type.pair.prod'(X262,X265),'const.pair.,_2'(s(X262,X264),s(X265,X267)))))) | s(bool,X268) = s(bool,happ(s(fun(X265,bool),X266),s(X265,X267))) | s(bool,'T_0') = s(bool,X268) | s(bool,'F_0') = s(bool,happ(s(fun(X262,bool),X263),s(X262,X264)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f64835,f37])).
fof(f64835,plain,(
  ( ! [X263,X265,X267,X262,X268,X264,X266] : (~p(s(bool,'T_0')) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X262,bool),X263),s(fun(X265,bool),X266),s('type.pair.prod'(X262,X265),'const.pair.,_2'(s(X262,X264),s(X265,X267)))))) | s(bool,X268) = s(bool,happ(s(fun(X265,bool),X266),s(X265,X267))) | s(bool,'T_0') = s(bool,X268) | s(bool,'F_0') = s(bool,happ(s(fun(X262,bool),X263),s(X262,X264)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(forward_demodulation,[],[f64742,f1085])).
fof(f64742,plain,(
  ( ! [X263,X265,X267,X262,X268,X264,X266] : (~p(s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X262,bool),X263),s(fun(X265,bool),X266),s('type.pair.prod'(X262,X265),'const.pair.,_2'(s(X262,X264),s(X265,X267)))))) | s(bool,X268) = s(bool,happ(s(fun(X265,bool),X266),s(X265,X267))) | s(bool,'T_0') = s(bool,X268) | s(bool,'F_0') = s(bool,happ(s(fun(X262,bool),X263),s(X262,X264)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f2763,f25792])).
fof(f64783,plain,(
  ( ! [X397,X399,X395,X396,X398,X394] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X395),s(fun(X396,bool),X397),s('type.pair.prod'(bool,X396),'const.pair.,_2'(s(bool,'T_0'),s(X396,X398)))))) | ~p(s(bool,happ(s(fun(bool,bool),X395),s(bool,X394)))) | s(bool,X399) = s(bool,happ(s(fun(X396,bool),X397),s(X396,X398))) | s(bool,'T_0') = s(bool,X399) | s(bool,'F_0') = s(bool,X394)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(forward_demodulation,[],[f64682,f49268])).
fof(f64682,plain,(
  ( ! [X397,X399,X395,X396,X398,X394] : (~p(s(bool,happ(s(fun(bool,bool),X395),s(bool,X394)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X395),s(fun(X396,bool),X397),s('type.pair.prod'(bool,X396),'const.pair.,_2'(s(bool,f(s(bool,'F_0'))),s(X396,X398)))))) | s(bool,X399) = s(bool,happ(s(fun(X396,bool),X397),s(X396,X398))) | s(bool,'T_0') = s(bool,X399) | s(bool,'F_0') = s(bool,X394)) ) | $spl176),
  inference(superposition,[],[f2763,f30750])).
fof(f64782,plain,(
  ( ! [X389,X391,X393,X388,X390,X392] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X389),s(fun(X390,bool),X391),s('type.pair.prod'(bool,X390),'const.pair.,_2'(s(bool,'T_0'),s(X390,X392)))))) | ~p(s(bool,happ(s(fun(bool,bool),X389),s(bool,X388)))) | s(bool,X393) = s(bool,happ(s(fun(X390,bool),X391),s(X390,X392))) | s(bool,'T_0') = s(bool,X393) | s(bool,'T_0') = s(bool,f(s(bool,X388)))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(forward_demodulation,[],[f64681,f49268])).
fof(f64681,plain,(
  ( ! [X389,X391,X393,X388,X390,X392] : (~p(s(bool,happ(s(fun(bool,bool),X389),s(bool,X388)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X389),s(fun(X390,bool),X391),s('type.pair.prod'(bool,X390),'const.pair.,_2'(s(bool,f(s(bool,'F_0'))),s(X390,X392)))))) | s(bool,X393) = s(bool,happ(s(fun(X390,bool),X391),s(X390,X392))) | s(bool,'T_0') = s(bool,X393) | s(bool,'T_0') = s(bool,f(s(bool,X388)))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(superposition,[],[f2763,f49702])).
fof(f64673,plain,(
  ( ! [X344,X341,X343,X339,X345,X340,X342] : (~p(s(bool,happ(s(fun(bool,bool),X340),s(bool,X339)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X340),s(fun(X341,bool),X342),s('type.pair.prod'(bool,X341),'const.pair.,_2'(s(bool,'F_0'),s(X341,X343)))))) | s(bool,X344) = s(bool,happ(s(fun(X341,bool),X342),s(X341,X343))) | s(bool,'T_0') = s(bool,X344) | s(bool,X339) = s(bool,X345) | s(bool,'F_0') = s(bool,X345)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f2763,f28992])).
fof(f64772,plain,(
  ( ! [X269,X265,X267,X268,X264,X266] : (~p(s(bool,happ(s(fun(bool,bool),X265),s(bool,'T_0')))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X265),s(fun(X266,bool),X267),s('type.pair.prod'(bool,X266),'const.pair.,_2'(s(bool,X264),s(X266,X268)))))) | s(bool,X269) = s(bool,happ(s(fun(X266,bool),X267),s(X266,X268))) | s(bool,'T_0') = s(bool,X269) | s(bool,'T_0') = s(bool,f(s(bool,X264)))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(forward_demodulation,[],[f64663,f49268])).
fof(f64663,plain,(
  ( ! [X269,X265,X267,X268,X264,X266] : (~p(s(bool,happ(s(fun(bool,bool),X265),s(bool,f(s(bool,'F_0')))))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X265),s(fun(X266,bool),X267),s('type.pair.prod'(bool,X266),'const.pair.,_2'(s(bool,X264),s(X266,X268)))))) | s(bool,X269) = s(bool,happ(s(fun(X266,bool),X267),s(X266,X268))) | s(bool,'T_0') = s(bool,X269) | s(bool,'T_0') = s(bool,f(s(bool,X264)))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(superposition,[],[f2763,f49702])).
fof(f64658,plain,(
  ( ! [X237,X235,X233,X238,X236,X234,X232] : (~p(s(bool,happ(s(fun(bool,bool),X234),s(bool,X233)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X234),s(fun(X235,bool),X236),s('type.pair.prod'(bool,X235),'const.pair.,_2'(s(bool,X232),s(X235,X237)))))) | s(bool,X238) = s(bool,happ(s(fun(X235,bool),X236),s(X235,X237))) | s(bool,'T_0') = s(bool,X238) | s(bool,'F_0') = s(bool,X233) | s(bool,'F_0') = s(bool,X232)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f2763,f28992])).
fof(f64656,plain,(
  ( ! [X218,X222,X220,X219,X223,X221,X224] : (~p(s(bool,happ(s(fun(bool,bool),X219),s(bool,'F_0')))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X219),s(fun(X220,bool),X221),s('type.pair.prod'(bool,X220),'const.pair.,_2'(s(bool,X218),s(X220,X222)))))) | s(bool,X223) = s(bool,happ(s(fun(X220,bool),X221),s(X220,X222))) | s(bool,'T_0') = s(bool,X223) | s(bool,X218) = s(bool,X224) | s(bool,'F_0') = s(bool,X224)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f2763,f28992])).
fof(f64768,plain,(
  ( ! [X212,X216,X215,X213,X217,X214] : (~p(s(bool,happ(s(fun(bool,bool),X213),s(bool,'T_0')))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X213),s(fun(X214,bool),X215),s('type.pair.prod'(bool,X214),'const.pair.,_2'(s(bool,X212),s(X214,X216)))))) | s(bool,X217) = s(bool,happ(s(fun(X214,bool),X215),s(X214,X216))) | s(bool,'T_0') = s(bool,X217) | s(bool,'F_0') = s(bool,X212)) ) | ($spl4 | $spl9 | $spl36)),
  inference(forward_demodulation,[],[f64655,f1085])).
fof(f64655,plain,(
  ( ! [X212,X216,X215,X213,X217,X214] : (~p(s(bool,happ(s(fun(bool,bool),X213),s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X213),s(fun(X214,bool),X215),s('type.pair.prod'(bool,X214),'const.pair.,_2'(s(bool,X212),s(X214,X216)))))) | s(bool,X217) = s(bool,happ(s(fun(X214,bool),X215),s(X214,X216))) | s(bool,'T_0') = s(bool,X217) | s(bool,'F_0') = s(bool,X212)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f2763,f25792])).
fof(f64463,plain,(
  ( ! [X28,X26,X27] : (s(bool,X27) != s(bool,f(s(bool,X26))) | s(bool,'T_0') = s(bool,f(s(bool,X27))) | s(bool,X28) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X26) = s(bool,X28)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(superposition,[],[f64321,f22295])).
fof(f64461,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s(bool,X17) != s(bool,f(s(bool,X16))) | s(bool,'T_0') = s(bool,f(s(bool,X17))) | s(bool,X18) = s(bool,X19) | s(bool,X16) = s(bool,X19) | s(bool,X18) = s(bool,X20) | s(bool,'T_0') = s(bool,X20)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(superposition,[],[f64321,f146])).
fof(f64460,plain,(
  ( ! [X14,X12,X15,X13,X11] : (s(bool,X12) != s(bool,f(s(bool,X11))) | s(bool,'T_0') = s(bool,f(s(bool,X12))) | s(bool,X11) = s(bool,X13) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X14,bool),'const.pair.,_2'(s(X14,X15),s(bool,X13)))))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(superposition,[],[f64321,f133])).
fof(f64459,plain,(
  ( ! [X6,X10,X8,X7,X9] : (s(bool,X7) != s(bool,f(s(bool,X6))) | s(bool,'T_0') = s(bool,f(s(bool,X7))) | s(bool,X6) = s(bool,X8) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X9),'const.pair.,_2'(s(bool,X8),s(X9,X10)))))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(superposition,[],[f64321,f132])).
fof(f64457,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,f(s(bool,X1))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(superposition,[],[f64321,f56])).
fof(f64321,plain,(
  ( ! [X100] : (s(bool,X100) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X100)))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(duplicate_literal_removal,[],[f64317])).
fof(f64317,plain,(
  ( ! [X100] : (s(bool,X100) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X100))) | s(bool,'T_0') = s(bool,f(s(bool,X100)))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(superposition,[],[f63838,f41609])).
fof(f64324,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,X76) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X77,bool),'const.pair.,_2'(s(X77,X78),s(bool,X79))))) | s(bool,'T_0') = s(bool,f(s(bool,X76))) | s(bool,'T_0') = s(bool,X79)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(duplicate_literal_removal,[],[f64304])).
fof(f64304,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,X76) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X77,bool),'const.pair.,_2'(s(X77,X78),s(bool,X79))))) | s(bool,'T_0') = s(bool,f(s(bool,X76))) | s(bool,'T_0') = s(bool,X79) | s(bool,'T_0') = s(bool,f(s(bool,X76)))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(superposition,[],[f63838,f157])).
fof(f64325,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,X72) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X73),'const.pair.,_2'(s(bool,X74),s(X73,X75))))) | s(bool,'T_0') = s(bool,f(s(bool,X72))) | s(bool,'T_0') = s(bool,X74)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(duplicate_literal_removal,[],[f64303])).
fof(f64303,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,X72) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X73),'const.pair.,_2'(s(bool,X74),s(X73,X75))))) | s(bool,'T_0') = s(bool,f(s(bool,X72))) | s(bool,'T_0') = s(bool,X74) | s(bool,'T_0') = s(bool,f(s(bool,X72)))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(superposition,[],[f63838,f156])).
fof(f64188,plain,(
  ( ! [X218,X222,X220,X219,X223,X221,X224] : (s(bool,X222) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X220,bool),X223),s(fun(X218,bool),X219),s('type.pair.prod'(X220,X218),X221))) | s(bool,'T_0') = s(bool,X222) | s(bool,X224) = s(bool,happ(s(fun(X218,bool),X219),s(X218,'const.pair.SND_1'(s('type.pair.prod'(X220,X218),X221))))) | s(bool,'F_0') = s(bool,X224)) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f64131,f39])).
fof(f64131,plain,(
  ( ! [X218,X222,X220,X219,X223,X221,X224] : (p(s(bool,'F_0')) | s(bool,X222) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X220,bool),X223),s(fun(X218,bool),X219),s('type.pair.prod'(X220,X218),X221))) | s(bool,'T_0') = s(bool,X222) | s(bool,X224) = s(bool,happ(s(fun(X218,bool),X219),s(X218,'const.pair.SND_1'(s('type.pair.prod'(X220,X218),X221))))) | s(bool,'F_0') = s(bool,X224)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f550,f28992])).
fof(f64165,plain,(
  ( ! [X269,X271,X268,X270,X272] : (p(s(bool,happ(s(fun(bool,bool),X270),s(bool,'T_0')))) | s(bool,X271) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X268,bool),X272),s(fun(bool,bool),X270),s('type.pair.prod'(X268,bool),X269))) | s(bool,'T_0') = s(bool,X271) | s(bool,'T_0') = s(bool,f(s(bool,'const.pair.SND_1'(s('type.pair.prod'(X268,bool),X269)))))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(forward_demodulation,[],[f64098,f49268])).
fof(f64098,plain,(
  ( ! [X269,X271,X268,X270,X272] : (p(s(bool,happ(s(fun(bool,bool),X270),s(bool,f(s(bool,'F_0')))))) | s(bool,X271) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X268,bool),X272),s(fun(bool,bool),X270),s('type.pair.prod'(X268,bool),X269))) | s(bool,'T_0') = s(bool,X271) | s(bool,'T_0') = s(bool,f(s(bool,'const.pair.SND_1'(s('type.pair.prod'(X268,bool),X269)))))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(superposition,[],[f550,f49702])).
fof(f64093,plain,(
  ( ! [X243,X241,X246,X244,X242,X245] : (p(s(bool,happ(s(fun(bool,bool),X244),s(bool,X243)))) | s(bool,X245) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X241,bool),X246),s(fun(bool,bool),X244),s('type.pair.prod'(X241,bool),X242))) | s(bool,'T_0') = s(bool,X245) | s(bool,'F_0') = s(bool,X243) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X241,bool),X242)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f550,f28992])).
fof(f64091,plain,(
  ( ! [X231,X229,X233,X230,X234,X232] : (p(s(bool,happ(s(fun(bool,bool),X231),s(bool,'F_0')))) | s(bool,X232) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X229,bool),X233),s(fun(bool,bool),X231),s('type.pair.prod'(X229,bool),X230))) | s(bool,'T_0') = s(bool,X232) | s(bool,X234) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X229,bool),X230))) | s(bool,'F_0') = s(bool,X234)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f550,f28992])).
fof(f64161,plain,(
  ( ! [X227,X225,X228,X226,X224] : (p(s(bool,happ(s(fun(bool,bool),X226),s(bool,'T_0')))) | s(bool,X227) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X224,bool),X228),s(fun(bool,bool),X226),s('type.pair.prod'(X224,bool),X225))) | s(bool,'T_0') = s(bool,X227) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X224,bool),X225)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(forward_demodulation,[],[f64090,f1085])).
fof(f64090,plain,(
  ( ! [X227,X225,X228,X226,X224] : (p(s(bool,happ(s(fun(bool,bool),X226),s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))))) | s(bool,X227) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X224,bool),X228),s(fun(bool,bool),X226),s('type.pair.prod'(X224,bool),X225))) | s(bool,'T_0') = s(bool,X227) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X224,bool),X225)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f550,f25792])).
fof(f63869,plain,(
  ( ! [X116,X117,X115] : (s(bool,X115) = s(bool,X116) | s(bool,'T_0') = s(bool,X116) | s(bool,X117) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X117) = s(bool,f(s(bool,X115)))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(subsumption_resolution,[],[f63764,f56])).
fof(f63764,plain,(
  ( ! [X116,X117,X115] : (s(bool,'T_0') != s(bool,X115) | s(bool,X115) = s(bool,X116) | s(bool,'T_0') = s(bool,X116) | s(bool,X117) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X117) = s(bool,f(s(bool,X115)))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(superposition,[],[f50029,f22295])).
fof(f63865,plain,(
  ( ! [X109,X107,X105,X108,X106] : (s(bool,X105) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X106,bool),'const.pair.,_2'(s(X106,X107),s(bool,X108))))) | s(bool,X105) = s(bool,X109) | s(bool,'T_0') = s(bool,X108) | s(bool,'T_0') = s(bool,f(s(bool,X105)))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(subsumption_resolution,[],[f63761,f36506])).
fof(f63761,plain,(
  ( ! [X109,X107,X105,X108,X106] : (s(bool,X105) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X106,bool),'const.pair.,_2'(s(X106,X107),s(bool,X108))))) | s(bool,X105) = s(bool,X109) | s(bool,'T_0') = s(bool,X109) | s(bool,'T_0') = s(bool,X108) | s(bool,'T_0') = s(bool,f(s(bool,X105)))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(superposition,[],[f50029,f157])).
fof(f63864,plain,(
  ( ! [X103,X101,X102,X100,X104] : (s(bool,X100) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X101),'const.pair.,_2'(s(bool,X102),s(X101,X103))))) | s(bool,X100) = s(bool,X104) | s(bool,'T_0') = s(bool,X102) | s(bool,'T_0') = s(bool,f(s(bool,X100)))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(subsumption_resolution,[],[f63760,f36506])).
fof(f63760,plain,(
  ( ! [X103,X101,X102,X100,X104] : (s(bool,X100) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X101),'const.pair.,_2'(s(bool,X102),s(X101,X103))))) | s(bool,X100) = s(bool,X104) | s(bool,'T_0') = s(bool,X104) | s(bool,'T_0') = s(bool,X102) | s(bool,'T_0') = s(bool,f(s(bool,X100)))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(superposition,[],[f50029,f156])).
fof(f63862,plain,(
  ( ! [X70,X74,X72,X71,X73] : (s(bool,X70) = s(bool,X71) | s(bool,'T_0') = s(bool,X71) | s(bool,X72) = s(bool,X73) | s(bool,X73) = s(bool,f(s(bool,X70))) | s(bool,X72) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(subsumption_resolution,[],[f63754,f56])).
fof(f63754,plain,(
  ( ! [X70,X74,X72,X71,X73] : (s(bool,'T_0') != s(bool,X70) | s(bool,X70) = s(bool,X71) | s(bool,'T_0') = s(bool,X71) | s(bool,X72) = s(bool,X73) | s(bool,X73) = s(bool,f(s(bool,X70))) | s(bool,X72) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(superposition,[],[f50029,f146])).
fof(f63855,plain,(
  ( ! [X59,X57,X58,X56,X55] : (s(bool,X55) = s(bool,X56) | s(bool,'T_0') = s(bool,X56) | s(bool,X57) = s(bool,f(s(bool,X55))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X58,bool),'const.pair.,_2'(s(X58,X59),s(bool,X57)))))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(subsumption_resolution,[],[f63751,f56])).
fof(f63751,plain,(
  ( ! [X59,X57,X58,X56,X55] : (s(bool,'T_0') != s(bool,X55) | s(bool,X55) = s(bool,X56) | s(bool,'T_0') = s(bool,X56) | s(bool,X57) = s(bool,f(s(bool,X55))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X58,bool),'const.pair.,_2'(s(X58,X59),s(bool,X57)))))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(superposition,[],[f50029,f133])).
fof(f63853,plain,(
  ( ! [X43,X41,X44,X42,X40] : (s(bool,X40) = s(bool,X41) | s(bool,'T_0') = s(bool,X41) | s(bool,X42) = s(bool,f(s(bool,X40))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X43),'const.pair.,_2'(s(bool,X42),s(X43,X44)))))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(subsumption_resolution,[],[f63748,f56])).
fof(f63748,plain,(
  ( ! [X43,X41,X44,X42,X40] : (s(bool,'T_0') != s(bool,X40) | s(bool,X40) = s(bool,X41) | s(bool,'T_0') = s(bool,X41) | s(bool,X42) = s(bool,f(s(bool,X40))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X43),'const.pair.,_2'(s(bool,X42),s(X43,X44)))))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(superposition,[],[f50029,f132])).
fof(f63731,plain,(
  ( ! [X271,X273,X270,X272] : (s(bool,X272) != s(bool,f(s(bool,X272))) | s(bool,X273) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X270,bool),'const.pair.,_2'(s(X270,X271),s(bool,X272))))) | s(bool,'T_0') = s(bool,X273)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(superposition,[],[f50029,f44])).
fof(f63730,plain,(
  ( ! [X269,X267,X268,X266] : (s(bool,X267) != s(bool,f(s(bool,X267))) | s(bool,X269) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X266),'const.pair.,_2'(s(bool,X267),s(X266,X268))))) | s(bool,'T_0') = s(bool,X269)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(superposition,[],[f50029,f43])).
fof(f63838,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,f(s(bool,X0)))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(condensation,[],[f63837])).
fof(f63837,plain,(
  ( ! [X180,X181] : (s(bool,X180) != s(bool,f(s(bool,X180))) | s(bool,'T_0') = s(bool,X181) | s(bool,'T_0') = s(bool,f(s(bool,X180)))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(subsumption_resolution,[],[f63704,f25786])).
fof(f63704,plain,(
  ( ! [X180,X181] : (s(bool,X180) != s(bool,f(s(bool,X180))) | s(bool,X181) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X181) | s(bool,'T_0') = s(bool,f(s(bool,X180)))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(superposition,[],[f50029,f49702])).
fof(f63641,plain,(
  ( ! [X218,X222,X220,X219,X223,X221,X224] : (s(bool,X222) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X218,bool),X219),s(fun(X220,bool),X223),s('type.pair.prod'(X218,X220),X221))) | s(bool,'T_0') = s(bool,X222) | s(bool,X224) = s(bool,happ(s(fun(X218,bool),X219),s(X218,'const.pair.FST_1'(s('type.pair.prod'(X218,X220),X221))))) | s(bool,'F_0') = s(bool,X224)) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f63584,f39])).
fof(f63584,plain,(
  ( ! [X218,X222,X220,X219,X223,X221,X224] : (p(s(bool,'F_0')) | s(bool,X222) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X218,bool),X219),s(fun(X220,bool),X223),s('type.pair.prod'(X218,X220),X221))) | s(bool,'T_0') = s(bool,X222) | s(bool,X224) = s(bool,happ(s(fun(X218,bool),X219),s(X218,'const.pair.FST_1'(s('type.pair.prod'(X218,X220),X221))))) | s(bool,'F_0') = s(bool,X224)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f475,f28992])).
fof(f63618,plain,(
  ( ! [X269,X267,X268,X270,X266] : (p(s(bool,happ(s(fun(bool,bool),X268),s(bool,'T_0')))) | s(bool,X269) = s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X268),s(fun(X266,bool),X270),s('type.pair.prod'(bool,X266),X267))) | s(bool,'T_0') = s(bool,X269) | s(bool,'T_0') = s(bool,f(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X266),X267)))))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(forward_demodulation,[],[f63551,f49268])).
fof(f63551,plain,(
  ( ! [X269,X267,X268,X270,X266] : (p(s(bool,happ(s(fun(bool,bool),X268),s(bool,f(s(bool,'F_0')))))) | s(bool,X269) = s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X268),s(fun(X266,bool),X270),s('type.pair.prod'(bool,X266),X267))) | s(bool,'T_0') = s(bool,X269) | s(bool,'T_0') = s(bool,f(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X266),X267)))))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(superposition,[],[f475,f49702])).
fof(f63546,plain,(
  ( ! [X243,X241,X239,X244,X242,X240] : (p(s(bool,happ(s(fun(bool,bool),X242),s(bool,X241)))) | s(bool,X243) = s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X242),s(fun(X239,bool),X244),s('type.pair.prod'(bool,X239),X240))) | s(bool,'T_0') = s(bool,X243) | s(bool,'F_0') = s(bool,X241) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X239),X240)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f475,f28992])).
fof(f63544,plain,(
  ( ! [X231,X229,X227,X230,X228,X232] : (p(s(bool,happ(s(fun(bool,bool),X229),s(bool,'F_0')))) | s(bool,X230) = s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X229),s(fun(X227,bool),X231),s('type.pair.prod'(bool,X227),X228))) | s(bool,'T_0') = s(bool,X230) | s(bool,X232) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X227),X228))) | s(bool,'F_0') = s(bool,X232)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f475,f28992])).
fof(f63614,plain,(
  ( ! [X222,X225,X223,X226,X224] : (p(s(bool,happ(s(fun(bool,bool),X224),s(bool,'T_0')))) | s(bool,X225) = s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X224),s(fun(X222,bool),X226),s('type.pair.prod'(bool,X222),X223))) | s(bool,'T_0') = s(bool,X225) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X222),X223)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(forward_demodulation,[],[f63543,f1085])).
fof(f63543,plain,(
  ( ! [X222,X225,X223,X226,X224] : (p(s(bool,happ(s(fun(bool,bool),X224),s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))))) | s(bool,X225) = s(bool,'const.pair.PROD_ALL_3'(s(fun(bool,bool),X224),s(fun(X222,bool),X226),s('type.pair.prod'(bool,X222),X223))) | s(bool,'T_0') = s(bool,X225) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X222),X223)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f475,f25792])).
fof(f60404,plain,(
  ( ! [X165,X166] : (s(bool,'T_0') = s(bool,f(s(bool,X166))) | s(bool,X166) = s(bool,f(s(bool,X165))) | s(bool,X165) = s(bool,f(s(bool,'F_0')))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(subsumption_resolution,[],[f60325,f37])).
fof(f60325,plain,(
  ( ! [X165,X166] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,f(s(bool,X166))) | s(bool,X166) = s(bool,f(s(bool,X165))) | s(bool,X165) = s(bool,f(s(bool,'F_0')))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(superposition,[],[f36504,f49702])).
fof(f11201,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,X16) != s(bool,f(s(bool,X15))) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X17),'const.pair.,_2'(s(bool,X15),s(X17,X18)))))) ) | ($spl9 | $spl36)),
  inference(subsumption_resolution,[],[f11169,f71])).
fof(f11169,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,X16) != s(bool,f(s(bool,X15))) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X17),'const.pair.,_2'(s(bool,X15),s(X17,X18)))))) ) | $spl36),
  inference(superposition,[],[f2455,f132])).
fof(f13853,plain,(
  ( ! [X308,X306,X309,X307] : (s(bool,X307) != s(bool,f(s(bool,X306))) | s(bool,'T_0') = s(bool,X307) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X308,bool),'const.pair.,_2'(s(X308,X309),s(bool,X306)))))) ) | ($spl9 | $spl36)),
  inference(subsumption_resolution,[],[f13097,f71])).
fof(f13097,plain,(
  ( ! [X308,X306,X309,X307] : (s(bool,X307) != s(bool,f(s(bool,X306))) | s(bool,'T_0') = s(bool,X307) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X308,bool),'const.pair.,_2'(s(X308,X309),s(bool,X306)))))) ) | $spl36),
  inference(superposition,[],[f2455,f133])).
fof(f13854,plain,(
  ( ! [X341,X340,X342] : (s(bool,'T_0') = s(bool,f(s(bool,X340))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X341,bool),'const.pair.,_2'(s(X341,X342),s(bool,X340)))))) ) | ($spl9 | $spl36)),
  inference(subsumption_resolution,[],[f13105,f71])).
fof(f13105,plain,(
  ( ! [X341,X340,X342] : (s(bool,'T_0') = s(bool,f(s(bool,X340))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X341,bool),'const.pair.,_2'(s(X341,X342),s(bool,X340)))))) ) | $spl36),
  inference(superposition,[],[f11203,f133])).
fof(f31902,plain,(
  ( ! [X24,X23,X21,X22] : (s(bool,X22) = s(bool,f(s(bool,X21))) | s(bool,'F_0') = s(bool,X22) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X23),'const.pair.,_2'(s(bool,X21),s(X23,X24)))))) ) | ($spl9 | $spl176)),
  inference(subsumption_resolution,[],[f31630,f71])).
fof(f31630,plain,(
  ( ! [X24,X23,X21,X22] : (s(bool,X22) = s(bool,f(s(bool,X21))) | s(bool,'F_0') = s(bool,X22) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X23),'const.pair.,_2'(s(bool,X21),s(X23,X24)))))) ) | $spl176),
  inference(superposition,[],[f30750,f132])).
fof(f31903,plain,(
  ( ! [X35,X33,X36,X34] : (s(bool,X34) = s(bool,f(s(bool,X33))) | s(bool,'F_0') = s(bool,X34) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X35,bool),'const.pair.,_2'(s(X35,X36),s(bool,X33)))))) ) | ($spl9 | $spl176)),
  inference(subsumption_resolution,[],[f31633,f71])).
fof(f31633,plain,(
  ( ! [X35,X33,X36,X34] : (s(bool,X34) = s(bool,f(s(bool,X33))) | s(bool,'F_0') = s(bool,X34) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X35,bool),'const.pair.,_2'(s(X35,X36),s(bool,X33)))))) ) | $spl176),
  inference(superposition,[],[f30750,f133])).
fof(f35840,plain,(
  ( ! [X261,X263,X262,X264] : (s(bool,'T_0') = s(bool,f(s(bool,X262))) | s(bool,X264) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X261),'const.pair.,_2'(s(bool,X262),s(X261,X263))))) | s(bool,'F_0') = s(bool,X264)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30029,f43])).
fof(f35841,plain,(
  ( ! [X265,X267,X268,X266] : (s(bool,'T_0') = s(bool,f(s(bool,X267))) | s(bool,X268) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X265,bool),'const.pair.,_2'(s(X265,X266),s(bool,X267))))) | s(bool,'F_0') = s(bool,X268)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30029,f44])).
fof(f36184,plain,(
  ( ! [X269,X271,X267,X268,X270,X266] : (~p(s(bool,happ(s(fun(bool,bool),X267),s(bool,'T_0')))) | ~p(s(bool,happ(s(fun(X268,bool),X269),s(X268,X270)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X268,bool),X269),s(fun(bool,bool),X267),s('type.pair.prod'(X268,bool),'const.pair.,_2'(s(X268,X270),s(bool,f(s(bool,X266)))))))) | s(bool,X266) = s(bool,X271) | s(bool,'F_0') = s(bool,X271)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f785,f30029])).
fof(f36331,plain,(
  ( ! [X492,X494,X496,X493,X495,X497] : (~p(s(bool,happ(s(fun(bool,bool),X494),s(bool,X493)))) | ~p(s(bool,happ(s(fun(X495,bool),X496),s(X495,X497)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X495,bool),X496),s(fun(bool,bool),X494),s('type.pair.prod'(X495,bool),'const.pair.,_2'(s(X495,X497),s(bool,X492)))))) | s(bool,'T_0') = s(bool,f(s(bool,X492))) | s(bool,'F_0') = s(bool,X493)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f785,f30029])).
fof(f36488,plain,(
  ( ! [X492,X494,X496,X493,X495,X497] : (~p(s(bool,happ(s(fun(bool,bool),X494),s(bool,X493)))) | ~p(s(bool,happ(s(fun(X495,bool),X496),s(X495,X497)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X495,bool),X496),s(fun(bool,bool),X494),s('type.pair.prod'(X495,bool),'const.pair.,_2'(s(X495,X497),s(bool,X492)))))) | s(bool,'T_0') = s(bool,f(s(bool,X493))) | s(bool,'F_0') = s(bool,X492)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f785,f30029])).
fof(f36500,plain,(
  ( ! [X557,X553,X555,X556,X554] : (~p(s(bool,happ(s(fun(bool,bool),X556),s(bool,X555)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X553,bool),X557),s(fun(bool,bool),X556),s('type.pair.prod'(X553,bool),X554)))) | ~p(s(bool,happ(s(fun(X553,bool),X557),s(X553,'const.pair.FST_1'(s('type.pair.prod'(X553,bool),X554)))))) | s(bool,'T_0') = s(bool,f(s(bool,X555))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X553,bool),X554)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f48,f30029])).
fof(f36504,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,f(s(bool,X1))) | s(bool,X0) = s(bool,X1)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f39,f30029])).
fof(f36575,plain,(
  ( ! [X292,X294,X290,X293,X289,X291] : (~p(s(bool,happ(s(fun(bool,bool),X290),s(bool,X289)))) | ~p(s(bool,happ(s(fun(X291,bool),X292),s(X291,X293)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X291,bool),X292),s(fun(bool,bool),X290),s('type.pair.prod'(X291,bool),'const.pair.,_2'(s(X291,X293),s(bool,'F_0')))))) | s(bool,'T_0') = s(bool,f(s(bool,X294))) | s(bool,X289) = s(bool,X294)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f785,f30029])).
fof(f36696,plain,(
  ( ! [X414,X416,X418,X413,X415,X417] : (~p(s(bool,happ(s(fun(bool,bool),X414),s(bool,'F_0')))) | ~p(s(bool,happ(s(fun(X415,bool),X416),s(X415,X417)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X415,bool),X416),s(fun(bool,bool),X414),s('type.pair.prod'(X415,bool),'const.pair.,_2'(s(X415,X417),s(bool,X413)))))) | s(bool,'T_0') = s(bool,f(s(bool,X418))) | s(bool,X413) = s(bool,X418)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f785,f30029])).
fof(f36708,plain,(
  ( ! [X472,X474,X471,X473,X470] : (~p(s(bool,happ(s(fun(bool,bool),X472),s(bool,'F_0')))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X470,bool),X473),s(fun(bool,bool),X472),s('type.pair.prod'(X470,bool),X471)))) | ~p(s(bool,happ(s(fun(X470,bool),X473),s(X470,'const.pair.FST_1'(s('type.pair.prod'(X470,bool),X471)))))) | s(bool,'T_0') = s(bool,f(s(bool,X474))) | s(bool,X474) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X470,bool),X471)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f48,f30029])).
fof(f36872,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,f(s(bool,X2))))))) | s(bool,X2) = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) ) | ($spl4 | $spl9 | $spl36)),
  inference(condensation,[],[f36871])).
fof(f36871,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,'T_0') = s(bool,X26) | s(bool,X25) = s(bool,X27) | s(bool,'F_0') = s(bool,X27) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X28,bool),'const.pair.,_2'(s(X28,X29),s(bool,f(s(bool,X25)))))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(duplicate_literal_removal,[],[f35849])).
fof(f35849,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,'T_0') = s(bool,X26) | s(bool,X25) = s(bool,X27) | s(bool,'F_0') = s(bool,X27) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X28,bool),'const.pair.,_2'(s(X28,X29),s(bool,f(s(bool,X25)))))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30029,f133])).
fof(f36874,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,f(s(bool,X1))),s(X0,X2))))) | s(bool,X1) = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) ) | ($spl4 | $spl9 | $spl36)),
  inference(condensation,[],[f36873])).
fof(f36873,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,'T_0') = s(bool,X15) | s(bool,X14) = s(bool,X16) | s(bool,'F_0') = s(bool,X16) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X17),'const.pair.,_2'(s(bool,f(s(bool,X14))),s(X17,X18)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(duplicate_literal_removal,[],[f35847])).
fof(f35847,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,'T_0') = s(bool,X15) | s(bool,X14) = s(bool,X16) | s(bool,'F_0') = s(bool,X16) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X17),'const.pair.,_2'(s(bool,f(s(bool,X14))),s(X17,X18)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30029,f132])).
fof(f45391,plain,(
  ( ! [X24,X23,X21,X22] : (s(bool,X22) = s(bool,f(s(bool,X21))) | s(bool,'T_0') = s(bool,f(s(bool,X22))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X23),'const.pair.,_2'(s(bool,X21),s(X23,X24)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f44869,f71])).
fof(f44869,plain,(
  ( ! [X24,X23,X21,X22] : (s(bool,X22) = s(bool,f(s(bool,X21))) | s(bool,'T_0') = s(bool,f(s(bool,X22))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X23),'const.pair.,_2'(s(bool,X21),s(X23,X24)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f37032,f132])).
fof(f45392,plain,(
  ( ! [X35,X33,X36,X34] : (s(bool,X34) = s(bool,f(s(bool,X33))) | s(bool,'T_0') = s(bool,f(s(bool,X34))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X35,bool),'const.pair.,_2'(s(X35,X36),s(bool,X33)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f44872,f71])).
fof(f44872,plain,(
  ( ! [X35,X33,X36,X34] : (s(bool,X34) = s(bool,f(s(bool,X33))) | s(bool,'T_0') = s(bool,f(s(bool,X34))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X35,bool),'const.pair.,_2'(s(X35,X36),s(bool,X33)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f37032,f133])).
fof(f45454,plain,(
  ( ! [X191,X194,X192,X195,X193] : (s(bool,'T_0') = s(bool,X191) | ~p(s(bool,happ(s(fun(bool,bool),X192),s(bool,'T_0')))) | ~p(s(bool,happ(s(fun(X193,bool),X194),s(X193,X195)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X193,bool),X194),s(fun(bool,bool),X192),s('type.pair.prod'(X193,bool),'const.pair.,_2'(s(X193,X195),s(bool,f(s(bool,X191))))))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(forward_demodulation,[],[f45001,f11203])).
fof(f45001,plain,(
  ( ! [X191,X194,X192,X195,X193] : (~p(s(bool,happ(s(fun(bool,bool),X192),s(bool,'T_0')))) | ~p(s(bool,happ(s(fun(X193,bool),X194),s(X193,X195)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X193,bool),X194),s(fun(bool,bool),X192),s('type.pair.prod'(X193,bool),'const.pair.,_2'(s(X193,X195),s(bool,f(s(bool,X191)))))))) | s(bool,X191) = s(bool,f(s(bool,'F_0')))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f785,f37032])).
fof(f45538,plain,(
  ( ! [X271,X273,X275,X272,X274] : (~p(s(bool,happ(s(fun(bool,bool),X272),s(bool,'T_0')))) | ~p(s(bool,happ(s(fun(X273,bool),X274),s(X273,X275)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X273,bool),X274),s(fun(bool,bool),X272),s('type.pair.prod'(X273,bool),'const.pair.,_2'(s(X273,X275),s(bool,X271)))))) | s(bool,'T_0') = s(bool,f(s(bool,X271)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(forward_demodulation,[],[f45115,f11203])).
fof(f45115,plain,(
  ( ! [X271,X273,X275,X272,X274] : (~p(s(bool,happ(s(fun(bool,bool),X272),s(bool,f(s(bool,'F_0')))))) | ~p(s(bool,happ(s(fun(X273,bool),X274),s(X273,X275)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X273,bool),X274),s(fun(bool,bool),X272),s('type.pair.prod'(X273,bool),'const.pair.,_2'(s(X273,X275),s(bool,X271)))))) | s(bool,'T_0') = s(bool,f(s(bool,X271)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f785,f37032])).
fof(f45588,plain,(
  ( ! [X206,X204,X207,X205,X203] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X205,bool),X206),s(fun(bool,bool),X204),s('type.pair.prod'(X205,bool),'const.pair.,_2'(s(X205,X207),s(bool,'T_0')))))) | ~p(s(bool,happ(s(fun(bool,bool),X204),s(bool,X203)))) | ~p(s(bool,happ(s(fun(X205,bool),X206),s(X205,X207)))) | s(bool,'T_0') = s(bool,f(s(bool,X203)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(forward_demodulation,[],[f45200,f11203])).
fof(f45200,plain,(
  ( ! [X206,X204,X207,X205,X203] : (~p(s(bool,happ(s(fun(bool,bool),X204),s(bool,X203)))) | ~p(s(bool,happ(s(fun(X205,bool),X206),s(X205,X207)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X205,bool),X206),s(fun(bool,bool),X204),s('type.pair.prod'(X205,bool),'const.pair.,_2'(s(X205,X207),s(bool,f(s(bool,'F_0')))))))) | s(bool,'T_0') = s(bool,f(s(bool,X203)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f785,f37032])).
fof(f46456,plain,(
  ( ! [X315,X314,X311] : (s(bool,'T_0') = s(bool,X311) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X314,bool),'const.pair.,_2'(s(X314,X315),s(bool,f(s(bool,X311)))))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(forward_demodulation,[],[f46455,f11203])).
fof(f46455,plain,(
  ( ! [X315,X314,X311] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X314,bool),'const.pair.,_2'(s(X314,X315),s(bool,f(s(bool,X311))))))) | s(bool,X311) = s(bool,f(s(bool,'F_0')))) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f46454,f71])).
fof(f46454,plain,(
  ( ! [X315,X314,X311] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X314,bool),'const.pair.,_2'(s(X314,X315),s(bool,f(s(bool,X311))))))) | s(bool,X311) = s(bool,f(s(bool,'F_0')))) ) | ($spl4 | $spl9 | $spl36)),
  inference(forward_demodulation,[],[f45884,f44])).
fof(f45884,plain,(
  ( ! [X313,X315,X312,X314,X311] : (s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X312,bool),'const.pair.,_2'(s(X312,X313),s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X314,bool),'const.pair.,_2'(s(X314,X315),s(bool,f(s(bool,X311))))))) | s(bool,X311) = s(bool,f(s(bool,'F_0')))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f1972,f37032])).
fof(f45544,plain,(
  ( ! [X308,X310,X309,X311] : (~p(s(bool,happ(s(fun(bool,bool),X310),s(bool,'T_0')))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X308,bool),X311),s(fun(bool,bool),X310),s('type.pair.prod'(X308,bool),X309)))) | ~p(s(bool,happ(s(fun(X308,bool),X311),s(X308,'const.pair.FST_1'(s('type.pair.prod'(X308,bool),X309)))))) | s(bool,'T_0') = s(bool,f(s(bool,'const.pair.SND_1'(s('type.pair.prod'(X308,bool),X309)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(forward_demodulation,[],[f45125,f11203])).
fof(f45125,plain,(
  ( ! [X308,X310,X309,X311] : (~p(s(bool,happ(s(fun(bool,bool),X310),s(bool,f(s(bool,'F_0')))))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X308,bool),X311),s(fun(bool,bool),X310),s('type.pair.prod'(X308,bool),X309)))) | ~p(s(bool,happ(s(fun(X308,bool),X311),s(X308,'const.pair.FST_1'(s('type.pair.prod'(X308,bool),X309)))))) | s(bool,'T_0') = s(bool,f(s(bool,'const.pair.SND_1'(s('type.pair.prod'(X308,bool),X309)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f48,f37032])).
fof(f36343,plain,(
  ( ! [X557,X553,X555,X556,X554] : (~p(s(bool,happ(s(fun(bool,bool),X556),s(bool,X555)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X553,bool),X557),s(fun(bool,bool),X556),s('type.pair.prod'(X553,bool),X554)))) | ~p(s(bool,happ(s(fun(X553,bool),X557),s(X553,'const.pair.FST_1'(s('type.pair.prod'(X553,bool),X554)))))) | s(bool,'T_0') = s(bool,f(s(bool,'const.pair.SND_1'(s('type.pair.prod'(X553,bool),X554))))) | s(bool,'F_0') = s(bool,X555)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f48,f30029])).
fof(f25786,plain,(
  ( ! [X694] : (s(bool,X694) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X694)) ) | ($spl4 | $spl36)),
  inference(global_subsumption,[],[f25785,f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f64,f132,f133,f146,f156,f157,f200,f361,f363,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f3832,f3890,f3878,f2550,f2549,f3652,f3732,f3828,f3727,f2700,f763,f2874,f2698,f761,f3651,f786,f550,f475,f48,f47,f46,f762,f2699,f2704,f785,f2763,f4154,f3827,f3726,f3646,f368,f366,f3759,f3758,f2455,f1935,f1480,f1972,f1473,f39,f82,f104,f365,f370,f2706,f2559,f11126,f11125,f11188,f11161,f11162,f11203,f11222,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13826,f13033,f13038,f13169,f13223,f13224,f13225,f13345,f13350,f13427,f13415,f15352,f15341,f15429,f15069,f17276,f17266,f17264,f17226,f17356,f16976,f22573,f18154,f18155,f22303,f22293,f22291,f22279,f22277,f22275,f22219,f22217,f18360,f18362,f18363,f22711,f18790,f18797,f18801,f18802,f22898,f19153,f19158,f21203,f21173,f19331,f19338,f19400,f19401,f19402,f19403,f19406,f22928,f22929,f22930,f22931,f19558,f19565,f19569,f19570,f22984,f20200])).
fof(f29241,plain,(
  ( ! [X134,X138,X136,X135,X137] : (p(s(bool,happ(s(fun(X136,bool),X137),s(X136,'const.pair.SND_1'(s('type.pair.prod'(X134,X136),X138)))))) | s(bool,'F_0') = s(bool,'const.pair.PROD_ALL_3'(s(fun(X134,bool),X135),s(fun(X136,bool),X137),s('type.pair.prod'(X134,X136),X138)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f29240,f37])).
fof(f29240,plain,(
  ( ! [X134,X138,X136,X135,X137] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X136,bool),X137),s(X136,'const.pair.SND_1'(s('type.pair.prod'(X134,X136),X138)))))) | s(bool,'F_0') = s(bool,'const.pair.PROD_ALL_3'(s(fun(X134,bool),X135),s(fun(X136,bool),X137),s('type.pair.prod'(X134,X136),X138)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(forward_demodulation,[],[f29069,f1085])).
fof(f29069,plain,(
  ( ! [X134,X138,X136,X135,X137] : (~p(s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) | p(s(bool,happ(s(fun(X136,bool),X137),s(X136,'const.pair.SND_1'(s('type.pair.prod'(X134,X136),X138)))))) | s(bool,'F_0') = s(bool,'const.pair.PROD_ALL_3'(s(fun(X134,bool),X135),s(fun(X136,bool),X137),s('type.pair.prod'(X134,X136),X138)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f47,f25792])).
fof(f29243,plain,(
  ( ! [X142,X140,X143,X141,X139] : (p(s(bool,happ(s(fun(X139,bool),X140),s(X139,'const.pair.FST_1'(s('type.pair.prod'(X139,X141),X143)))))) | s(bool,'F_0') = s(bool,'const.pair.PROD_ALL_3'(s(fun(X139,bool),X140),s(fun(X141,bool),X142),s('type.pair.prod'(X139,X141),X143)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f29242,f37])).
fof(f29242,plain,(
  ( ! [X142,X140,X143,X141,X139] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X139,bool),X140),s(X139,'const.pair.FST_1'(s('type.pair.prod'(X139,X141),X143)))))) | s(bool,'F_0') = s(bool,'const.pair.PROD_ALL_3'(s(fun(X139,bool),X140),s(fun(X141,bool),X142),s('type.pair.prod'(X139,X141),X143)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(forward_demodulation,[],[f29070,f1085])).
fof(f29070,plain,(
  ( ! [X142,X140,X143,X141,X139] : (~p(s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) | p(s(bool,happ(s(fun(X139,bool),X140),s(X139,'const.pair.FST_1'(s('type.pair.prod'(X139,X141),X143)))))) | s(bool,'F_0') = s(bool,'const.pair.PROD_ALL_3'(s(fun(X139,bool),X140),s(fun(X141,bool),X142),s('type.pair.prod'(X139,X141),X143)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f46,f25792])).
fof(f29245,plain,(
  ( ! [X146,X144,X147,X145,X148] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X146,bool),X148),s(fun(X144,bool),X145),s('type.pair.prod'(X146,X144),X147)))) | ~p(s(bool,happ(s(fun(X146,bool),X148),s(X146,'const.pair.FST_1'(s('type.pair.prod'(X146,X144),X147)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X144,bool),X145),s(X144,'const.pair.SND_1'(s('type.pair.prod'(X146,X144),X147)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f29244,f37])).
fof(f29244,plain,(
  ( ! [X146,X144,X147,X145,X148] : (~p(s(bool,'T_0')) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X146,bool),X148),s(fun(X144,bool),X145),s('type.pair.prod'(X146,X144),X147)))) | ~p(s(bool,happ(s(fun(X146,bool),X148),s(X146,'const.pair.FST_1'(s('type.pair.prod'(X146,X144),X147)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X144,bool),X145),s(X144,'const.pair.SND_1'(s('type.pair.prod'(X146,X144),X147)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(forward_demodulation,[],[f29071,f1085])).
fof(f29071,plain,(
  ( ! [X146,X144,X147,X145,X148] : (~p(s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X146,bool),X148),s(fun(X144,bool),X145),s('type.pair.prod'(X146,X144),X147)))) | ~p(s(bool,happ(s(fun(X146,bool),X148),s(X146,'const.pair.FST_1'(s('type.pair.prod'(X146,X144),X147)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X144,bool),X145),s(X144,'const.pair.SND_1'(s('type.pair.prod'(X146,X144),X147)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f48,f25792])).
fof(f29247,plain,(
  ( ! [X154,X152,X151,X149,X153,X150] : (~p(s(bool,happ(s(fun(X152,bool),X153),s(X152,X154)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X152,bool),X153),s(fun(X149,bool),X150),s('type.pair.prod'(X152,X149),'const.pair.,_2'(s(X152,X154),s(X149,X151)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X149,bool),X150),s(X149,X151)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f29246,f37])).
fof(f29246,plain,(
  ( ! [X154,X152,X151,X149,X153,X150] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X152,bool),X153),s(X152,X154)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X152,bool),X153),s(fun(X149,bool),X150),s('type.pair.prod'(X152,X149),'const.pair.,_2'(s(X152,X154),s(X149,X151)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X149,bool),X150),s(X149,X151)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(forward_demodulation,[],[f29072,f1085])).
fof(f29072,plain,(
  ( ! [X154,X152,X151,X149,X153,X150] : (~p(s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) | ~p(s(bool,happ(s(fun(X152,bool),X153),s(X152,X154)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X152,bool),X153),s(fun(X149,bool),X150),s('type.pair.prod'(X152,X149),'const.pair.,_2'(s(X152,X154),s(X149,X151)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X149,bool),X150),s(X149,X151)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f785,f25792])).
fof(f29252,plain,(
  ( ! [X222,X220,X221] : (s(bool,'T_0') != s(bool,f(s(X220,sK6(X221,X222,bool,X220)))) | s(fun(X220,bool),X221) = s(fun(X220,bool),X222) | s(bool,'F_0') = s(bool,g(s(X220,sK6(X221,X222,bool,X220))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(forward_demodulation,[],[f29095,f1085])).
fof(f29095,plain,(
  ( ! [X222,X220,X221] : (s(bool,f(s(X220,sK6(X221,X222,bool,X220)))) != s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(fun(X220,bool),X221) = s(fun(X220,bool),X222) | s(bool,'F_0') = s(bool,g(s(X220,sK6(X221,X222,bool,X220))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f45,f25792])).
fof(f29255,plain,(
  ( ! [X231,X229,X230,X232] : (~p(s(bool,happ(s(fun(bool,bool),X231),s(bool,'T_0')))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X229,bool),X232),s(fun(bool,bool),X231),s('type.pair.prod'(X229,bool),X230)))) | ~p(s(bool,happ(s(fun(X229,bool),X232),s(X229,'const.pair.FST_1'(s('type.pair.prod'(X229,bool),X230)))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X229,bool),X230)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(forward_demodulation,[],[f29098,f1085])).
fof(f29098,plain,(
  ( ! [X231,X229,X230,X232] : (~p(s(bool,happ(s(fun(bool,bool),X231),s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X229,bool),X232),s(fun(bool,bool),X231),s('type.pair.prod'(X229,bool),X230)))) | ~p(s(bool,happ(s(fun(X229,bool),X232),s(X229,'const.pair.FST_1'(s('type.pair.prod'(X229,bool),X230)))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X229,bool),X230)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f48,f25792])).
fof(f29256,plain,(
  ( ! [X235,X233,X234] : (s(X235,f(s(bool,'T_0'))) != s(X235,g(s(bool,'T_0'))) | s(fun(bool,X235),X233) = s(fun(bool,X235),X234) | s(bool,'F_0') = s(bool,sK6(X233,X234,X235,bool))) ) | ($spl4 | $spl9 | $spl36)),
  inference(forward_demodulation,[],[f29099,f1085])).
fof(f29099,plain,(
  ( ! [X235,X233,X234] : (s(X235,f(s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))))) != s(X235,g(s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))))) | s(fun(bool,X235),X233) = s(fun(bool,X235),X234) | s(bool,'F_0') = s(bool,sK6(X233,X234,X235,bool))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f45,f25792])).
fof(f29283,plain,(
  ( ! [X167,X165,X163,X166,X164] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X165,bool),X166),s(fun(bool,bool),X164),s('type.pair.prod'(X165,bool),'const.pair.,_2'(s(X165,X167),s(bool,'T_0')))))) | ~p(s(bool,happ(s(fun(bool,bool),X164),s(bool,X163)))) | ~p(s(bool,happ(s(fun(X165,bool),X166),s(X165,X167)))) | s(bool,'F_0') = s(bool,X163)) ) | ($spl4 | $spl9 | $spl36)),
  inference(forward_demodulation,[],[f29155,f1085])).
fof(f29155,plain,(
  ( ! [X167,X165,X163,X166,X164] : (~p(s(bool,happ(s(fun(bool,bool),X164),s(bool,X163)))) | ~p(s(bool,happ(s(fun(X165,bool),X166),s(X165,X167)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X165,bool),X166),s(fun(bool,bool),X164),s('type.pair.prod'(X165,bool),'const.pair.,_2'(s(X165,X167),s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))))))) | s(bool,'F_0') = s(bool,X163)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f785,f25792])).
fof(f29580,plain,(
  ( ! [X208,X206,X207,X205] : (s(bool,X206) = s(bool,X208) | s(bool,'F_0') = s(bool,X208) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X205),'const.pair.,_2'(s(bool,X206),s(X205,X207)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f28992,f43])).
fof(f29581,plain,(
  ( ! [X212,X210,X211,X209] : (s(bool,X211) = s(bool,X212) | s(bool,'F_0') = s(bool,X212) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X209,bool),'const.pair.,_2'(s(X209,X210),s(bool,X211)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f28992,f44])).
fof(f29714,plain,(
  ( ! [X175,X173,X176,X174] : (s(bool,'F_0') = s(bool,X174) | s(bool,X176) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X173),'const.pair.,_2'(s(bool,X174),s(X173,X175))))) | s(bool,'F_0') = s(bool,X176)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f28992,f43])).
fof(f29715,plain,(
  ( ! [X177,X180,X178,X179] : (s(bool,'F_0') = s(bool,X179) | s(bool,X180) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X177,bool),'const.pair.,_2'(s(X177,X178),s(bool,X179))))) | s(bool,'F_0') = s(bool,X180)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f28992,f44])).
fof(f29898,plain,(
  ( ! [X377,X379,X374,X376,X378,X375] : (~p(s(bool,happ(s(fun(bool,bool),X376),s(bool,X375)))) | ~p(s(bool,happ(s(fun(X377,bool),X378),s(X377,X379)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X377,bool),X378),s(fun(bool,bool),X376),s('type.pair.prod'(X377,bool),'const.pair.,_2'(s(X377,X379),s(bool,X374)))))) | s(bool,'F_0') = s(bool,X375) | s(bool,'F_0') = s(bool,X374)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f785,f28992])).
fof(f29900,plain,(
  ( ! [X389,X391,X390,X392] : (s(bool,X392) != s(bool,f(s(X389,sK6(X390,X391,bool,X389)))) | s(fun(X389,bool),X390) = s(fun(X389,bool),X391) | s(bool,'F_0') = s(bool,X392) | s(bool,'F_0') = s(bool,g(s(X389,sK6(X390,X391,bool,X389))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f45,f28992])).
fof(f29904,plain,(
  ( ! [X408,X405,X407,X404,X406] : (~p(s(bool,happ(s(fun(bool,bool),X407),s(bool,X406)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X404,bool),X408),s(fun(bool,bool),X407),s('type.pair.prod'(X404,bool),X405)))) | ~p(s(bool,happ(s(fun(X404,bool),X408),s(X404,'const.pair.FST_1'(s('type.pair.prod'(X404,bool),X405)))))) | s(bool,'F_0') = s(bool,X406) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X404,bool),X405)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f48,f28992])).
fof(f29905,plain,(
  ( ! [X412,X410,X409,X411] : (s(X411,f(s(bool,X412))) != s(X411,g(s(bool,X412))) | s(fun(bool,X411),X409) = s(fun(bool,X411),X410) | s(bool,'F_0') = s(bool,X412) | s(bool,'F_0') = s(bool,sK6(X409,X410,X411,bool))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f45,f28992])).
fof(f30024,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f39,f28992])).
fof(f30027,plain,(
  ( ! [X10,X8,X7,X11,X9] : (s(X8,f(s(bool,X7))) != s(X8,g(s(bool,X7))) | s(fun(bool,X8),X9) = s(fun(bool,X8),X10) | s(bool,'T_0') = s(bool,sK6(X9,X10,X8,bool)) | s(bool,X7) = s(bool,X11) | s(bool,'F_0') = s(bool,X11)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f365,f28992])).
fof(f30030,plain,(
  ( ! [X24,X23,X22] : (s(bool,X23) != s(bool,f(s(bool,X22))) | s(bool,'T_0') = s(bool,X23) | s(bool,X22) = s(bool,X24) | s(bool,'F_0') = s(bool,X24)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f25786,f28992])).
fof(f30080,plain,(
  ( ! [X231,X229,X233,X230,X228,X232] : (~p(s(bool,happ(s(fun(bool,bool),X229),s(bool,X228)))) | ~p(s(bool,happ(s(fun(X230,bool),X231),s(X230,X232)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X230,bool),X231),s(fun(bool,bool),X229),s('type.pair.prod'(X230,bool),'const.pair.,_2'(s(X230,X232),s(bool,'F_0')))))) | s(bool,X228) = s(bool,X233) | s(bool,'F_0') = s(bool,X233)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f785,f28992])).
fof(f30172,plain,(
  ( ! [X321,X316,X318,X320,X317,X319] : (~p(s(bool,happ(s(fun(bool,bool),X317),s(bool,'F_0')))) | ~p(s(bool,happ(s(fun(X318,bool),X319),s(X318,X320)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X318,bool),X319),s(fun(bool,bool),X317),s('type.pair.prod'(X318,bool),'const.pair.,_2'(s(X318,X320),s(bool,X316)))))) | s(bool,X316) = s(bool,X321) | s(bool,'F_0') = s(bool,X321)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f785,f28992])).
fof(f30174,plain,(
  ( ! [X329,X331,X332,X330] : (s(bool,'F_0') != s(bool,f(s(X329,sK6(X330,X331,bool,X329)))) | s(fun(X329,bool),X330) = s(fun(X329,bool),X331) | s(bool,X332) = s(bool,g(s(X329,sK6(X330,X331,bool,X329)))) | s(bool,'F_0') = s(bool,X332)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f45,f28992])).
fof(f30178,plain,(
  ( ! [X344,X346,X345,X347,X348] : (~p(s(bool,happ(s(fun(bool,bool),X346),s(bool,'F_0')))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X344,bool),X347),s(fun(bool,bool),X346),s('type.pair.prod'(X344,bool),X345)))) | ~p(s(bool,happ(s(fun(X344,bool),X347),s(X344,'const.pair.FST_1'(s('type.pair.prod'(X344,bool),X345)))))) | s(bool,X348) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X344,bool),X345))) | s(bool,'F_0') = s(bool,X348)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f48,f28992])).
fof(f30179,plain,(
  ( ! [X350,X352,X349,X351] : (s(X351,f(s(bool,'F_0'))) != s(X351,g(s(bool,'F_0'))) | s(fun(bool,X351),X349) = s(fun(bool,X351),X350) | s(bool,X352) = s(bool,sK6(X349,X350,X351,bool)) | s(bool,'F_0') = s(bool,X352)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f45,f28992])).
fof(f30404,plain,(
  ( ! [X206,X204,X205,X203] : (s(fun(X203,bool),X204) = s(fun(X203,bool),X205) | s(bool,'T_0') = s(bool,g(s(X203,sK6(X204,X205,bool,X203)))) | s(bool,X206) = s(bool,f(s(X203,sK6(X204,X205,bool,X203)))) | s(bool,'F_0') = s(bool,X206)) ) | ($spl4 | $spl9 | $spl36)),
  inference(trivial_inequality_removal,[],[f30141])).
fof(f30141,plain,(
  ( ! [X206,X204,X205,X203] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(X203,bool),X204) = s(fun(X203,bool),X205) | s(bool,'T_0') = s(bool,g(s(X203,sK6(X204,X205,bool,X203)))) | s(bool,X206) = s(bool,f(s(X203,sK6(X204,X205,bool,X203)))) | s(bool,'F_0') = s(bool,X206)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f370,f28992])).
fof(f31001,plain,(
  ( ! [X231,X235,X233,X234,X232] : (s(bool,X231) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X231),s(bool,X232))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X233),'const.pair.,_2'(s(bool,X232),s(X233,X234))))) | s(bool,X231) = s(bool,X235) | s(bool,'F_0') = s(bool,X235)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f1473,f28992])).
fof(f31621,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) = s(bool,f(s(bool,X0))) | s(bool,'F_0') = s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'F_0') = s(bool,X2)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(superposition,[],[f30750,f28992])).
fof(f31907,plain,(
  ( ! [X61,X64,X62,X63] : (s(bool,X62) = s(bool,f(s(bool,X61))) | s(bool,'F_0') = s(bool,X62) | s(bool,X61) = s(bool,X63) | s(bool,X63) = s(bool,X64) | s(bool,'T_0') = s(bool,X64)) ) | ($spl9 | $spl176)),
  inference(subsumption_resolution,[],[f31640,f71])).
fof(f31640,plain,(
  ( ! [X61,X64,X62,X63] : (s(bool,X62) = s(bool,f(s(bool,X61))) | s(bool,'F_0') = s(bool,X62) | s(bool,X61) = s(bool,X63) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X63) = s(bool,X64) | s(bool,'T_0') = s(bool,X64)) ) | $spl176),
  inference(superposition,[],[f30750,f146])).
fof(f31911,plain,(
  ( ! [X81,X82] : (s(bool,X82) = s(bool,f(s(bool,X81))) | s(bool,'F_0') = s(bool,X82) | s(bool,X81) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl9 | $spl176)),
  inference(subsumption_resolution,[],[f31644,f71])).
fof(f31644,plain,(
  ( ! [X81,X82] : (s(bool,X82) = s(bool,f(s(bool,X81))) | s(bool,'F_0') = s(bool,X82) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X81) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl176)),
  inference(superposition,[],[f30750,f22295])).
fof(f32101,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,X24) = s(bool,X25) | s(bool,'F_0') = s(bool,X25) | s(bool,X24) = s(bool,X26) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X27),'const.pair.,_2'(s(bool,X26),s(X27,X28)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f32023,f37])).
fof(f32023,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,'T_0')) | s(bool,X24) = s(bool,X25) | s(bool,'F_0') = s(bool,X25) | s(bool,X24) = s(bool,X26) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X27),'const.pair.,_2'(s(bool,X26),s(X27,X28)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30024,f132])).
fof(f32105,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s(bool,X39) = s(bool,X40) | s(bool,'F_0') = s(bool,X40) | s(bool,X39) = s(bool,X41) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X42,bool),'const.pair.,_2'(s(X42,X43),s(bool,X41)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f32026,f37])).
fof(f32026,plain,(
  ( ! [X39,X43,X41,X42,X40] : (~p(s(bool,'T_0')) | s(bool,X39) = s(bool,X40) | s(bool,'F_0') = s(bool,X40) | s(bool,X39) = s(bool,X41) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X42,bool),'const.pair.,_2'(s(X42,X43),s(bool,X41)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30024,f133])).
fof(f32126,plain,(
  ( ! [X101,X99,X100] : (s(bool,X99) = s(bool,X100) | s(bool,'F_0') = s(bool,X100) | s(bool,X101) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X99) = s(bool,X101)) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f32039,f37])).
fof(f32039,plain,(
  ( ! [X101,X99,X100] : (~p(s(bool,'T_0')) | s(bool,X99) = s(bool,X100) | s(bool,'F_0') = s(bool,X100) | s(bool,X101) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X99) = s(bool,X101)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30024,f22295])).
fof(f32151,plain,(
  ( ! [X185,X182,X184,X183,X181] : (s(bool,X184) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X181),'const.pair.,_2'(s(bool,X182),s(X181,X183))))) | s(bool,'F_0') = s(bool,X184) | s(bool,X182) = s(bool,X185) | s(bool,'T_0') = s(bool,X185)) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f32070,f37])).
fof(f32070,plain,(
  ( ! [X185,X182,X184,X183,X181] : (~p(s(bool,'T_0')) | s(bool,X184) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X181),'const.pair.,_2'(s(bool,X182),s(X181,X183))))) | s(bool,'F_0') = s(bool,X184) | s(bool,X182) = s(bool,X185) | s(bool,'T_0') = s(bool,X185)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30024,f132])).
fof(f32152,plain,(
  ( ! [X191,X196,X194,X192,X195,X193] : (s(bool,X194) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X191),'const.pair.,_2'(s(bool,X192),s(X191,X193))))) | s(bool,'F_0') = s(bool,X194) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X195,bool),'const.pair.,_2'(s(X195,X196),s(bool,X192)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f32072,f37])).
fof(f32072,plain,(
  ( ! [X191,X196,X194,X192,X195,X193] : (~p(s(bool,'T_0')) | s(bool,X194) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X191),'const.pair.,_2'(s(bool,X192),s(X191,X193))))) | s(bool,'F_0') = s(bool,X194) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X195,bool),'const.pair.,_2'(s(X195,X196),s(bool,X192)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30024,f1473])).
fof(f32154,plain,(
  ( ! [X218,X222,X220,X219,X221] : (s(bool,X221) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X218,bool),'const.pair.,_2'(s(X218,X219),s(bool,X220))))) | s(bool,'F_0') = s(bool,X221) | s(bool,X220) = s(bool,X222) | s(bool,'T_0') = s(bool,X222)) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f32077,f37])).
fof(f32077,plain,(
  ( ! [X218,X222,X220,X219,X221] : (~p(s(bool,'T_0')) | s(bool,X221) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X218,bool),'const.pair.,_2'(s(X218,X219),s(bool,X220))))) | s(bool,'F_0') = s(bool,X221) | s(bool,X220) = s(bool,X222) | s(bool,'T_0') = s(bool,X222)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30024,f133])).
fof(f32361,plain,(
  ( ! [X239,X237,X235,X238,X236] : (s(bool,X235) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X235),s(bool,X236))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X237,bool),'const.pair.,_2'(s(X237,X238),s(bool,X236))))) | s(bool,X235) = s(bool,X239) | s(bool,'F_0') = s(bool,X239)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f1972,f28992])).
fof(f32991,plain,(
  ( ! [X358,X360,X362,X359,X361,X363] : (s(bool,X361) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X358,bool),'const.pair.,_2'(s(X358,X359),s(bool,X360))))) | s(bool,'F_0') = s(bool,X361) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X362,bool),'const.pair.,_2'(s(X362,X363),s(bool,X360)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f32539,f37])).
fof(f32539,plain,(
  ( ! [X358,X360,X362,X359,X361,X363] : (~p(s(bool,'T_0')) | s(bool,X361) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X358,bool),'const.pair.,_2'(s(X358,X359),s(bool,X360))))) | s(bool,'F_0') = s(bool,X361) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X362,bool),'const.pair.,_2'(s(X362,X363),s(bool,X360)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30024,f1972])).
fof(f33551,plain,(
  ( ! [X239,X237,X240,X238] : (s(bool,X238) != s(bool,X240) | s(bool,'T_0') = s(bool,X240) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X237),'const.pair.,_2'(s(bool,X238),s(X237,X239)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f29101,f43])).
fof(f33552,plain,(
  ( ! [X243,X241,X244,X242] : (s(bool,X243) != s(bool,X244) | s(bool,'T_0') = s(bool,X244) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X241,bool),'const.pair.,_2'(s(X241,X242),s(bool,X243)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f29101,f44])).
fof(f33627,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X2,bool),'const.pair.,_2'(s(X2,X3),s(bool,X0))))) | s(bool,'F_0') = s(bool,X1)) ) | ($spl4 | $spl9 | $spl36)),
  inference(condensation,[],[f33626])).
fof(f33626,plain,(
  ( ! [X92,X90,X93,X91,X89] : (s(bool,X93) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X90,bool),'const.pair.,_2'(s(X90,X91),s(bool,X92))))) | s(bool,'T_0') = s(bool,X89) | s(bool,'F_0') = s(bool,X93) | s(bool,'T_0') = s(bool,X92)) ) | ($spl4 | $spl9 | $spl36)),
  inference(duplicate_literal_removal,[],[f33573])).
fof(f33573,plain,(
  ( ! [X92,X90,X93,X91,X89] : (s(bool,X93) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X90,bool),'const.pair.,_2'(s(X90,X91),s(bool,X92))))) | s(bool,'T_0') = s(bool,X89) | s(bool,'F_0') = s(bool,X93) | s(bool,'T_0') = s(bool,X92) | s(bool,'T_0') = s(bool,X89)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f29101,f157])).
fof(f33629,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),'const.pair.,_2'(s(bool,X0),s(X2,X3))))) | s(bool,'F_0') = s(bool,X1)) ) | ($spl4 | $spl9 | $spl36)),
  inference(condensation,[],[f33628])).
fof(f33628,plain,(
  ( ! [X88,X87,X85,X86,X84] : (s(bool,X88) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X85),'const.pair.,_2'(s(bool,X86),s(X85,X87))))) | s(bool,'T_0') = s(bool,X84) | s(bool,'F_0') = s(bool,X88) | s(bool,'T_0') = s(bool,X86)) ) | ($spl4 | $spl9 | $spl36)),
  inference(duplicate_literal_removal,[],[f33572])).
fof(f33572,plain,(
  ( ! [X88,X87,X85,X86,X84] : (s(bool,X88) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X85),'const.pair.,_2'(s(bool,X86),s(X85,X87))))) | s(bool,'T_0') = s(bool,X84) | s(bool,'F_0') = s(bool,X88) | s(bool,'T_0') = s(bool,X86) | s(bool,'T_0') = s(bool,X84)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f29101,f156])).
fof(f33643,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,X2))))) | s(bool,X2) != s(bool,X3) | s(bool,'F_0') = s(bool,X3)) ) | ($spl4 | $spl9 | $spl36)),
  inference(condensation,[],[f33642])).
fof(f33642,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(bool,X35) != s(bool,X36) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X36) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X37,bool),'const.pair.,_2'(s(X37,X38),s(bool,X35)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(duplicate_literal_removal,[],[f33562])).
fof(f33562,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(bool,X35) != s(bool,X36) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,X36) | s(bool,'T_0') = s(bool,X34) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X37,bool),'const.pair.,_2'(s(X37,X38),s(bool,X35)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f29101,f133])).
fof(f33645,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,X1),s(X0,X2))))) | s(bool,X1) != s(bool,X3) | s(bool,'F_0') = s(bool,X3)) ) | ($spl4 | $spl9 | $spl36)),
  inference(condensation,[],[f33644])).
fof(f33644,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s(bool,X20) != s(bool,X21) | s(bool,'T_0') = s(bool,X19) | s(bool,'F_0') = s(bool,X21) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X22),'const.pair.,_2'(s(bool,X20),s(X22,X23)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(duplicate_literal_removal,[],[f33559])).
fof(f33559,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s(bool,X20) != s(bool,X21) | s(bool,'T_0') = s(bool,X19) | s(bool,'F_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X22),'const.pair.,_2'(s(bool,X20),s(X22,X23)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f29101,f132])).
fof(f34115,plain,(
  ( ! [X12,X10,X8,X13,X11,X9] : (s(bool,X8) != s(bool,X9) | s(bool,X9) = s(bool,X10) | s(bool,'F_0') = s(bool,X10) | s(bool,X8) = s(bool,X11) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X12),'const.pair.,_2'(s(bool,X11),s(X12,X13)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30026,f132])).
fof(f34116,plain,(
  ( ! [X14,X19,X17,X15,X18,X16] : (s(bool,X14) != s(bool,X15) | s(bool,X15) = s(bool,X16) | s(bool,'F_0') = s(bool,X16) | s(bool,X14) = s(bool,X17) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X18,bool),'const.pair.,_2'(s(X18,X19),s(bool,X17)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30026,f133])).
fof(f34119,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,X32) != s(bool,X33) | s(bool,X33) = s(bool,X34) | s(bool,'F_0') = s(bool,X34) | s(bool,X35) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X32) = s(bool,X35)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30026,f22295])).
fof(f34207,plain,(
  ( ! [X237,X235,X238,X236] : (s(bool,'T_0') != s(bool,X236) | s(bool,X238) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X235),'const.pair.,_2'(s(bool,X236),s(X235,X237))))) | s(bool,'F_0') = s(bool,X238)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30026,f43])).
fof(f34208,plain,(
  ( ! [X241,X239,X242,X240] : (s(bool,'T_0') != s(bool,X241) | s(bool,X242) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X239,bool),'const.pair.,_2'(s(X239,X240),s(bool,X241))))) | s(bool,'F_0') = s(bool,X242)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30026,f44])).
fof(f34518,plain,(
  ( ! [X271,X273,X270,X272] : (s(bool,X273) = s(bool,g(s(X270,sK6(X271,X272,bool,X270)))) | s(bool,'F_0') = s(bool,X273) | s(fun(X270,bool),X271) = s(fun(X270,bool),X272) | s(bool,'T_0') = s(bool,f(s(X270,sK6(X271,X272,bool,X270))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f34412,f37])).
fof(f34412,plain,(
  ( ! [X271,X273,X270,X272] : (~p(s(bool,'T_0')) | s(bool,X273) = s(bool,g(s(X270,sK6(X271,X272,bool,X270)))) | s(bool,'F_0') = s(bool,X273) | s(fun(X270,bool),X271) = s(fun(X270,bool),X272) | s(bool,'T_0') = s(bool,f(s(X270,sK6(X271,X272,bool,X270))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30024,f2559])).
fof(f36022,plain,(
  ( ! [X261,X259,X260,X258] : (s(bool,X259) = s(bool,X261) | s(bool,'T_0') = s(bool,f(s(bool,X261))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X258),'const.pair.,_2'(s(bool,X259),s(X258,X260)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30029,f43])).
fof(f36023,plain,(
  ( ! [X263,X265,X262,X264] : (s(bool,X264) = s(bool,X265) | s(bool,'T_0') = s(bool,f(s(bool,X265))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X262,bool),'const.pair.,_2'(s(X262,X263),s(bool,X264)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30029,f44])).
fof(f36025,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,f(s(bool,X6))) | s(bool,X5) = s(bool,X6) | s(bool,X4) = s(bool,X7) | s(bool,'F_0') = s(bool,X7)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30029,f28992])).
fof(f36110,plain,(
  ( ! [X218,X216,X219,X217] : (s(bool,'F_0') = s(bool,X217) | s(bool,'T_0') = s(bool,f(s(bool,X219))) | s(bool,X219) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X216),'const.pair.,_2'(s(bool,X217),s(X216,X218)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30029,f43])).
fof(f36111,plain,(
  ( ! [X222,X220,X223,X221] : (s(bool,'F_0') = s(bool,X222) | s(bool,'T_0') = s(bool,f(s(bool,X223))) | s(bool,X223) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X220,bool),'const.pair.,_2'(s(X220,X221),s(bool,X222)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30029,f44])).
fof(f36337,plain,(
  ( ! [X532,X530,X533,X531] : (s(bool,X533) != s(bool,f(s(X530,sK6(X531,X532,bool,X530)))) | s(fun(X530,bool),X531) = s(fun(X530,bool),X532) | s(bool,'T_0') = s(bool,f(s(bool,g(s(X530,sK6(X531,X532,bool,X530)))))) | s(bool,'F_0') = s(bool,X533)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f45,f30029])).
fof(f36344,plain,(
  ( ! [X561,X559,X560,X558] : (s(X560,f(s(bool,X561))) != s(X560,g(s(bool,X561))) | s(fun(bool,X560),X558) = s(fun(bool,X560),X559) | s(bool,'T_0') = s(bool,f(s(bool,sK6(X558,X559,X560,bool)))) | s(bool,'F_0') = s(bool,X561)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f45,f30029])).
fof(f36494,plain,(
  ( ! [X532,X530,X533,X531] : (s(bool,X533) != s(bool,f(s(X530,sK6(X531,X532,bool,X530)))) | s(fun(X530,bool),X531) = s(fun(X530,bool),X532) | s(bool,'T_0') = s(bool,f(s(bool,X533))) | s(bool,'F_0') = s(bool,g(s(X530,sK6(X531,X532,bool,X530))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f45,f30029])).
fof(f36501,plain,(
  ( ! [X561,X559,X560,X558] : (s(X560,f(s(bool,X561))) != s(X560,g(s(bool,X561))) | s(fun(bool,X560),X558) = s(fun(bool,X560),X559) | s(bool,'T_0') = s(bool,f(s(bool,X561))) | s(bool,'F_0') = s(bool,sK6(X558,X559,X560,bool))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f45,f30029])).
fof(f36513,plain,(
  ( ! [X35,X33,X34] : (s(bool,X34) = s(bool,f(s(bool,X33))) | s(bool,'F_0') = s(bool,X34) | s(bool,'T_0') = s(bool,f(s(bool,X35))) | s(bool,X33) = s(bool,X35)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(superposition,[],[f30750,f30029])).
fof(f36702,plain,(
  ( ! [X447,X449,X448,X450] : (s(bool,'F_0') != s(bool,f(s(X447,sK6(X448,X449,bool,X447)))) | s(fun(X447,bool),X448) = s(fun(X447,bool),X449) | s(bool,'T_0') = s(bool,f(s(bool,X450))) | s(bool,X450) = s(bool,g(s(X447,sK6(X448,X449,bool,X447))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f45,f30029])).
fof(f36709,plain,(
  ( ! [X476,X478,X475,X477] : (s(X477,f(s(bool,'F_0'))) != s(X477,g(s(bool,'F_0'))) | s(fun(bool,X477),X475) = s(fun(bool,X477),X476) | s(bool,'T_0') = s(bool,f(s(bool,X478))) | s(bool,X478) = s(bool,sK6(X475,X476,X477,bool))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f45,f30029])).
fof(f36870,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X3) = s(bool,f(s(bool,X1))) | s(bool,X0) = s(bool,X3)) ) | ($spl4 | $spl9 | $spl36)),
  inference(condensation,[],[f36869])).
fof(f36869,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(bool,'T_0') = s(bool,X49) | s(bool,X48) = s(bool,X50) | s(bool,'F_0') = s(bool,X50) | s(bool,X51) = s(bool,f(s(bool,X48))) | s(bool,X51) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) ) | ($spl4 | $spl9 | $spl36)),
  inference(duplicate_literal_removal,[],[f35853])).
fof(f35853,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(bool,'T_0') = s(bool,X49) | s(bool,X48) = s(bool,X50) | s(bool,'F_0') = s(bool,X50) | s(bool,X51) = s(bool,f(s(bool,X48))) | s(bool,'T_0') = s(bool,X49) | s(bool,X51) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30029,f146])).
fof(f36953,plain,(
  ( ! [X114,X112,X113] : (s(bool,X112) = s(bool,X113) | s(bool,'F_0') = s(bool,X113) | s(bool,'T_0') = s(bool,f(s(bool,X114))) | s(bool,X114) = s(bool,f(s(bool,X112)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f35871,f71])).
fof(f35871,plain,(
  ( ! [X114,X112,X113] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X112) = s(bool,X113) | s(bool,'F_0') = s(bool,X113) | s(bool,'T_0') = s(bool,f(s(bool,X114))) | s(bool,X114) = s(bool,f(s(bool,X112)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30029,f30029])).
fof(f37038,plain,(
  ( ! [X177,X175,X180,X178,X176,X179] : (s(bool,'T_0') = s(bool,f(s(bool,X178))) | s(bool,X178) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X175),'const.pair.,_2'(s(bool,X176),s(X175,X177))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X179,bool),'const.pair.,_2'(s(X179,X180),s(bool,X176)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f36103,f71])).
fof(f36103,plain,(
  ( ! [X177,X175,X180,X178,X176,X179] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(bool,X178))) | s(bool,X178) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X175),'const.pair.,_2'(s(bool,X176),s(X175,X177))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X179,bool),'const.pair.,_2'(s(X179,X180),s(bool,X176)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30029,f1473])).
fof(f37040,plain,(
  ( ! [X198,X196,X194,X199,X197,X195] : (s(bool,'T_0') = s(bool,f(s(bool,X197))) | s(bool,X197) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X194,bool),'const.pair.,_2'(s(X194,X195),s(bool,X196))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X198,bool),'const.pair.,_2'(s(X198,X199),s(bool,X196)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f36106,f71])).
fof(f36106,plain,(
  ( ! [X198,X196,X194,X199,X197,X195] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(bool,X197))) | s(bool,X197) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X194,bool),'const.pair.,_2'(s(X194,X195),s(bool,X196))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X198,bool),'const.pair.,_2'(s(X198,X199),s(bool,X196)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30029,f1972])).
fof(f37211,plain,(
  ( ! [X243,X244,X245] : (s(bool,X243) != s(bool,X244) | s(bool,'F_0') = s(bool,X244) | s(bool,'T_0') = s(bool,f(s(bool,X245))) | s(bool,X243) = s(bool,X245)) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f36562,f71])).
fof(f36562,plain,(
  ( ! [X243,X244,X245] : (s(bool,X243) != s(bool,X244) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X244) | s(bool,'T_0') = s(bool,f(s(bool,X245))) | s(bool,X243) = s(bool,X245)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f29101,f30029])).
fof(f37594,plain,(
  ( ! [X167,X165,X166,X164] : (s(fun(X164,bool),X165) = s(fun(X164,bool),X166) | s(bool,X167) = s(bool,g(s(X164,sK6(X165,X166,bool,X164)))) | s(bool,'T_0') = s(bool,X167) | s(bool,'F_0') = s(bool,f(s(X164,sK6(X165,X166,bool,X164))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f37563,f1085])).
fof(f37563,plain,(
  ( ! [X167,X165,X166,X164] : (s(bool,'T_0') != s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(fun(X164,bool),X165) = s(fun(X164,bool),X166) | s(bool,X167) = s(bool,g(s(X164,sK6(X165,X166,bool,X164)))) | s(bool,'T_0') = s(bool,X167) | s(bool,'F_0') = s(bool,f(s(X164,sK6(X165,X166,bool,X164))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f366,f25792])).
fof(f45254,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,f(s(bool,'F_0'))) | s(bool,X2) = s(bool,f(s(bool,X1))) | s(bool,X0) = s(bool,X2)) ) | ($spl4 | $spl9 | $spl36)),
  inference(condensation,[],[f45253])).
fof(f45253,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,'T_0') = s(bool,X41) | s(bool,X40) = s(bool,f(s(bool,'F_0'))) | s(bool,X42) = s(bool,f(s(bool,X40))) | s(bool,X42) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) ) | ($spl4 | $spl9 | $spl36)),
  inference(duplicate_literal_removal,[],[f44787])).
fof(f44787,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,'T_0') = s(bool,X41) | s(bool,X40) = s(bool,f(s(bool,'F_0'))) | s(bool,X42) = s(bool,f(s(bool,X40))) | s(bool,'T_0') = s(bool,X41) | s(bool,X42) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f37032,f146])).
fof(f45453,plain,(
  ( ! [X152,X151] : (s(bool,X152) = s(bool,f(s(bool,X151))) | s(bool,'F_0') = s(bool,X152) | s(bool,X151) = s(bool,f(s(bool,'F_0')))) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f44982,f37])).
fof(f44982,plain,(
  ( ! [X152,X151] : (~p(s(bool,'T_0')) | s(bool,X152) = s(bool,f(s(bool,X151))) | s(bool,'F_0') = s(bool,X152) | s(bool,X151) = s(bool,f(s(bool,'F_0')))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30024,f37032])).
fof(f49037,plain,(
  ( ! [X191,X192] : (s(bool,X191) != s(bool,f(s(bool,X191))) | s(bool,X191) = s(bool,X192) | s(bool,'F_0') = s(bool,X192)) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f30070,f71])).
fof(f30070,plain,(
  ( ! [X191,X192] : (s(bool,X191) != s(bool,f(s(bool,X191))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X191) = s(bool,X192) | s(bool,'F_0') = s(bool,X192)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f25786,f28992])).
fof(f49265,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,X41) = s(bool,f(s(bool,'F_0'))) | s(bool,X42) = s(bool,X43) | s(bool,X41) = s(bool,X43) | s(bool,X42) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(subsumption_resolution,[],[f44717,f49260])).
fof(f44717,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,X41) = s(bool,f(s(bool,'F_0'))) | s(bool,X42) = s(bool,X43) | s(bool,X41) = s(bool,X43) | s(bool,X42) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f37032,f146])).
fof(f36728,plain,(
  ( ! [X269,X271,X270,X272] : (s(fun(X269,bool),X270) = s(fun(X269,bool),X271) | s(bool,'T_0') = s(bool,g(s(X269,sK6(X270,X271,bool,X269)))) | s(bool,'T_0') = s(bool,f(s(bool,X272))) | s(bool,X272) = s(bool,f(s(X269,sK6(X270,X271,bool,X269))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(trivial_inequality_removal,[],[f36658])).
fof(f36658,plain,(
  ( ! [X269,X271,X270,X272] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(X269,bool),X270) = s(fun(X269,bool),X271) | s(bool,'T_0') = s(bool,g(s(X269,sK6(X270,X271,bool,X269)))) | s(bool,'T_0') = s(bool,f(s(bool,X272))) | s(bool,X272) = s(bool,f(s(X269,sK6(X270,X271,bool,X269))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f370,f30029])).
fof(f37033,plain,(
  ( ! [X152,X151,X153,X150] : (s(bool,'T_0') = s(bool,f(s(bool,X153))) | s(bool,X153) = s(bool,g(s(X150,sK6(X151,X152,bool,X150)))) | s(fun(X150,bool),X151) = s(fun(X150,bool),X152) | s(bool,'T_0') = s(bool,f(s(X150,sK6(X151,X152,bool,X150))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f36094,f71])).
fof(f36094,plain,(
  ( ! [X152,X151,X153,X150] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(bool,X153))) | s(bool,X153) = s(bool,g(s(X150,sK6(X151,X152,bool,X150)))) | s(fun(X150,bool),X151) = s(fun(X150,bool),X152) | s(bool,'T_0') = s(bool,f(s(X150,sK6(X151,X152,bool,X150))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30029,f2559])).
fof(f45539,plain,(
  ( ! [X292,X293,X291] : (s(bool,'T_0') != s(bool,f(s(X291,sK6(X292,X293,bool,X291)))) | s(fun(X291,bool),X292) = s(fun(X291,bool),X293) | s(bool,'T_0') = s(bool,f(s(bool,g(s(X291,sK6(X292,X293,bool,X291))))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(forward_demodulation,[],[f45119,f11203])).
fof(f45119,plain,(
  ( ! [X292,X293,X291] : (s(bool,f(s(bool,'F_0'))) != s(bool,f(s(X291,sK6(X292,X293,bool,X291)))) | s(fun(X291,bool),X292) = s(fun(X291,bool),X293) | s(bool,'T_0') = s(bool,f(s(bool,g(s(X291,sK6(X292,X293,bool,X291))))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f45,f37032])).
fof(f45545,plain,(
  ( ! [X313,X312,X314] : (s(X314,f(s(bool,'T_0'))) != s(X314,g(s(bool,'T_0'))) | s(fun(bool,X314),X312) = s(fun(bool,X314),X313) | s(bool,'T_0') = s(bool,f(s(bool,sK6(X312,X313,X314,bool))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(forward_demodulation,[],[f45126,f11203])).
fof(f45126,plain,(
  ( ! [X313,X312,X314] : (s(X314,f(s(bool,f(s(bool,'F_0'))))) != s(X314,g(s(bool,f(s(bool,'F_0'))))) | s(fun(bool,X314),X312) = s(fun(bool,X314),X313) | s(bool,'T_0') = s(bool,f(s(bool,sK6(X312,X313,X314,bool))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f45,f37032])).
fof(f50035,plain,(
  ( ! [X21,X20] : (s(bool,X20) != s(bool,f(s(bool,X20))) | s(bool,X21) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X20) = s(bool,X21)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(superposition,[],[f49260,f22295])).
fof(f50033,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) != s(bool,f(s(bool,X12))) | s(bool,X13) = s(bool,X14) | s(bool,X12) = s(bool,X14) | s(bool,X13) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(superposition,[],[f49260,f146])).
fof(f50032,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) != s(bool,f(s(bool,X8))) | s(bool,X8) = s(bool,X9) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X10,bool),'const.pair.,_2'(s(X10,X11),s(bool,X9)))))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(superposition,[],[f49260,f133])).
fof(f50031,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X4) != s(bool,f(s(bool,X4))) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X6),'const.pair.,_2'(s(bool,X5),s(X6,X7)))))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(superposition,[],[f49260,f132])).
fof(f50029,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,f(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(superposition,[],[f49260,f56])).
fof(f49281,plain,(
  ( ! [X163,X161,X164,X162] : (s(bool,f(s(bool,'F_0'))) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X161,bool),'const.pair.,_2'(s(X161,X162),s(bool,X163))))) | s(bool,X163) = s(bool,X164) | s(bool,'T_0') = s(bool,X164)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(subsumption_resolution,[],[f44769,f49260])).
fof(f44769,plain,(
  ( ! [X163,X161,X164,X162] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,f(s(bool,'F_0'))) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X161,bool),'const.pair.,_2'(s(X161,X162),s(bool,X163))))) | s(bool,X163) = s(bool,X164) | s(bool,'T_0') = s(bool,X164)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f37032,f133])).
fof(f49697,plain,(
  ( ! [X142,X140,X143,X141,X139] : (s(bool,f(s(bool,'F_0'))) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X139),'const.pair.,_2'(s(bool,X140),s(X139,X141))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X142,bool),'const.pair.,_2'(s(X142,X143),s(bool,X140)))))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(subsumption_resolution,[],[f44764,f49260])).
fof(f44764,plain,(
  ( ! [X142,X140,X143,X141,X139] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,f(s(bool,'F_0'))) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X139),'const.pair.,_2'(s(bool,X140),s(X139,X141))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X142,bool),'const.pair.,_2'(s(X142,X143),s(bool,X140)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f37032,f1473])).
fof(f49276,plain,(
  ( ! [X134,X132,X133,X131] : (s(bool,f(s(bool,'F_0'))) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X131),'const.pair.,_2'(s(bool,X132),s(X131,X133))))) | s(bool,X132) = s(bool,X134) | s(bool,'T_0') = s(bool,X134)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(subsumption_resolution,[],[f44762,f49260])).
fof(f44762,plain,(
  ( ! [X134,X132,X133,X131] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,f(s(bool,'F_0'))) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X131),'const.pair.,_2'(s(bool,X132),s(X131,X133))))) | s(bool,X132) = s(bool,X134) | s(bool,'T_0') = s(bool,X134)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f37032,f132])).
fof(f49266,plain,(
  ( ! [X76,X77] : (s(bool,X76) = s(bool,f(s(bool,'F_0'))) | s(bool,X77) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X76) = s(bool,X77)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(subsumption_resolution,[],[f44727,f49260])).
fof(f44727,plain,(
  ( ! [X76,X77] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,X76) = s(bool,f(s(bool,'F_0'))) | s(bool,X77) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X76) = s(bool,X77)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f37032,f22295])).
fof(f49263,plain,(
  ( ! [X30,X31,X29,X32] : (s(bool,X29) = s(bool,f(s(bool,'F_0'))) | s(bool,X29) = s(bool,X30) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X31,bool),'const.pair.,_2'(s(X31,X32),s(bool,X30)))))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(subsumption_resolution,[],[f44714,f49260])).
fof(f44714,plain,(
  ( ! [X30,X31,X29,X32] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,X29) = s(bool,f(s(bool,'F_0'))) | s(bool,X29) = s(bool,X30) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X31,bool),'const.pair.,_2'(s(X31,X32),s(bool,X30)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f37032,f133])).
fof(f49261,plain,(
  ( ! [X19,X17,X20,X18] : (s(bool,X17) = s(bool,f(s(bool,'F_0'))) | s(bool,X17) = s(bool,X18) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X19),'const.pair.,_2'(s(bool,X18),s(X19,X20)))))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(subsumption_resolution,[],[f44711,f49260])).
fof(f44711,plain,(
  ( ! [X19,X17,X20,X18] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,X17) = s(bool,f(s(bool,'F_0'))) | s(bool,X17) = s(bool,X18) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X19),'const.pair.,_2'(s(bool,X18),s(X19,X20)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f37032,f132])).
fof(f49695,plain,(
  ( ! [X401,X403,X404,X400,X402] : (s(bool,f(s(bool,'F_0'))) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X400,bool),'const.pair.,_2'(s(X400,X401),s(bool,X402))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X403,bool),'const.pair.,_2'(s(X403,X404),s(bool,X402)))))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(subsumption_resolution,[],[f46000,f49260])).
fof(f46000,plain,(
  ( ! [X401,X403,X404,X400,X402] : (s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,f(s(bool,'F_0'))) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X400,bool),'const.pair.,_2'(s(X400,X401),s(bool,X402))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X403,bool),'const.pair.,_2'(s(X403,X404),s(bool,X402)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f37032,f1972])).
fof(f36225,plain,(
  ( ! [X163,X161,X159,X164,X162,X160] : (s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X161,bool),'const.pair.,_2'(s(X161,X162),s(bool,X160))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X163,bool),'const.pair.,_2'(s(X163,X164),s(bool,X159))))) | s(bool,'T_0') = s(bool,f(s(bool,X159))) | s(bool,'F_0') = s(bool,X160)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f1972,f30029])).
fof(f36220,plain,(
  ( ! [X134,X132,X136,X135,X133,X137] : (s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X134,bool),'const.pair.,_2'(s(X134,X135),s(bool,X133))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X136),'const.pair.,_2'(s(bool,X132),s(X136,X137))))) | s(bool,'T_0') = s(bool,f(s(bool,X132))) | s(bool,'F_0') = s(bool,X133)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f1473,f30029])).
fof(f44862,plain,(
  ( ! [X6,X7] : (s(bool,X7) = s(bool,f(s(bool,X6))) | s(bool,'T_0') = s(bool,f(s(bool,X7))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f37032,f40])).
fof(f45330,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,f(s(bool,X2)))))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(forward_demodulation,[],[f45256,f11203])).
fof(f45256,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,f(s(bool,X2))))))) | s(bool,X2) = s(bool,f(s(bool,'F_0')))) ) | ($spl4 | $spl9 | $spl36)),
  inference(condensation,[],[f45255])).
fof(f45255,plain,(
  ( ! [X24,X23,X21,X22] : (s(bool,'T_0') = s(bool,X22) | s(bool,X21) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X23,bool),'const.pair.,_2'(s(X23,X24),s(bool,f(s(bool,X21)))))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(duplicate_literal_removal,[],[f44783])).
fof(f44783,plain,(
  ( ! [X24,X23,X21,X22] : (s(bool,'T_0') = s(bool,X22) | s(bool,X21) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X23,bool),'const.pair.,_2'(s(X23,X24),s(bool,f(s(bool,X21)))))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f37032,f133])).
fof(f45327,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,f(s(bool,X1))),s(X0,X2)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(forward_demodulation,[],[f45258,f11203])).
fof(f45258,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,f(s(bool,X1))),s(X0,X2))))) | s(bool,X1) = s(bool,f(s(bool,'F_0')))) ) | ($spl4 | $spl9 | $spl36)),
  inference(condensation,[],[f45257])).
fof(f45257,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'T_0') = s(bool,X13) | s(bool,X12) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X14),'const.pair.,_2'(s(bool,f(s(bool,X12))),s(X14,X15)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(duplicate_literal_removal,[],[f44781])).
fof(f44781,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'T_0') = s(bool,X13) | s(bool,X12) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X14),'const.pair.,_2'(s(bool,f(s(bool,X12))),s(X14,X15)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f37032,f132])).
fof(f49078,plain,(
  ( ! [X5] : (s(bool,X5) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X5)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f44778,f71])).
fof(f44778,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X5) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X5)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f37032,f40])).
fof(f49070,plain,(
  ( ! [X83,X84] : (s(bool,X84) = s(bool,f(s(bool,X83))) | s(bool,'F_0') = s(bool,X84) | s(bool,'T_0') = s(bool,X83)) ) | ($spl4 | $spl9 | $spl176)),
  inference(subsumption_resolution,[],[f49069,f71])).
fof(f49069,plain,(
  ( ! [X83,X84] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X84) = s(bool,f(s(bool,X83))) | s(bool,'F_0') = s(bool,X84) | s(bool,'T_0') = s(bool,X83)) ) | ($spl4 | $spl176)),
  inference(forward_demodulation,[],[f39751,f1085])).
fof(f39751,plain,(
  ( ! [X83,X84] : (s(bool,X84) = s(bool,f(s(bool,X83))) | s(bool,'F_0') = s(bool,X84) | s(bool,'T_0') = s(bool,X83) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl176)),
  inference(superposition,[],[f30750,f22295])).
fof(f39992,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,f(s(bool,X1))) | s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(duplicate_literal_removal,[],[f39991])).
fof(f39991,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,f(s(bool,X1))) | s(bool,'F_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(condensation,[],[f39990])).
fof(f39990,plain,(
  ( ! [X68,X66,X67,X65] : (s(bool,X66) = s(bool,f(s(bool,X65))) | s(bool,'F_0') = s(bool,X66) | s(bool,'F_0') = s(bool,X67) | s(bool,X67) = s(bool,X68) | s(bool,'T_0') = s(bool,X65)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(subsumption_resolution,[],[f39747,f30026])).
fof(f39747,plain,(
  ( ! [X68,X66,X67,X65] : (s(bool,X66) = s(bool,f(s(bool,X65))) | s(bool,'F_0') = s(bool,X66) | s(bool,'F_0') = s(bool,X67) | s(bool,X67) = s(bool,X68) | s(bool,'T_0') = s(bool,X68) | s(bool,'T_0') = s(bool,X65)) ) | $spl176),
  inference(superposition,[],[f30750,f146])).
fof(f39989,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,f(s(bool,X1))) | s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(duplicate_literal_removal,[],[f39988])).
fof(f39988,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,f(s(bool,X1))) | s(bool,'F_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(condensation,[],[f39987])).
fof(f39987,plain,(
  ( ! [X59,X57,X60,X58] : (s(bool,X58) = s(bool,f(s(bool,X57))) | s(bool,'F_0') = s(bool,X58) | s(bool,'F_0') = s(bool,X59) | s(bool,'T_0') = s(bool,X57) | s(bool,X59) = s(bool,X60)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(subsumption_resolution,[],[f39745,f30026])).
fof(f39745,plain,(
  ( ! [X59,X57,X60,X58] : (s(bool,X58) = s(bool,f(s(bool,X57))) | s(bool,'F_0') = s(bool,X58) | s(bool,'F_0') = s(bool,X59) | s(bool,'T_0') = s(bool,X57) | s(bool,X59) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) ) | $spl176),
  inference(superposition,[],[f30750,f146])).
fof(f39738,plain,(
  ( ! [X30,X31,X29,X32] : (s(bool,X30) = s(bool,f(s(bool,X29))) | s(bool,'F_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X31,bool),'const.pair.,_2'(s(X31,X32),s(bool,'F_0')))))) ) | $spl176),
  inference(superposition,[],[f30750,f133])).
fof(f39735,plain,(
  ( ! [X19,X17,X20,X18] : (s(bool,X18) = s(bool,f(s(bool,X17))) | s(bool,'F_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X17) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X19),'const.pair.,_2'(s(bool,'F_0'),s(X19,X20)))))) ) | $spl176),
  inference(superposition,[],[f30750,f132])).
fof(f49063,plain,(
  ( ! [X14,X13] : (s(bool,X14) = s(bool,f(s(bool,X13))) | s(bool,'F_0') = s(bool,X14) | s(bool,'T_0') = s(bool,X13)) ) | ($spl9 | $spl176)),
  inference(subsumption_resolution,[],[f39733,f71])).
fof(f39733,plain,(
  ( ! [X14,X13] : (s(bool,X14) = s(bool,f(s(bool,X13))) | s(bool,'F_0') = s(bool,X14) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X13)) ) | $spl176),
  inference(superposition,[],[f30750,f56])).
fof(f49061,plain,(
  ( ! [X12,X11] : (s(bool,X12) = s(bool,f(s(bool,X11))) | s(bool,'F_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X11)) ) | ($spl9 | $spl176)),
  inference(subsumption_resolution,[],[f39732,f71])).
fof(f39732,plain,(
  ( ! [X12,X11] : (s(bool,X12) = s(bool,f(s(bool,X11))) | s(bool,'F_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,'F_0')) ) | $spl176),
  inference(superposition,[],[f30750,f56])).
fof(f39729,plain,(
  ( ! [X6,X7] : (s(bool,X7) = s(bool,f(s(bool,X6))) | s(bool,'F_0') = s(bool,X7) | s(bool,'T_0') = s(bool,X6)) ) | $spl176),
  inference(superposition,[],[f30750,f40])).
fof(f49053,plain,(
  ( ! [X231,X232] : (s(bool,X231) != s(bool,f(s(bool,X231))) | s(bool,'T_0') = s(bool,f(s(bool,X232))) | s(bool,X231) = s(bool,X232)) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f36558,f71])).
fof(f36558,plain,(
  ( ! [X231,X232] : (s(bool,X231) != s(bool,f(s(bool,X231))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(bool,X232))) | s(bool,X231) = s(bool,X232)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f25786,f30029])).
fof(f49051,plain,(
  ( ! [X300,X299] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X299,bool),'const.pair.,_2'(s(X299,X300),s(bool,f(s(bool,'F_0')))))))) ) | ($spl9 | $spl36)),
  inference(subsumption_resolution,[],[f32925,f71])).
fof(f32925,plain,(
  ( ! [X300,X299] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X299,bool),'const.pair.,_2'(s(X299,X300),s(bool,f(s(bool,'F_0')))))))) ) | $spl36),
  inference(forward_demodulation,[],[f32433,f44])).
fof(f32433,plain,(
  ( ! [X300,X298,X297,X299] : (s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X297,bool),'const.pair.,_2'(s(X297,X298),s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X299,bool),'const.pair.,_2'(s(X299,X300),s(bool,f(s(bool,'F_0')))))))) ) | $spl36),
  inference(superposition,[],[f1972,f11203])).
fof(f49049,plain,(
  ( ! [X83,X84] : (s(bool,X84) = s(bool,f(s(bool,X83))) | s(bool,'F_0') = s(bool,X84) | s(bool,'T_0') = s(bool,X83)) ) | ($spl4 | $spl9 | $spl176)),
  inference(subsumption_resolution,[],[f49048,f71])).
fof(f49048,plain,(
  ( ! [X83,X84] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X84) = s(bool,f(s(bool,X83))) | s(bool,'F_0') = s(bool,X84) | s(bool,'T_0') = s(bool,X83)) ) | ($spl4 | $spl176)),
  inference(forward_demodulation,[],[f31645,f1085])).
fof(f31645,plain,(
  ( ! [X83,X84] : (s(bool,X84) = s(bool,f(s(bool,X83))) | s(bool,'F_0') = s(bool,X84) | s(bool,'T_0') = s(bool,X83) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl176)),
  inference(superposition,[],[f30750,f22295])).
fof(f31910,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,f(s(bool,X1))) | s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(duplicate_literal_removal,[],[f31909])).
fof(f31909,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,f(s(bool,X1))) | s(bool,'F_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(condensation,[],[f31908])).
fof(f31908,plain,(
  ( ! [X68,X66,X67,X65] : (s(bool,X66) = s(bool,f(s(bool,X65))) | s(bool,'F_0') = s(bool,X66) | s(bool,'F_0') = s(bool,X67) | s(bool,X67) = s(bool,X68) | s(bool,'T_0') = s(bool,X65)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(subsumption_resolution,[],[f31641,f30026])).
fof(f31641,plain,(
  ( ! [X68,X66,X67,X65] : (s(bool,X66) = s(bool,f(s(bool,X65))) | s(bool,'F_0') = s(bool,X66) | s(bool,'F_0') = s(bool,X67) | s(bool,X67) = s(bool,X68) | s(bool,'T_0') = s(bool,X68) | s(bool,'T_0') = s(bool,X65)) ) | $spl176),
  inference(superposition,[],[f30750,f146])).
fof(f31906,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,f(s(bool,X1))) | s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(duplicate_literal_removal,[],[f31905])).
fof(f31905,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,f(s(bool,X1))) | s(bool,'F_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(condensation,[],[f31904])).
fof(f31904,plain,(
  ( ! [X59,X57,X60,X58] : (s(bool,X58) = s(bool,f(s(bool,X57))) | s(bool,'F_0') = s(bool,X58) | s(bool,'F_0') = s(bool,X59) | s(bool,'T_0') = s(bool,X57) | s(bool,X59) = s(bool,X60)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(subsumption_resolution,[],[f31639,f30026])).
fof(f31639,plain,(
  ( ! [X59,X57,X60,X58] : (s(bool,X58) = s(bool,f(s(bool,X57))) | s(bool,'F_0') = s(bool,X58) | s(bool,'F_0') = s(bool,X59) | s(bool,'T_0') = s(bool,X57) | s(bool,X59) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) ) | $spl176),
  inference(superposition,[],[f30750,f146])).
fof(f31632,plain,(
  ( ! [X30,X31,X29,X32] : (s(bool,X30) = s(bool,f(s(bool,X29))) | s(bool,'F_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X31,bool),'const.pair.,_2'(s(X31,X32),s(bool,'F_0')))))) ) | $spl176),
  inference(superposition,[],[f30750,f133])).
fof(f31629,plain,(
  ( ! [X19,X17,X20,X18] : (s(bool,X18) = s(bool,f(s(bool,X17))) | s(bool,'F_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X17) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X19),'const.pair.,_2'(s(bool,'F_0'),s(X19,X20)))))) ) | $spl176),
  inference(superposition,[],[f30750,f132])).
fof(f49042,plain,(
  ( ! [X14,X13] : (s(bool,X14) = s(bool,f(s(bool,X13))) | s(bool,'F_0') = s(bool,X14) | s(bool,'T_0') = s(bool,X13)) ) | ($spl9 | $spl176)),
  inference(subsumption_resolution,[],[f31627,f71])).
fof(f31627,plain,(
  ( ! [X14,X13] : (s(bool,X14) = s(bool,f(s(bool,X13))) | s(bool,'F_0') = s(bool,X14) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X13)) ) | $spl176),
  inference(superposition,[],[f30750,f56])).
fof(f49040,plain,(
  ( ! [X12,X11] : (s(bool,X12) = s(bool,f(s(bool,X11))) | s(bool,'F_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X11)) ) | ($spl9 | $spl176)),
  inference(subsumption_resolution,[],[f31626,f71])).
fof(f31626,plain,(
  ( ! [X12,X11] : (s(bool,X12) = s(bool,f(s(bool,X11))) | s(bool,'F_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,'F_0')) ) | $spl176),
  inference(superposition,[],[f30750,f56])).
fof(f49038,plain,(
  ( ! [X290,X289] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X289),'const.pair.,_2'(s(bool,f(s(bool,'F_0'))),s(X289,X290)))))) ) | ($spl9 | $spl36)),
  inference(subsumption_resolution,[],[f31491,f71])).
fof(f31491,plain,(
  ( ! [X290,X289] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X289),'const.pair.,_2'(s(bool,f(s(bool,'F_0'))),s(X289,X290)))))) ) | $spl36),
  inference(forward_demodulation,[],[f31068,f44])).
fof(f31068,plain,(
  ( ! [X288,X290,X287,X289] : (s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X287,bool),'const.pair.,_2'(s(X287,X288),s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X289),'const.pair.,_2'(s(bool,f(s(bool,'F_0'))),s(X289,X290)))))) ) | $spl36),
  inference(superposition,[],[f1473,f11203])).
fof(f45400,plain,(
  ( ! [X81,X82] : (s(bool,X82) = s(bool,f(s(bool,X81))) | s(bool,'T_0') = s(bool,f(s(bool,X82))) | s(bool,X81) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f44883,f71])).
fof(f44883,plain,(
  ( ! [X81,X82] : (s(bool,X82) = s(bool,f(s(bool,X81))) | s(bool,'T_0') = s(bool,f(s(bool,X82))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X81) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f37032,f22295])).
fof(f45396,plain,(
  ( ! [X61,X64,X62,X63] : (s(bool,X62) = s(bool,f(s(bool,X61))) | s(bool,'T_0') = s(bool,f(s(bool,X62))) | s(bool,X61) = s(bool,X63) | s(bool,X63) = s(bool,X64) | s(bool,'T_0') = s(bool,X64)) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f44879,f71])).
fof(f44879,plain,(
  ( ! [X61,X64,X62,X63] : (s(bool,X62) = s(bool,f(s(bool,X61))) | s(bool,'T_0') = s(bool,f(s(bool,X62))) | s(bool,X61) = s(bool,X63) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X63) = s(bool,X64) | s(bool,'T_0') = s(bool,X64)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f37032,f146])).
fof(f41687,plain,(
  ( ! [X28,X26,X27] : (s(bool,X27) = s(bool,f(s(bool,X26))) | s(bool,'T_0') = s(bool,X27) | s(bool,X28) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X26) = s(bool,X28)) ) | ($spl4 | $spl9 | $spl36 | $spl40)),
  inference(superposition,[],[f41609,f22295])).
fof(f41685,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s(bool,X17) = s(bool,f(s(bool,X16))) | s(bool,'T_0') = s(bool,X17) | s(bool,X18) = s(bool,X19) | s(bool,X16) = s(bool,X19) | s(bool,X18) = s(bool,X20) | s(bool,'T_0') = s(bool,X20)) ) | ($spl4 | $spl9 | $spl36 | $spl40)),
  inference(superposition,[],[f41609,f146])).
fof(f41684,plain,(
  ( ! [X14,X12,X15,X13,X11] : (s(bool,X12) = s(bool,f(s(bool,X11))) | s(bool,'T_0') = s(bool,X12) | s(bool,X11) = s(bool,X13) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X14,bool),'const.pair.,_2'(s(X14,X15),s(bool,X13)))))) ) | ($spl4 | $spl9 | $spl36 | $spl40)),
  inference(superposition,[],[f41609,f133])).
fof(f41683,plain,(
  ( ! [X6,X10,X8,X7,X9] : (s(bool,X7) = s(bool,f(s(bool,X6))) | s(bool,'T_0') = s(bool,X7) | s(bool,X6) = s(bool,X8) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X9),'const.pair.,_2'(s(bool,X8),s(X9,X10)))))) ) | ($spl4 | $spl9 | $spl36 | $spl40)),
  inference(superposition,[],[f41609,f132])).
fof(f41681,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl4 | $spl9 | $spl36 | $spl40)),
  inference(superposition,[],[f41609,f56])).
fof(f42049,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,f(s(bool,'T_0'))) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X2)) ) | ($spl4 | $spl9 | $spl36 | $spl40)),
  inference(condensation,[],[f42048])).
fof(f42048,plain,(
  ( ! [X54,X56,X55,X53] : (s(bool,X54) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X53) | s(bool,X54) = s(bool,X55) | s(bool,X55) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) ) | ($spl4 | $spl9 | $spl36 | $spl40)),
  inference(duplicate_literal_removal,[],[f41640])).
fof(f41640,plain,(
  ( ! [X54,X56,X55,X53] : (s(bool,X54) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X53) | s(bool,X54) = s(bool,X55) | s(bool,'T_0') = s(bool,X53) | s(bool,X55) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) ) | ($spl4 | $spl9 | $spl36 | $spl40)),
  inference(superposition,[],[f41609,f146])).
fof(f11293,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') = s(bool,f(s(bool,X8))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X9),'const.pair.,_2'(s(bool,X8),s(X9,X10)))))) ) | ($spl9 | $spl36)),
  inference(subsumption_resolution,[],[f11229,f71])).
fof(f11229,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') = s(bool,f(s(bool,X8))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X9),'const.pair.,_2'(s(bool,X8),s(X9,X10)))))) ) | $spl36),
  inference(superposition,[],[f11203,f132])).
fof(f36580,plain,(
  ( ! [X315,X316,X318,X314,X317] : (s(bool,X314) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X314),s(bool,X315))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X316,bool),'const.pair.,_2'(s(X316,X317),s(bool,X315))))) | s(bool,'T_0') = s(bool,f(s(bool,X318))) | s(bool,X314) = s(bool,X318)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f1972,f30029])).
fof(f36577,plain,(
  ( ! [X300,X302,X304,X301,X303] : (s(bool,X300) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X300),s(bool,X301))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X302),'const.pair.,_2'(s(bool,X301),s(X302,X303))))) | s(bool,'T_0') = s(bool,f(s(bool,X304))) | s(bool,X300) = s(bool,X304)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f1473,f30029])).
fof(f37203,plain,(
  ( ! [X144,X142,X140,X143,X141] : (s(bool,X141) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X142,bool),'const.pair.,_2'(s(X142,X143),s(bool,X140))))) | s(bool,'T_0') = s(bool,X141) | s(bool,'T_0') = s(bool,f(s(bool,X144))) | s(bool,X140) = s(bool,X144)) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f36536,f71])).
fof(f36536,plain,(
  ( ! [X144,X142,X140,X143,X141] : (s(bool,X141) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X142,bool),'const.pair.,_2'(s(X142,X143),s(bool,X140))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X141) | s(bool,'T_0') = s(bool,f(s(bool,X144))) | s(bool,X140) = s(bool,X144)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f157,f30029])).
fof(f37202,plain,(
  ( ! [X134,X132,X135,X133,X131] : (s(bool,X132) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X133),'const.pair.,_2'(s(bool,X131),s(X133,X134))))) | s(bool,'T_0') = s(bool,X132) | s(bool,'T_0') = s(bool,f(s(bool,X135))) | s(bool,X131) = s(bool,X135)) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f36535,f71])).
fof(f36535,plain,(
  ( ! [X134,X132,X135,X133,X131] : (s(bool,X132) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X133),'const.pair.,_2'(s(bool,X131),s(X133,X134))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X132) | s(bool,'T_0') = s(bool,f(s(bool,X135))) | s(bool,X131) = s(bool,X135)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f156,f30029])).
fof(f36510,plain,(
  ( ! [X24,X23,X22] : (s(bool,X23) != s(bool,f(s(bool,X22))) | s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,f(s(bool,X24))) | s(bool,X22) = s(bool,X24)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f25786,f30029])).
fof(f36509,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') = s(bool,f(s(bool,X20))) | s(bool,'T_0') = s(bool,f(s(bool,X21))) | s(bool,X20) = s(bool,X21)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f11203,f30029])).
fof(f36507,plain,(
  ( ! [X10,X8,X7,X11,X9] : (s(X8,f(s(bool,X7))) != s(X8,g(s(bool,X7))) | s(fun(bool,X8),X9) = s(fun(bool,X8),X10) | s(bool,'T_0') = s(bool,sK6(X9,X10,X8,bool)) | s(bool,'T_0') = s(bool,f(s(bool,X11))) | s(bool,X7) = s(bool,X11)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f365,f30029])).
fof(f36506,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') != s(bool,X5) | s(bool,'T_0') = s(bool,f(s(bool,X6))) | s(bool,X5) = s(bool,X6)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f71,f30029])).
fof(f37191,plain,(
  ( ! [X335,X336] : (s(bool,X335) != s(bool,X336) | s(bool,'T_0') = s(bool,X336) | s(bool,'T_0') = s(bool,f(s(bool,X335)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f37190,f71])).
fof(f37190,plain,(
  ( ! [X335,X336] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X335) != s(bool,X336) | s(bool,'T_0') = s(bool,X336) | s(bool,'T_0') = s(bool,f(s(bool,X335)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(forward_demodulation,[],[f36443,f11203])).
fof(f36443,plain,(
  ( ! [X335,X336] : (s(bool,X335) != s(bool,X336) | s(bool,'T_0') = s(bool,X336) | s(bool,'T_0') = s(bool,f(s(bool,X335))) | s(bool,'F_0') = s(bool,f(s(bool,'F_0')))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f25786,f30029])).
fof(f37187,plain,(
  ( ! [X296,X297] : (s(bool,X296) != s(bool,X297) | s(bool,X297) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,'T_0') = s(bool,f(s(bool,X296)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f36426,f71])).
fof(f36426,plain,(
  ( ! [X296,X297] : (s(bool,X296) != s(bool,X297) | s(bool,X297) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,'T_0') = s(bool,f(s(bool,X296))) | s(bool,'T_0') = s(bool,'F_0')) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f25261,f30029])).
fof(f37183,plain,(
  ( ! [X271,X273,X275,X272,X274] : (s(X272,f(s(bool,X271))) != s(X272,g(s(bool,X271))) | s(fun(bool,X272),X273) = s(fun(bool,X272),X274) | s(bool,X275) = s(bool,sK6(X273,X274,X272,bool)) | s(bool,'T_0') = s(bool,X275) | s(bool,'T_0') = s(bool,f(s(bool,X271)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f36420,f71])).
fof(f36420,plain,(
  ( ! [X271,X273,X275,X272,X274] : (s(X272,f(s(bool,X271))) != s(X272,g(s(bool,X271))) | s(fun(bool,X272),X273) = s(fun(bool,X272),X274) | s(bool,X275) = s(bool,sK6(X273,X274,X272,bool)) | s(bool,'T_0') = s(bool,X275) | s(bool,'T_0') = s(bool,f(s(bool,X271))) | s(bool,'T_0') = s(bool,'F_0')) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f361,f30029])).
fof(f37041,plain,(
  ( ! [X212,X210,X208,X211,X209] : (s(bool,'T_0') = s(bool,f(s(bool,X211))) | s(bool,X211) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X208,bool),'const.pair.,_2'(s(X208,X209),s(bool,X210))))) | s(bool,X210) = s(bool,X212) | s(bool,'T_0') = s(bool,X212)) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f36108,f71])).
fof(f36108,plain,(
  ( ! [X212,X210,X208,X211,X209] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(bool,X211))) | s(bool,X211) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X208,bool),'const.pair.,_2'(s(X208,X209),s(bool,X210))))) | s(bool,X210) = s(bool,X212) | s(bool,'T_0') = s(bool,X212)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30029,f133])).
fof(f37037,plain,(
  ( ! [X167,X165,X169,X166,X168] : (s(bool,'T_0') = s(bool,f(s(bool,X168))) | s(bool,X168) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X165),'const.pair.,_2'(s(bool,X166),s(X165,X167))))) | s(bool,X166) = s(bool,X169) | s(bool,'T_0') = s(bool,X169)) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f36101,f71])).
fof(f36101,plain,(
  ( ! [X167,X165,X169,X166,X168] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(bool,X168))) | s(bool,X168) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X165),'const.pair.,_2'(s(bool,X166),s(X165,X167))))) | s(bool,X166) = s(bool,X169) | s(bool,'T_0') = s(bool,X169)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30029,f132])).
fof(f36024,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,f(s(bool,X2))) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,f(s(bool,X3))) | s(bool,X0) = s(bool,X3)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30029,f30029])).
fof(f37021,plain,(
  ( ! [X165,X163,X164] : (s(bool,X163) = s(bool,X164) | s(bool,'T_0') = s(bool,f(s(bool,X164))) | s(bool,X165) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X163) = s(bool,X165)) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f35992,f71])).
fof(f35992,plain,(
  ( ! [X165,X163,X164] : (s(bool,X163) = s(bool,X164) | s(bool,'T_0') = s(bool,f(s(bool,X164))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X165) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X163) = s(bool,X165)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30029,f22295])).
fof(f37020,plain,(
  ( ! [X156,X154,X157,X155,X153] : (s(bool,X153) = s(bool,X154) | s(bool,'T_0') = s(bool,f(s(bool,X154))) | s(bool,X155) = s(bool,X156) | s(bool,X153) = s(bool,X156) | s(bool,X155) = s(bool,X157) | s(bool,'T_0') = s(bool,X157)) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f35990,f71])).
fof(f35990,plain,(
  ( ! [X156,X154,X157,X155,X153] : (s(bool,X153) = s(bool,X154) | s(bool,'T_0') = s(bool,f(s(bool,X154))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X155) = s(bool,X156) | s(bool,X153) = s(bool,X156) | s(bool,X155) = s(bool,X157) | s(bool,'T_0') = s(bool,X157)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30029,f146])).
fof(f37019,plain,(
  ( ! [X152,X151,X149,X150,X148] : (s(bool,X148) = s(bool,X149) | s(bool,'T_0') = s(bool,f(s(bool,X149))) | s(bool,X148) = s(bool,X150) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X151,bool),'const.pair.,_2'(s(X151,X152),s(bool,X150)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f35989,f71])).
fof(f35989,plain,(
  ( ! [X152,X151,X149,X150,X148] : (s(bool,X148) = s(bool,X149) | s(bool,'T_0') = s(bool,f(s(bool,X149))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X148) = s(bool,X150) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X151,bool),'const.pair.,_2'(s(X151,X152),s(bool,X150)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30029,f133])).
fof(f37018,plain,(
  ( ! [X146,X144,X147,X145,X143] : (s(bool,X143) = s(bool,X144) | s(bool,'T_0') = s(bool,f(s(bool,X144))) | s(bool,X143) = s(bool,X145) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X146),'const.pair.,_2'(s(bool,X145),s(X146,X147)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f35988,f71])).
fof(f35988,plain,(
  ( ! [X146,X144,X147,X145,X143] : (s(bool,X143) = s(bool,X144) | s(bool,'T_0') = s(bool,f(s(bool,X144))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X143) = s(bool,X145) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X146),'const.pair.,_2'(s(bool,X145),s(X146,X147)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30029,f132])).
fof(f34117,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (s(bool,X20) != s(bool,X21) | s(bool,X21) = s(bool,X22) | s(bool,'F_0') = s(bool,X22) | s(bool,X23) = s(bool,X24) | s(bool,X20) = s(bool,X24) | s(bool,X23) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30026,f146])).
fof(f34113,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30026,f56])).
fof(f33751,plain,(
  ( ! [X21,X20] : (s(bool,'F_0') = s(bool,f(s(bool,X20))) | s(bool,X21) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X20) = s(bool,X21)) ) | ($spl4 | $spl40)),
  inference(superposition,[],[f2561,f22295])).
fof(f33637,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X1) = s(bool,X3) | s(bool,X0) = s(bool,X3)) ) | ($spl4 | $spl9 | $spl36)),
  inference(condensation,[],[f33636])).
fof(f33636,plain,(
  ( ! [X70,X72,X71,X69,X73] : (s(bool,X70) != s(bool,X71) | s(bool,'T_0') = s(bool,X69) | s(bool,'F_0') = s(bool,X71) | s(bool,X70) = s(bool,X72) | s(bool,X72) = s(bool,X73) | s(bool,'T_0') = s(bool,X73)) ) | ($spl4 | $spl9 | $spl36)),
  inference(duplicate_literal_removal,[],[f33569])).
fof(f33569,plain,(
  ( ! [X70,X72,X71,X69,X73] : (s(bool,X70) != s(bool,X71) | s(bool,'T_0') = s(bool,X69) | s(bool,'F_0') = s(bool,X71) | s(bool,X70) = s(bool,X72) | s(bool,'T_0') = s(bool,X69) | s(bool,X72) = s(bool,X73) | s(bool,'T_0') = s(bool,X73)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f29101,f146])).
fof(f33700,plain,(
  ( ! [X156,X154,X155] : (s(bool,X154) != s(bool,X155) | s(bool,'T_0') = s(bool,X155) | s(bool,X156) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X154) = s(bool,X156)) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f33525,f71])).
fof(f33525,plain,(
  ( ! [X156,X154,X155] : (s(bool,X154) != s(bool,X155) | s(bool,'T_0') = s(bool,X155) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X156) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X154) = s(bool,X156)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f29101,f22295])).
fof(f33699,plain,(
  ( ! [X146,X144,X147,X145,X148] : (s(bool,X144) != s(bool,X145) | s(bool,'T_0') = s(bool,X145) | s(bool,X146) = s(bool,X147) | s(bool,X144) = s(bool,X147) | s(bool,X146) = s(bool,X148) | s(bool,'T_0') = s(bool,X148)) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f33523,f71])).
fof(f33523,plain,(
  ( ! [X146,X144,X147,X145,X148] : (s(bool,X144) != s(bool,X145) | s(bool,'T_0') = s(bool,X145) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X146) = s(bool,X147) | s(bool,X144) = s(bool,X147) | s(bool,X146) = s(bool,X148) | s(bool,'T_0') = s(bool,X148)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f29101,f146])).
fof(f32112,plain,(
  ( ! [X57,X54,X58,X56,X55] : (s(bool,X54) = s(bool,X55) | s(bool,'F_0') = s(bool,X55) | s(bool,X56) = s(bool,X57) | s(bool,X54) = s(bool,X57) | s(bool,X56) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f32029,f37])).
fof(f32029,plain,(
  ( ! [X57,X54,X58,X56,X55] : (~p(s(bool,'T_0')) | s(bool,X54) = s(bool,X55) | s(bool,'F_0') = s(bool,X55) | s(bool,X56) = s(bool,X57) | s(bool,X54) = s(bool,X57) | s(bool,X56) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30024,f146])).
fof(f29642,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f28992,f28992])).
fof(f29101,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f24431,f25792])).
fof(f28801,plain,(
  ( ! [X28,X26,X27] : (s(bool,X26) != s(bool,X27) | s(bool,X27) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X28) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X26) = s(bool,X28)) ) | ($spl4 | $spl9)),
  inference(superposition,[],[f25261,f22295])).
fof(f28799,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s(bool,X16) != s(bool,X17) | s(bool,X17) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X18) = s(bool,X19) | s(bool,X16) = s(bool,X19) | s(bool,X18) = s(bool,X20) | s(bool,'T_0') = s(bool,X20)) ) | ($spl4 | $spl9)),
  inference(superposition,[],[f25261,f146])).
fof(f28929,plain,(
  ( ! [X14,X12,X15,X13,X11] : (s(bool,'T_0') = s(bool,X12) | s(bool,X11) != s(bool,X12) | s(bool,X11) = s(bool,X13) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X14,bool),'const.pair.,_2'(s(X14,X15),s(bool,X13)))))) ) | ($spl4 | $spl9)),
  inference(forward_demodulation,[],[f28798,f1085])).
fof(f28798,plain,(
  ( ! [X14,X12,X15,X13,X11] : (s(bool,X11) != s(bool,X12) | s(bool,X12) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X11) = s(bool,X13) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X14,bool),'const.pair.,_2'(s(X14,X15),s(bool,X13)))))) ) | ($spl4 | $spl9)),
  inference(superposition,[],[f25261,f133])).
fof(f28928,plain,(
  ( ! [X6,X10,X8,X7,X9] : (s(bool,'T_0') = s(bool,X7) | s(bool,X6) != s(bool,X7) | s(bool,X6) = s(bool,X8) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X9),'const.pair.,_2'(s(bool,X8),s(X9,X10)))))) ) | ($spl4 | $spl9)),
  inference(forward_demodulation,[],[f28797,f1085])).
fof(f28797,plain,(
  ( ! [X6,X10,X8,X7,X9] : (s(bool,X6) != s(bool,X7) | s(bool,X7) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X6) = s(bool,X8) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X9),'const.pair.,_2'(s(bool,X8),s(X9,X10)))))) ) | ($spl4 | $spl9)),
  inference(superposition,[],[f25261,f132])).
fof(f28795,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl4 | $spl9)),
  inference(superposition,[],[f25261,f56])).
fof(f25787,plain,(
  ( ! [X694] : (s(bool,X694) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X694) != s(bool,f(s(bool,'F_0')))) ) | ($spl4 | $spl9 | $spl36)),
  inference(global_subsumption,[],[f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f67,f64,f132,f133,f146,f156,f157,f200,f361,f363,f71,f54,f1085,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f4170,f4171,f3832,f3890,f3878,f2550,f2549,f3652,f3732,f3828,f3727,f2700,f763,f2874,f2698,f761,f3651,f786,f550,f475,f48,f47,f46,f762,f2762,f2699,f2704,f785,f2763,f4154,f3827,f3726,f3646,f368,f366,f3759,f3758,f2455,f1935,f1480,f1972,f1473,f39,f82,f104,f365,f370,f2706,f2559,f11126,f11125,f11188,f11200,f11161,f11162,f11201,f11203,f11222,f11293,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13826,f13033,f13038,f13852,f13853,f13854,f13169,f13223,f13224,f13225,f13345,f13350,f13427,f13415,f15352,f15341,f15429,f15069,f17276,f17266,f17264,f17226,f17356,f16976,f22573,f18154,f18155,f22303,f22295,f22293,f22291,f22279,f22277,f22275,f22219,f22217,f18360,f18362,f18363,f22711,f18790,f18797,f18801,f18802,f22898,f19153,f19158,f21203,f21173,f22901,f22902,f22905,f19331,f19338,f19400,f19401,f19402,f19403,f19406,f22928,f22929,f22930,f22931,f19558,f19565,f19569,f19570,f22984,f20200,f23229,f23233,f23235,f24931,f24930,f24854,f24822,f24800,f23782,f23787,f25217,f25261,f25262,f25265,f25280,f24214,f25784,f25785,f25786,f24220,f24215])).
fof(f25280,plain,(
  ( ! [X383] : (s(bool,'T_0') = s(bool,f(s(bool,X383))) | s(bool,X383) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f24129,f71])).
fof(f24129,plain,(
  ( ! [X383] : (s(bool,'T_0') = s(bool,f(s(bool,X383))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X383) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl36)),
  inference(superposition,[],[f11203,f22295])).
fof(f25265,plain,(
  ( ! [X360,X359] : (s(bool,X360) != s(bool,f(s(bool,X359))) | s(bool,'T_0') = s(bool,X360) | s(bool,X359) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f24119,f71])).
fof(f24119,plain,(
  ( ! [X360,X359] : (s(bool,X360) != s(bool,f(s(bool,X359))) | s(bool,'T_0') = s(bool,X360) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X359) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl36)),
  inference(superposition,[],[f2455,f22295])).
fof(f22905,plain,(
  ( ! [X490,X489,X491] : (s(bool,'T_0') = s(bool,f(s(bool,X489))) | s(bool,X489) = s(bool,X490) | s(bool,X490) = s(bool,X491) | s(bool,'T_0') = s(bool,X491)) ) | ($spl9 | $spl36)),
  inference(subsumption_resolution,[],[f19246,f71])).
fof(f19246,plain,(
  ( ! [X490,X489,X491] : (s(bool,'T_0') = s(bool,f(s(bool,X489))) | s(bool,X489) = s(bool,X490) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X490) = s(bool,X491) | s(bool,'T_0') = s(bool,X491)) ) | $spl36),
  inference(superposition,[],[f11203,f146])).
fof(f22902,plain,(
  ( ! [X445,X447,X446,X448] : (s(bool,X446) != s(bool,f(s(bool,X445))) | s(bool,'T_0') = s(bool,X446) | s(bool,X445) = s(bool,X447) | s(bool,X447) = s(bool,X448) | s(bool,'T_0') = s(bool,X448)) ) | ($spl9 | $spl36)),
  inference(subsumption_resolution,[],[f19236,f71])).
fof(f19236,plain,(
  ( ! [X445,X447,X446,X448] : (s(bool,X446) != s(bool,f(s(bool,X445))) | s(bool,'T_0') = s(bool,X446) | s(bool,X445) = s(bool,X447) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X447) = s(bool,X448) | s(bool,'T_0') = s(bool,X448)) ) | $spl36),
  inference(superposition,[],[f2455,f146])).
fof(f11200,plain,(
  ( ! [X31] : (s(bool,X31) != s(bool,f(s(bool,X31))) | s(bool,'T_0') = s(bool,X31)) ) | ($spl9 | $spl36)),
  inference(subsumption_resolution,[],[f11139,f71])).
fof(f11139,plain,(
  ( ! [X31] : (s(bool,X31) != s(bool,f(s(bool,X31))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X31)) ) | $spl36),
  inference(superposition,[],[f2455,f40])).
fof(f496,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2))))) | s(bool,'T_0') = s(bool,X2)) ) | $spl13),
  inference(trivial_inequality_removal,[],[f492])).
fof(f492,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X2) = s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2))))) | s(bool,'T_0') = s(bool,X2)) ) | $spl13),
  inference(superposition,[],[f439,f56])).
fof(f502,plain,(
  s(bool,'F_0') = s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2))))) | $spl13),
  inference(unit_resulting_resolution,[],[f39,f439,f200])).
fof(f74388,plain,(
  ( ! [X243,X244] : (p(s(bool,f(s(bool,X244)))) | s(bool,'F_0') = s(bool,X243) | s(bool,'T_0') = s(bool,X244)) ) | ($spl172 | $spl176)),
  inference(duplicate_literal_removal,[],[f74072])).
fof(f74072,plain,(
  ( ! [X243,X244] : (p(s(bool,f(s(bool,X244)))) | s(bool,'F_0') = s(bool,X243) | s(bool,'F_0') = s(bool,X243) | s(bool,'T_0') = s(bool,X244)) ) | ($spl172 | $spl176)),
  inference(superposition,[],[f29278,f31623])).
fof(f74387,plain,(
  ( ! [X257,X259,X258] : (s(bool,X259) = s(bool,f(s(bool,X258))) | s(bool,'T_0') = s(bool,f(s(bool,X259))) | s(bool,'F_0') = s(bool,X257) | s(bool,'T_0') = s(bool,X258)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(duplicate_literal_removal,[],[f74077])).
fof(f74077,plain,(
  ( ! [X257,X259,X258] : (s(bool,X259) = s(bool,f(s(bool,X258))) | s(bool,'T_0') = s(bool,f(s(bool,X259))) | s(bool,'F_0') = s(bool,X257) | s(bool,'F_0') = s(bool,X257) | s(bool,'T_0') = s(bool,X258)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(superposition,[],[f30029,f31623])).
fof(f74748,plain,(
  ( ! [X305] : (s(bool,f(s(bool,X305))) != s(bool,f(s(bool,f(s(bool,X305))))) | s(bool,'T_0') = s(bool,X305)) ) | $spl350),
  inference(cnf_transformation,[],[f74748_D])).
fof(f74748_D,plain,(
  ( ! [X305] : (s(bool,f(s(bool,X305))) != s(bool,f(s(bool,f(s(bool,X305))))) | s(bool,'T_0') = s(bool,X305)) ) <=> ~$spl350),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl350])])).
fof(f74751,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,f(s(bool,X1))) != s(bool,f(s(bool,f(s(bool,X1))))) | s(bool,'T_0') = s(bool,X1)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl176)),
  inference(condensation,[],[f74750])).
fof(f74750,plain,(
  ( ! [X292,X293,X291] : (s(bool,f(s(bool,X292))) != s(bool,f(s(bool,f(s(bool,X292))))) | s(bool,X291) = s(bool,X293) | s(bool,'F_0') = s(bool,X291) | s(bool,'T_0') = s(bool,X292)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl176)),
  inference(subsumption_resolution,[],[f74092,f30026])).
fof(f74092,plain,(
  ( ! [X292,X293,X291] : (s(bool,f(s(bool,X292))) != s(bool,f(s(bool,f(s(bool,X292))))) | s(bool,X291) = s(bool,X293) | s(bool,'T_0') = s(bool,X293) | s(bool,'F_0') = s(bool,X291) | s(bool,'T_0') = s(bool,X292)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl176)),
  inference(superposition,[],[f50029,f31623])).
fof(f74644,plain,(
  ( ! [X310] : (s(bool,'T_0') = s(bool,f(s(bool,X310)))) ) | $spl334),
  inference(cnf_transformation,[],[f74644_D])).
fof(f74644_D,plain,(
  ( ! [X310] : (s(bool,'T_0') = s(bool,f(s(bool,X310)))) ) <=> ~$spl334),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl334])])).
fof(f74749,plain,(
  ( ! [X302,X303] : (s(bool,f(s(bool,X303))) != s(bool,f(s(bool,f(s(bool,X303))))) | s(bool,'T_0') = s(bool,f(s(bool,X302))) | s(bool,'T_0') = s(bool,X303)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl176 | $spl250)),
  inference(subsumption_resolution,[],[f74097,f49698])).
fof(f74097,plain,(
  ( ! [X302,X303] : (s(bool,f(s(bool,X303))) != s(bool,f(s(bool,f(s(bool,X303))))) | s(bool,'T_0') = s(bool,f(s(bool,X302))) | s(bool,'F_0') = s(bool,X302) | s(bool,'T_0') = s(bool,X303)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl176)),
  inference(superposition,[],[f63838,f31623])).
fof(f74379,plain,(
  ( ! [X304,X305] : (s(bool,f(s(bool,X305))) != s(bool,f(s(bool,f(s(bool,X305))))) | s(bool,'F_0') = s(bool,X304) | s(bool,'T_0') = s(bool,X305)) ) | ($spl176 | $spl280)),
  inference(duplicate_literal_removal,[],[f74098])).
fof(f74098,plain,(
  ( ! [X304,X305] : (s(bool,f(s(bool,X305))) != s(bool,f(s(bool,f(s(bool,X305))))) | s(bool,'F_0') = s(bool,X304) | s(bool,'F_0') = s(bool,X304) | s(bool,'T_0') = s(bool,X305)) ) | ($spl176 | $spl280)),
  inference(superposition,[],[f63840,f31623])).
fof(f74746,plain,(
  ( ! [X472,X474,X469,X471,X475,X470] : (s(bool,X474) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X471,bool),X475),s(fun(X469,bool),X470),s('type.pair.prod'(X471,X469),X472))) | s(bool,'F_0') = s(bool,happ(s(fun(X469,bool),X470),s(X469,'const.pair.SND_1'(s('type.pair.prod'(X471,X469),X472))))) | s(bool,'T_0') = s(bool,X474)) ) | $spl348),
  inference(cnf_transformation,[],[f74746_D])).
fof(f74746_D,plain,(
  ( ! [X472,X474,X469,X471,X475,X470] : (s(bool,X474) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X471,bool),X475),s(fun(X469,bool),X470),s('type.pair.prod'(X471,X469),X472))) | s(bool,'F_0') = s(bool,happ(s(fun(X469,bool),X470),s(X469,'const.pair.SND_1'(s('type.pair.prod'(X471,X469),X472))))) | s(bool,'T_0') = s(bool,X474)) ) <=> ~$spl348),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl348])])).
fof(f74148,plain,(
  ( ! [X472,X474,X469,X471,X475,X473,X470] : (p(s(bool,f(s(bool,X473)))) | s(bool,X474) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X471,bool),X475),s(fun(X469,bool),X470),s('type.pair.prod'(X471,X469),X472))) | s(bool,'T_0') = s(bool,X474) | s(bool,'F_0') = s(bool,happ(s(fun(X469,bool),X470),s(X469,'const.pair.SND_1'(s('type.pair.prod'(X471,X469),X472))))) | s(bool,'T_0') = s(bool,X473)) ) | $spl176),
  inference(superposition,[],[f550,f31623])).
fof(f74744,plain,(
  ( ! [X509,X505,X507,X510,X504,X506] : (s(bool,X509) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X504,bool),X505),s(fun(X506,bool),X510),s('type.pair.prod'(X504,X506),X507))) | s(bool,'F_0') = s(bool,happ(s(fun(X504,bool),X505),s(X504,'const.pair.FST_1'(s('type.pair.prod'(X504,X506),X507))))) | s(bool,'T_0') = s(bool,X509)) ) | $spl344),
  inference(cnf_transformation,[],[f74744_D])).
fof(f74744_D,plain,(
  ( ! [X509,X505,X507,X510,X504,X506] : (s(bool,X509) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X504,bool),X505),s(fun(X506,bool),X510),s('type.pair.prod'(X504,X506),X507))) | s(bool,'F_0') = s(bool,happ(s(fun(X504,bool),X505),s(X504,'const.pair.FST_1'(s('type.pair.prod'(X504,X506),X507))))) | s(bool,'T_0') = s(bool,X509)) ) <=> ~$spl344),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl344])])).
fof(f74153,plain,(
  ( ! [X509,X505,X507,X508,X510,X504,X506] : (p(s(bool,f(s(bool,X508)))) | s(bool,X509) = s(bool,'const.pair.PROD_ALL_3'(s(fun(X504,bool),X505),s(fun(X506,bool),X510),s('type.pair.prod'(X504,X506),X507))) | s(bool,'T_0') = s(bool,X509) | s(bool,'F_0') = s(bool,happ(s(fun(X504,bool),X505),s(X504,'const.pair.FST_1'(s('type.pair.prod'(X504,X506),X507))))) | s(bool,'T_0') = s(bool,X508)) ) | $spl176),
  inference(superposition,[],[f475,f31623])).
fof(f74741,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,X1) = s(bool,f(s(bool,X2))) | s(bool,X1) = s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(condensation,[],[f74740])).
fof(f74740,plain,(
  ( ! [X127,X130,X128,X131,X129] : (s(bool,X128) = s(bool,X129) | s(bool,X130) = s(bool,f(s(bool,X127))) | s(bool,X130) = s(bool,X131) | s(bool,'T_0') = s(bool,X131) | s(bool,'F_0') = s(bool,X128) | s(bool,'T_0') = s(bool,X127)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(subsumption_resolution,[],[f74254,f30026])).
fof(f74254,plain,(
  ( ! [X127,X130,X128,X131,X129] : (s(bool,X128) = s(bool,X129) | s(bool,X130) = s(bool,f(s(bool,X127))) | s(bool,'T_0') = s(bool,X129) | s(bool,X130) = s(bool,X131) | s(bool,'T_0') = s(bool,X131) | s(bool,'F_0') = s(bool,X128) | s(bool,'T_0') = s(bool,X127)) ) | $spl176),
  inference(superposition,[],[f146,f31623])).
fof(f74739,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,X1) = s(bool,f(s(bool,X2))) | s(bool,X1) = s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(condensation,[],[f74738])).
fof(f74738,plain,(
  ( ! [X140,X138,X141,X139,X137] : (s(bool,X138) = s(bool,X139) | s(bool,X140) = s(bool,f(s(bool,X137))) | s(bool,X140) = s(bool,X141) | s(bool,'T_0') = s(bool,X141) | s(bool,'F_0') = s(bool,X138) | s(bool,'T_0') = s(bool,X137)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(subsumption_resolution,[],[f74256,f30026])).
fof(f74256,plain,(
  ( ! [X140,X138,X141,X139,X137] : (s(bool,X138) = s(bool,X139) | s(bool,X140) = s(bool,f(s(bool,X137))) | s(bool,X140) = s(bool,X141) | s(bool,'T_0') = s(bool,X141) | s(bool,'T_0') = s(bool,X139) | s(bool,'F_0') = s(bool,X138) | s(bool,'T_0') = s(bool,X137)) ) | $spl176),
  inference(superposition,[],[f146,f31623])).
fof(f74731,plain,(
  ( ! [X265,X267,X264,X266] : (s(bool,'T_0') = s(bool,X266) | s(bool,X267) = s(bool,f(s(bool,X264))) | s(bool,X266) = s(bool,X267) | s(bool,'F_0') = s(bool,X265) | s(bool,'T_0') = s(bool,X264)) ) | ($spl13 | $spl22 | $spl36 | $spl176)),
  inference(subsumption_resolution,[],[f74286,f40])).
fof(f74286,plain,(
  ( ! [X265,X267,X264,X266] : (s(bool,'T_0') != s(bool,X265) | s(bool,'T_0') = s(bool,X266) | s(bool,X267) = s(bool,f(s(bool,X264))) | s(bool,X266) = s(bool,X267) | s(bool,'F_0') = s(bool,X265) | s(bool,'T_0') = s(bool,X264)) ) | ($spl13 | $spl22 | $spl36 | $spl176)),
  inference(superposition,[],[f22900,f31623])).
fof(f22900,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X2)) ) | ($spl13 | $spl22 | $spl36)),
  inference(condensation,[],[f22899])).
fof(f22899,plain,(
  ( ! [X212,X208,X213,X209] : (s(bool,'T_0') != s(bool,X209) | s(bool,X209) = s(bool,X212) | s(bool,'T_0') = s(bool,X208) | s(bool,X212) = s(bool,X213) | s(bool,'T_0') = s(bool,X213)) ) | ($spl13 | $spl22 | $spl36)),
  inference(subsumption_resolution,[],[f19197,f17356])).
fof(f19197,plain,(
  ( ! [X212,X210,X208,X213,X211,X209] : (s(bool,'T_0') != s(bool,X209) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X210,bool),'const.pair.,_2'(s(X210,X211),s(bool,X208))))) | s(bool,X209) = s(bool,X212) | s(bool,'T_0') = s(bool,X208) | s(bool,X212) = s(bool,X213) | s(bool,'T_0') = s(bool,X213)) ) | ($spl13 | $spl22)),
  inference(superposition,[],[f7282,f146])).
fof(f7282,plain,(
  ( ! [X59,X60,X58] : (s(bool,'T_0') != s(bool,X60) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X58,bool),'const.pair.,_2'(s(X58,X59),s(bool,X60)))))) ) | ($spl13 | $spl22)),
  inference(forward_demodulation,[],[f7226,f43])).
fof(f7226,plain,(
  ( ! [X59,X60,X58] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X60),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X58,bool),'const.pair.,_2'(s(X58,X59),s(bool,X60)))))) ) | ($spl13 | $spl22)),
  inference(superposition,[],[f7061,f44])).
fof(f7061,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X0),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl13 | $spl22)),
  inference(condensation,[],[f7060])).
fof(f7060,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X1),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1)) ) | ($spl13 | $spl22)),
  inference(duplicate_literal_removal,[],[f6979])).
fof(f6979,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X1),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | ($spl13 | $spl22)),
  inference(superposition,[],[f6527,f40])).
fof(f6527,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X0),s(bool,X1))))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | ($spl13 | $spl22)),
  inference(superposition,[],[f6026,f40])).
fof(f6026,plain,(
  ( ! [X85] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,X85))))) | s(bool,'T_0') = s(bool,X85)) ) | ($spl13 | $spl22)),
  inference(subsumption_resolution,[],[f5852,f439])).
fof(f5852,plain,(
  ( ! [X85] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,X85))))) | s(bool,'T_0') = s(bool,X85) | s(bool,'T_0') = s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2)))))) ) | ($spl13 | $spl22)),
  inference(superposition,[],[f439,f5520])).
fof(f74360,plain,(
  ( ! [X313,X312,X311] : (s(bool,'F_0') = s(bool,X312) | s(bool,'T_0') = s(bool,f(s(bool,X313))) | s(bool,X313) = s(bool,f(s(bool,X311))) | s(bool,'T_0') = s(bool,X311)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(duplicate_literal_removal,[],[f74303])).
fof(f74303,plain,(
  ( ! [X313,X312,X311] : (s(bool,'F_0') = s(bool,X312) | s(bool,'T_0') = s(bool,f(s(bool,X313))) | s(bool,X313) = s(bool,f(s(bool,X311))) | s(bool,'F_0') = s(bool,X312) | s(bool,'T_0') = s(bool,X311)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(superposition,[],[f30029,f31623])).
fof(f74729,plain,(
  ( ! [X321,X320,X322] : (s(bool,'T_0') = s(bool,f(s(bool,X322))) | s(bool,X322) = s(bool,f(s(bool,X320))) | s(bool,'F_0') = s(bool,X321) | s(bool,'T_0') = s(bool,X320)) ) | ($spl4 | $spl9 | $spl36 | $spl172 | $spl176)),
  inference(subsumption_resolution,[],[f74307,f29278])).
fof(f74307,plain,(
  ( ! [X321,X320,X322] : (~p(s(bool,X321)) | s(bool,'T_0') = s(bool,f(s(bool,X322))) | s(bool,X322) = s(bool,f(s(bool,X320))) | s(bool,'F_0') = s(bool,X321) | s(bool,'T_0') = s(bool,X320)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(superposition,[],[f36504,f31623])).
fof(f74667,plain,(
  ( ! [X648,X646,X647] : (s(bool,'F_0') = s(bool,X646) | s(bool,'T_0') = s(bool,X647) | ~p(s(bool,happ(s(fun(bool,bool),X648),s(bool,f(s(bool,X647))))))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl176 | $spl290)),
  inference(global_subsumption,[],[f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f67,f64,f132,f133,f146,f156,f157,f200,f361,f363,f502,f496,f71,f54,f439,f1085,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f4170,f4171,f3832,f3890,f3878,f2550,f2549,f3652,f3732,f3828,f3727,f2700,f763,f2698,f761,f2455,f1935,f1480,f82,f104,f2706,f11126,f11125,f11188,f11200,f11161,f11162,f11222,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13852,f13225,f13350,f13427,f13415,f15429,f17356,f22573,f18154,f18155,f22295,f22291,f22279,f22275,f22219,f22217,f18360,f18362,f18363,f22711,f18802,f19158,f21203,f22901,f22902,f22905,f19338,f19402,f19403,f19570,f22984,f20200,f23229,f23233,f23235,f24931,f24930,f24854,f24822,f24800,f23787,f25217,f25261,f25262,f25265,f25280,f25784,f25785,f24220,f24215,f25787,f25788,f23788,f23783,f25790,f23712,f20161,f20156,f25794,f19927,f19749,f19571,f19566,f25798,f19337,f19159,f18981,f18976,f18803,f18798,f25801,f25802,f13351,f13346,f25804,f13292,f13175,f13170,f13039,f11230,f9980,f9978,f9977,f9974,f6734,f6732,f6731,f6728,f24431,f25850,f25854,f28579,f28795,f28928,f28929,f28799,f28801,f28992,f29101,f29277,f29642,f30026,f30749,f30802,f30902,f31517,f31518,f31523,f32112,f32190,f32954,f32955,f32960,f33699,f33700,f33637,f33624,f33745,f33747,f33748,f33749,f33751,f33888,f33897,f34113,f34117,f35168,f35169,f35170,f35172,f36886,f36851,f36833,f36832,f37018,f37019,f37020,f37021,f36024,f36821,f36820,f37037,f37041,f36264,f36272,f36746,f37183,f37187,f37191,f36506,f36507,f36509,f36510,f37202,f37203,f36577,f36580,f37595,f11203,f11293,f36263,f41608,f41609,f42049,f41681,f41683,f41684,f41685,f41687,f42116,f42119,f41950,f42132,f45396,f45400,f49038,f49040,f49042,f31629,f31632,f31906,f31910,f49049,f49051,f37194,f49053,f36723,f39729,f49061,f49063,f39735,f39738,f39989,f39992,f49070,f49078,f45327,f45330,f44862,f36220,f49246,f36225,f43667,f43893,f43894,f43907,f45693,f45866,f45867,f45883,f49260,f44738,f37032,f49695,f49248,f49261,f49263,f49266,f49269,f49276,f49697,f49281,f36718,f50029,f50031,f50032,f50033,f50035,f45545,f13038,f3651,f15352,f22930,f22929,f22928,f786,f550,f475,f48,f47,f46,f25855,f19925,f19747,f19569,f18801,f762,f2762,f11302,f2704,f19406,f25856,f19914,f19736,f19558,f18790,f2699,f785,f2763,f22931,f45539,f37580,f37578,f37577,f37576,f37574,f37033,f36728,f34476,f34474,f34472,f34425,f24214,f23782,f22898,f19565,f19331,f19153,f18797,f15069,f13826,f13345,f13169,f13033,f4154,f3827,f3726,f3646,f2559,f368,f366,f49702,f49268,f49265,f49037,f45453,f45254,f37594,f37211,f37040,f37038,f36953,f36870,f36709,f36702,f36513,f36501,f36494,f36344,f36337,f36111,f36110,f36025,f36023,f36022,f34518,f34208,f34207,f34119,f34116,f34115,f33645,f33643,f33629,f33627,f33552,f33551,f32991,f32792,f32361,f32194,f32193,f32192,f32154,f32152,f32151,f32126,f32105,f32101,f31911,f31907,f31623,f31621,f31431,f31357,f31001,f30906,f30905,f30904,f30750,f30404,f30179,f30178,f30174,f30172,f30080,f30030,f30029,f30027,f30024,f29905,f29904,f29900,f29898,f29715,f29714,f29581,f29580,f29283,f29256,f29255,f29252,f29247,f29245,f29243,f29241,f25810,f25806,f25792,f25786,f22303,f22293,f22277,f21173,f19401,f19400,f17276,f17226,f15341,f13224,f13223,f3759,f3758,f2874,f1972,f1473,f370,f365,f39,f36343,f45544,f2561,f46456,f45588,f45538,f45454,f45392,f45391,f36874,f36872,f36708,f36696,f36575,f36504,f36500,f36488,f36331,f36184,f35841,f35840,f31903,f31902,f13854,f13853,f11201,f17266,f17264,f16976,f60404,f63512,f63513,f63598,f63517,f63518,f63519,f63528,f63530,f63532,f63533,f63614,f63544,f63546,f63618,f63554,f63621,f63623,f63625,f63627,f63641,f63838,f63839,f63845,f63730,f63731,f63853,f63855,f63862,f63864,f63865,f63869,f64145,f64146,f64062,f64063,f64064,f64065,f64066,f64075,f64077,f64079,f64080,f64161,f64091,f64093,f64165,f64101,f64168,f64170,f64172,f64174,f64188,f64325,f64324,f64321,f64457,f64459,f64460,f64461,f64463,f64629,f64630,f64631,f64640,f64642,f64644,f64645,f64768,f64656,f64658,f64772,f64666,f64670,f64673,f64675,f64782,f64783,f64689,f64820,f64821,f64828,f64836,f64839,f65344,f66452,f65529,f65535,f65537,f65728,f65734,f65736,f65754,f65830,f66588,f66589,f65985,f65991,f65993,f66956,f66957,f66958,f66967,f66969,f66971,f66972,f67066,f66983,f66985,f66987,f66988,f66989,f67070,f67074,f67081,f67089,f67094,f67781,f67872,f67873,f67874,f67881,f67889,f67894,f67866,f68252,f68256,f68413,f68414,f68415,f68419,f68423,f68430,f68434,f68646,f68653,f68661,f68666,f69129,f69131,f69133,f69134,f69138,f69246,f69154,f69156,f69169,f69170,f69181,f69187,f69280,f69281,f69288,f69296,f69301,f69521,f69523,f69525,f69532,f69536,f69542,f70246,f70249,f70256,f70263,f70275,f70160,f70161,f70322,f70198,f70206,f74463,f74466,f74438,f73853,f73854,f74556,f74562,f74376,f74112,f74190])).
fof(f74649,plain,(
  ( ! [X326] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,X326))))) | s(bool,'T_0') = s(bool,X326)) ) | $spl340),
  inference(cnf_transformation,[],[f74649_D])).
fof(f74649_D,plain,(
  ( ! [X326] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,X326))))) | s(bool,'T_0') = s(bool,X326)) ) <=> ~$spl340),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl340])])).
fof(f74373,plain,(
  ( ! [X325,X326] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,X326))))) | s(bool,'F_0') = s(bool,X325) | s(bool,'T_0') = s(bool,X326)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl166 | $spl176 | $spl204 | $spl234)),
  inference(duplicate_literal_removal,[],[f74107])).
fof(f74107,plain,(
  ( ! [X325,X326] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,X326))))) | s(bool,'F_0') = s(bool,X325) | s(bool,'F_0') = s(bool,X325) | s(bool,'T_0') = s(bool,X326)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl166 | $spl176 | $spl204 | $spl234)),
  inference(superposition,[],[f70304,f31623])).
fof(f70304,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(bool,X0))) | s(bool,'F_0') = s(bool,X0)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl166 | $spl204 | $spl234)),
  inference(condensation,[],[f70303])).
fof(f70303,plain,(
  ( ! [X17,X18] : (s(bool,'T_0') != s(bool,f(s(bool,X17))) | s(bool,'F_0') = s(bool,X18) | s(bool,'F_0') = s(bool,X17)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl166 | $spl204 | $spl234)),
  inference(subsumption_resolution,[],[f70171,f66448])).
fof(f66448,plain,(
  ( ! [X341,X340] : (s(bool,X341) != s(bool,f(s(bool,X340))) | s(bool,'F_0') = s(bool,X341) | s(bool,'F_0') = s(bool,X340)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl234)),
  inference(subsumption_resolution,[],[f65440,f49260])).
fof(f65440,plain,(
  ( ! [X341,X340] : (s(bool,X341) != s(bool,f(s(bool,X340))) | s(bool,'F_0') = s(bool,X341) | s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X340)) ) | ($spl4 | $spl9 | $spl36 | $spl234)),
  inference(superposition,[],[f42127,f30029])).
fof(f70171,plain,(
  ( ! [X17,X18] : (s(bool,'T_0') != s(bool,f(s(bool,X17))) | s(bool,X18) = s(bool,f(s(bool,X17))) | s(bool,'F_0') = s(bool,X18) | s(bool,'F_0') = s(bool,X17)) ) | ($spl4 | $spl9 | $spl36 | $spl166 | $spl204)),
  inference(superposition,[],[f49037,f40393])).
fof(f40393,plain,(
  ( ! [X150] : (s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,X150))))) | s(bool,'F_0') = s(bool,X150)) ) | ($spl4 | $spl9 | $spl36 | $spl166 | $spl204)),
  inference(subsumption_resolution,[],[f40301,f39])).
fof(f40301,plain,(
  ( ! [X150] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,X150))))) | s(bool,'F_0') = s(bool,X150)) ) | ($spl4 | $spl9 | $spl36 | $spl166 | $spl204)),
  inference(superposition,[],[f37178,f37102])).
fof(f37178,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') = s(bool,f(s(bool,X0)))) ) | ($spl4 | $spl9 | $spl36 | $spl166)),
  inference(condensation,[],[f37177])).
fof(f37177,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | s(bool,'T_0') = s(bool,f(s(bool,X1)))) ) | ($spl4 | $spl9 | $spl36 | $spl166)),
  inference(condensation,[],[f37176])).
fof(f37176,plain,(
  ( ! [X118,X116,X117] : (p(s(bool,X117)) | s(bool,X116) = s(bool,X118) | s(bool,'T_0') = s(bool,X118) | s(bool,'T_0') = s(bool,f(s(bool,X117)))) ) | ($spl4 | $spl9 | $spl36 | $spl166)),
  inference(subsumption_resolution,[],[f36373,f25851])).
fof(f36373,plain,(
  ( ! [X118,X116,X117] : (p(s(bool,X117)) | s(bool,X116) = s(bool,X118) | s(bool,'T_0') = s(bool,X118) | s(bool,'T_0') = s(bool,f(s(bool,X117))) | s(bool,'F_0') = s(bool,X116)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f200,f30029])).
fof(f74648,plain,(
  ( ! [X321,X320,X319] : (~p(s(bool,happ(s(fun(bool,bool),X321),s(bool,f(s(bool,X320)))))) | s(bool,'T_0') = s(bool,f(s(bool,X319))) | s(bool,'T_0') = s(bool,X320)) ) | ($spl4 | $spl9 | $spl36 | $spl176 | $spl250 | $spl290)),
  inference(subsumption_resolution,[],[f74105,f49698])).
fof(f74105,plain,(
  ( ! [X321,X320,X319] : (~p(s(bool,happ(s(fun(bool,bool),X321),s(bool,f(s(bool,X320)))))) | s(bool,'T_0') = s(bool,f(s(bool,X319))) | s(bool,'F_0') = s(bool,X319) | s(bool,'T_0') = s(bool,X320)) ) | ($spl4 | $spl9 | $spl36 | $spl176 | $spl290)),
  inference(superposition,[],[f68414,f31623])).
fof(f74376,plain,(
  ( ! [X316,X318,X317] : (~p(s(bool,happ(s(fun(bool,bool),X318),s(bool,f(s(bool,X317)))))) | s(bool,'F_0') = s(bool,X316) | s(bool,'T_0') = s(bool,X317)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl176 | $spl290)),
  inference(duplicate_literal_removal,[],[f74104])).
fof(f74104,plain,(
  ( ! [X316,X318,X317] : (~p(s(bool,happ(s(fun(bool,bool),X318),s(bool,f(s(bool,X317)))))) | s(bool,'F_0') = s(bool,X316) | s(bool,'F_0') = s(bool,X316) | s(bool,'T_0') = s(bool,X317)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl176 | $spl290)),
  inference(superposition,[],[f68413,f31623])).
fof(f74645,plain,(
  ( ! [X313] : (s(bool,'T_0') = s(bool,g(s(bool,f(s(bool,X313))))) | s(bool,'T_0') = s(bool,X313)) ) | $spl336),
  inference(cnf_transformation,[],[f74645_D])).
fof(f74645_D,plain,(
  ( ! [X313] : (s(bool,'T_0') = s(bool,g(s(bool,f(s(bool,X313))))) | s(bool,'T_0') = s(bool,X313)) ) <=> ~$spl336),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl336])])).
fof(f74646,plain,(
  ( ! [X315,X314] : (s(bool,'T_0') = s(bool,g(s(bool,f(s(bool,X315))))) | s(bool,'T_0') = s(bool,f(s(bool,X314))) | s(bool,'T_0') = s(bool,X315)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl46 | $spl176 | $spl250)),
  inference(subsumption_resolution,[],[f74103,f49698])).
fof(f74103,plain,(
  ( ! [X315,X314] : (s(bool,'T_0') = s(bool,g(s(bool,f(s(bool,X315))))) | s(bool,'T_0') = s(bool,f(s(bool,X314))) | s(bool,'F_0') = s(bool,X314) | s(bool,'T_0') = s(bool,X315)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl46 | $spl176 | $spl250)),
  inference(superposition,[],[f66941,f31623])).
fof(f66941,plain,(
  ( ! [X197] : (s(bool,'T_0') = s(bool,g(s(bool,X197))) | s(bool,'T_0') = s(bool,f(s(bool,X197)))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl46 | $spl250)),
  inference(forward_demodulation,[],[f66940,f49268])).
fof(f66940,plain,(
  ( ! [X197] : (s(bool,'T_0') = s(bool,f(s(bool,X197))) | s(bool,f(s(bool,'F_0'))) = s(bool,g(s(bool,X197)))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl46 | $spl250)),
  inference(subsumption_resolution,[],[f66804,f49698])).
fof(f66804,plain,(
  ( ! [X197] : (s(bool,'T_0') = s(bool,f(s(bool,X197))) | s(bool,f(s(bool,'F_0'))) = s(bool,g(s(bool,X197))) | s(bool,'F_0') = s(bool,X197)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl46)),
  inference(superposition,[],[f49702,f66447])).
fof(f66447,plain,(
  ( ! [X313] : (s(bool,X313) = s(bool,g(s(bool,X313))) | s(bool,'F_0') = s(bool,X313)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl46)),
  inference(subsumption_resolution,[],[f65433,f49260])).
fof(f65433,plain,(
  ( ! [X313] : (s(bool,X313) = s(bool,g(s(bool,X313))) | s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X313)) ) | ($spl4 | $spl9 | $spl36 | $spl46)),
  inference(superposition,[],[f3303,f30029])).
fof(f3303,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl46),
  inference(cnf_transformation,[],[f3303_D])).
fof(f3303_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f74377,plain,(
  ( ! [X313,X312] : (s(bool,'T_0') = s(bool,g(s(bool,f(s(bool,X313))))) | s(bool,'F_0') = s(bool,X312) | s(bool,'T_0') = s(bool,X313)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl46 | $spl176)),
  inference(duplicate_literal_removal,[],[f74102])).
fof(f74102,plain,(
  ( ! [X313,X312] : (s(bool,'T_0') = s(bool,g(s(bool,f(s(bool,X313))))) | s(bool,'F_0') = s(bool,X312) | s(bool,'F_0') = s(bool,X312) | s(bool,'T_0') = s(bool,X313)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl46 | $spl176)),
  inference(superposition,[],[f66840,f31623])).
fof(f66840,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,g(s(bool,X5))) | s(bool,'F_0') = s(bool,X5)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl46)),
  inference(duplicate_literal_removal,[],[f66687])).
fof(f66687,plain,(
  ( ! [X5] : (s(bool,'F_0') = s(bool,X5) | s(bool,'F_0') = s(bool,X5) | s(bool,'T_0') = s(bool,g(s(bool,X5)))) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl46)),
  inference(superposition,[],[f66447,f40])).
fof(f74642,plain,(
  ( ! [X309] : (s(bool,f(s(bool,X309))) = s(bool,g(s(bool,f(s(bool,X309))))) | s(bool,'T_0') = s(bool,X309)) ) | $spl332),
  inference(cnf_transformation,[],[f74642_D])).
fof(f74642_D,plain,(
  ( ! [X309] : (s(bool,f(s(bool,X309))) = s(bool,g(s(bool,f(s(bool,X309))))) | s(bool,'T_0') = s(bool,X309)) ) <=> ~$spl332),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl332])])).
fof(f74643,plain,(
  ( ! [X310,X311] : (s(bool,f(s(bool,X311))) = s(bool,g(s(bool,f(s(bool,X311))))) | s(bool,'T_0') = s(bool,f(s(bool,X310))) | s(bool,'T_0') = s(bool,X311)) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl176 | $spl250)),
  inference(subsumption_resolution,[],[f74101,f49698])).
fof(f74101,plain,(
  ( ! [X310,X311] : (s(bool,f(s(bool,X311))) = s(bool,g(s(bool,f(s(bool,X311))))) | s(bool,'T_0') = s(bool,f(s(bool,X310))) | s(bool,'F_0') = s(bool,X310) | s(bool,'T_0') = s(bool,X311)) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl176)),
  inference(superposition,[],[f66540,f31623])).
fof(f66540,plain,(
  ( ! [X313] : (s(bool,X313) = s(bool,g(s(bool,X313))) | s(bool,'T_0') = s(bool,f(s(bool,X313)))) ) | ($spl4 | $spl9 | $spl36 | $spl46)),
  inference(subsumption_resolution,[],[f65632,f71])).
fof(f65632,plain,(
  ( ! [X313] : (s(bool,X313) = s(bool,g(s(bool,X313))) | s(bool,'T_0') = s(bool,f(s(bool,X313))) | s(bool,'T_0') = s(bool,'F_0')) ) | ($spl4 | $spl9 | $spl36 | $spl46)),
  inference(superposition,[],[f3303,f30029])).
fof(f74378,plain,(
  ( ! [X308,X309] : (s(bool,f(s(bool,X309))) = s(bool,g(s(bool,f(s(bool,X309))))) | s(bool,'F_0') = s(bool,X308) | s(bool,'T_0') = s(bool,X309)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl46 | $spl176)),
  inference(duplicate_literal_removal,[],[f74100])).
fof(f74100,plain,(
  ( ! [X308,X309] : (s(bool,f(s(bool,X309))) = s(bool,g(s(bool,f(s(bool,X309))))) | s(bool,'F_0') = s(bool,X308) | s(bool,'F_0') = s(bool,X308) | s(bool,'T_0') = s(bool,X309)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl46 | $spl176)),
  inference(superposition,[],[f66447,f31623])).
fof(f74627,plain,(
  ( ! [X277] : (s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,X277))))))) | s(bool,'T_0') = s(bool,X277)) ) | $spl330),
  inference(cnf_transformation,[],[f74627_D])).
fof(f74627_D,plain,(
  ( ! [X277] : (s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,X277))))))) | s(bool,'T_0') = s(bool,X277)) ) <=> ~$spl330),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl330])])).
fof(f74382,plain,(
  ( ! [X277,X276] : (s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,X277))))))) | s(bool,'F_0') = s(bool,X276) | s(bool,'T_0') = s(bool,X277)) ) | ($spl4 | $spl9 | $spl36 | $spl166 | $spl176 | $spl204)),
  inference(duplicate_literal_removal,[],[f74085])).
fof(f74085,plain,(
  ( ! [X277,X276] : (s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,X277))))))) | s(bool,'F_0') = s(bool,X276) | s(bool,'F_0') = s(bool,X276) | s(bool,'T_0') = s(bool,X277)) ) | ($spl4 | $spl9 | $spl36 | $spl166 | $spl176 | $spl204)),
  inference(superposition,[],[f40393,f31623])).
fof(f74562,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1) | s(bool,X2) = s(bool,f(s(bool,X1))) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(condensation,[],[f74561])).
fof(f74561,plain,(
  ( ! [X99,X97,X100,X98,X96] : (s(bool,X97) = s(bool,X98) | s(bool,'F_0') = s(bool,X98) | s(bool,'T_0') = s(bool,X96) | s(bool,X99) = s(bool,f(s(bool,X96))) | s(bool,X99) = s(bool,X100) | s(bool,'T_0') = s(bool,X100)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(subsumption_resolution,[],[f73983,f30026])).
fof(f73983,plain,(
  ( ! [X99,X97,X100,X98,X96] : (s(bool,X97) = s(bool,X98) | s(bool,'F_0') = s(bool,X98) | s(bool,'T_0') = s(bool,X96) | s(bool,X99) = s(bool,f(s(bool,X96))) | s(bool,X99) = s(bool,X100) | s(bool,'T_0') = s(bool,X100) | s(bool,'T_0') = s(bool,X97)) ) | $spl176),
  inference(superposition,[],[f31623,f146])).
fof(f74556,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1) | s(bool,X2) = s(bool,f(s(bool,X1))) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(condensation,[],[f74555])).
fof(f74555,plain,(
  ( ! [X90,X88,X87,X89,X86] : (s(bool,X87) = s(bool,X88) | s(bool,'F_0') = s(bool,X88) | s(bool,'T_0') = s(bool,X86) | s(bool,X89) = s(bool,f(s(bool,X86))) | s(bool,X89) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(subsumption_resolution,[],[f73981,f30026])).
fof(f73981,plain,(
  ( ! [X90,X88,X87,X89,X86] : (s(bool,X87) = s(bool,X88) | s(bool,'F_0') = s(bool,X88) | s(bool,'T_0') = s(bool,X86) | s(bool,X89) = s(bool,f(s(bool,X86))) | s(bool,'T_0') = s(bool,X87) | s(bool,X89) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) ) | $spl176),
  inference(superposition,[],[f31623,f146])).
fof(f74479,plain,(
  ( ! [X146,X147] : (s(bool,f(s(bool,X146))) = s(bool,f(s(bool,X147))) | s(bool,'T_0') = s(bool,X146) | s(bool,'T_0') = s(bool,X147)) ) | $spl328),
  inference(cnf_transformation,[],[f74479_D])).
fof(f74479_D,plain,(
  ( ! [X146,X147] : (s(bool,f(s(bool,X146))) = s(bool,f(s(bool,X147))) | s(bool,'T_0') = s(bool,X146) | s(bool,'T_0') = s(bool,X147)) ) <=> ~$spl328),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl328])])).
fof(f74438,plain,(
  ( ! [X146,X147,X145] : (s(bool,f(s(bool,X146))) = s(bool,f(s(bool,X147))) | s(bool,'F_0') = s(bool,X145) | s(bool,'T_0') = s(bool,X147) | s(bool,'T_0') = s(bool,X146)) ) | $spl176),
  inference(duplicate_literal_removal,[],[f73806])).
fof(f73806,plain,(
  ( ! [X146,X147,X145] : (s(bool,f(s(bool,X146))) = s(bool,f(s(bool,X147))) | s(bool,'F_0') = s(bool,X145) | s(bool,'T_0') = s(bool,X147) | s(bool,'F_0') = s(bool,X145) | s(bool,'T_0') = s(bool,X146)) ) | $spl176),
  inference(superposition,[],[f31623,f31623])).
fof(f74466,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,X1) = s(bool,f(s(bool,X2))) | s(bool,'T_0') = s(bool,X2) | s(bool,X1) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(condensation,[],[f74465])).
fof(f74465,plain,(
  ( ! [X52,X50,X53,X51,X49] : (s(bool,X50) = s(bool,f(s(bool,X51))) | s(bool,'F_0') = s(bool,X49) | s(bool,'T_0') = s(bool,X51) | s(bool,X49) = s(bool,X52) | s(bool,X50) = s(bool,X53) | s(bool,'T_0') = s(bool,X53)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(subsumption_resolution,[],[f73779,f30026])).
fof(f73779,plain,(
  ( ! [X52,X50,X53,X51,X49] : (s(bool,X50) = s(bool,f(s(bool,X51))) | s(bool,'F_0') = s(bool,X49) | s(bool,'T_0') = s(bool,X51) | s(bool,X49) = s(bool,X52) | s(bool,'T_0') = s(bool,X52) | s(bool,X50) = s(bool,X53) | s(bool,'T_0') = s(bool,X53)) ) | $spl176),
  inference(superposition,[],[f31623,f146])).
fof(f74463,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,X1) = s(bool,f(s(bool,X2))) | s(bool,'T_0') = s(bool,X2) | s(bool,X1) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(condensation,[],[f74462])).
fof(f74462,plain,(
  ( ! [X47,X45,X48,X46,X44] : (s(bool,X45) = s(bool,f(s(bool,X46))) | s(bool,'F_0') = s(bool,X44) | s(bool,'T_0') = s(bool,X46) | s(bool,X45) = s(bool,X47) | s(bool,'T_0') = s(bool,X47) | s(bool,X44) = s(bool,X48)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(subsumption_resolution,[],[f73778,f30026])).
fof(f73778,plain,(
  ( ! [X47,X45,X48,X46,X44] : (s(bool,X45) = s(bool,f(s(bool,X46))) | s(bool,'F_0') = s(bool,X44) | s(bool,'T_0') = s(bool,X46) | s(bool,X45) = s(bool,X47) | s(bool,'T_0') = s(bool,X47) | s(bool,X44) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) ) | $spl176),
  inference(superposition,[],[f31623,f146])).
fof(f73207,plain,(
  ( ! [X186] : (s(bool,'T_0') != s(bool,X186) | s(bool,X186) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) ) | $spl320),
  inference(cnf_transformation,[],[f73207_D])).
fof(f73207_D,plain,(
  ( ! [X186] : (s(bool,'T_0') != s(bool,X186) | s(bool,X186) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) ) <=> ~$spl320),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl320])])).
fof(f73242,plain,(
  ( ! [X285,X284] : (s(bool,'T_0') != s(bool,X285) | s(bool,'T_0') = s(bool,X284) | s(bool,X285) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) ) | ($spl2 | $spl46 | $spl276)),
  inference(forward_demodulation,[],[f72123,f3303])).
fof(f72123,plain,(
  ( ! [X285,X284] : (s(bool,X285) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X284) | s(bool,X285) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) ) | ($spl2 | $spl276)),
  inference(duplicate_literal_removal,[],[f71630])).
fof(f71630,plain,(
  ( ! [X285,X284] : (s(bool,X285) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X284) | s(bool,'T_0') = s(bool,X284) | s(bool,X285) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) ) | ($spl2 | $spl276)),
  inference(superposition,[],[f57575,f22297])).
fof(f73239,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,f(s(bool,X1))) | s(bool,X1) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) ) | $spl324),
  inference(cnf_transformation,[],[f73239_D])).
fof(f73239_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,f(s(bool,X1))) | s(bool,X1) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) ) <=> ~$spl324),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl324])])).
fof(f73238,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,X1) != s(bool,f(s(bool,X1))) | s(bool,X1) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) ) | ($spl2 | $spl4 | $spl9 | $spl36)),
  inference(condensation,[],[f73237])).
fof(f73237,plain,(
  ( ! [X271,X270,X272] : (s(bool,X271) != s(bool,f(s(bool,X271))) | s(bool,X270) = s(bool,X272) | s(bool,'F_0') = s(bool,X272) | s(bool,X271) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) ) | ($spl2 | $spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f71624,f30026])).
fof(f71624,plain,(
  ( ! [X271,X270,X272] : (s(bool,X271) != s(bool,f(s(bool,X271))) | s(bool,X270) = s(bool,X272) | s(bool,'F_0') = s(bool,X272) | s(bool,'T_0') = s(bool,X270) | s(bool,X271) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) ) | ($spl2 | $spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f49037,f22297])).
fof(f73228,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') != s(bool,X1) | s(bool,X1) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) ) | ($spl2 | $spl4 | $spl9 | $spl36)),
  inference(condensation,[],[f73227])).
fof(f73227,plain,(
  ( ! [X237,X238,X236] : (s(bool,'T_0') != s(bool,X237) | s(bool,X236) = s(bool,X238) | s(bool,'F_0') = s(bool,X238) | s(bool,X237) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) ) | ($spl2 | $spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f71610,f30026])).
fof(f71610,plain,(
  ( ! [X237,X238,X236] : (s(bool,'T_0') != s(bool,X237) | s(bool,X236) = s(bool,X238) | s(bool,'F_0') = s(bool,X238) | s(bool,'T_0') = s(bool,X236) | s(bool,X237) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) ) | ($spl2 | $spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30026,f22297])).
fof(f73223,plain,(
  ( ! [X229,X230] : (s(bool,X229) != s(bool,X230) | s(bool,X229) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3))) | s(bool,'F_0') = s(bool,X230)) ) | $spl322),
  inference(cnf_transformation,[],[f73223_D])).
fof(f73223_D,plain,(
  ( ! [X229,X230] : (s(bool,X229) != s(bool,X230) | s(bool,X229) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3))) | s(bool,'F_0') = s(bool,X230)) ) <=> ~$spl322),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl322])])).
fof(f72127,plain,(
  ( ! [X229,X230,X228] : (s(bool,X229) != s(bool,X230) | s(bool,'T_0') = s(bool,X228) | s(bool,'F_0') = s(bool,X230) | s(bool,X229) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) ) | ($spl2 | $spl4 | $spl9 | $spl36)),
  inference(duplicate_literal_removal,[],[f71607])).
fof(f71607,plain,(
  ( ! [X229,X230,X228] : (s(bool,X229) != s(bool,X230) | s(bool,'T_0') = s(bool,X228) | s(bool,'F_0') = s(bool,X230) | s(bool,'T_0') = s(bool,X228) | s(bool,X229) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) ) | ($spl2 | $spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f29101,f22297])).
fof(f73214,plain,(
  ( ! [X210,X209] : (s(bool,'T_0') != s(bool,X210) | s(bool,'T_0') = s(bool,X209) | s(bool,X210) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) ) | ($spl2 | $spl4 | $spl9 | $spl36)),
  inference(forward_demodulation,[],[f72128,f49268])).
fof(f72128,plain,(
  ( ! [X210,X209] : (s(bool,X210) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X209) | s(bool,X210) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) ) | ($spl2 | $spl4 | $spl36)),
  inference(duplicate_literal_removal,[],[f71600])).
fof(f71600,plain,(
  ( ! [X210,X209] : (s(bool,X210) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X209) | s(bool,'T_0') = s(bool,X209) | s(bool,X210) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) ) | ($spl2 | $spl4 | $spl36)),
  inference(superposition,[],[f25786,f22297])).
fof(f72131,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') != s(bool,X3) | s(bool,X3) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) ) | ($spl2 | $spl13 | $spl22 | $spl36)),
  inference(duplicate_literal_removal,[],[f72130])).
fof(f72130,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') != s(bool,X3) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X3) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) ) | ($spl2 | $spl13 | $spl22 | $spl36)),
  inference(condensation,[],[f71597])).
fof(f71597,plain,(
  ( ! [X204,X202,X203,X201] : (s(bool,'T_0') != s(bool,X202) | s(bool,'T_0') = s(bool,X203) | s(bool,X201) = s(bool,X204) | s(bool,X203) = s(bool,X204) | s(bool,'T_0') = s(bool,X201) | s(bool,X202) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) ) | ($spl2 | $spl13 | $spl22 | $spl36)),
  inference(superposition,[],[f22900,f22297])).
fof(f73206,plain,(
  ( ! [X185,X186] : (s(bool,'T_0') != s(bool,X186) | s(bool,'T_0') = s(bool,X185) | s(bool,X186) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) ) | ($spl2 | $spl13 | $spl22 | $spl36)),
  inference(subsumption_resolution,[],[f71592,f17356])).
fof(f71592,plain,(
  ( ! [X187,X185,X188,X186] : (s(bool,'T_0') != s(bool,X186) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X187,bool),'const.pair.,_2'(s(X187,X188),s(bool,X185))))) | s(bool,'T_0') = s(bool,X185) | s(bool,X186) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) ) | ($spl2 | $spl13 | $spl22)),
  inference(superposition,[],[f7282,f22297])).
fof(f73075,plain,(
  ( ! [X321,X322] : (s(bool,X321) != s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3))) | s(bool,'F_0') = s(bool,X322) | s(bool,X321) = s(bool,X322)) ) | $spl318),
  inference(cnf_transformation,[],[f73075_D])).
fof(f73075_D,plain,(
  ( ! [X321,X322] : (s(bool,X321) != s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3))) | s(bool,'F_0') = s(bool,X322) | s(bool,X321) = s(bool,X322)) ) <=> ~$spl318),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl318])])).
fof(f73074,plain,(
  ( ! [X321,X323,X322] : (s(bool,X321) != s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3))) | s(bool,X321) = s(bool,X322) | s(bool,'F_0') = s(bool,X322) | s(bool,'T_0') = s(bool,X323)) ) | ($spl2 | $spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f71360,f30030])).
fof(f71360,plain,(
  ( ! [X321,X323,X322] : (s(bool,X321) != s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3))) | s(bool,X321) = s(bool,X322) | s(bool,'F_0') = s(bool,X322) | s(bool,'T_0') = s(bool,X323) | s(bool,X323) = s(bool,f(s(bool,X321)))) ) | ($spl2 | $spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f49037,f22297])).
fof(f72593,plain,(
  ( ! [X95,X93] : (s(bool,X95) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3))) | s(bool,X93) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3))) | s(bool,X93) = s(bool,X95)) ) | $spl316),
  inference(cnf_transformation,[],[f72593_D])).
fof(f72593_D,plain,(
  ( ! [X95,X93] : (s(bool,X95) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3))) | s(bool,X93) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3))) | s(bool,X93) = s(bool,X95)) ) <=> ~$spl316),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl316])])).
fof(f72402,plain,(
  ( ! [X111,X112,X113] : (s(bool,'T_0') = s(bool,X112) | s(bool,X111) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3))) | s(bool,X113) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3))) | s(bool,X111) = s(bool,X113)) ) | $spl2),
  inference(duplicate_literal_removal,[],[f70879])).
fof(f70879,plain,(
  ( ! [X111,X112,X113] : (s(bool,'T_0') = s(bool,X112) | s(bool,'T_0') = s(bool,X112) | s(bool,X111) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3))) | s(bool,X113) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3))) | s(bool,X111) = s(bool,X113)) ) | $spl2),
  inference(superposition,[],[f22297,f22297])).
fof(f72428,plain,(
  ( ! [X116,X114,X115] : (s(bool,X115) = s(bool,X116) | s(bool,'T_0') = s(bool,X114) | s(bool,X116) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3))) | s(bool,X115) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) ) | $spl2),
  inference(duplicate_literal_removal,[],[f70782])).
fof(f70782,plain,(
  ( ! [X116,X114,X115] : (s(bool,X115) = s(bool,X116) | s(bool,'T_0') = s(bool,X114) | s(bool,X116) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X114) | s(bool,X115) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) ) | $spl2),
  inference(superposition,[],[f22297,f22297])).
fof(f72532,plain,(
  ( ! [X94,X95,X93] : (s(bool,'T_0') = s(bool,X94) | s(bool,X95) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3))) | s(bool,X93) = s(bool,X95) | s(bool,X93) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) ) | $spl2),
  inference(duplicate_literal_removal,[],[f70595])).
fof(f70595,plain,(
  ( ! [X94,X95,X93] : (s(bool,'T_0') = s(bool,X94) | s(bool,X95) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3))) | s(bool,X93) = s(bool,X95) | s(bool,'T_0') = s(bool,X94) | s(bool,X93) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) ) | $spl2),
  inference(superposition,[],[f22297,f22297])).
fof(f70504,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,X2) | s(bool,X1) = s(bool,X2)) ) | $spl314),
  inference(cnf_transformation,[],[f70504_D])).
fof(f70504_D,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,X2) | s(bool,X1) = s(bool,X2)) ) <=> ~$spl314),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl314])])).
fof(f70506,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') != s(bool,f(s(bool,X1))) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl166 | $spl204 | $spl234)),
  inference(condensation,[],[f70505])).
fof(f70505,plain,(
  ( ! [X39,X37,X38,X40] : (s(bool,'T_0') != s(bool,f(s(bool,X38))) | s(bool,'F_0') = s(bool,X37) | s(bool,X37) = s(bool,X39) | s(bool,X38) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl166 | $spl204 | $spl234)),
  inference(subsumption_resolution,[],[f70346,f30026])).
fof(f70346,plain,(
  ( ! [X39,X37,X38,X40] : (s(bool,'T_0') != s(bool,f(s(bool,X38))) | s(bool,'F_0') = s(bool,X37) | s(bool,X37) = s(bool,X39) | s(bool,'T_0') = s(bool,X39) | s(bool,X38) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl166 | $spl204 | $spl234)),
  inference(superposition,[],[f70304,f146])).
fof(f70503,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') != s(bool,f(s(bool,X1))) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl166 | $spl204 | $spl234)),
  inference(condensation,[],[f70502])).
fof(f70502,plain,(
  ( ! [X35,X33,X36,X34] : (s(bool,'T_0') != s(bool,f(s(bool,X34))) | s(bool,'F_0') = s(bool,X33) | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35) | s(bool,X33) = s(bool,X36)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl166 | $spl204 | $spl234)),
  inference(subsumption_resolution,[],[f70345,f30026])).
fof(f70345,plain,(
  ( ! [X35,X33,X36,X34] : (s(bool,'T_0') != s(bool,f(s(bool,X34))) | s(bool,'F_0') = s(bool,X33) | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35) | s(bool,X33) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl166 | $spl204 | $spl234)),
  inference(superposition,[],[f70304,f146])).
fof(f70323,plain,(
  ( ! [X116,X115] : (s(bool,X115) != s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,'F_0') = s(bool,X116) | s(bool,X115) = s(bool,X116)) ) | $spl312),
  inference(cnf_transformation,[],[f70323_D])).
fof(f70323_D,plain,(
  ( ! [X116,X115] : (s(bool,X115) != s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,'F_0') = s(bool,X116) | s(bool,X115) = s(bool,X116)) ) <=> ~$spl312),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl312])])).
fof(f70322,plain,(
  ( ! [X116,X117,X115] : (s(bool,X115) != s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X115) = s(bool,X116) | s(bool,'F_0') = s(bool,X116) | s(bool,'T_0') = s(bool,X117)) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f70194,f30030])).
fof(f70194,plain,(
  ( ! [X116,X117,X115] : (s(bool,X115) != s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X115) = s(bool,X116) | s(bool,'F_0') = s(bool,X116) | s(bool,'T_0') = s(bool,X117) | s(bool,X117) = s(bool,f(s(bool,X115)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f49037,f22295])).
fof(f70276,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,f(s(bool,X1))) | s(bool,X1) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | $spl310),
  inference(cnf_transformation,[],[f70276_D])).
fof(f70276_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,f(s(bool,X1))) | s(bool,X1) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) <=> ~$spl310),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl310])])).
fof(f70275,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,X1) != s(bool,f(s(bool,X1))) | s(bool,X1) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(condensation,[],[f70274])).
fof(f70274,plain,(
  ( ! [X103,X102,X104] : (s(bool,X103) != s(bool,f(s(bool,X103))) | s(bool,X102) = s(bool,X104) | s(bool,'F_0') = s(bool,X104) | s(bool,X103) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f70094,f30026])).
fof(f70094,plain,(
  ( ! [X103,X102,X104] : (s(bool,X103) != s(bool,f(s(bool,X103))) | s(bool,X102) = s(bool,X104) | s(bool,'F_0') = s(bool,X104) | s(bool,'T_0') = s(bool,X102) | s(bool,X103) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f49037,f22295])).
fof(f70257,plain,(
  ( ! [X2,X3,X1] : (s(bool,X1) != s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,X3) | s(bool,X2) = s(bool,X3) | s(bool,X1) = s(bool,X2)) ) | $spl308),
  inference(cnf_transformation,[],[f70257_D])).
fof(f70257_D,plain,(
  ( ! [X2,X3,X1] : (s(bool,X1) != s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,X3) | s(bool,X2) = s(bool,X3) | s(bool,X1) = s(bool,X2)) ) <=> ~$spl308),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl308])])).
fof(f70263,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,X1) != s(bool,f(s(bool,X1))) | s(bool,X1) = s(bool,X2) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | ($spl4 | $spl9 | $spl36)),
  inference(condensation,[],[f70262])).
fof(f70262,plain,(
  ( ! [X80,X83,X81,X79,X82] : (s(bool,X80) != s(bool,f(s(bool,X80))) | s(bool,X79) = s(bool,X81) | s(bool,'F_0') = s(bool,X81) | s(bool,X80) = s(bool,X82) | s(bool,X82) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f70088,f30026])).
fof(f70088,plain,(
  ( ! [X80,X83,X81,X79,X82] : (s(bool,X80) != s(bool,f(s(bool,X80))) | s(bool,X79) = s(bool,X81) | s(bool,'F_0') = s(bool,X81) | s(bool,X80) = s(bool,X82) | s(bool,X82) = s(bool,X83) | s(bool,'T_0') = s(bool,X83) | s(bool,'T_0') = s(bool,X79)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f49037,f146])).
fof(f70256,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,X1) != s(bool,f(s(bool,X1))) | s(bool,X1) = s(bool,X2) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | ($spl4 | $spl9 | $spl36)),
  inference(condensation,[],[f70255])).
fof(f70255,plain,(
  ( ! [X70,X72,X71,X69,X73] : (s(bool,X70) != s(bool,f(s(bool,X70))) | s(bool,X69) = s(bool,X71) | s(bool,'F_0') = s(bool,X71) | s(bool,X70) = s(bool,X72) | s(bool,X72) = s(bool,X73) | s(bool,'T_0') = s(bool,X73)) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f70086,f30026])).
fof(f70086,plain,(
  ( ! [X70,X72,X71,X69,X73] : (s(bool,X70) != s(bool,f(s(bool,X70))) | s(bool,X69) = s(bool,X71) | s(bool,'F_0') = s(bool,X71) | s(bool,X70) = s(bool,X72) | s(bool,'T_0') = s(bool,X69) | s(bool,X72) = s(bool,X73) | s(bool,'T_0') = s(bool,X73)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f49037,f146])).
fof(f70250,plain,(
  ( ! [X2,X3,X1] : (s(bool,X1) != s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X2,bool),'const.pair.,_2'(s(X2,X3),s(bool,X1)))))) ) | $spl306),
  inference(cnf_transformation,[],[f70250_D])).
fof(f70250_D,plain,(
  ( ! [X2,X3,X1] : (s(bool,X1) != s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X2,bool),'const.pair.,_2'(s(X2,X3),s(bool,X1)))))) ) <=> ~$spl306),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl306])])).
fof(f70249,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,X1) != s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X2,bool),'const.pair.,_2'(s(X2,X3),s(bool,X1)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(condensation,[],[f70248])).
fof(f70248,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(bool,X35) != s(bool,f(s(bool,X35))) | s(bool,X34) = s(bool,X36) | s(bool,'F_0') = s(bool,X36) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X37,bool),'const.pair.,_2'(s(X37,X38),s(bool,X35)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f70079,f30026])).
fof(f70079,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(bool,X35) != s(bool,f(s(bool,X35))) | s(bool,X34) = s(bool,X36) | s(bool,'F_0') = s(bool,X36) | s(bool,'T_0') = s(bool,X34) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X37,bool),'const.pair.,_2'(s(X37,X38),s(bool,X35)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f49037,f133])).
fof(f70247,plain,(
  ( ! [X2,X3,X1] : (s(bool,X1) != s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),'const.pair.,_2'(s(bool,X1),s(X2,X3)))))) ) | $spl304),
  inference(cnf_transformation,[],[f70247_D])).
fof(f70247_D,plain,(
  ( ! [X2,X3,X1] : (s(bool,X1) != s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),'const.pair.,_2'(s(bool,X1),s(X2,X3)))))) ) <=> ~$spl304),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl304])])).
fof(f70246,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,X1) != s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),'const.pair.,_2'(s(bool,X1),s(X2,X3)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(condensation,[],[f70245])).
fof(f70245,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s(bool,X20) != s(bool,f(s(bool,X20))) | s(bool,X19) = s(bool,X21) | s(bool,'F_0') = s(bool,X21) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X22),'const.pair.,_2'(s(bool,X20),s(X22,X23)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f70076,f30026])).
fof(f70076,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s(bool,X20) != s(bool,f(s(bool,X20))) | s(bool,X19) = s(bool,X21) | s(bool,'F_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X22),'const.pair.,_2'(s(bool,X20),s(X22,X23)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f49037,f132])).
fof(f70051,plain,(
  $false | ($spl1 | $spl9 | $spl13 | $spl60 | $spl276)),
  inference(subsumption_resolution,[],[f70050,f439])).
fof(f70050,plain,(
  s(bool,'T_0') = s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2))))) | ($spl1 | $spl9 | $spl60 | $spl276)),
  inference(subsumption_resolution,[],[f69975,f4900])).
fof(f69975,plain,(
  s(bool,'F_0') != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2))))) | ($spl1 | $spl9 | $spl276)),
  inference(superposition,[],[f57575,f1103])).
fof(f69851,plain,(
  ( ! [X177,X180,X176,X181] : (s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X180,bool),'const.pair.,_2'(s(X180,X181),s(bool,happ(s(fun(bool,bool),X176),s(bool,X177))))))) | s(bool,'F_0') = s(bool,X177)) ) | $spl300),
  inference(cnf_transformation,[],[f69851_D])).
fof(f69851_D,plain,(
  ( ! [X177,X180,X176,X181] : (s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X180,bool),'const.pair.,_2'(s(X180,X181),s(bool,happ(s(fun(bool,bool),X176),s(bool,X177))))))) | s(bool,'F_0') = s(bool,X177)) ) <=> ~$spl300),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl300])])).
fof(f69719,plain,(
  ( ! [X208,X206,X204,X207,X205,X203] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X205,bool),'const.pair.,_2'(s(X205,X206),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X207,bool),'const.pair.,_2'(s(X207,X208),s(bool,happ(s(fun(bool,bool),X203),s(bool,X204))))))) | s(bool,'F_0') = s(bool,X204)) ) | ($spl4 | $spl9 | $spl13 | $spl22 | $spl36 | $spl40 | $spl290)),
  inference(superposition,[],[f1972,f68868])).
fof(f68868,plain,(
  ( ! [X105,X104] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X105),s(bool,X104))) | s(bool,'F_0') = s(bool,X104)) ) | ($spl4 | $spl9 | $spl13 | $spl22 | $spl36 | $spl40 | $spl290)),
  inference(subsumption_resolution,[],[f68703,f49260])).
fof(f68703,plain,(
  ( ! [X105,X104] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X105),s(bool,X104))) | s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X104)) ) | ($spl4 | $spl9 | $spl13 | $spl22 | $spl36 | $spl290)),
  inference(superposition,[],[f68336,f30029])).
fof(f68336,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0')))) ) | ($spl13 | $spl22 | $spl290)),
  inference(forward_demodulation,[],[f68247,f43])).
fof(f68247,plain,(
  ( ! [X0] : (s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0'))) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,'F_0')))))) ) | ($spl13 | $spl22 | $spl290)),
  inference(unit_resulting_resolution,[],[f6406,f67866,f200])).
fof(f69852,plain,(
  ( ! [X198,X202,X199] : (s(bool,X202) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,happ(s(fun(bool,bool),X198),s(bool,X199))),s(bool,X202))))) | s(bool,'F_0') = s(bool,X199) | s(bool,'T_0') = s(bool,X202)) ) | $spl302),
  inference(cnf_transformation,[],[f69852_D])).
fof(f69852_D,plain,(
  ( ! [X198,X202,X199] : (s(bool,X202) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,happ(s(fun(bool,bool),X198),s(bool,X199))),s(bool,X202))))) | s(bool,'F_0') = s(bool,X199) | s(bool,'T_0') = s(bool,X202)) ) <=> ~$spl302),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl302])])).
fof(f69718,plain,(
  ( ! [X198,X202,X200,X199,X201] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X200),'const.pair.,_2'(s(bool,'F_0'),s(X200,X201))))) | s(bool,X202) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,happ(s(fun(bool,bool),X198),s(bool,X199))),s(bool,X202))))) | s(bool,'T_0') = s(bool,X202) | s(bool,'F_0') = s(bool,X199)) ) | ($spl4 | $spl9 | $spl13 | $spl16 | $spl22 | $spl36 | $spl40 | $spl290)),
  inference(superposition,[],[f1557,f68868])).
fof(f69714,plain,(
  ( ! [X177,X180,X178,X176,X181,X179] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X178),'const.pair.,_2'(s(bool,'F_0'),s(X178,X179))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X180,bool),'const.pair.,_2'(s(X180,X181),s(bool,happ(s(fun(bool,bool),X176),s(bool,X177))))))) | s(bool,'F_0') = s(bool,X177)) ) | ($spl4 | $spl9 | $spl13 | $spl22 | $spl36 | $spl40 | $spl290)),
  inference(superposition,[],[f1473,f68868])).
fof(f69539,plain,(
  ( ! [X123,X124] : (s(bool,'T_0') = s(bool,f(s(bool,X124))) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(bool,bool),X123),s(bool,X124)))))) ) | $spl298),
  inference(cnf_transformation,[],[f69539_D])).
fof(f69539_D,plain,(
  ( ! [X123,X124] : (s(bool,'T_0') = s(bool,f(s(bool,X124))) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(bool,bool),X123),s(bool,X124)))))) ) <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f69538,plain,(
  ( ! [X125,X123,X124] : (s(bool,'T_0') = s(bool,f(s(bool,X124))) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(bool,bool),X123),s(bool,X124))))) | s(bool,'F_0') = s(bool,X125)) ) | ($spl4 | $spl9 | $spl36 | $spl172 | $spl290)),
  inference(subsumption_resolution,[],[f69429,f29278])).
fof(f69429,plain,(
  ( ! [X125,X123,X124] : (~p(s(bool,X125)) | s(bool,'T_0') = s(bool,f(s(bool,X124))) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(bool,bool),X123),s(bool,X124))))) | s(bool,'F_0') = s(bool,X125)) ) | ($spl4 | $spl9 | $spl36 | $spl290)),
  inference(superposition,[],[f68414,f30029])).
fof(f69298,plain,(
  ( ! [X216,X220,X215,X219,X217] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X215,bool),X216),s(fun(bool,bool),X219),s('type.pair.prod'(X215,bool),X217)))) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(X215,bool),X216),s(X215,'const.pair.FST_1'(s('type.pair.prod'(X215,bool),X217))))))) | s(bool,'T_0') = s(bool,X220) | s(bool,X220) = s(bool,happ(s(fun(bool,bool),X219),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X215,bool),X217)))) ) | $spl296),
  inference(cnf_transformation,[],[f69298_D])).
fof(f69298_D,plain,(
  ( ! [X216,X220,X215,X219,X217] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X215,bool),X216),s(fun(bool,bool),X219),s('type.pair.prod'(X215,bool),X217)))) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(X215,bool),X216),s(X215,'const.pair.FST_1'(s('type.pair.prod'(X215,bool),X217))))))) | s(bool,'T_0') = s(bool,X220) | s(bool,X220) = s(bool,happ(s(fun(bool,bool),X219),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X215,bool),X217)))) ) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f69297,plain,(
  ( ! [X218,X216,X220,X215,X219,X217] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X215,bool),X216),s(fun(bool,bool),X219),s('type.pair.prod'(X215,bool),X217)))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X215,bool),X217))) | s(bool,X220) = s(bool,happ(s(fun(bool,bool),X219),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X220) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(X215,bool),X216),s(X215,'const.pair.FST_1'(s('type.pair.prod'(X215,bool),X217))))))) | s(bool,'F_0') = s(bool,X218)) ) | ($spl4 | $spl9 | $spl36 | $spl172)),
  inference(subsumption_resolution,[],[f69228,f29278])).
fof(f69228,plain,(
  ( ! [X218,X216,X220,X215,X219,X217] : (~p(s(bool,X218)) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X215,bool),X216),s(fun(bool,bool),X219),s('type.pair.prod'(X215,bool),X217)))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X215,bool),X217))) | s(bool,X220) = s(bool,happ(s(fun(bool,bool),X219),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X220) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(X215,bool),X216),s(X215,'const.pair.FST_1'(s('type.pair.prod'(X215,bool),X217))))))) | s(bool,'F_0') = s(bool,X218)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f2874,f30029])).
fof(f68663,plain,(
  ( ! [X251,X249,X246,X252,X250,X247] : (~p(s(bool,happ(s(fun(X249,bool),X250),s(X249,X251)))) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(bool,bool),X246),s(bool,X247))))) | s(bool,X247) = s(bool,X252) | s(bool,'T_0') = s(bool,X252) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X249,bool),X250),s(fun(bool,bool),X246),s('type.pair.prod'(X249,bool),'const.pair.,_2'(s(X249,X251),s(bool,'T_0'))))))) ) | $spl294),
  inference(cnf_transformation,[],[f68663_D])).
fof(f68663_D,plain,(
  ( ! [X251,X249,X246,X252,X250,X247] : (~p(s(bool,happ(s(fun(X249,bool),X250),s(X249,X251)))) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(bool,bool),X246),s(bool,X247))))) | s(bool,X247) = s(bool,X252) | s(bool,'T_0') = s(bool,X252) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X249,bool),X250),s(fun(bool,bool),X246),s('type.pair.prod'(X249,bool),'const.pair.,_2'(s(X249,X251),s(bool,'T_0'))))))) ) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f68662,plain,(
  ( ! [X251,X249,X246,X252,X250,X248,X247] : (~p(s(bool,happ(s(fun(X249,bool),X250),s(X249,X251)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X249,bool),X250),s(fun(bool,bool),X246),s('type.pair.prod'(X249,bool),'const.pair.,_2'(s(X249,X251),s(bool,'T_0')))))) | s(bool,X247) = s(bool,X252) | s(bool,'T_0') = s(bool,X252) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(bool,bool),X246),s(bool,X247))))) | s(bool,'F_0') = s(bool,X248)) ) | ($spl4 | $spl9 | $spl36 | $spl172)),
  inference(subsumption_resolution,[],[f68574,f29278])).
fof(f68574,plain,(
  ( ! [X251,X249,X246,X252,X250,X248,X247] : (~p(s(bool,X248)) | ~p(s(bool,happ(s(fun(X249,bool),X250),s(X249,X251)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X249,bool),X250),s(fun(bool,bool),X246),s('type.pair.prod'(X249,bool),'const.pair.,_2'(s(X249,X251),s(bool,'T_0')))))) | s(bool,X247) = s(bool,X252) | s(bool,'T_0') = s(bool,X252) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(bool,bool),X246),s(bool,X247))))) | s(bool,'F_0') = s(bool,X248)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f2704,f30029])).
fof(f67891,plain,(
  ( ! [X251,X249,X246,X252,X250,X248] : (~p(s(bool,happ(s(fun(X248,bool),X249),s(X248,X250)))) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(bool,bool),X246),s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,X252) | s(bool,X251) = s(bool,X252) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X248,bool),X249),s(fun(bool,bool),X246),s('type.pair.prod'(X248,bool),'const.pair.,_2'(s(X248,X250),s(bool,X251))))))) ) | $spl292),
  inference(cnf_transformation,[],[f67891_D])).
fof(f67891_D,plain,(
  ( ! [X251,X249,X246,X252,X250,X248] : (~p(s(bool,happ(s(fun(X248,bool),X249),s(X248,X250)))) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(bool,bool),X246),s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,X252) | s(bool,X251) = s(bool,X252) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X248,bool),X249),s(fun(bool,bool),X246),s('type.pair.prod'(X248,bool),'const.pair.,_2'(s(X248,X250),s(bool,X251))))))) ) <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f67890,plain,(
  ( ! [X251,X249,X246,X252,X250,X248,X247] : (~p(s(bool,happ(s(fun(X248,bool),X249),s(X248,X250)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X248,bool),X249),s(fun(bool,bool),X246),s('type.pair.prod'(X248,bool),'const.pair.,_2'(s(X248,X250),s(bool,X251)))))) | s(bool,X251) = s(bool,X252) | s(bool,'T_0') = s(bool,X252) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(bool,bool),X246),s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,X247)) ) | ($spl4 | $spl9 | $spl36 | $spl172)),
  inference(subsumption_resolution,[],[f67840,f29278])).
fof(f67840,plain,(
  ( ! [X251,X249,X246,X252,X250,X248,X247] : (~p(s(bool,X247)) | ~p(s(bool,happ(s(fun(X248,bool),X249),s(X248,X250)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X248,bool),X249),s(fun(bool,bool),X246),s('type.pair.prod'(X248,bool),'const.pair.,_2'(s(X248,X250),s(bool,X251)))))) | s(bool,X251) = s(bool,X252) | s(bool,'T_0') = s(bool,X252) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(bool,bool),X246),s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,X247)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f2699,f30029])).
fof(f67865,plain,(
  ( ! [X6,X4,X2,X0,X5,X3] : (~p(s(bool,happ(s(fun(X2,bool),X3),s(X2,X4)))) | s(bool,'T_0') = s(bool,X6) | s(bool,X5) = s(bool,X6) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X2,bool),X3),s(fun(bool,bool),X0),s('type.pair.prod'(X2,bool),'const.pair.,_2'(s(X2,X4),s(bool,X5))))))) ) | $spl288),
  inference(cnf_transformation,[],[f67865_D])).
fof(f67865_D,plain,(
  ( ! [X6,X4,X2,X0,X5,X3] : (~p(s(bool,happ(s(fun(X2,bool),X3),s(X2,X4)))) | s(bool,'T_0') = s(bool,X6) | s(bool,X5) = s(bool,X6) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X2,bool),X3),s(fun(bool,bool),X0),s('type.pair.prod'(X2,bool),'const.pair.,_2'(s(X2,X4),s(bool,X5))))))) ) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f67866,plain,(
  ( ! [X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0'))))) ) | $spl290),
  inference(cnf_transformation,[],[f67866_D])).
fof(f67866_D,plain,(
  ( ! [X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0'))))) ) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f67774,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (~p(s(bool,happ(s(fun(bool,bool),X8),s(bool,'T_0')))) | ~p(s(bool,happ(s(fun(X9,bool),X10),s(X9,X11)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X9,bool),X10),s(fun(bool,bool),X7),s('type.pair.prod'(X9,bool),'const.pair.,_2'(s(X9,X11),s(bool,X12)))))) | s(bool,X12) = s(bool,X13) | s(bool,'T_0') = s(bool,X13)) ) | $spl44),
  inference(superposition,[],[f2699,f3302])).
fof(f67773,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) | ~p(s(bool,happ(s(fun(X2,bool),X3),s(X2,X4)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X2,bool),X3),s(fun(bool,bool),X0),s('type.pair.prod'(X2,bool),'const.pair.,_2'(s(X2,X4),s(bool,X5)))))) | s(bool,X5) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) ) | $spl44),
  inference(superposition,[],[f2699,f3302])).
fof(f67091,plain,(
  ( ! [X251,X249,X246,X252,X248,X247] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X246,bool),X247),s(fun(X248,bool),X251),s('type.pair.prod'(X246,X248),X249)))) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(X246,bool),X247),s(X246,'const.pair.FST_1'(s('type.pair.prod'(X246,X248),X249))))))) | s(bool,'T_0') = s(bool,X252) | s(bool,X252) = s(bool,happ(s(fun(X248,bool),X251),s(X248,'const.pair.SND_1'(s('type.pair.prod'(X246,X248),X249)))))) ) | $spl286),
  inference(cnf_transformation,[],[f67091_D])).
fof(f67091_D,plain,(
  ( ! [X251,X249,X246,X252,X248,X247] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X246,bool),X247),s(fun(X248,bool),X251),s('type.pair.prod'(X246,X248),X249)))) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(X246,bool),X247),s(X246,'const.pair.FST_1'(s('type.pair.prod'(X246,X248),X249))))))) | s(bool,'T_0') = s(bool,X252) | s(bool,X252) = s(bool,happ(s(fun(X248,bool),X251),s(X248,'const.pair.SND_1'(s('type.pair.prod'(X246,X248),X249)))))) ) <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f67090,plain,(
  ( ! [X251,X249,X246,X252,X250,X248,X247] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X246,bool),X247),s(fun(X248,bool),X251),s('type.pair.prod'(X246,X248),X249)))) | s(bool,X252) = s(bool,happ(s(fun(X248,bool),X251),s(X248,'const.pair.SND_1'(s('type.pair.prod'(X246,X248),X249))))) | s(bool,'T_0') = s(bool,X252) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(X246,bool),X247),s(X246,'const.pair.FST_1'(s('type.pair.prod'(X246,X248),X249))))))) | s(bool,'F_0') = s(bool,X250)) ) | ($spl4 | $spl9 | $spl36 | $spl172)),
  inference(subsumption_resolution,[],[f67030,f29278])).
fof(f67030,plain,(
  ( ! [X251,X249,X246,X252,X250,X248,X247] : (~p(s(bool,X250)) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X246,bool),X247),s(fun(X248,bool),X251),s('type.pair.prod'(X246,X248),X249)))) | s(bool,X252) = s(bool,happ(s(fun(X248,bool),X251),s(X248,'const.pair.SND_1'(s('type.pair.prod'(X246,X248),X249))))) | s(bool,'T_0') = s(bool,X252) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(X246,bool),X247),s(X246,'const.pair.FST_1'(s('type.pair.prod'(X246,X248),X249))))))) | s(bool,'F_0') = s(bool,X250)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f786,f30029])).
fof(f64817,plain,(
  ( ! [X12,X10,X13,X11] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(sK1,bool),sK4),s(fun(X10,bool),X11),s('type.pair.prod'(sK1,X10),'const.pair.,_2'(s(sK1,sK2),s(X10,X12)))))) | s(bool,'T_0') = s(bool,X13) | s(bool,X13) = s(bool,happ(s(fun(X10,bool),X11),s(X10,X12)))) ) | $spl284),
  inference(cnf_transformation,[],[f64817_D])).
fof(f64817_D,plain,(
  ( ! [X12,X10,X13,X11] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(sK1,bool),sK4),s(fun(X10,bool),X11),s('type.pair.prod'(sK1,X10),'const.pair.,_2'(s(sK1,sK2),s(X10,X12)))))) | s(bool,'T_0') = s(bool,X13) | s(bool,X13) = s(bool,happ(s(fun(X10,bool),X11),s(X10,X12)))) ) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f64816,plain,(
  ( ! [X12,X10,X13,X11,X9] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(sK1,bool),sK4),s(fun(X10,bool),X11),s('type.pair.prod'(sK1,X10),'const.pair.,_2'(s(sK1,sK2),s(X10,X12)))))) | s(bool,X13) = s(bool,happ(s(fun(X10,bool),X11),s(X10,X12))) | s(bool,'T_0') = s(bool,X13) | s(bool,'F_0') = s(bool,X9)) ) | ($spl4 | $spl9 | $spl36 | $spl172)),
  inference(subsumption_resolution,[],[f64713,f29278])).
fof(f64713,plain,(
  ( ! [X12,X10,X13,X11,X9] : (~p(s(bool,X9)) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(sK1,bool),sK4),s(fun(X10,bool),X11),s('type.pair.prod'(sK1,X10),'const.pair.,_2'(s(sK1,sK2),s(X10,X12)))))) | s(bool,X13) = s(bool,happ(s(fun(X10,bool),X11),s(X10,X12))) | s(bool,'T_0') = s(bool,X13) | s(bool,'F_0') = s(bool,X9)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f2763,f25792])).
fof(f64814,plain,(
  ( ! [X4,X2,X3,X1] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(X1,bool),X2),s('type.pair.prod'(sK0,X1),'const.pair.,_2'(s(sK0,sK3),s(X1,X3)))))) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,happ(s(fun(X1,bool),X2),s(X1,X3)))) ) | $spl282),
  inference(cnf_transformation,[],[f64814_D])).
fof(f64814_D,plain,(
  ( ! [X4,X2,X3,X1] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(X1,bool),X2),s('type.pair.prod'(sK0,X1),'const.pair.,_2'(s(sK0,sK3),s(X1,X3)))))) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,happ(s(fun(X1,bool),X2),s(X1,X3)))) ) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f64813,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(X1,bool),X2),s('type.pair.prod'(sK0,X1),'const.pair.,_2'(s(sK0,sK3),s(X1,X3)))))) | s(bool,X4) = s(bool,happ(s(fun(X1,bool),X2),s(X1,X3))) | s(bool,'T_0') = s(bool,X4) | s(bool,'F_0') = s(bool,X0)) ) | ($spl2 | $spl4 | $spl9 | $spl36 | $spl172)),
  inference(subsumption_resolution,[],[f64711,f29278])).
fof(f64711,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,X0)) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(X1,bool),X2),s('type.pair.prod'(sK0,X1),'const.pair.,_2'(s(sK0,sK3),s(X1,X3)))))) | s(bool,X4) = s(bool,happ(s(fun(X1,bool),X2),s(X1,X3))) | s(bool,'T_0') = s(bool,X4) | s(bool,'F_0') = s(bool,X0)) ) | ($spl2 | $spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f2763,f32150])).
fof(f32150,plain,(
  ( ! [X174] : (s(bool,X174) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3))) | s(bool,'F_0') = s(bool,X174)) ) | ($spl2 | $spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f32066,f37])).
fof(f32066,plain,(
  ( ! [X174] : (~p(s(bool,'T_0')) | s(bool,X174) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3))) | s(bool,'F_0') = s(bool,X174)) ) | ($spl2 | $spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30024,f1074])).
fof(f49427,plain,(
  ( ! [X693] : (s(bool,X693) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | $spl270),
  inference(cnf_transformation,[],[f49427_D])).
fof(f49427_D,plain,(
  ( ! [X693] : (s(bool,X693) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f63840,plain,(
  ( ! [X182] : (s(bool,X182) != s(bool,f(s(bool,X182))) | s(bool,'F_0') = s(bool,X182)) ) | $spl280),
  inference(cnf_transformation,[],[f63840_D])).
fof(f63840_D,plain,(
  ( ! [X182] : (s(bool,X182) != s(bool,f(s(bool,X182))) | s(bool,'F_0') = s(bool,X182)) ) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f63845,plain,(
  ( ! [X204,X203] : (s(bool,X203) != s(bool,f(s(bool,X203))) | s(bool,X204) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,'F_0') = s(bool,X203)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(subsumption_resolution,[],[f63716,f25261])).
fof(f63716,plain,(
  ( ! [X204,X203] : (s(bool,X203) != s(bool,f(s(bool,X203))) | s(bool,X204) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,'T_0') = s(bool,X204) | s(bool,'F_0') = s(bool,X203)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(superposition,[],[f50029,f25792])).
fof(f63844,plain,(
  ( ! [X200,X201] : (s(bool,X200) != s(bool,f(s(bool,X200))) | s(bool,'T_0') = s(bool,X201) | s(bool,'F_0') = s(bool,X200)) ) | ($spl2 | $spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl194)),
  inference(subsumption_resolution,[],[f63714,f34914])).
fof(f63714,plain,(
  ( ! [X200,X201] : (s(bool,X200) != s(bool,f(s(bool,X200))) | s(bool,X201) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X201) | s(bool,'F_0') = s(bool,X200)) ) | ($spl2 | $spl4 | $spl9 | $spl13 | $spl36 | $spl40)),
  inference(superposition,[],[f50029,f32150])).
fof(f63843,plain,(
  ( ! [X191,X192] : (s(bool,X191) != s(bool,f(s(bool,X191))) | s(bool,'T_0') = s(bool,X192) | s(bool,'F_0') = s(bool,X191)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl46 | $spl276)),
  inference(subsumption_resolution,[],[f63710,f57575])).
fof(f63710,plain,(
  ( ! [X191,X192] : (s(bool,X191) != s(bool,f(s(bool,X191))) | s(bool,X192) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X192) | s(bool,'F_0') = s(bool,X191)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl46)),
  inference(superposition,[],[f50029,f56332])).
fof(f56332,plain,(
  ( ! [X77] : (s(bool,X77) = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X77)) ) | ($spl4 | $spl9 | $spl36 | $spl46)),
  inference(subsumption_resolution,[],[f56210,f37])).
fof(f56210,plain,(
  ( ! [X77] : (~p(s(bool,'T_0')) | s(bool,X77) = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X77)) ) | ($spl4 | $spl9 | $spl36 | $spl46)),
  inference(superposition,[],[f30024,f3303])).
fof(f63839,plain,(
  ( ! [X182,X183] : (s(bool,X182) != s(bool,f(s(bool,X182))) | s(bool,'T_0') = s(bool,X183) | s(bool,'F_0') = s(bool,X182)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl176)),
  inference(subsumption_resolution,[],[f63705,f25786])).
fof(f63705,plain,(
  ( ! [X182,X183] : (s(bool,X182) != s(bool,f(s(bool,X182))) | s(bool,X183) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X183) | s(bool,'F_0') = s(bool,X182)) ) | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl176)),
  inference(superposition,[],[f50029,f30750])).
fof(f60407,plain,(
  ( ! [X180] : (s(bool,'T_0') = s(bool,f(s(bool,X180))) | s(bool,X180) = s(bool,g(s(bool,'T_0')))) ) | $spl278),
  inference(cnf_transformation,[],[f60407_D])).
fof(f60407_D,plain,(
  ( ! [X180] : (s(bool,'T_0') = s(bool,f(s(bool,X180))) | s(bool,X180) = s(bool,g(s(bool,'T_0')))) ) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f60406,plain,(
  ( ! [X180,X179] : (s(bool,'T_0') = s(bool,f(s(bool,X180))) | s(bool,X180) = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X179)) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl172)),
  inference(subsumption_resolution,[],[f60333,f29278])).
fof(f60333,plain,(
  ( ! [X180,X179] : (~p(s(bool,X179)) | s(bool,'T_0') = s(bool,f(s(bool,X180))) | s(bool,X180) = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X179)) ) | ($spl4 | $spl9 | $spl36 | $spl46)),
  inference(superposition,[],[f36504,f56332])).
fof(f57575,plain,(
  ( ! [X150] : (s(bool,X150) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X150)) ) | $spl276),
  inference(cnf_transformation,[],[f57575_D])).
fof(f57575_D,plain,(
  ( ! [X150] : (s(bool,X150) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X150)) ) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f57455,plain,(
  ( ! [X149,X150] : (s(bool,X150) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X150) | s(bool,'F_0') = s(bool,X149)) ) | ($spl4 | $spl9 | $spl36 | $spl46)),
  inference(duplicate_literal_removal,[],[f57311])).
fof(f57311,plain,(
  ( ! [X149,X150] : (s(bool,X150) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X150) | s(bool,'F_0') = s(bool,X149) | s(bool,'F_0') = s(bool,X149)) ) | ($spl4 | $spl9 | $spl36 | $spl46)),
  inference(superposition,[],[f29101,f56332])).
fof(f57515,plain,(
  $false | ($spl4 | $spl9 | $spl36 | $spl46 | $spl230)),
  inference(subsumption_resolution,[],[f57514,f71])).
fof(f57514,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl4 | $spl9 | $spl36 | $spl46 | $spl230)),
  inference(condensation,[],[f57513])).
fof(f57513,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl230)),
  inference(condensation,[],[f57512])).
fof(f57512,plain,(
  ( ! [X103,X102] : (s(bool,'F_0') = s(bool,X103) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X102)) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl230)),
  inference(subsumption_resolution,[],[f57229,f38765])).
fof(f57229,plain,(
  ( ! [X103,X102] : (s(bool,X103) = s(bool,g(s(bool,X102))) | s(bool,'F_0') = s(bool,X103) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X102)) ) | ($spl4 | $spl9 | $spl36 | $spl46)),
  inference(superposition,[],[f56332,f28992])).
fof(f57511,plain,(
  $false | ($spl4 | $spl9 | $spl36 | $spl46 | $spl230)),
  inference(subsumption_resolution,[],[f57510,f71])).
fof(f57510,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl4 | $spl9 | $spl36 | $spl46 | $spl230)),
  inference(condensation,[],[f57509])).
fof(f57509,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,'F_0')) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl230)),
  inference(condensation,[],[f57508])).
fof(f57508,plain,(
  ( ! [X101,X100] : (s(bool,'F_0') = s(bool,X101) | s(bool,'F_0') = s(bool,X100) | s(bool,'T_0') = s(bool,'F_0')) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl230)),
  inference(subsumption_resolution,[],[f57228,f38765])).
fof(f57228,plain,(
  ( ! [X101,X100] : (s(bool,X101) = s(bool,g(s(bool,X100))) | s(bool,'F_0') = s(bool,X101) | s(bool,'F_0') = s(bool,X100) | s(bool,'T_0') = s(bool,'F_0')) ) | ($spl4 | $spl9 | $spl36 | $spl46)),
  inference(superposition,[],[f56332,f28992])).
fof(f56335,plain,(
  $false | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl46 | $spl218)),
  inference(subsumption_resolution,[],[f56334,f49260])).
fof(f56334,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | ($spl4 | $spl9 | $spl36 | $spl40 | $spl46 | $spl218)),
  inference(subsumption_resolution,[],[f56215,f37182])).
fof(f56215,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | ($spl4 | $spl9 | $spl36 | $spl40 | $spl46)),
  inference(superposition,[],[f41609,f3303])).
fof(f56311,plain,(
  $false | ($spl4 | $spl9 | $spl13 | $spl36 | $spl40 | $spl46 | $spl218)),
  inference(subsumption_resolution,[],[f56310,f49260])).
fof(f56310,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | ($spl4 | $spl9 | $spl36 | $spl40 | $spl46 | $spl218)),
  inference(subsumption_resolution,[],[f56170,f37182])).
fof(f56170,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | ($spl4 | $spl9 | $spl36 | $spl40 | $spl46)),
  inference(superposition,[],[f3303,f41609])).
fof(f53059,plain,(
  $false | ($spl4 | $spl9 | $spl13 | $spl22 | $spl36 | $spl40 | $spl42 | $spl234 | $spl250 | $spl258)),
  inference(subsumption_resolution,[],[f52793,f52765])).
fof(f52765,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl4 | $spl9 | $spl36 | $spl40 | $spl42 | $spl234 | $spl250 | $spl258)),
  inference(subsumption_resolution,[],[f52764,f3890])).
fof(f52764,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'T_0'),s(bool,X0)))))) ) | ($spl4 | $spl9 | $spl36 | $spl40 | $spl42 | $spl234 | $spl250 | $spl258)),
  inference(condensation,[],[f52707])).
fof(f52707,plain,(
  ( ! [X64,X63] : (s(bool,X64) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'T_0'),s(bool,X64))))) | s(bool,'T_0') = s(bool,X63) | s(bool,'T_0') = s(bool,X64)) ) | ($spl4 | $spl9 | $spl36 | $spl40 | $spl42 | $spl234 | $spl250 | $spl258)),
  inference(backward_demodulation,[],[f52706,f3170])).
fof(f52706,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,f(s(bool,X2)))) ) | ($spl4 | $spl9 | $spl36 | $spl40 | $spl234 | $spl250 | $spl258)),
  inference(forward_demodulation,[],[f52382,f44])).
fof(f52382,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,f(s(bool,'const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,X2)))))))) ) | ($spl4 | $spl9 | $spl36 | $spl40 | $spl234 | $spl250 | $spl258)),
  inference(unit_resulting_resolution,[],[f51575,f49698])).
fof(f51575,plain,(
  ( ! [X152,X151,X153] : (s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X151,bool),'const.pair.,_2'(s(X151,X152),s(bool,X153)))))) ) | ($spl40 | $spl234 | $spl258)),
  inference(subsumption_resolution,[],[f51574,f49250])).
fof(f51574,plain,(
  ( ! [X152,X151,X153] : (s(bool,'F_0') != s(bool,X153) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X151,bool),'const.pair.,_2'(s(X151,X152),s(bool,X153)))))) ) | ($spl40 | $spl234)),
  inference(forward_demodulation,[],[f51478,f2561])).
fof(f51478,plain,(
  ( ! [X152,X151,X153] : (s(bool,X153) != s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X151,bool),'const.pair.,_2'(s(X151,X152),s(bool,X153)))))) ) | $spl234),
  inference(superposition,[],[f42127,f44])).
fof(f52793,plain,(
  s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'T_0'),s(bool,'T_0'))))) | ($spl4 | $spl9 | $spl13 | $spl22 | $spl36 | $spl40 | $spl42 | $spl234 | $spl250 | $spl258)),
  inference(backward_demodulation,[],[f52765,f6406])).
fof(f53057,plain,(
  $false | ($spl4 | $spl9 | $spl36 | $spl40 | $spl42 | $spl234 | $spl250 | $spl258)),
  inference(subsumption_resolution,[],[f52766,f37])).
fof(f52766,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl9 | $spl36 | $spl40 | $spl42 | $spl234 | $spl250 | $spl258)),
  inference(backward_demodulation,[],[f52765,f39])).
fof(f49696,plain,(
  ( ! [X163,X159,X164] : (s(bool,'T_0') = s(bool,f(s(bool,X159))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X163,bool),'const.pair.,_2'(s(X163,X164),s(bool,X159)))))) ) | $spl274),
  inference(cnf_transformation,[],[f49696_D])).
fof(f49696_D,plain,(
  ( ! [X163,X159,X164] : (s(bool,'T_0') = s(bool,f(s(bool,X159))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X163,bool),'const.pair.,_2'(s(X163,X164),s(bool,X159)))))) ) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f49248,plain,(
  ( ! [X163,X159,X164,X160] : (s(bool,'T_0') = s(bool,f(s(bool,X159))) | s(bool,'F_0') = s(bool,X160) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X163,bool),'const.pair.,_2'(s(X163,X164),s(bool,X159)))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(global_subsumption,[],[f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f67,f64,f132,f133,f146,f156,f157,f200,f361,f363,f71,f54,f1085,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f4170,f4171,f3832,f3890,f3878,f2550,f2549,f3652,f3732,f3828,f3727,f2700,f763,f2698,f761,f2455,f1935,f1480,f82,f104,f2706,f11126,f11125,f11188,f11200,f11161,f11162,f11222,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13852,f13225,f13350,f13427,f13415,f15429,f17356,f22573,f18154,f18155,f22295,f22291,f22279,f22275,f22219,f22217,f18360,f18362,f18363,f22711,f18802,f19158,f21203,f22901,f22902,f22905,f19338,f19402,f19403,f19570,f22984,f20200,f23229,f23233,f23235,f24931,f24930,f24854,f24822,f24800,f23787,f25217,f25261,f25262,f25265,f25280,f25784,f25785,f24220,f24215,f25787,f25788,f23788,f23783,f25790,f23712,f20161,f20156,f25794,f19927,f19749,f19571,f19566,f25798,f19337,f19159,f18981,f18976,f18803,f18798,f25801,f25802,f13351,f13346,f25804,f13292,f13175,f13170,f13039,f11230,f9980,f9978,f9977,f9974,f6734,f6732,f6731,f6728,f24431,f25850,f25854,f28579,f28795,f28928,f28929,f28799,f28801,f28992,f29101,f29277,f29642,f30026,f30749,f30802,f30902,f31517,f31518,f31523,f32112,f32190,f32954,f32955,f32960,f33699,f33700,f33637,f33624,f34113,f34117,f35168,f35169,f35170,f35172,f36886,f36851,f36833,f37018,f37019,f37020,f37021,f36024,f36821,f37037,f37041,f36272,f36746,f37183,f37187,f37191,f36506,f36507,f36509,f36510,f37202,f37203,f36577,f36580,f37595,f13038,f3651,f17264,f17266,f15352,f22930,f22929,f22928,f786,f550,f475,f48,f47,f46,f25855,f19925,f19747,f19569,f18801,f762,f2762,f11302,f2704,f19406,f25856,f19914,f19736,f19558,f18790,f2699,f785,f2763,f22931,f37580,f37578,f37577,f37576,f37574,f37033,f36728,f34476,f34474,f34472,f34425,f24214,f23782,f22898,f19565,f19331,f19153,f18797,f16976,f15069,f13826,f13345,f13169,f13033,f4154,f3827,f3726,f3646,f2559,f368,f366,f37594,f37211,f37040,f37038,f37032,f36953,f36870,f36709,f36702,f36501,f36494,f36344,f36337,f36111,f36110,f36025,f36023,f36022,f34518,f34208,f34207,f34119,f34116,f34115,f33645,f33643,f33629,f33627,f33552,f33551,f32991,f32792,f32361,f32194,f32193,f32192,f32154,f32152,f32151,f32126,f32105,f32101,f31431,f31357,f31001,f30906,f30905,f30904,f30404,f30179,f30178,f30174,f30172,f30080,f30030,f30029,f30027,f30024,f29905,f29904,f29900,f29898,f29715,f29714,f29581,f29580,f29283,f29256,f29255,f29252,f29247,f29245,f29243,f29241,f25810,f25806,f25792,f25786,f22303,f22293,f22277,f21173,f19401,f19400,f17276,f17226,f15341,f13224,f13223,f11203,f3759,f3758,f2874,f1972,f1473,f370,f365,f39,f36343,f36874,f36872,f36708,f36696,f36575,f36504,f36500,f36488,f36331,f36184,f35841,f35840,f13854,f13853,f11293,f11201,f36263,f45254,f45391,f45392,f45396,f45400,f45453,f45454,f45538,f45539,f45544,f45545,f45588,f49037,f49038,f49051,f37194,f49053,f36718,f36723,f49078,f45327,f45330,f44862,f36220,f49246,f36225])).
fof(f49570,plain,(
  $false | ($spl13 | $spl22 | $spl42 | $spl250 | $spl252)),
  inference(subsumption_resolution,[],[f49549,f49541])).
fof(f49541,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl42 | $spl250 | $spl252)),
  inference(subsumption_resolution,[],[f49540,f3890])).
fof(f49540,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'T_0'),s(bool,X0)))))) ) | ($spl42 | $spl250 | $spl252)),
  inference(condensation,[],[f49539])).
fof(f49539,plain,(
  ( ! [X64,X63] : (s(bool,X64) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'T_0'),s(bool,X64))))) | s(bool,'T_0') = s(bool,X63) | s(bool,'T_0') = s(bool,X64)) ) | ($spl42 | $spl250 | $spl252)),
  inference(forward_demodulation,[],[f3170,f49513])).
fof(f49513,plain,(
  ( ! [X98] : (s(bool,'T_0') = s(bool,f(s(bool,X98)))) ) | ($spl250 | $spl252)),
  inference(subsumption_resolution,[],[f49056,f49052])).
fof(f49549,plain,(
  s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'T_0'),s(bool,'T_0'))))) | ($spl13 | $spl22 | $spl42 | $spl250 | $spl252)),
  inference(backward_demodulation,[],[f49541,f6406])).
fof(f49569,plain,(
  $false | ($spl42 | $spl250 | $spl252)),
  inference(subsumption_resolution,[],[f49542,f37])).
fof(f49542,plain,(
  ~p(s(bool,'T_0')) | ($spl42 | $spl250 | $spl252)),
  inference(backward_demodulation,[],[f49541,f39])).
fof(f49445,plain,(
  ( ! [X4,X2,X3] : (s(fun(X2,bool),X3) = s(fun(X2,bool),X4)) ) | $spl272),
  inference(cnf_transformation,[],[f49445_D])).
fof(f49445_D,plain,(
  ( ! [X4,X2,X3] : (s(fun(X2,bool),X3) = s(fun(X2,bool),X4)) ) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f49450,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,'F_0'),s(X0,X4))))) | s(fun(X1,bool),X2) = s(fun(X1,bool),X3)) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl228)),
  inference(forward_demodulation,[],[f49449,f49258])).
fof(f49258,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl228)),
  inference(subsumption_resolution,[],[f49257,f3303])).
fof(f49257,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') != s(bool,g(s(bool,'T_0')))) ) | ($spl4 | $spl9 | $spl36 | $spl228)),
  inference(condensation,[],[f49256])).
fof(f49256,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2)) ) | ($spl4 | $spl9 | $spl36 | $spl228)),
  inference(subsumption_resolution,[],[f40642,f30026])).
fof(f40642,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl228),
  inference(superposition,[],[f38764,f56])).
fof(f49449,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,g(s(X1,sK6(X2,X3,bool,X1)))),s(X0,X4))))) | s(fun(X1,bool),X2) = s(fun(X1,bool),X3)) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl228)),
  inference(condensation,[],[f49448])).
fof(f49448,plain,(
  ( ! [X412,X414,X410,X413,X409,X411] : (s(fun(X409,bool),X410) = s(fun(X409,bool),X411) | s(bool,'T_0') = s(bool,X412) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X413),'const.pair.,_2'(s(bool,g(s(X409,sK6(X410,X411,bool,X409)))),s(X413,X414)))))) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl228)),
  inference(subsumption_resolution,[],[f49447,f40])).
fof(f49447,plain,(
  ( ! [X412,X414,X410,X413,X409,X411] : (s(bool,'F_0') != s(bool,X412) | s(fun(X409,bool),X410) = s(fun(X409,bool),X411) | s(bool,'T_0') = s(bool,X412) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X413),'const.pair.,_2'(s(bool,g(s(X409,sK6(X410,X411,bool,X409)))),s(X413,X414)))))) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl228)),
  inference(forward_demodulation,[],[f3646,f49258])).
fof(f49444,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,'F_0'))))) | s(fun(X2,bool),X3) = s(fun(X2,bool),X4)) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl228)),
  inference(forward_demodulation,[],[f49443,f49258])).
fof(f49443,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,g(s(X2,sK6(X3,X4,bool,X2)))))))) | s(fun(X2,bool),X3) = s(fun(X2,bool),X4)) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl228)),
  inference(condensation,[],[f49442])).
fof(f49442,plain,(
  ( ! [X705,X700,X702,X704,X701,X703] : (s(fun(X700,bool),X701) = s(fun(X700,bool),X702) | s(bool,'T_0') = s(bool,X703) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X704,bool),'const.pair.,_2'(s(X704,X705),s(bool,g(s(X700,sK6(X701,X702,bool,X700))))))))) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl228)),
  inference(subsumption_resolution,[],[f49441,f40])).
fof(f49441,plain,(
  ( ! [X705,X700,X702,X704,X701,X703] : (s(bool,'F_0') != s(bool,X703) | s(fun(X700,bool),X701) = s(fun(X700,bool),X702) | s(bool,'T_0') = s(bool,X703) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X704,bool),'const.pair.,_2'(s(X704,X705),s(bool,g(s(X700,sK6(X701,X702,bool,X700))))))))) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl228)),
  inference(forward_demodulation,[],[f13033,f49258])).
fof(f25852,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK6(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | $spl168),
  inference(cnf_transformation,[],[f25852_D])).
fof(f25852_D,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK6(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f49426,plain,(
  ( ! [X693,X691,X692,X690] : (s(fun(X690,bool),X691) = s(fun(X690,bool),X692) | s(bool,'T_0') = s(bool,g(s(X690,sK6(X691,X692,bool,X690)))) | s(bool,X693) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl228)),
  inference(subsumption_resolution,[],[f49425,f49258])).
fof(f49425,plain,(
  ( ! [X693,X691,X692,X690] : (s(bool,'F_0') != s(bool,X693) | s(fun(X690,bool),X691) = s(fun(X690,bool),X692) | s(bool,'T_0') = s(bool,g(s(X690,sK6(X691,X692,bool,X690)))) | s(bool,X693) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl228)),
  inference(forward_demodulation,[],[f24214,f49258])).
fof(f49286,plain,(
  ( ! [X732,X730,X733,X731] : (s(X732,f(s(bool,X733))) != s(X732,g(s(bool,X733))) | s(bool,'T_0') = s(bool,X733) | s(fun(bool,X732),X730) = s(fun(bool,X732),X731)) ) | $spl268),
  inference(cnf_transformation,[],[f49286_D])).
fof(f49286_D,plain,(
  ( ! [X732,X730,X733,X731] : (s(X732,f(s(bool,X733))) != s(X732,g(s(bool,X733))) | s(bool,'T_0') = s(bool,X733) | s(fun(bool,X732),X730) = s(fun(bool,X732),X731)) ) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f49287,plain,(
  ( ! [X441,X443,X444,X440,X442,X439] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X443),'const.pair.,_2'(s(bool,'F_0'),s(X443,X444))))) | s(X441,f(s(bool,X442))) != s(X441,g(s(bool,X442))) | s(fun(bool,X441),X439) = s(fun(bool,X441),X440) | s(bool,'T_0') = s(bool,X442)) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl228)),
  inference(forward_demodulation,[],[f3651,f49258])).
fof(f49285,plain,(
  ( ! [X732,X734,X730,X733,X735,X731] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X734,bool),'const.pair.,_2'(s(X734,X735),s(bool,'F_0'))))) | s(X732,f(s(bool,X733))) != s(X732,g(s(bool,X733))) | s(fun(bool,X732),X730) = s(fun(bool,X732),X731) | s(bool,'T_0') = s(bool,X733)) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl228)),
  inference(forward_demodulation,[],[f13038,f49258])).
fof(f49255,plain,(
  $false | ($spl9 | $spl46 | $spl228)),
  inference(subsumption_resolution,[],[f40640,f3303])).
fof(f40640,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | ($spl9 | $spl228)),
  inference(unit_resulting_resolution,[],[f71,f38764])).
fof(f49253,plain,(
  ( ! [X216,X217] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X216,bool),'const.pair.,_2'(s(X216,X217),s(bool,f(s(bool,'F_0')))))))) ) | $spl264),
  inference(cnf_transformation,[],[f49253_D])).
fof(f49253_D,plain,(
  ( ! [X216,X217] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X216,bool),'const.pair.,_2'(s(X216,X217),s(bool,f(s(bool,'F_0')))))))) ) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f49250,plain,(
  ( ! [X218,X215,X219] : (s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X218,bool),'const.pair.,_2'(s(X218,X219),s(bool,X215))))) | s(bool,'F_0') = s(bool,X215)) ) | $spl258),
  inference(cnf_transformation,[],[f49250_D])).
fof(f49250_D,plain,(
  ( ! [X218,X215,X219] : (s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X218,bool),'const.pair.,_2'(s(X218,X219),s(bool,X215))))) | s(bool,'F_0') = s(bool,X215)) ) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f45883,plain,(
  ( ! [X308,X310,X306,X309,X307] : (s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X307,bool),'const.pair.,_2'(s(X307,X308),s(bool,X306))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X309,bool),'const.pair.,_2'(s(X309,X310),s(bool,f(s(bool,'F_0'))))))) | s(bool,'F_0') = s(bool,X306)) ) | $spl176),
  inference(superposition,[],[f1972,f30750])).
fof(f49254,plain,(
  ( ! [X208,X207,X203] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X207,bool),'const.pair.,_2'(s(X207,X208),s(bool,X203))))) | s(bool,'F_0') = s(bool,X203)) ) | $spl266),
  inference(cnf_transformation,[],[f49254_D])).
fof(f49254_D,plain,(
  ( ! [X208,X207,X203] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X207,bool),'const.pair.,_2'(s(X207,X208),s(bool,X203))))) | s(bool,'F_0') = s(bool,X203)) ) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f45867,plain,(
  ( ! [X212,X210,X213,X211,X209,X214] : (s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X211,bool),'const.pair.,_2'(s(X211,X212),s(bool,X210))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X213,bool),'const.pair.,_2'(s(X213,X214),s(bool,X209))))) | s(bool,'F_0') = s(bool,X209) | s(bool,'F_0') = s(bool,X210)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f1972,f28992])).
fof(f45866,plain,(
  ( ! [X208,X206,X204,X207,X205,X203] : (s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X205,bool),'const.pair.,_2'(s(X205,X206),s(bool,X204))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X207,bool),'const.pair.,_2'(s(X207,X208),s(bool,X203))))) | s(bool,'F_0') = s(bool,X204) | s(bool,'F_0') = s(bool,X203)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f1972,f28992])).
fof(f45693,plain,(
  ( ! [X218,X216,X215,X219,X217] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X216,bool),'const.pair.,_2'(s(X216,X217),s(bool,f(s(bool,'F_0'))))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X218,bool),'const.pair.,_2'(s(X218,X219),s(bool,X215))))) | s(bool,'F_0') = s(bool,X215)) ) | $spl176),
  inference(superposition,[],[f1972,f30750])).
fof(f49251,plain,(
  ( ! [X216,X217] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X216),'const.pair.,_2'(s(bool,f(s(bool,'F_0'))),s(X216,X217)))))) ) | $spl260),
  inference(cnf_transformation,[],[f49251_D])).
fof(f49251_D,plain,(
  ( ! [X216,X217] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X216),'const.pair.,_2'(s(bool,f(s(bool,'F_0'))),s(X216,X217)))))) ) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f43907,plain,(
  ( ! [X292,X294,X296,X293,X295] : (s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X293,bool),'const.pair.,_2'(s(X293,X294),s(bool,X292))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X295),'const.pair.,_2'(s(bool,f(s(bool,'F_0'))),s(X295,X296))))) | s(bool,'F_0') = s(bool,X292)) ) | $spl176),
  inference(superposition,[],[f1473,f30750])).
fof(f49252,plain,(
  ( ! [X208,X207,X203] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X207),'const.pair.,_2'(s(bool,X203),s(X207,X208))))) | s(bool,'F_0') = s(bool,X203)) ) | $spl262),
  inference(cnf_transformation,[],[f49252_D])).
fof(f49252_D,plain,(
  ( ! [X208,X207,X203] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X207),'const.pair.,_2'(s(bool,X203),s(X207,X208))))) | s(bool,'F_0') = s(bool,X203)) ) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f43894,plain,(
  ( ! [X212,X210,X213,X211,X209,X214] : (s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X211,bool),'const.pair.,_2'(s(X211,X212),s(bool,X210))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X213),'const.pair.,_2'(s(bool,X209),s(X213,X214))))) | s(bool,'F_0') = s(bool,X209) | s(bool,'F_0') = s(bool,X210)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f1473,f28992])).
fof(f43893,plain,(
  ( ! [X208,X206,X204,X207,X205,X203] : (s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X205,bool),'const.pair.,_2'(s(X205,X206),s(bool,X204))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X207),'const.pair.,_2'(s(bool,X203),s(X207,X208))))) | s(bool,'F_0') = s(bool,X204) | s(bool,'F_0') = s(bool,X203)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f1473,f28992])).
fof(f43667,plain,(
  ( ! [X218,X216,X215,X219,X217] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X216),'const.pair.,_2'(s(bool,f(s(bool,'F_0'))),s(X216,X217))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X218,bool),'const.pair.,_2'(s(X218,X219),s(bool,X215))))) | s(bool,'F_0') = s(bool,X215)) ) | $spl176),
  inference(superposition,[],[f1473,f30750])).
fof(f7901,plain,(
  ( ! [X14,X15,X16] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X15,bool),'const.pair.,_2'(s(X15,X16),s(bool,X14)))))) ) | $spl100),
  inference(cnf_transformation,[],[f7901_D])).
fof(f7901_D,plain,(
  ( ! [X14,X15,X16] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X15,bool),'const.pair.,_2'(s(X15,X16),s(bool,X14)))))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f49249,plain,(
  ( ! [X163,X159,X164,X160] : (s(bool,'F_0') = s(bool,X160) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X163,bool),'const.pair.,_2'(s(X163,X164),s(bool,X159)))))) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl228)),
  inference(global_subsumption,[],[f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f67,f64,f132,f133,f146,f156,f157,f200,f361,f363,f71,f54,f1085,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f4170,f4171,f3832,f3890,f3878,f2550,f2549,f3652,f3732,f3828,f3727,f2700,f763,f2698,f761,f5160,f5161,f2455,f1935,f1480,f82,f104,f2706,f11126,f11125,f11188,f11200,f11161,f11162,f11222,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13852,f13181,f13225,f13350,f13427,f13415,f15429,f17356,f22573,f18154,f18155,f22299,f22295,f22291,f22279,f22275,f22219,f22217,f18360,f18362,f18363,f22711,f18802,f19158,f21203,f22901,f22902,f22905,f19338,f19344,f19402,f19403,f19570,f22984,f20200,f23229,f23233,f23235,f24931,f24930,f24854,f24822,f24800,f23543,f23787,f25217,f25261,f25262,f25265,f25280,f25784,f25785,f24220,f24215,f25787,f25788,f23788,f23783,f25790,f23712,f20161,f20156,f25794,f19927,f19749,f19571,f19566,f25798,f19337,f19159,f18981,f18976,f18803,f18798,f25801,f25802,f13351,f13346,f25804,f13292,f13175,f13170,f13039,f11230,f9980,f9978,f9977,f9974,f6734,f6732,f6731,f6728,f24431,f25850,f25854,f28579,f28795,f28928,f28929,f28799,f28801,f28992,f29101,f29277,f29642,f30026,f30749,f30802,f30902,f31517,f31518,f31523,f32112,f32190,f32954,f32955,f32960,f33699,f33700,f33637,f33624,f34113,f34117,f35168,f35169,f35170,f35172,f36886,f36851,f36833,f37018,f37019,f37020,f37021,f36024,f36821,f37037,f37041,f36272,f36746,f37183,f37187,f37191,f36506,f36507,f36509,f36510,f37202,f37203,f36577,f36580,f37595,f13038,f3651,f17264,f17266,f15352,f22930,f22929,f22928,f786,f550,f475,f48,f47,f46,f25855,f19925,f19747,f19569,f18801,f762,f2762,f11302,f2704,f19406,f25856,f19914,f19736,f19558,f18790,f2699,f785,f2763,f22931,f37580,f37578,f37577,f37576,f37574,f37033,f36728,f34476,f34474,f34472,f34425,f24214,f23782,f22898,f19565,f19331,f19153,f18797,f16976,f15069,f13826,f13345,f13169,f13033,f4154,f3827,f3726,f3646,f2559,f368,f366,f37594,f37211,f37040,f37038,f37032,f36953,f36870,f36709,f36702,f36501,f36494,f36344,f36337,f36111,f36110,f36025,f36023,f36022,f34518,f34208,f34207,f34119,f34116,f34115,f33645,f33643,f33629,f33627,f33552,f33551,f32991,f32792,f32361,f32194,f32193,f32192,f32154,f32152,f32151,f32126,f32105,f32101,f31431,f31357,f31001,f30906,f30905,f30904,f30404,f30179,f30178,f30174,f30172,f30080,f30030,f30029,f30027,f30024,f29905,f29904,f29900,f29898,f29715,f29714,f29581,f29580,f29283,f29256,f29255,f29252,f29247,f29245,f29243,f29241,f25810,f25806,f25792,f25786,f22303,f22293,f22277,f21173,f19401,f19400,f17276,f17226,f15341,f13224,f13223,f11203,f3759,f3758,f2874,f1972,f1473,f370,f365,f39,f36343,f36874,f36872,f36708,f36696,f36575,f36504,f36500,f36488,f36331,f36184,f35841,f35840,f13854,f13853,f11293,f11201,f36263,f38764,f40698,f40699,f45254,f45391,f45392,f45396,f45400,f45453,f45454,f45538,f45539,f45544,f45545,f45588,f48458,f48459,f46149,f48467,f48475,f3303,f49037,f49038,f49051,f37194,f49053,f36718,f36723,f49078,f45327,f45330,f44862,f49239,f40747,f40748,f40749,f49240,f40752,f45083,f36220,f49246,f36225,f49248])).
fof(f45083,plain,(
  ( ! [X181] : (s(bool,X181) != s(bool,f(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X181) | s(bool,'T_0') = s(bool,f(s(bool,g(s(bool,X181)))))) ) | ($spl4 | $spl9 | $spl36 | $spl228)),
  inference(superposition,[],[f38764,f37032])).
fof(f40752,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'F_0') = s(bool,g(s(bool,X0)))) ) | ($spl4 | $spl9 | $spl36 | $spl228)),
  inference(condensation,[],[f40751])).
fof(f40751,plain,(
  ( ! [X88,X89] : (s(bool,'F_0') = s(bool,X88) | s(bool,'F_0') = s(bool,X89) | s(bool,'F_0') = s(bool,g(s(bool,X88)))) ) | ($spl4 | $spl9 | $spl36 | $spl228)),
  inference(subsumption_resolution,[],[f40731,f28992])).
fof(f40731,plain,(
  ( ! [X88,X89] : (s(bool,X88) != s(bool,X89) | s(bool,'F_0') = s(bool,X88) | s(bool,'F_0') = s(bool,X89) | s(bool,'F_0') = s(bool,g(s(bool,X88)))) ) | ($spl4 | $spl9 | $spl36 | $spl228)),
  inference(superposition,[],[f38764,f28992])).
fof(f49240,plain,(
  ( ! [X78,X77] : (s(bool,'F_0') = s(bool,X77) | s(bool,X78) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X78) = s(bool,g(s(bool,X77)))) ) | ($spl4 | $spl228)),
  inference(subsumption_resolution,[],[f40725,f40])).
fof(f40725,plain,(
  ( ! [X78,X77] : (s(bool,'T_0') != s(bool,X77) | s(bool,'F_0') = s(bool,X77) | s(bool,X78) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X78) = s(bool,g(s(bool,X77)))) ) | ($spl4 | $spl228)),
  inference(superposition,[],[f38764,f22295])).
fof(f40749,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,'F_0') = s(bool,X42) | s(bool,X43) = s(bool,X44) | s(bool,X44) = s(bool,g(s(bool,X42))) | s(bool,X43) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) ) | $spl228),
  inference(subsumption_resolution,[],[f40715,f40])).
fof(f40715,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,'T_0') != s(bool,X42) | s(bool,'F_0') = s(bool,X42) | s(bool,X43) = s(bool,X44) | s(bool,X44) = s(bool,g(s(bool,X42))) | s(bool,X43) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) ) | $spl228),
  inference(superposition,[],[f38764,f146])).
fof(f40748,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,'F_0') = s(bool,X30) | s(bool,X31) = s(bool,g(s(bool,X30))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X32,bool),'const.pair.,_2'(s(X32,X33),s(bool,X31)))))) ) | $spl228),
  inference(subsumption_resolution,[],[f40712,f40])).
fof(f40712,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,'T_0') != s(bool,X30) | s(bool,'F_0') = s(bool,X30) | s(bool,X31) = s(bool,g(s(bool,X30))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X32,bool),'const.pair.,_2'(s(X32,X33),s(bool,X31)))))) ) | $spl228),
  inference(superposition,[],[f38764,f133])).
fof(f40747,plain,(
  ( ! [X21,X19,X20,X18] : (s(bool,'F_0') = s(bool,X18) | s(bool,X19) = s(bool,g(s(bool,X18))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X20),'const.pair.,_2'(s(bool,X19),s(X20,X21)))))) ) | $spl228),
  inference(subsumption_resolution,[],[f40709,f40])).
fof(f40709,plain,(
  ( ! [X21,X19,X20,X18] : (s(bool,'T_0') != s(bool,X18) | s(bool,'F_0') = s(bool,X18) | s(bool,X19) = s(bool,g(s(bool,X18))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X20),'const.pair.,_2'(s(bool,X19),s(X20,X21)))))) ) | $spl228),
  inference(superposition,[],[f38764,f132])).
fof(f49239,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') = s(bool,X2) | s(bool,X3) = s(bool,g(s(bool,X2))) | s(bool,'T_0') = s(bool,X3)) ) | $spl228),
  inference(subsumption_resolution,[],[f40703,f40])).
fof(f40703,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X3) = s(bool,g(s(bool,X2))) | s(bool,'T_0') = s(bool,X3)) ) | $spl228),
  inference(superposition,[],[f38764,f56])).
fof(f48475,plain,(
  ( ! [X174] : (s(bool,X174) != s(bool,g(s(bool,X174))) | s(bool,'T_0') = s(bool,f(s(bool,X174)))) ) | ($spl4 | $spl9 | $spl36 | $spl228)),
  inference(subsumption_resolution,[],[f48474,f71])).
fof(f48474,plain,(
  ( ! [X174] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X174) != s(bool,g(s(bool,X174))) | s(bool,'T_0') = s(bool,f(s(bool,X174)))) ) | ($spl4 | $spl9 | $spl36 | $spl228)),
  inference(forward_demodulation,[],[f45189,f11203])).
fof(f45189,plain,(
  ( ! [X174] : (s(bool,X174) != s(bool,g(s(bool,X174))) | s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X174)))) ) | ($spl4 | $spl9 | $spl36 | $spl228)),
  inference(superposition,[],[f38764,f37032])).
fof(f48467,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,X2)))))) ) | ($spl46 | $spl228)),
  inference(global_subsumption,[],[f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f64,f132,f133,f146,f156,f157,f200,f361,f363,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f3832,f3890,f3878,f2550,f2549,f3303,f3652,f3732,f3828,f3727,f2700,f763,f2698,f761,f5160,f5161,f1935,f1480,f82,f104,f2706,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13181,f13225,f13350,f13427,f13415,f22573,f18154,f18155,f22299,f22291,f22279,f22275,f22219,f22217,f18360,f18362,f18363,f18802,f19158,f21203,f19338,f19344,f19402,f19403,f19570,f22984,f20200,f20161,f20156,f19927,f19749,f19571,f19566,f19337,f19159,f18981,f18976,f18803,f18798,f25801,f25802,f13351,f13346,f13175,f13170,f13039,f9980,f9978,f9977,f9974,f6734,f6732,f6731,f6728,f25850,f25854,f30902,f32190,f35168,f35169,f35170,f13038,f3651,f17264,f17266,f15352,f22930,f22929,f22928,f786,f550,f475,f48,f47,f46,f19925,f19747,f19569,f18801,f762,f2704,f19406,f19914,f19736,f19558,f18790,f2699,f785,f2763,f22931,f37580,f37578,f37577,f37576,f34476,f34474,f34472,f34425,f22898,f19565,f19331,f19153,f18797,f16976,f15069,f13826,f13345,f13169,f13033,f4154,f3827,f3726,f3646,f2559,f368,f366,f32792,f32194,f32193,f32192,f31357,f30906,f30905,f30904,f22293,f22277,f19401,f19400,f17276,f17226,f15341,f13224,f13223,f3759,f3758,f2874,f1972,f1473,f370,f365,f39,f38764,f40698,f40699,f46149])).
fof(f46149,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,X2))))) | s(bool,'T_0') != s(bool,g(s(bool,'T_0')))) ) | $spl228),
  inference(condensation,[],[f46005])).
fof(f46005,plain,(
  ( ! [X428,X426,X429,X425,X427] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X425,bool),'const.pair.,_2'(s(X425,X426),s(bool,X427))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X428,bool),'const.pair.,_2'(s(X428,X429),s(bool,X427)))))) ) | $spl228),
  inference(superposition,[],[f38764,f1972])).
fof(f48459,plain,(
  ( ! [X313,X312,X314] : (s(bool,'F_0') = s(bool,g(s(X312,sK6(X313,X314,bool,X312)))) | s(fun(X312,bool),X313) = s(fun(X312,bool),X314)) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl228)),
  inference(global_subsumption,[],[f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f67,f64,f132,f133,f146,f156,f157,f200,f361,f363,f71,f54,f1085,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f4170,f4171,f3832,f3890,f3878,f2550,f2549,f3303,f3652,f3732,f3828,f3727,f2700,f763,f2698,f761,f5160,f5161,f2455,f1935,f1480,f82,f104,f2706,f11126,f11125,f11188,f11200,f11161,f11162,f11222,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13852,f13181,f13225,f13350,f13427,f13415,f15429,f17356,f22573,f18154,f18155,f22299,f22295,f22291,f22279,f22275,f22219,f22217,f18360,f18362,f18363,f22711,f18802,f19158,f21203,f22901,f22902,f22905,f19338,f19344,f19402,f19403,f19570,f22984,f20200,f23229,f23233,f23235,f24931,f24930,f24854,f24822,f24800,f23543,f23787,f25217,f25261,f25262,f25265,f25280,f25784,f25785,f24220,f24215,f25787,f25788,f23788,f23783,f25790,f23712,f20161,f20156,f25794,f19927,f19749,f19571,f19566,f25798,f19337,f19159,f18981,f18976,f18803,f18798,f25801,f25802,f13351,f13346,f25804,f13292,f13175,f13170,f13039,f11230,f9980,f9978,f9977,f9974,f6734,f6732,f6731,f6728,f24431,f25850,f25854,f28579,f28795,f28928,f28929,f28799,f28801,f28992,f29101,f29277,f29642,f30026,f30749,f30802,f30902,f31517,f31518,f31523,f32112,f32190,f32954,f32955,f32960,f33699,f33700,f33637,f33624,f34113,f34117,f35168,f35169,f35170,f35172,f36886,f36851,f36833,f37018,f37019,f37020,f37021,f36024,f36821,f37037,f37041,f36272,f36746,f37183,f37187,f37191,f36506,f36507,f36509,f36510,f37202,f37203,f36577,f36580,f37595,f13038,f3651,f17264,f17266,f15352,f22930,f22929,f22928,f786,f550,f475,f48,f47,f46,f25855,f19925,f19747,f19569,f18801,f762,f2762,f11302,f2704,f19406,f25856,f19914,f19736,f19558,f18790,f2699,f785,f2763,f22931,f37580,f37578,f37577,f37576,f37574,f37033,f36728,f34476,f34474,f34472,f34425,f24214,f23782,f22898,f19565,f19331,f19153,f18797,f16976,f15069,f13826,f13345,f13169,f13033,f4154,f3827,f3726,f3646,f2559,f368,f366,f37594,f37211,f37040,f37038,f37032,f36953,f36870,f36709,f36702,f36501,f36494,f36344,f36337,f36111,f36110,f36025,f36023,f36022,f34518,f34208,f34207,f34119,f34116,f34115,f33645,f33643,f33629,f33627,f33552,f33551,f32991,f32792,f32361,f32194,f32193,f32192,f32154,f32152,f32151,f32126,f32105,f32101,f31431,f31357,f31001,f30906,f30905,f30904,f30404,f30179,f30178,f30174,f30172,f30080,f30030,f30029,f30027,f30024,f29905,f29904,f29900,f29898,f29715,f29714,f29581,f29580,f29283,f29256,f29255,f29252,f29247,f29245,f29243,f29241,f25810,f25806,f25792,f25786,f22303,f22293,f22277,f21173,f19401,f19400,f17276,f17226,f15341,f13224,f13223,f11203,f3759,f3758,f2874,f1972,f1473,f370,f365,f39,f36343,f36874,f36872,f36708,f36696,f36575,f36504,f36500,f36488,f36331,f36184,f35841,f35840,f13854,f13853,f11293,f11201,f36263,f38764,f40698,f40699,f45254,f45391,f45392,f45396,f45400,f45453,f45454,f45538,f45539,f45544,f45545,f45588,f48458])).
fof(f48458,plain,(
  ( ! [X313,X312,X314] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,g(s(X312,sK6(X313,X314,bool,X312)))) | s(fun(X312,bool),X313) = s(fun(X312,bool),X314)) ) | ($spl4 | $spl9 | $spl36 | $spl228)),
  inference(subsumption_resolution,[],[f47233,f29252])).
fof(f47233,plain,(
  ( ! [X313,X312,X314] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,g(s(X312,sK6(X313,X314,bool,X312)))) | s(fun(X312,bool),X313) = s(fun(X312,bool),X314) | s(bool,'T_0') = s(bool,f(s(X312,sK6(X313,X314,bool,X312))))) ) | $spl228),
  inference(superposition,[],[f38764,f2559])).
fof(f40699,plain,(
  ( ! [X154,X152,X153] : (s(bool,X154) != s(bool,g(s(bool,X154))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X152,bool),'const.pair.,_2'(s(X152,X153),s(bool,X154)))))) ) | $spl228),
  inference(superposition,[],[f38764,f44])).
fof(f40698,plain,(
  ( ! [X151,X149,X150] : (s(bool,X150) != s(bool,g(s(bool,X150))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X149),'const.pair.,_2'(s(bool,X150),s(X149,X151)))))) ) | $spl228),
  inference(superposition,[],[f38764,f43])).
fof(f23543,plain,(
  ( ! [X39,X38] : (s(bool,X38) = s(bool,g(s(bool,X38))) | s(bool,X39) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X38) = s(bool,X39)) ) | ($spl4 | $spl46)),
  inference(superposition,[],[f3303,f22295])).
fof(f19344,plain,(
  ( ! [X39,X41,X38,X40] : (s(bool,X38) = s(bool,g(s(bool,X38))) | s(bool,X39) = s(bool,X40) | s(bool,X38) = s(bool,X40) | s(bool,X39) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) ) | $spl46),
  inference(superposition,[],[f3303,f146])).
fof(f22299,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,g(s(bool,'T_0'))) | s(bool,X0) = s(bool,X1)) ) | $spl46),
  inference(condensation,[],[f22298])).
fof(f22298,plain,(
  ( ! [X210,X208,X209] : (s(bool,'T_0') = s(bool,X208) | s(bool,X209) = s(bool,g(s(bool,'T_0'))) | s(bool,X209) = s(bool,X210) | s(bool,'T_0') = s(bool,X210)) ) | $spl46),
  inference(duplicate_literal_removal,[],[f18259])).
fof(f18259,plain,(
  ( ! [X210,X208,X209] : (s(bool,'T_0') = s(bool,X208) | s(bool,X209) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X208) | s(bool,X209) = s(bool,X210) | s(bool,'T_0') = s(bool,X210)) ) | $spl46),
  inference(superposition,[],[f146,f3303])).
fof(f13181,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,X30) = s(bool,g(s(bool,X30))) | s(bool,X30) = s(bool,X31) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X32,bool),'const.pair.,_2'(s(X32,X33),s(bool,X31)))))) ) | $spl46),
  inference(superposition,[],[f3303,f133])).
fof(f5161,plain,(
  ( ! [X4,X5] : (s(bool,X4) = s(bool,g(s(bool,X4))) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) ) | $spl46),
  inference(superposition,[],[f3303,f56])).
fof(f5160,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,g(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),'const.pair.,_2'(s(bool,X1),s(X2,X3)))))) ) | $spl46),
  inference(superposition,[],[f3303,f132])).
fof(f49247,plain,(
  ( ! [X132,X136,X137] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X136),'const.pair.,_2'(s(bool,X132),s(X136,X137))))) | s(bool,'T_0') = s(bool,f(s(bool,X132)))) ) | $spl256),
  inference(cnf_transformation,[],[f49247_D])).
fof(f49247_D,plain,(
  ( ! [X132,X136,X137] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X136),'const.pair.,_2'(s(bool,X132),s(X136,X137))))) | s(bool,'T_0') = s(bool,f(s(bool,X132)))) ) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f49246,plain,(
  ( ! [X132,X136,X133,X137] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X136),'const.pair.,_2'(s(bool,X132),s(X136,X137))))) | s(bool,'F_0') = s(bool,X133) | s(bool,'T_0') = s(bool,f(s(bool,X132)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(global_subsumption,[],[f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f67,f64,f132,f133,f146,f156,f157,f200,f361,f363,f71,f54,f1085,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f4170,f4171,f3832,f3890,f3878,f2550,f2549,f3652,f3732,f3828,f3727,f2700,f763,f2698,f761,f2455,f1935,f1480,f82,f104,f2706,f11126,f11125,f11188,f11200,f11161,f11162,f11222,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13852,f13225,f13350,f13427,f13415,f15429,f17356,f22573,f18154,f18155,f22295,f22291,f22279,f22275,f22219,f22217,f18360,f18362,f18363,f22711,f18802,f19158,f21203,f22901,f22902,f22905,f19338,f19402,f19403,f19570,f22984,f20200,f23229,f23233,f23235,f24931,f24930,f24854,f24822,f24800,f23787,f25217,f25261,f25262,f25265,f25280,f25784,f25785,f24220,f24215,f25787,f25788,f23788,f23783,f25790,f23712,f20161,f20156,f25794,f19927,f19749,f19571,f19566,f25798,f19337,f19159,f18981,f18976,f18803,f18798,f25801,f25802,f13351,f13346,f25804,f13292,f13175,f13170,f13039,f11230,f9980,f9978,f9977,f9974,f6734,f6732,f6731,f6728,f24431,f25850,f25854,f28579,f28795,f28928,f28929,f28799,f28801,f28992,f29101,f29277,f29642,f30026,f30749,f30802,f30902,f31517,f31518,f31523,f32112,f32190,f32954,f32955,f32960,f33699,f33700,f33637,f33624,f34113,f34117,f35168,f35169,f35170,f35172,f36886,f36851,f36833,f37018,f37019,f37020,f37021,f36024,f36821,f37037,f37041,f36272,f36746,f37183,f37187,f37191,f36506,f36507,f36509,f36510,f37202,f37203,f36577,f36580,f37595,f13038,f3651,f17264,f17266,f15352,f22930,f22929,f22928,f786,f550,f475,f48,f47,f46,f25855,f19925,f19747,f19569,f18801,f762,f2762,f11302,f2704,f19406,f25856,f19914,f19736,f19558,f18790,f2699,f785,f2763,f22931,f37580,f37578,f37577,f37576,f37574,f37033,f36728,f34476,f34474,f34472,f34425,f24214,f23782,f22898,f19565,f19331,f19153,f18797,f16976,f15069,f13826,f13345,f13169,f13033,f4154,f3827,f3726,f3646,f2559,f368,f366,f37594,f37211,f37040,f37038,f37032,f36953,f36870,f36709,f36702,f36501,f36494,f36344,f36337,f36111,f36110,f36025,f36023,f36022,f34518,f34208,f34207,f34119,f34116,f34115,f33645,f33643,f33629,f33627,f33552,f33551,f32991,f32792,f32361,f32194,f32193,f32192,f32154,f32152,f32151,f32126,f32105,f32101,f31431,f31357,f31001,f30906,f30905,f30904,f30404,f30179,f30178,f30174,f30172,f30080,f30030,f30029,f30027,f30024,f29905,f29904,f29900,f29898,f29715,f29714,f29581,f29580,f29283,f29256,f29255,f29252,f29247,f29245,f29243,f29241,f25810,f25806,f25792,f25786,f22303,f22293,f22277,f21173,f19401,f19400,f17276,f17226,f15341,f13224,f13223,f11203,f3759,f3758,f2874,f1972,f1473,f370,f365,f39,f36343,f36874,f36872,f36708,f36696,f36575,f36504,f36500,f36488,f36331,f36184,f35841,f35840,f13854,f13853,f11293,f11201,f36263,f45254,f45391,f45392,f45396,f45400,f45453,f45454,f45538,f45539,f45544,f45545,f45588,f49037,f49038,f49051,f37194,f49053,f36718,f36723,f49078,f45327,f45330,f44862,f36220])).
fof(f49238,plain,(
  $false | ($spl4 | $spl9 | $spl36 | $spl42 | $spl46 | $spl186 | $spl188)),
  inference(subsumption_resolution,[],[f33915,f49103])).
fof(f49103,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl4 | $spl9 | $spl36 | $spl42 | $spl46 | $spl188)),
  inference(subsumption_resolution,[],[f49102,f3890])).
fof(f49102,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'T_0'),s(bool,X0)))))) ) | ($spl4 | $spl9 | $spl36 | $spl42 | $spl46 | $spl188)),
  inference(condensation,[],[f49101])).
fof(f49101,plain,(
  ( ! [X64,X63] : (s(bool,X64) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'T_0'),s(bool,X64))))) | s(bool,'T_0') = s(bool,X63) | s(bool,'T_0') = s(bool,X64)) ) | ($spl4 | $spl9 | $spl36 | $spl42 | $spl46 | $spl188)),
  inference(forward_demodulation,[],[f3170,f49080])).
fof(f49080,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,f(s(bool,X0)))) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl188)),
  inference(condensation,[],[f49079])).
fof(f49079,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,f(s(bool,X7))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl188)),
  inference(global_subsumption,[],[f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f67,f64,f132,f133,f146,f156,f157,f200,f361,f363,f71,f54,f1085,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f4170,f4171,f3832,f3890,f3878,f2550,f2549,f3652,f3732,f3828,f3727,f2700,f763,f2698,f761,f5160,f5161,f2455,f1935,f1480,f82,f104,f2706,f11126,f11125,f11188,f11200,f11161,f11162,f11222,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13852,f13181,f13225,f13350,f13427,f13415,f15429,f17356,f22573,f18154,f18155,f22299,f22295,f22291,f22279,f22275,f22219,f22217,f18360,f18362,f18363,f22711,f18802,f19158,f21203,f22901,f22902,f22905,f19338,f19344,f19402,f19403,f19570,f22984,f20200,f23229,f23233,f23235,f24931,f24930,f24854,f24822,f24800,f23543,f23787,f25217,f25261,f25262,f25265,f25280,f25784,f25785,f24220,f24215,f25787,f25788,f23788,f23783,f25790,f23712,f20161,f20156,f25794,f19927,f19749,f19571,f19566,f25798,f19337,f19159,f18981,f18976,f18803,f18798,f25801,f25802,f13351,f13346,f25804,f13292,f13175,f13170,f13039,f11230,f9980,f9978,f9977,f9974,f6734,f6732,f6731,f6728,f24431,f25850,f25854,f28579,f28795,f28928,f28929,f28799,f28801,f28992,f29101,f29277,f29642,f30026,f30749,f30802,f30902,f31517,f31518,f31523,f32112,f32190,f32954,f32955,f32960,f33699,f33700,f33637,f33624,f34071,f34113,f34117,f35168,f35169,f35170,f35172,f36886,f36851,f36833,f37018,f37019,f37020,f37021,f36024,f36821,f37037,f37041,f36272,f36746,f37183,f37187,f37191,f36506,f36507,f36509,f36510,f37202,f37203,f36577,f36580,f37595,f13038,f3651,f17264,f17266,f15352,f22930,f22929,f22928,f786,f550,f475,f48,f47,f46,f25855,f19925,f19747,f19569,f18801,f762,f2762,f11302,f2704,f19406,f25856,f19914,f19736,f19558,f18790,f2699,f785,f2763,f22931,f37580,f37578,f37577,f37576,f37574,f37033,f36728,f34476,f34474,f34472,f34425,f24214,f23782,f22898,f19565,f19331,f19153,f18797,f16976,f15069,f13826,f13345,f13169,f13033,f4154,f3827,f3726,f3646,f2559,f368,f366,f37594,f37211,f37040,f37038,f37032,f36953,f36870,f36709,f36702,f36501,f36494,f36344,f36337,f36111,f36110,f36025,f36023,f36022,f34518,f34208,f34207,f34119,f34116,f34115,f33645,f33643,f33629,f33627,f33552,f33551,f32991,f32792,f32361,f32194,f32193,f32192,f32154,f32152,f32151,f32126,f32105,f32101,f31431,f31357,f31001,f30906,f30905,f30904,f30404,f30179,f30178,f30174,f30172,f30080,f30030,f30029,f30027,f30024,f29905,f29904,f29900,f29898,f29715,f29714,f29581,f29580,f29283,f29256,f29255,f29252,f29247,f29245,f29243,f29241,f25810,f25806,f25792,f25786,f22303,f22293,f22277,f21173,f19401,f19400,f17276,f17226,f15341,f13224,f13223,f11203,f3759,f3758,f2874,f1972,f1473,f370,f365,f39,f36343,f36874,f36872,f36708,f36696,f36575,f36504,f36500,f36488,f36331,f36184,f35841,f35840,f13854,f13853,f11293,f11201,f36263,f45254,f45391,f45392,f45396,f45400,f45453,f45454,f45538,f45539,f45544,f45545,f45588,f3303,f49037,f49038,f49051,f37194,f49053,f36718,f36723,f49078,f45327,f45330,f44862])).
fof(f33915,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | ($spl9 | $spl186)),
  inference(unit_resulting_resolution,[],[f71,f33724])).
fof(f49237,plain,(
  $false | ($spl4 | $spl9 | $spl36 | $spl42 | $spl46 | $spl188 | $spl228)),
  inference(subsumption_resolution,[],[f40640,f49103])).
fof(f49236,plain,(
  $false | ($spl4 | $spl9 | $spl36 | $spl42 | $spl46 | $spl186 | $spl188)),
  inference(subsumption_resolution,[],[f41131,f49103])).
fof(f41131,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | ($spl9 | $spl186)),
  inference(unit_resulting_resolution,[],[f71,f33724])).
fof(f49234,plain,(
  $false | ($spl4 | $spl9 | $spl36 | $spl42 | $spl46 | $spl188 | $spl203)),
  inference(subsumption_resolution,[],[f37100,f49103])).
fof(f49151,plain,(
  $false | ($spl4 | $spl9 | $spl13 | $spl22 | $spl36 | $spl42 | $spl46 | $spl188)),
  inference(subsumption_resolution,[],[f6406,f49103])).
fof(f49143,plain,(
  $false | ($spl4 | $spl9 | $spl36 | $spl42 | $spl46 | $spl188)),
  inference(subsumption_resolution,[],[f49142,f37])).
fof(f49142,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl9 | $spl36 | $spl42 | $spl46 | $spl188)),
  inference(forward_demodulation,[],[f39,f49103])).
fof(f49074,plain,(
  ( ! [X80,X81] : (s(bool,'T_0') = s(bool,X81) | s(bool,'F_0') = s(bool,X80)) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl188 | $spl204)),
  inference(global_subsumption,[],[f40182,f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f67,f64,f132,f133,f146,f156,f157,f200,f361,f363,f71,f54,f1085,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f4170,f4171,f3832,f3890,f3878,f2550,f2549,f3652,f3732,f3828,f3727,f2700,f763,f2698,f761,f5160,f5161,f2455,f1935,f1480,f82,f104,f2706,f11126,f11125,f11188,f11200,f11161,f11162,f11222,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13852,f13181,f13225,f13350,f13427,f13415,f15429,f17356,f22573,f18154,f18155,f22299,f22295,f22291,f22279,f22275,f22219,f22217,f18360,f18362,f18363,f22711,f18802,f19158,f21203,f22901,f22902,f22905,f19338,f19344,f19402,f19403,f19570,f22984,f20200,f23229,f23233,f23235,f24931,f24930,f24854,f24822,f24800,f23543,f23787,f25217,f25261,f25262,f25265,f25280,f25784,f25785,f24220,f24215,f25787,f25788,f23788,f23783,f25790,f23712,f20161,f20156,f25794,f19927,f19749,f19571,f19566,f25798,f19337,f19159,f18981,f18976,f18803,f18798,f25801,f25802,f13351,f13346,f25804,f13292,f13175,f13170,f13039,f11230,f9980,f9978,f9977,f9974,f6734,f6732,f6731,f6728,f24431,f25850,f25854,f28579,f28795,f28928,f28929,f28799,f28801,f28992,f29101,f29277,f29642,f30026,f30749,f30802,f30902,f31517,f31518,f31523,f32112,f32190,f32954,f32955,f32960,f33699,f33700,f33637,f33624,f34071,f34113,f34117,f35168,f35169,f35170,f35172,f36886,f36851,f36833,f37018,f37019,f37020,f37021,f36024,f36821,f37037,f37041,f36272,f36746,f37183,f37187,f37191,f36506,f36507,f36509,f36510,f37202,f37203,f36577,f36580,f37595,f13038,f3651,f17264,f17266,f15352,f22930,f22929,f22928,f786,f550,f475,f48,f47,f46,f25855,f19925,f19747,f19569,f18801,f762,f2762,f11302,f2704,f19406,f25856,f19914,f19736,f19558,f18790,f2699,f785,f2763,f22931,f37580,f37578,f37577,f37576,f37574,f37033,f36728,f34476,f34474,f34472,f34425,f24214,f23782,f22898,f19565,f19331,f19153,f18797,f16976,f15069,f13826,f13345,f13169,f13033,f4154,f3827,f3726,f3646,f2559,f368,f366,f37594,f37211,f37040,f37038,f37032,f36953,f36870,f36709,f36702,f36501,f36494,f36344,f36337,f36111,f36110,f36025,f36023,f36022,f34518,f34208,f34207,f34119,f34116,f34115,f33645,f33643,f33629,f33627,f33552,f33551,f32991,f32792,f32361,f32194,f32193,f32192,f32154,f32152,f32151,f32126,f32105,f32101,f31431,f31357,f31001,f30906,f30905,f30904,f30404,f30179,f30178,f30174,f30172,f30080,f30030,f30029,f30027,f30024,f29905,f29904,f29900,f29898,f29715,f29714,f29581,f29580,f29283,f29256,f29255,f29252,f29247,f29245,f29243,f29241,f25810,f25806,f25792,f25786,f22303,f22293,f22277,f21173,f19401,f19400,f17276,f17226,f15341,f13224,f13223,f11203,f3759,f3758,f2874,f1972,f1473,f370,f365,f39,f36343,f36874,f36872,f36708,f36696,f36575,f36504,f36500,f36488,f36331,f36184,f35841,f35840,f13854,f13853,f11293,f11201,f36263,f45254,f45391,f45392,f45396,f45400,f45453,f45454,f45538,f45539,f45544,f45545,f45588,f3303,f49037,f49038,f49051,f37194,f49053,f36718,f36723])).
fof(f40182,plain,(
  ( ! [X80,X81] : (s(bool,'F_0') = s(bool,f(s(bool,X81))) | s(bool,'F_0') = s(bool,X80) | s(bool,'T_0') = s(bool,X81) | s(bool,X80) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl204)),
  inference(superposition,[],[f37102,f22295])).
fof(f49073,plain,(
  ( ! [X21,X22] : (s(bool,'F_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X22)) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl188 | $spl204)),
  inference(global_subsumption,[],[f40165,f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f67,f64,f132,f133,f146,f156,f157,f200,f361,f363,f71,f54,f1085,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f4170,f4171,f3832,f3890,f3878,f2550,f2549,f3652,f3732,f3828,f3727,f2700,f763,f2698,f761,f5160,f5161,f2455,f1935,f1480,f82,f104,f2706,f11126,f11125,f11188,f11200,f11161,f11162,f11222,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13852,f13181,f13225,f13350,f13427,f13415,f15429,f17356,f22573,f18154,f18155,f22299,f22295,f22291,f22279,f22275,f22219,f22217,f18360,f18362,f18363,f22711,f18802,f19158,f21203,f22901,f22902,f22905,f19338,f19344,f19402,f19403,f19570,f22984,f20200,f23229,f23233,f23235,f24931,f24930,f24854,f24822,f24800,f23543,f23787,f25217,f25261,f25262,f25265,f25280,f25784,f25785,f24220,f24215,f25787,f25788,f23788,f23783,f25790,f23712,f20161,f20156,f25794,f19927,f19749,f19571,f19566,f25798,f19337,f19159,f18981,f18976,f18803,f18798,f25801,f25802,f13351,f13346,f25804,f13292,f13175,f13170,f13039,f11230,f9980,f9978,f9977,f9974,f6734,f6732,f6731,f6728,f24431,f25850,f25854,f28579,f28795,f28928,f28929,f28799,f28801,f28992,f29101,f29277,f29642,f30026,f30749,f30802,f30902,f31517,f31518,f31523,f32112,f32190,f32954,f32955,f32960,f33699,f33700,f33637,f33624,f34071,f34113,f34117,f35168,f35169,f35170,f35172,f36886,f36851,f36833,f37018,f37019,f37020,f37021,f36024,f36821,f37037,f37041,f36272,f36746,f37183,f37187,f37191,f36506,f36507,f36509,f36510,f37202,f37203,f36577,f36580,f37595,f13038,f3651,f17264,f17266,f15352,f22930,f22929,f22928,f786,f550,f475,f48,f47,f46,f25855,f19925,f19747,f19569,f18801,f762,f2762,f11302,f2704,f19406,f25856,f19914,f19736,f19558,f18790,f2699,f785,f2763,f22931,f37580,f37578,f37577,f37576,f37574,f37033,f36728,f34476,f34474,f34472,f34425,f24214,f23782,f22898,f19565,f19331,f19153,f18797,f16976,f15069,f13826,f13345,f13169,f13033,f4154,f3827,f3726,f3646,f2559,f368,f366,f37594,f37211,f37040,f37038,f37032,f36953,f36870,f36709,f36702,f36501,f36494,f36344,f36337,f36111,f36110,f36025,f36023,f36022,f34518,f34208,f34207,f34119,f34116,f34115,f33645,f33643,f33629,f33627,f33552,f33551,f32991,f32792,f32361,f32194,f32193,f32192,f32154,f32152,f32151,f32126,f32105,f32101,f31431,f31357,f31001,f30906,f30905,f30904,f30404,f30179,f30178,f30174,f30172,f30080,f30030,f30029,f30027,f30024,f29905,f29904,f29900,f29898,f29715,f29714,f29581,f29580,f29283,f29256,f29255,f29252,f29247,f29245,f29243,f29241,f25810,f25806,f25792,f25786,f22303,f22293,f22277,f21173,f19401,f19400,f17276,f17226,f15341,f13224,f13223,f11203,f3759,f3758,f2874,f1972,f1473,f370,f365,f39,f36343,f36874,f36872,f36708,f36696,f36575,f36504,f36500,f36488,f36331,f36184,f35841,f35840,f13854,f13853,f11293,f11201,f36263,f45254,f45391,f45392,f45396,f45400,f45453,f45454,f45538,f45539,f45544,f45545,f45588,f3303,f49037,f49038,f49051,f37194,f49053,f36718,f36723])).
fof(f40165,plain,(
  ( ! [X24,X23,X21,X22] : (s(bool,'F_0') = s(bool,f(s(bool,X22))) | s(bool,'F_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X23,bool),'const.pair.,_2'(s(X23,X24),s(bool,X21)))))) ) | $spl204),
  inference(superposition,[],[f37102,f133])).
fof(f49072,plain,(
  ( ! [X10,X9] : (s(bool,'F_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X10)) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl188 | $spl204)),
  inference(global_subsumption,[],[f40162,f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f67,f64,f132,f133,f146,f156,f157,f200,f361,f363,f71,f54,f1085,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f4170,f4171,f3832,f3890,f3878,f2550,f2549,f3652,f3732,f3828,f3727,f2700,f763,f2698,f761,f5160,f5161,f2455,f1935,f1480,f82,f104,f2706,f11126,f11125,f11188,f11200,f11161,f11162,f11222,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13852,f13181,f13225,f13350,f13427,f13415,f15429,f17356,f22573,f18154,f18155,f22299,f22295,f22291,f22279,f22275,f22219,f22217,f18360,f18362,f18363,f22711,f18802,f19158,f21203,f22901,f22902,f22905,f19338,f19344,f19402,f19403,f19570,f22984,f20200,f23229,f23233,f23235,f24931,f24930,f24854,f24822,f24800,f23543,f23787,f25217,f25261,f25262,f25265,f25280,f25784,f25785,f24220,f24215,f25787,f25788,f23788,f23783,f25790,f23712,f20161,f20156,f25794,f19927,f19749,f19571,f19566,f25798,f19337,f19159,f18981,f18976,f18803,f18798,f25801,f25802,f13351,f13346,f25804,f13292,f13175,f13170,f13039,f11230,f9980,f9978,f9977,f9974,f6734,f6732,f6731,f6728,f24431,f25850,f25854,f28579,f28795,f28928,f28929,f28799,f28801,f28992,f29101,f29277,f29642,f30026,f30749,f30802,f30902,f31517,f31518,f31523,f32112,f32190,f32954,f32955,f32960,f33699,f33700,f33637,f33624,f34071,f34113,f34117,f35168,f35169,f35170,f35172,f36886,f36851,f36833,f37018,f37019,f37020,f37021,f36024,f36821,f37037,f37041,f36272,f36746,f37183,f37187,f37191,f36506,f36507,f36509,f36510,f37202,f37203,f36577,f36580,f37595,f13038,f3651,f17264,f17266,f15352,f22930,f22929,f22928,f786,f550,f475,f48,f47,f46,f25855,f19925,f19747,f19569,f18801,f762,f2762,f11302,f2704,f19406,f25856,f19914,f19736,f19558,f18790,f2699,f785,f2763,f22931,f37580,f37578,f37577,f37576,f37574,f37033,f36728,f34476,f34474,f34472,f34425,f24214,f23782,f22898,f19565,f19331,f19153,f18797,f16976,f15069,f13826,f13345,f13169,f13033,f4154,f3827,f3726,f3646,f2559,f368,f366,f37594,f37211,f37040,f37038,f37032,f36953,f36870,f36709,f36702,f36501,f36494,f36344,f36337,f36111,f36110,f36025,f36023,f36022,f34518,f34208,f34207,f34119,f34116,f34115,f33645,f33643,f33629,f33627,f33552,f33551,f32991,f32792,f32361,f32194,f32193,f32192,f32154,f32152,f32151,f32126,f32105,f32101,f31431,f31357,f31001,f30906,f30905,f30904,f30404,f30179,f30178,f30174,f30172,f30080,f30030,f30029,f30027,f30024,f29905,f29904,f29900,f29898,f29715,f29714,f29581,f29580,f29283,f29256,f29255,f29252,f29247,f29245,f29243,f29241,f25810,f25806,f25792,f25786,f22303,f22293,f22277,f21173,f19401,f19400,f17276,f17226,f15341,f13224,f13223,f11203,f3759,f3758,f2874,f1972,f1473,f370,f365,f39,f36343,f36874,f36872,f36708,f36696,f36575,f36504,f36500,f36488,f36331,f36184,f35841,f35840,f13854,f13853,f11293,f11201,f36263,f45254,f45391,f45392,f45396,f45400,f45453,f45454,f45538,f45539,f45544,f45545,f45588,f3303,f49037,f49038,f49051,f37194,f49053,f36718,f36723])).
fof(f40162,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,'F_0') = s(bool,f(s(bool,X10))) | s(bool,'F_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X11),'const.pair.,_2'(s(bool,X9),s(X11,X12)))))) ) | $spl204),
  inference(superposition,[],[f37102,f132])).
fof(f49056,plain,(
  ( ! [X98] : (s(bool,'T_0') = s(bool,f(s(bool,X98))) | s(bool,'F_0') = s(bool,X98)) ) | $spl252),
  inference(cnf_transformation,[],[f49056_D])).
fof(f49056_D,plain,(
  ( ! [X98] : (s(bool,'T_0') = s(bool,f(s(bool,X98))) | s(bool,'F_0') = s(bool,X98)) ) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f49058,plain,(
  p(s(bool,f(s(bool,'F_0')))) | $spl254),
  inference(cnf_transformation,[],[f49058_D])).
fof(f49058_D,plain,(
  p(s(bool,f(s(bool,'F_0')))) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f39839,plain,(
  ( ! [X149] : (p(s(bool,f(s(bool,'F_0')))) | s(bool,'T_0') = s(bool,f(s(bool,X149))) | s(bool,'F_0') = s(bool,X149)) ) | ($spl4 | $spl9 | $spl36 | $spl166 | $spl176)),
  inference(superposition,[],[f37178,f30750])).
fof(f49071,plain,(
  ( ! [X83,X84] : (s(bool,'F_0') = s(bool,X84) | s(bool,'T_0') = s(bool,X83)) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl176 | $spl188)),
  inference(global_subsumption,[],[f49070,f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f67,f64,f132,f133,f146,f156,f157,f200,f361,f363,f71,f54,f1085,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f4170,f4171,f3832,f3890,f3878,f2550,f2549,f3652,f3732,f3828,f3727,f2700,f763,f2698,f761,f5160,f5161,f2455,f1935,f1480,f82,f104,f2706,f11126,f11125,f11188,f11200,f11161,f11162,f11222,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13852,f13181,f13225,f13350,f13427,f13415,f15429,f17356,f22573,f18154,f18155,f22299,f22295,f22291,f22279,f22275,f22219,f22217,f18360,f18362,f18363,f22711,f18802,f19158,f21203,f22901,f22902,f22905,f19338,f19344,f19402,f19403,f19570,f22984,f20200,f23229,f23233,f23235,f24931,f24930,f24854,f24822,f24800,f23543,f23787,f25217,f25261,f25262,f25265,f25280,f25784,f25785,f24220,f24215,f25787,f25788,f23788,f23783,f25790,f23712,f20161,f20156,f25794,f19927,f19749,f19571,f19566,f25798,f19337,f19159,f18981,f18976,f18803,f18798,f25801,f25802,f13351,f13346,f25804,f13292,f13175,f13170,f13039,f11230,f9980,f9978,f9977,f9974,f6734,f6732,f6731,f6728,f24431,f25850,f25854,f28579,f28795,f28928,f28929,f28799,f28801,f28992,f29101,f29277,f29642,f30026,f30749,f30802,f30902,f31517,f31518,f31523,f32112,f32190,f32954,f32955,f32960,f33699,f33700,f33637,f33624,f34071,f34113,f34117,f35168,f35169,f35170,f35172,f36886,f36851,f36833,f37018,f37019,f37020,f37021,f36024,f36821,f37037,f37041,f36272,f36746,f37183,f37187,f37191,f36506,f36507,f36509,f36510,f37202,f37203,f36577,f36580,f37595,f13038,f3651,f17264,f17266,f15352,f22930,f22929,f22928,f786,f550,f475,f48,f47,f46,f25855,f19925,f19747,f19569,f18801,f762,f2762,f11302,f2704,f19406,f25856,f19914,f19736,f19558,f18790,f2699,f785,f2763,f22931,f37580,f37578,f37577,f37576,f37574,f37033,f36728,f34476,f34474,f34472,f34425,f24214,f23782,f22898,f19565,f19331,f19153,f18797,f16976,f15069,f13826,f13345,f13169,f13033,f4154,f3827,f3726,f3646,f2559,f368,f366,f37594,f37211,f37040,f37038,f37032,f36953,f36870,f36709,f36702,f36501,f36494,f36344,f36337,f36111,f36110,f36025,f36023,f36022,f34518,f34208,f34207,f34119,f34116,f34115,f33645,f33643,f33629,f33627,f33552,f33551,f32991,f32792,f32361,f32194,f32193,f32192,f32154,f32152,f32151,f32126,f32105,f32101,f31431,f31357,f31001,f30906,f30905,f30904,f30404,f30179,f30178,f30174,f30172,f30080,f30030,f30029,f30027,f30024,f29905,f29904,f29900,f29898,f29715,f29714,f29581,f29580,f29283,f29256,f29255,f29252,f29247,f29245,f29243,f29241,f25810,f25806,f25792,f25786,f22303,f22293,f22277,f21173,f19401,f19400,f17276,f17226,f15341,f13224,f13223,f11203,f3759,f3758,f2874,f1972,f1473,f370,f365,f39,f36343,f36874,f36872,f36708,f36696,f36575,f36504,f36500,f36488,f36331,f36184,f35841,f35840,f13854,f13853,f11293,f11201,f36263,f45254,f45391,f45392,f45396,f45400,f45453,f45454,f45538,f45539,f45544,f45545,f45588,f3303,f49037,f49038,f49051,f37194,f49053,f36718,f36723])).
fof(f49068,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl176 | $spl188)),
  inference(global_subsumption,[],[f39992,f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f67,f64,f132,f133,f146,f156,f157,f200,f361,f363,f71,f54,f1085,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f4170,f4171,f3832,f3890,f3878,f2550,f2549,f3652,f3732,f3828,f3727,f2700,f763,f2698,f761,f5160,f5161,f2455,f1935,f1480,f82,f104,f2706,f11126,f11125,f11188,f11200,f11161,f11162,f11222,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13852,f13181,f13225,f13350,f13427,f13415,f15429,f17356,f22573,f18154,f18155,f22299,f22295,f22291,f22279,f22275,f22219,f22217,f18360,f18362,f18363,f22711,f18802,f19158,f21203,f22901,f22902,f22905,f19338,f19344,f19402,f19403,f19570,f22984,f20200,f23229,f23233,f23235,f24931,f24930,f24854,f24822,f24800,f23543,f23787,f25217,f25261,f25262,f25265,f25280,f25784,f25785,f24220,f24215,f25787,f25788,f23788,f23783,f25790,f23712,f20161,f20156,f25794,f19927,f19749,f19571,f19566,f25798,f19337,f19159,f18981,f18976,f18803,f18798,f25801,f25802,f13351,f13346,f25804,f13292,f13175,f13170,f13039,f11230,f9980,f9978,f9977,f9974,f6734,f6732,f6731,f6728,f24431,f25850,f25854,f28579,f28795,f28928,f28929,f28799,f28801,f28992,f29101,f29277,f29642,f30026,f30749,f30802,f30902,f31517,f31518,f31523,f32112,f32190,f32954,f32955,f32960,f33699,f33700,f33637,f33624,f34071,f34113,f34117,f35168,f35169,f35170,f35172,f36886,f36851,f36833,f37018,f37019,f37020,f37021,f36024,f36821,f37037,f37041,f36272,f36746,f37183,f37187,f37191,f36506,f36507,f36509,f36510,f37202,f37203,f36577,f36580,f37595,f13038,f3651,f17264,f17266,f15352,f22930,f22929,f22928,f786,f550,f475,f48,f47,f46,f25855,f19925,f19747,f19569,f18801,f762,f2762,f11302,f2704,f19406,f25856,f19914,f19736,f19558,f18790,f2699,f785,f2763,f22931,f37580,f37578,f37577,f37576,f37574,f37033,f36728,f34476,f34474,f34472,f34425,f24214,f23782,f22898,f19565,f19331,f19153,f18797,f16976,f15069,f13826,f13345,f13169,f13033,f4154,f3827,f3726,f3646,f2559,f368,f366,f37594,f37211,f37040,f37038,f37032,f36953,f36870,f36709,f36702,f36501,f36494,f36344,f36337,f36111,f36110,f36025,f36023,f36022,f34518,f34208,f34207,f34119,f34116,f34115,f33645,f33643,f33629,f33627,f33552,f33551,f32991,f32792,f32361,f32194,f32193,f32192,f32154,f32152,f32151,f32126,f32105,f32101,f31431,f31357,f31001,f30906,f30905,f30904,f30404,f30179,f30178,f30174,f30172,f30080,f30030,f30029,f30027,f30024,f29905,f29904,f29900,f29898,f29715,f29714,f29581,f29580,f29283,f29256,f29255,f29252,f29247,f29245,f29243,f29241,f25810,f25806,f25792,f25786,f22303,f22293,f22277,f21173,f19401,f19400,f17276,f17226,f15341,f13224,f13223,f11203,f3759,f3758,f2874,f1972,f1473,f370,f365,f39,f36343,f36874,f36872,f36708,f36696,f36575,f36504,f36500,f36488,f36331,f36184,f35841,f35840,f13854,f13853,f11293,f11201,f36263,f45254,f45391,f45392,f45396,f45400,f45453,f45454,f45538,f45539,f45544,f45545,f45588,f3303,f49037,f49038,f49051,f37194,f49053,f36718,f36723])).
fof(f49067,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1)) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl176 | $spl188)),
  inference(global_subsumption,[],[f39989,f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f67,f64,f132,f133,f146,f156,f157,f200,f361,f363,f71,f54,f1085,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f4170,f4171,f3832,f3890,f3878,f2550,f2549,f3652,f3732,f3828,f3727,f2700,f763,f2698,f761,f5160,f5161,f2455,f1935,f1480,f82,f104,f2706,f11126,f11125,f11188,f11200,f11161,f11162,f11222,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13852,f13181,f13225,f13350,f13427,f13415,f15429,f17356,f22573,f18154,f18155,f22299,f22295,f22291,f22279,f22275,f22219,f22217,f18360,f18362,f18363,f22711,f18802,f19158,f21203,f22901,f22902,f22905,f19338,f19344,f19402,f19403,f19570,f22984,f20200,f23229,f23233,f23235,f24931,f24930,f24854,f24822,f24800,f23543,f23787,f25217,f25261,f25262,f25265,f25280,f25784,f25785,f24220,f24215,f25787,f25788,f23788,f23783,f25790,f23712,f20161,f20156,f25794,f19927,f19749,f19571,f19566,f25798,f19337,f19159,f18981,f18976,f18803,f18798,f25801,f25802,f13351,f13346,f25804,f13292,f13175,f13170,f13039,f11230,f9980,f9978,f9977,f9974,f6734,f6732,f6731,f6728,f24431,f25850,f25854,f28579,f28795,f28928,f28929,f28799,f28801,f28992,f29101,f29277,f29642,f30026,f30749,f30802,f30902,f31517,f31518,f31523,f32112,f32190,f32954,f32955,f32960,f33699,f33700,f33637,f33624,f34071,f34113,f34117,f35168,f35169,f35170,f35172,f36886,f36851,f36833,f37018,f37019,f37020,f37021,f36024,f36821,f37037,f37041,f36272,f36746,f37183,f37187,f37191,f36506,f36507,f36509,f36510,f37202,f37203,f36577,f36580,f37595,f13038,f3651,f17264,f17266,f15352,f22930,f22929,f22928,f786,f550,f475,f48,f47,f46,f25855,f19925,f19747,f19569,f18801,f762,f2762,f11302,f2704,f19406,f25856,f19914,f19736,f19558,f18790,f2699,f785,f2763,f22931,f37580,f37578,f37577,f37576,f37574,f37033,f36728,f34476,f34474,f34472,f34425,f24214,f23782,f22898,f19565,f19331,f19153,f18797,f16976,f15069,f13826,f13345,f13169,f13033,f4154,f3827,f3726,f3646,f2559,f368,f366,f37594,f37211,f37040,f37038,f37032,f36953,f36870,f36709,f36702,f36501,f36494,f36344,f36337,f36111,f36110,f36025,f36023,f36022,f34518,f34208,f34207,f34119,f34116,f34115,f33645,f33643,f33629,f33627,f33552,f33551,f32991,f32792,f32361,f32194,f32193,f32192,f32154,f32152,f32151,f32126,f32105,f32101,f31431,f31357,f31001,f30906,f30905,f30904,f30404,f30179,f30178,f30174,f30172,f30080,f30030,f30029,f30027,f30024,f29905,f29904,f29900,f29898,f29715,f29714,f29581,f29580,f29283,f29256,f29255,f29252,f29247,f29245,f29243,f29241,f25810,f25806,f25792,f25786,f22303,f22293,f22277,f21173,f19401,f19400,f17276,f17226,f15341,f13224,f13223,f11203,f3759,f3758,f2874,f1972,f1473,f370,f365,f39,f36343,f36874,f36872,f36708,f36696,f36575,f36504,f36500,f36488,f36331,f36184,f35841,f35840,f13854,f13853,f11293,f11201,f36263,f45254,f45391,f45392,f45396,f45400,f45453,f45454,f45538,f45539,f45544,f45545,f45588,f3303,f49037,f49038,f49051,f37194,f49053,f36718,f36723])).
fof(f49066,plain,(
  ( ! [X30,X29] : (s(bool,'F_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X29)) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl176 | $spl188)),
  inference(global_subsumption,[],[f39738,f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f67,f64,f132,f133,f146,f156,f157,f200,f361,f363,f71,f54,f1085,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f4170,f4171,f3832,f3890,f3878,f2550,f2549,f3652,f3732,f3828,f3727,f2700,f763,f2698,f761,f5160,f5161,f2455,f1935,f1480,f82,f104,f2706,f11126,f11125,f11188,f11200,f11161,f11162,f11222,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13852,f13181,f13225,f13350,f13427,f13415,f15429,f17356,f22573,f18154,f18155,f22299,f22295,f22291,f22279,f22275,f22219,f22217,f18360,f18362,f18363,f22711,f18802,f19158,f21203,f22901,f22902,f22905,f19338,f19344,f19402,f19403,f19570,f22984,f20200,f23229,f23233,f23235,f24931,f24930,f24854,f24822,f24800,f23543,f23787,f25217,f25261,f25262,f25265,f25280,f25784,f25785,f24220,f24215,f25787,f25788,f23788,f23783,f25790,f23712,f20161,f20156,f25794,f19927,f19749,f19571,f19566,f25798,f19337,f19159,f18981,f18976,f18803,f18798,f25801,f25802,f13351,f13346,f25804,f13292,f13175,f13170,f13039,f11230,f9980,f9978,f9977,f9974,f6734,f6732,f6731,f6728,f24431,f25850,f25854,f28579,f28795,f28928,f28929,f28799,f28801,f28992,f29101,f29277,f29642,f30026,f30749,f30802,f30902,f31517,f31518,f31523,f32112,f32190,f32954,f32955,f32960,f33699,f33700,f33637,f33624,f34071,f34113,f34117,f35168,f35169,f35170,f35172,f36886,f36851,f36833,f37018,f37019,f37020,f37021,f36024,f36821,f37037,f37041,f36272,f36746,f37183,f37187,f37191,f36506,f36507,f36509,f36510,f37202,f37203,f36577,f36580,f37595,f13038,f3651,f17264,f17266,f15352,f22930,f22929,f22928,f786,f550,f475,f48,f47,f46,f25855,f19925,f19747,f19569,f18801,f762,f2762,f11302,f2704,f19406,f25856,f19914,f19736,f19558,f18790,f2699,f785,f2763,f22931,f37580,f37578,f37577,f37576,f37574,f37033,f36728,f34476,f34474,f34472,f34425,f24214,f23782,f22898,f19565,f19331,f19153,f18797,f16976,f15069,f13826,f13345,f13169,f13033,f4154,f3827,f3726,f3646,f2559,f368,f366,f37594,f37211,f37040,f37038,f37032,f36953,f36870,f36709,f36702,f36501,f36494,f36344,f36337,f36111,f36110,f36025,f36023,f36022,f34518,f34208,f34207,f34119,f34116,f34115,f33645,f33643,f33629,f33627,f33552,f33551,f32991,f32792,f32361,f32194,f32193,f32192,f32154,f32152,f32151,f32126,f32105,f32101,f31431,f31357,f31001,f30906,f30905,f30904,f30404,f30179,f30178,f30174,f30172,f30080,f30030,f30029,f30027,f30024,f29905,f29904,f29900,f29898,f29715,f29714,f29581,f29580,f29283,f29256,f29255,f29252,f29247,f29245,f29243,f29241,f25810,f25806,f25792,f25786,f22303,f22293,f22277,f21173,f19401,f19400,f17276,f17226,f15341,f13224,f13223,f11203,f3759,f3758,f2874,f1972,f1473,f370,f365,f39,f36343,f36874,f36872,f36708,f36696,f36575,f36504,f36500,f36488,f36331,f36184,f35841,f35840,f13854,f13853,f11293,f11201,f36263,f45254,f45391,f45392,f45396,f45400,f45453,f45454,f45538,f45539,f45544,f45545,f45588,f3303,f49037,f49038,f49051,f37194,f49053,f36718,f36723])).
fof(f49065,plain,(
  ( ! [X17,X18] : (s(bool,'F_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X17)) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl176 | $spl188)),
  inference(global_subsumption,[],[f39735,f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f67,f64,f132,f133,f146,f156,f157,f200,f361,f363,f71,f54,f1085,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f4170,f4171,f3832,f3890,f3878,f2550,f2549,f3652,f3732,f3828,f3727,f2700,f763,f2698,f761,f5160,f5161,f2455,f1935,f1480,f82,f104,f2706,f11126,f11125,f11188,f11200,f11161,f11162,f11222,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13852,f13181,f13225,f13350,f13427,f13415,f15429,f17356,f22573,f18154,f18155,f22299,f22295,f22291,f22279,f22275,f22219,f22217,f18360,f18362,f18363,f22711,f18802,f19158,f21203,f22901,f22902,f22905,f19338,f19344,f19402,f19403,f19570,f22984,f20200,f23229,f23233,f23235,f24931,f24930,f24854,f24822,f24800,f23543,f23787,f25217,f25261,f25262,f25265,f25280,f25784,f25785,f24220,f24215,f25787,f25788,f23788,f23783,f25790,f23712,f20161,f20156,f25794,f19927,f19749,f19571,f19566,f25798,f19337,f19159,f18981,f18976,f18803,f18798,f25801,f25802,f13351,f13346,f25804,f13292,f13175,f13170,f13039,f11230,f9980,f9978,f9977,f9974,f6734,f6732,f6731,f6728,f24431,f25850,f25854,f28579,f28795,f28928,f28929,f28799,f28801,f28992,f29101,f29277,f29642,f30026,f30749,f30802,f30902,f31517,f31518,f31523,f32112,f32190,f32954,f32955,f32960,f33699,f33700,f33637,f33624,f34071,f34113,f34117,f35168,f35169,f35170,f35172,f36886,f36851,f36833,f37018,f37019,f37020,f37021,f36024,f36821,f37037,f37041,f36272,f36746,f37183,f37187,f37191,f36506,f36507,f36509,f36510,f37202,f37203,f36577,f36580,f37595,f13038,f3651,f17264,f17266,f15352,f22930,f22929,f22928,f786,f550,f475,f48,f47,f46,f25855,f19925,f19747,f19569,f18801,f762,f2762,f11302,f2704,f19406,f25856,f19914,f19736,f19558,f18790,f2699,f785,f2763,f22931,f37580,f37578,f37577,f37576,f37574,f37033,f36728,f34476,f34474,f34472,f34425,f24214,f23782,f22898,f19565,f19331,f19153,f18797,f16976,f15069,f13826,f13345,f13169,f13033,f4154,f3827,f3726,f3646,f2559,f368,f366,f37594,f37211,f37040,f37038,f37032,f36953,f36870,f36709,f36702,f36501,f36494,f36344,f36337,f36111,f36110,f36025,f36023,f36022,f34518,f34208,f34207,f34119,f34116,f34115,f33645,f33643,f33629,f33627,f33552,f33551,f32991,f32792,f32361,f32194,f32193,f32192,f32154,f32152,f32151,f32126,f32105,f32101,f31431,f31357,f31001,f30906,f30905,f30904,f30404,f30179,f30178,f30174,f30172,f30080,f30030,f30029,f30027,f30024,f29905,f29904,f29900,f29898,f29715,f29714,f29581,f29580,f29283,f29256,f29255,f29252,f29247,f29245,f29243,f29241,f25810,f25806,f25792,f25786,f22303,f22293,f22277,f21173,f19401,f19400,f17276,f17226,f15341,f13224,f13223,f11203,f3759,f3758,f2874,f1972,f1473,f370,f365,f39,f36343,f36874,f36872,f36708,f36696,f36575,f36504,f36500,f36488,f36331,f36184,f35841,f35840,f13854,f13853,f11293,f11201,f36263,f45254,f45391,f45392,f45396,f45400,f45453,f45454,f45538,f45539,f45544,f45545,f45588,f3303,f49037,f49038,f49051,f37194,f49053,f36718,f36723])).
fof(f49064,plain,(
  ( ! [X14,X13] : (s(bool,'T_0') = s(bool,X13) | s(bool,'F_0') = s(bool,X14)) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl176 | $spl188)),
  inference(global_subsumption,[],[f49063,f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f67,f64,f132,f133,f146,f156,f157,f200,f361,f363,f71,f54,f1085,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f4170,f4171,f3832,f3890,f3878,f2550,f2549,f3652,f3732,f3828,f3727,f2700,f763,f2698,f761,f5160,f5161,f2455,f1935,f1480,f82,f104,f2706,f11126,f11125,f11188,f11200,f11161,f11162,f11222,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13852,f13181,f13225,f13350,f13427,f13415,f15429,f17356,f22573,f18154,f18155,f22299,f22295,f22291,f22279,f22275,f22219,f22217,f18360,f18362,f18363,f22711,f18802,f19158,f21203,f22901,f22902,f22905,f19338,f19344,f19402,f19403,f19570,f22984,f20200,f23229,f23233,f23235,f24931,f24930,f24854,f24822,f24800,f23543,f23787,f25217,f25261,f25262,f25265,f25280,f25784,f25785,f24220,f24215,f25787,f25788,f23788,f23783,f25790,f23712,f20161,f20156,f25794,f19927,f19749,f19571,f19566,f25798,f19337,f19159,f18981,f18976,f18803,f18798,f25801,f25802,f13351,f13346,f25804,f13292,f13175,f13170,f13039,f11230,f9980,f9978,f9977,f9974,f6734,f6732,f6731,f6728,f24431,f25850,f25854,f28579,f28795,f28928,f28929,f28799,f28801,f28992,f29101,f29277,f29642,f30026,f30749,f30802,f30902,f31517,f31518,f31523,f32112,f32190,f32954,f32955,f32960,f33699,f33700,f33637,f33624,f34071,f34113,f34117,f35168,f35169,f35170,f35172,f36886,f36851,f36833,f37018,f37019,f37020,f37021,f36024,f36821,f37037,f37041,f36272,f36746,f37183,f37187,f37191,f36506,f36507,f36509,f36510,f37202,f37203,f36577,f36580,f37595,f13038,f3651,f17264,f17266,f15352,f22930,f22929,f22928,f786,f550,f475,f48,f47,f46,f25855,f19925,f19747,f19569,f18801,f762,f2762,f11302,f2704,f19406,f25856,f19914,f19736,f19558,f18790,f2699,f785,f2763,f22931,f37580,f37578,f37577,f37576,f37574,f37033,f36728,f34476,f34474,f34472,f34425,f24214,f23782,f22898,f19565,f19331,f19153,f18797,f16976,f15069,f13826,f13345,f13169,f13033,f4154,f3827,f3726,f3646,f2559,f368,f366,f37594,f37211,f37040,f37038,f37032,f36953,f36870,f36709,f36702,f36501,f36494,f36344,f36337,f36111,f36110,f36025,f36023,f36022,f34518,f34208,f34207,f34119,f34116,f34115,f33645,f33643,f33629,f33627,f33552,f33551,f32991,f32792,f32361,f32194,f32193,f32192,f32154,f32152,f32151,f32126,f32105,f32101,f31431,f31357,f31001,f30906,f30905,f30904,f30404,f30179,f30178,f30174,f30172,f30080,f30030,f30029,f30027,f30024,f29905,f29904,f29900,f29898,f29715,f29714,f29581,f29580,f29283,f29256,f29255,f29252,f29247,f29245,f29243,f29241,f25810,f25806,f25792,f25786,f22303,f22293,f22277,f21173,f19401,f19400,f17276,f17226,f15341,f13224,f13223,f11203,f3759,f3758,f2874,f1972,f1473,f370,f365,f39,f36343,f36874,f36872,f36708,f36696,f36575,f36504,f36500,f36488,f36331,f36184,f35841,f35840,f13854,f13853,f11293,f11201,f36263,f45254,f45391,f45392,f45396,f45400,f45453,f45454,f45538,f45539,f45544,f45545,f45588,f3303,f49037,f49038,f49051,f37194,f49053,f36718,f36723])).
fof(f49062,plain,(
  ( ! [X12,X11] : (s(bool,'T_0') = s(bool,X11) | s(bool,'F_0') = s(bool,X12)) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl176 | $spl188)),
  inference(global_subsumption,[],[f49061,f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f67,f64,f132,f133,f146,f156,f157,f200,f361,f363,f71,f54,f1085,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f4170,f4171,f3832,f3890,f3878,f2550,f2549,f3652,f3732,f3828,f3727,f2700,f763,f2698,f761,f5160,f5161,f2455,f1935,f1480,f82,f104,f2706,f11126,f11125,f11188,f11200,f11161,f11162,f11222,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13852,f13181,f13225,f13350,f13427,f13415,f15429,f17356,f22573,f18154,f18155,f22299,f22295,f22291,f22279,f22275,f22219,f22217,f18360,f18362,f18363,f22711,f18802,f19158,f21203,f22901,f22902,f22905,f19338,f19344,f19402,f19403,f19570,f22984,f20200,f23229,f23233,f23235,f24931,f24930,f24854,f24822,f24800,f23543,f23787,f25217,f25261,f25262,f25265,f25280,f25784,f25785,f24220,f24215,f25787,f25788,f23788,f23783,f25790,f23712,f20161,f20156,f25794,f19927,f19749,f19571,f19566,f25798,f19337,f19159,f18981,f18976,f18803,f18798,f25801,f25802,f13351,f13346,f25804,f13292,f13175,f13170,f13039,f11230,f9980,f9978,f9977,f9974,f6734,f6732,f6731,f6728,f24431,f25850,f25854,f28579,f28795,f28928,f28929,f28799,f28801,f28992,f29101,f29277,f29642,f30026,f30749,f30802,f30902,f31517,f31518,f31523,f32112,f32190,f32954,f32955,f32960,f33699,f33700,f33637,f33624,f34071,f34113,f34117,f35168,f35169,f35170,f35172,f36886,f36851,f36833,f37018,f37019,f37020,f37021,f36024,f36821,f37037,f37041,f36272,f36746,f37183,f37187,f37191,f36506,f36507,f36509,f36510,f37202,f37203,f36577,f36580,f37595,f13038,f3651,f17264,f17266,f15352,f22930,f22929,f22928,f786,f550,f475,f48,f47,f46,f25855,f19925,f19747,f19569,f18801,f762,f2762,f11302,f2704,f19406,f25856,f19914,f19736,f19558,f18790,f2699,f785,f2763,f22931,f37580,f37578,f37577,f37576,f37574,f37033,f36728,f34476,f34474,f34472,f34425,f24214,f23782,f22898,f19565,f19331,f19153,f18797,f16976,f15069,f13826,f13345,f13169,f13033,f4154,f3827,f3726,f3646,f2559,f368,f366,f37594,f37211,f37040,f37038,f37032,f36953,f36870,f36709,f36702,f36501,f36494,f36344,f36337,f36111,f36110,f36025,f36023,f36022,f34518,f34208,f34207,f34119,f34116,f34115,f33645,f33643,f33629,f33627,f33552,f33551,f32991,f32792,f32361,f32194,f32193,f32192,f32154,f32152,f32151,f32126,f32105,f32101,f31431,f31357,f31001,f30906,f30905,f30904,f30404,f30179,f30178,f30174,f30172,f30080,f30030,f30029,f30027,f30024,f29905,f29904,f29900,f29898,f29715,f29714,f29581,f29580,f29283,f29256,f29255,f29252,f29247,f29245,f29243,f29241,f25810,f25806,f25792,f25786,f22303,f22293,f22277,f21173,f19401,f19400,f17276,f17226,f15341,f13224,f13223,f11203,f3759,f3758,f2874,f1972,f1473,f370,f365,f39,f36343,f36874,f36872,f36708,f36696,f36575,f36504,f36500,f36488,f36331,f36184,f35841,f35840,f13854,f13853,f11293,f11201,f36263,f45254,f45391,f45392,f45396,f45400,f45453,f45454,f45538,f45539,f45544,f45545,f45588,f3303,f49037,f49038,f49051,f37194,f49053,f36718,f36723])).
fof(f49060,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,X6) | s(bool,'F_0') = s(bool,X7)) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl176 | $spl188)),
  inference(global_subsumption,[],[f39729,f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f67,f64,f132,f133,f146,f156,f157,f200,f361,f363,f71,f54,f1085,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f4170,f4171,f3832,f3890,f3878,f2550,f2549,f3652,f3732,f3828,f3727,f2700,f763,f2698,f761,f5160,f5161,f2455,f1935,f1480,f82,f104,f2706,f11126,f11125,f11188,f11200,f11161,f11162,f11222,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13852,f13181,f13225,f13350,f13427,f13415,f15429,f17356,f22573,f18154,f18155,f22299,f22295,f22291,f22279,f22275,f22219,f22217,f18360,f18362,f18363,f22711,f18802,f19158,f21203,f22901,f22902,f22905,f19338,f19344,f19402,f19403,f19570,f22984,f20200,f23229,f23233,f23235,f24931,f24930,f24854,f24822,f24800,f23543,f23787,f25217,f25261,f25262,f25265,f25280,f25784,f25785,f24220,f24215,f25787,f25788,f23788,f23783,f25790,f23712,f20161,f20156,f25794,f19927,f19749,f19571,f19566,f25798,f19337,f19159,f18981,f18976,f18803,f18798,f25801,f25802,f13351,f13346,f25804,f13292,f13175,f13170,f13039,f11230,f9980,f9978,f9977,f9974,f6734,f6732,f6731,f6728,f24431,f25850,f25854,f28579,f28795,f28928,f28929,f28799,f28801,f28992,f29101,f29277,f29642,f30026,f30749,f30802,f30902,f31517,f31518,f31523,f32112,f32190,f32954,f32955,f32960,f33699,f33700,f33637,f33624,f34071,f34113,f34117,f35168,f35169,f35170,f35172,f36886,f36851,f36833,f37018,f37019,f37020,f37021,f36024,f36821,f37037,f37041,f36272,f36746,f37183,f37187,f37191,f36506,f36507,f36509,f36510,f37202,f37203,f36577,f36580,f37595,f13038,f3651,f17264,f17266,f15352,f22930,f22929,f22928,f786,f550,f475,f48,f47,f46,f25855,f19925,f19747,f19569,f18801,f762,f2762,f11302,f2704,f19406,f25856,f19914,f19736,f19558,f18790,f2699,f785,f2763,f22931,f37580,f37578,f37577,f37576,f37574,f37033,f36728,f34476,f34474,f34472,f34425,f24214,f23782,f22898,f19565,f19331,f19153,f18797,f16976,f15069,f13826,f13345,f13169,f13033,f4154,f3827,f3726,f3646,f2559,f368,f366,f37594,f37211,f37040,f37038,f37032,f36953,f36870,f36709,f36702,f36501,f36494,f36344,f36337,f36111,f36110,f36025,f36023,f36022,f34518,f34208,f34207,f34119,f34116,f34115,f33645,f33643,f33629,f33627,f33552,f33551,f32991,f32792,f32361,f32194,f32193,f32192,f32154,f32152,f32151,f32126,f32105,f32101,f31431,f31357,f31001,f30906,f30905,f30904,f30404,f30179,f30178,f30174,f30172,f30080,f30030,f30029,f30027,f30024,f29905,f29904,f29900,f29898,f29715,f29714,f29581,f29580,f29283,f29256,f29255,f29252,f29247,f29245,f29243,f29241,f25810,f25806,f25792,f25786,f22303,f22293,f22277,f21173,f19401,f19400,f17276,f17226,f15341,f13224,f13223,f11203,f3759,f3758,f2874,f1972,f1473,f370,f365,f39,f36343,f36874,f36872,f36708,f36696,f36575,f36504,f36500,f36488,f36331,f36184,f35841,f35840,f13854,f13853,f11293,f11201,f36263,f45254,f45391,f45392,f45396,f45400,f45453,f45454,f45538,f45539,f45544,f45545,f45588,f3303,f49037,f49038,f49051,f37194,f49053,f36718,f36723])).
fof(f37633,plain,(
  ( ! [X98] : (p(s(bool,f(s(bool,'F_0')))) | s(bool,'T_0') = s(bool,f(s(bool,X98))) | s(bool,'F_0') = s(bool,X98)) ) | ($spl4 | $spl9 | $spl36 | $spl166 | $spl176)),
  inference(superposition,[],[f37178,f30750])).
fof(f49052,plain,(
  ( ! [X343] : (s(bool,'F_0') != s(bool,X343) | s(bool,'T_0') = s(bool,f(s(bool,X343)))) ) | $spl250),
  inference(cnf_transformation,[],[f49052_D])).
fof(f49052_D,plain,(
  ( ! [X343] : (s(bool,'F_0') != s(bool,X343) | s(bool,'T_0') = s(bool,f(s(bool,X343)))) ) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f36723,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,'T_0') = s(bool,f(s(bool,X1))) | s(bool,'F_0') = s(bool,X0)) ) | ($spl4 | $spl9 | $spl36)),
  inference(equality_factoring,[],[f30029])).
fof(f37194,plain,(
  ( ! [X341,X343,X340,X342] : (s(bool,'F_0') != s(bool,X343) | s(fun(X340,bool),X341) = s(fun(X340,bool),X342) | s(bool,'T_0') = s(bool,g(s(X340,sK6(X341,X342,bool,X340)))) | s(bool,'T_0') = s(bool,f(s(bool,X343)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f36446,f370])).
fof(f36446,plain,(
  ( ! [X341,X343,X340,X342] : (s(bool,'F_0') != s(bool,X343) | s(fun(X340,bool),X341) = s(fun(X340,bool),X342) | s(bool,'T_0') = s(bool,g(s(X340,sK6(X341,X342,bool,X340)))) | s(bool,'T_0') = s(bool,f(s(bool,X343))) | s(bool,'F_0') = s(bool,f(s(X340,sK6(X341,X342,bool,X340))))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f370,f30029])).
fof(f49050,plain,(
  ( ! [X83,X84] : (s(bool,'F_0') = s(bool,X84) | s(bool,'T_0') = s(bool,X83)) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl176 | $spl188)),
  inference(global_subsumption,[],[f49049,f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f67,f64,f132,f133,f146,f156,f157,f200,f361,f363,f71,f54,f1085,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f4170,f4171,f3832,f3890,f3878,f2550,f2549,f3652,f3732,f3828,f3727,f2700,f763,f2698,f761,f5160,f5161,f2455,f1935,f1480,f82,f104,f2706,f11126,f11125,f11188,f11200,f11161,f11162,f11222,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13852,f13181,f13225,f13350,f13427,f13415,f15429,f17356,f22573,f18154,f18155,f22299,f22295,f22291,f22279,f22275,f22219,f22217,f18360,f18362,f18363,f22711,f18802,f19158,f21203,f22901,f22902,f22905,f19338,f19344,f19402,f19403,f19570,f22984,f20200,f23229,f23233,f23235,f24931,f24930,f24854,f24822,f24800,f23543,f23787,f25217,f25261,f25262,f25265,f25280,f25784,f25785,f24220,f24215,f25787,f25788,f23788,f23783,f25790,f23712,f20161,f20156,f25794,f19927,f19749,f19571,f19566,f25798,f19337,f19159,f18981,f18976,f18803,f18798,f25801,f25802,f13351,f13346,f25804,f13292,f13175,f13170,f13039,f11230,f9980,f9978,f9977,f9974,f6734,f6732,f6731,f6728,f24431,f25850,f25854,f28579,f28795,f28928,f28929,f28799,f28801,f28992,f29101,f29277,f29642,f30026,f30749,f30802,f30902,f31517,f31518,f31523,f32112,f32190,f32954,f32955,f32960,f33699,f33700,f33637,f33624,f34071,f34113,f34117,f35168,f35169,f35170,f35172,f36886,f36851,f36833,f37018,f37019,f37020,f37021,f36024,f36821,f37037,f37041,f36272,f36746,f37183,f37187,f37191,f36506,f36507,f36509,f36510,f37202,f37203,f36577,f36580,f37595,f13038,f3651,f17264,f17266,f15352,f22930,f22929,f22928,f786,f550,f475,f48,f47,f46,f25855,f19925,f19747,f19569,f18801,f762,f2762,f11302,f2704,f19406,f25856,f19914,f19736,f19558,f18790,f2699,f785,f2763,f22931,f37580,f37578,f37577,f37576,f37574,f37033,f36728,f34476,f34474,f34472,f34425,f24214,f23782,f22898,f19565,f19331,f19153,f18797,f16976,f15069,f13826,f13345,f13169,f13033,f4154,f3827,f3726,f3646,f2559,f368,f366,f37594,f37211,f37040,f37038,f37032,f36953,f36870,f36709,f36702,f36501,f36494,f36344,f36337,f36111,f36110,f36025,f36023,f36022,f34518,f34208,f34207,f34119,f34116,f34115,f33645,f33643,f33629,f33627,f33552,f33551,f32991,f32792,f32361,f32194,f32193,f32192,f32154,f32152,f32151,f32126,f32105,f32101,f31431,f31357,f31001,f30906,f30905,f30904,f30404,f30179,f30178,f30174,f30172,f30080,f30030,f30029,f30027,f30024,f29905,f29904,f29900,f29898,f29715,f29714,f29581,f29580,f29283,f29256,f29255,f29252,f29247,f29245,f29243,f29241,f25810,f25806,f25792,f25786,f22303,f22293,f22277,f21173,f19401,f19400,f17276,f17226,f15341,f13224,f13223,f11203,f3759,f3758,f2874,f1972,f1473,f370,f365,f39,f36343,f36874,f36872,f36708,f36696,f36575,f36504,f36500,f36488,f36331,f36184,f35841,f35840,f13854,f13853,f11293,f11201,f36263,f45254,f45391,f45392,f45396,f45400,f45453,f45454,f45538,f45539,f45544,f45545,f45588,f3303,f49037,f49038])).
fof(f49047,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl176 | $spl188)),
  inference(global_subsumption,[],[f31910,f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f67,f64,f132,f133,f146,f156,f157,f200,f361,f363,f71,f54,f1085,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f4170,f4171,f3832,f3890,f3878,f2550,f2549,f3652,f3732,f3828,f3727,f2700,f763,f2698,f761,f5160,f5161,f2455,f1935,f1480,f82,f104,f2706,f11126,f11125,f11188,f11200,f11161,f11162,f11222,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13852,f13181,f13225,f13350,f13427,f13415,f15429,f17356,f22573,f18154,f18155,f22299,f22295,f22291,f22279,f22275,f22219,f22217,f18360,f18362,f18363,f22711,f18802,f19158,f21203,f22901,f22902,f22905,f19338,f19344,f19402,f19403,f19570,f22984,f20200,f23229,f23233,f23235,f24931,f24930,f24854,f24822,f24800,f23543,f23787,f25217,f25261,f25262,f25265,f25280,f25784,f25785,f24220,f24215,f25787,f25788,f23788,f23783,f25790,f23712,f20161,f20156,f25794,f19927,f19749,f19571,f19566,f25798,f19337,f19159,f18981,f18976,f18803,f18798,f25801,f25802,f13351,f13346,f25804,f13292,f13175,f13170,f13039,f11230,f9980,f9978,f9977,f9974,f6734,f6732,f6731,f6728,f24431,f25850,f25854,f28579,f28795,f28928,f28929,f28799,f28801,f28992,f29101,f29277,f29642,f30026,f30749,f30802,f30902,f31517,f31518,f31523,f32112,f32190,f32954,f32955,f32960,f33699,f33700,f33637,f33624,f34071,f34113,f34117,f35168,f35169,f35170,f35172,f36886,f36851,f36833,f37018,f37019,f37020,f37021,f36024,f36821,f37037,f37041,f36272,f36746,f37183,f37187,f37191,f36506,f36507,f36509,f36510,f37202,f37203,f36577,f36580,f37595,f13038,f3651,f17264,f17266,f15352,f22930,f22929,f22928,f786,f550,f475,f48,f47,f46,f25855,f19925,f19747,f19569,f18801,f762,f2762,f11302,f2704,f19406,f25856,f19914,f19736,f19558,f18790,f2699,f785,f2763,f22931,f37580,f37578,f37577,f37576,f37574,f37033,f36728,f34476,f34474,f34472,f34425,f24214,f23782,f22898,f19565,f19331,f19153,f18797,f16976,f15069,f13826,f13345,f13169,f13033,f4154,f3827,f3726,f3646,f2559,f368,f366,f37594,f37211,f37040,f37038,f37032,f36953,f36870,f36709,f36702,f36501,f36494,f36344,f36337,f36111,f36110,f36025,f36023,f36022,f34518,f34208,f34207,f34119,f34116,f34115,f33645,f33643,f33629,f33627,f33552,f33551,f32991,f32792,f32361,f32194,f32193,f32192,f32154,f32152,f32151,f32126,f32105,f32101,f31431,f31357,f31001,f30906,f30905,f30904,f30404,f30179,f30178,f30174,f30172,f30080,f30030,f30029,f30027,f30024,f29905,f29904,f29900,f29898,f29715,f29714,f29581,f29580,f29283,f29256,f29255,f29252,f29247,f29245,f29243,f29241,f25810,f25806,f25792,f25786,f22303,f22293,f22277,f21173,f19401,f19400,f17276,f17226,f15341,f13224,f13223,f11203,f3759,f3758,f2874,f1972,f1473,f370,f365,f39,f36343,f36874,f36872,f36708,f36696,f36575,f36504,f36500,f36488,f36331,f36184,f35841,f35840,f13854,f13853,f11293,f11201,f36263,f45254,f45391,f45392,f45396,f45400,f45453,f45454,f45538,f45539,f45544,f45545,f45588,f3303,f49037,f49038])).
fof(f49046,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1)) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl176 | $spl188)),
  inference(global_subsumption,[],[f31906,f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f67,f64,f132,f133,f146,f156,f157,f200,f361,f363,f71,f54,f1085,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f4170,f4171,f3832,f3890,f3878,f2550,f2549,f3652,f3732,f3828,f3727,f2700,f763,f2698,f761,f5160,f5161,f2455,f1935,f1480,f82,f104,f2706,f11126,f11125,f11188,f11200,f11161,f11162,f11222,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13852,f13181,f13225,f13350,f13427,f13415,f15429,f17356,f22573,f18154,f18155,f22299,f22295,f22291,f22279,f22275,f22219,f22217,f18360,f18362,f18363,f22711,f18802,f19158,f21203,f22901,f22902,f22905,f19338,f19344,f19402,f19403,f19570,f22984,f20200,f23229,f23233,f23235,f24931,f24930,f24854,f24822,f24800,f23543,f23787,f25217,f25261,f25262,f25265,f25280,f25784,f25785,f24220,f24215,f25787,f25788,f23788,f23783,f25790,f23712,f20161,f20156,f25794,f19927,f19749,f19571,f19566,f25798,f19337,f19159,f18981,f18976,f18803,f18798,f25801,f25802,f13351,f13346,f25804,f13292,f13175,f13170,f13039,f11230,f9980,f9978,f9977,f9974,f6734,f6732,f6731,f6728,f24431,f25850,f25854,f28579,f28795,f28928,f28929,f28799,f28801,f28992,f29101,f29277,f29642,f30026,f30749,f30802,f30902,f31517,f31518,f31523,f32112,f32190,f32954,f32955,f32960,f33699,f33700,f33637,f33624,f34071,f34113,f34117,f35168,f35169,f35170,f35172,f36886,f36851,f36833,f37018,f37019,f37020,f37021,f36024,f36821,f37037,f37041,f36272,f36746,f37183,f37187,f37191,f36506,f36507,f36509,f36510,f37202,f37203,f36577,f36580,f37595,f13038,f3651,f17264,f17266,f15352,f22930,f22929,f22928,f786,f550,f475,f48,f47,f46,f25855,f19925,f19747,f19569,f18801,f762,f2762,f11302,f2704,f19406,f25856,f19914,f19736,f19558,f18790,f2699,f785,f2763,f22931,f37580,f37578,f37577,f37576,f37574,f37033,f36728,f34476,f34474,f34472,f34425,f24214,f23782,f22898,f19565,f19331,f19153,f18797,f16976,f15069,f13826,f13345,f13169,f13033,f4154,f3827,f3726,f3646,f2559,f368,f366,f37594,f37211,f37040,f37038,f37032,f36953,f36870,f36709,f36702,f36501,f36494,f36344,f36337,f36111,f36110,f36025,f36023,f36022,f34518,f34208,f34207,f34119,f34116,f34115,f33645,f33643,f33629,f33627,f33552,f33551,f32991,f32792,f32361,f32194,f32193,f32192,f32154,f32152,f32151,f32126,f32105,f32101,f31431,f31357,f31001,f30906,f30905,f30904,f30404,f30179,f30178,f30174,f30172,f30080,f30030,f30029,f30027,f30024,f29905,f29904,f29900,f29898,f29715,f29714,f29581,f29580,f29283,f29256,f29255,f29252,f29247,f29245,f29243,f29241,f25810,f25806,f25792,f25786,f22303,f22293,f22277,f21173,f19401,f19400,f17276,f17226,f15341,f13224,f13223,f11203,f3759,f3758,f2874,f1972,f1473,f370,f365,f39,f36343,f36874,f36872,f36708,f36696,f36575,f36504,f36500,f36488,f36331,f36184,f35841,f35840,f13854,f13853,f11293,f11201,f36263,f45254,f45391,f45392,f45396,f45400,f45453,f45454,f45538,f45539,f45544,f45545,f45588,f3303,f49037,f49038])).
fof(f49045,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') = s(bool,X29) | s(bool,'F_0') = s(bool,X30)) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl176 | $spl188)),
  inference(global_subsumption,[],[f31632,f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f67,f64,f132,f133,f146,f156,f157,f200,f361,f363,f71,f54,f1085,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f4170,f4171,f3832,f3890,f3878,f2550,f2549,f3652,f3732,f3828,f3727,f2700,f763,f2698,f761,f5160,f5161,f2455,f1935,f1480,f82,f104,f2706,f11126,f11125,f11188,f11200,f11161,f11162,f11222,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13852,f13181,f13225,f13350,f13427,f13415,f15429,f17356,f22573,f18154,f18155,f22299,f22295,f22291,f22279,f22275,f22219,f22217,f18360,f18362,f18363,f22711,f18802,f19158,f21203,f22901,f22902,f22905,f19338,f19344,f19402,f19403,f19570,f22984,f20200,f23229,f23233,f23235,f24931,f24930,f24854,f24822,f24800,f23543,f23787,f25217,f25261,f25262,f25265,f25280,f25784,f25785,f24220,f24215,f25787,f25788,f23788,f23783,f25790,f23712,f20161,f20156,f25794,f19927,f19749,f19571,f19566,f25798,f19337,f19159,f18981,f18976,f18803,f18798,f25801,f25802,f13351,f13346,f25804,f13292,f13175,f13170,f13039,f11230,f9980,f9978,f9977,f9974,f6734,f6732,f6731,f6728,f24431,f25850,f25854,f28579,f28795,f28928,f28929,f28799,f28801,f28992,f29101,f29277,f29642,f30026,f30749,f30802,f30902,f31517,f31518,f31523,f32112,f32190,f32954,f32955,f32960,f33699,f33700,f33637,f33624,f34071,f34113,f34117,f35168,f35169,f35170,f35172,f36886,f36851,f36833,f37018,f37019,f37020,f37021,f36024,f36821,f37037,f37041,f36272,f36746,f37183,f37187,f37191,f36506,f36507,f36509,f36510,f37202,f37203,f36577,f36580,f37595,f13038,f3651,f17264,f17266,f15352,f22930,f22929,f22928,f786,f550,f475,f48,f47,f46,f25855,f19925,f19747,f19569,f18801,f762,f2762,f11302,f2704,f19406,f25856,f19914,f19736,f19558,f18790,f2699,f785,f2763,f22931,f37580,f37578,f37577,f37576,f37574,f37033,f36728,f34476,f34474,f34472,f34425,f24214,f23782,f22898,f19565,f19331,f19153,f18797,f16976,f15069,f13826,f13345,f13169,f13033,f4154,f3827,f3726,f3646,f2559,f368,f366,f37594,f37211,f37040,f37038,f37032,f36953,f36870,f36709,f36702,f36501,f36494,f36344,f36337,f36111,f36110,f36025,f36023,f36022,f34518,f34208,f34207,f34119,f34116,f34115,f33645,f33643,f33629,f33627,f33552,f33551,f32991,f32792,f32361,f32194,f32193,f32192,f32154,f32152,f32151,f32126,f32105,f32101,f31431,f31357,f31001,f30906,f30905,f30904,f30404,f30179,f30178,f30174,f30172,f30080,f30030,f30029,f30027,f30024,f29905,f29904,f29900,f29898,f29715,f29714,f29581,f29580,f29283,f29256,f29255,f29252,f29247,f29245,f29243,f29241,f25810,f25806,f25792,f25786,f22303,f22293,f22277,f21173,f19401,f19400,f17276,f17226,f15341,f13224,f13223,f11203,f3759,f3758,f2874,f1972,f1473,f370,f365,f39,f36343,f36874,f36872,f36708,f36696,f36575,f36504,f36500,f36488,f36331,f36184,f35841,f35840,f13854,f13853,f11293,f11201,f36263,f45254,f45391,f45392,f45396,f45400,f45453,f45454,f45538,f45539,f45544,f45545,f45588,f3303,f49037,f49038])).
fof(f49044,plain,(
  ( ! [X17,X18] : (s(bool,'F_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X17)) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl176 | $spl188)),
  inference(global_subsumption,[],[f31629,f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f67,f64,f132,f133,f146,f156,f157,f200,f361,f363,f71,f54,f1085,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f4170,f4171,f3832,f3890,f3878,f2550,f2549,f3652,f3732,f3828,f3727,f2700,f763,f2698,f761,f5160,f5161,f2455,f1935,f1480,f82,f104,f2706,f11126,f11125,f11188,f11200,f11161,f11162,f11222,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13852,f13181,f13225,f13350,f13427,f13415,f15429,f17356,f22573,f18154,f18155,f22299,f22295,f22291,f22279,f22275,f22219,f22217,f18360,f18362,f18363,f22711,f18802,f19158,f21203,f22901,f22902,f22905,f19338,f19344,f19402,f19403,f19570,f22984,f20200,f23229,f23233,f23235,f24931,f24930,f24854,f24822,f24800,f23543,f23787,f25217,f25261,f25262,f25265,f25280,f25784,f25785,f24220,f24215,f25787,f25788,f23788,f23783,f25790,f23712,f20161,f20156,f25794,f19927,f19749,f19571,f19566,f25798,f19337,f19159,f18981,f18976,f18803,f18798,f25801,f25802,f13351,f13346,f25804,f13292,f13175,f13170,f13039,f11230,f9980,f9978,f9977,f9974,f6734,f6732,f6731,f6728,f24431,f25850,f25854,f28579,f28795,f28928,f28929,f28799,f28801,f28992,f29101,f29277,f29642,f30026,f30749,f30802,f30902,f31517,f31518,f31523,f32112,f32190,f32954,f32955,f32960,f33699,f33700,f33637,f33624,f34071,f34113,f34117,f35168,f35169,f35170,f35172,f36886,f36851,f36833,f37018,f37019,f37020,f37021,f36024,f36821,f37037,f37041,f36272,f36746,f37183,f37187,f37191,f36506,f36507,f36509,f36510,f37202,f37203,f36577,f36580,f37595,f13038,f3651,f17264,f17266,f15352,f22930,f22929,f22928,f786,f550,f475,f48,f47,f46,f25855,f19925,f19747,f19569,f18801,f762,f2762,f11302,f2704,f19406,f25856,f19914,f19736,f19558,f18790,f2699,f785,f2763,f22931,f37580,f37578,f37577,f37576,f37574,f37033,f36728,f34476,f34474,f34472,f34425,f24214,f23782,f22898,f19565,f19331,f19153,f18797,f16976,f15069,f13826,f13345,f13169,f13033,f4154,f3827,f3726,f3646,f2559,f368,f366,f37594,f37211,f37040,f37038,f37032,f36953,f36870,f36709,f36702,f36501,f36494,f36344,f36337,f36111,f36110,f36025,f36023,f36022,f34518,f34208,f34207,f34119,f34116,f34115,f33645,f33643,f33629,f33627,f33552,f33551,f32991,f32792,f32361,f32194,f32193,f32192,f32154,f32152,f32151,f32126,f32105,f32101,f31431,f31357,f31001,f30906,f30905,f30904,f30404,f30179,f30178,f30174,f30172,f30080,f30030,f30029,f30027,f30024,f29905,f29904,f29900,f29898,f29715,f29714,f29581,f29580,f29283,f29256,f29255,f29252,f29247,f29245,f29243,f29241,f25810,f25806,f25792,f25786,f22303,f22293,f22277,f21173,f19401,f19400,f17276,f17226,f15341,f13224,f13223,f11203,f3759,f3758,f2874,f1972,f1473,f370,f365,f39,f36343,f36874,f36872,f36708,f36696,f36575,f36504,f36500,f36488,f36331,f36184,f35841,f35840,f13854,f13853,f11293,f11201,f36263,f45254,f45391,f45392,f45396,f45400,f45453,f45454,f45538,f45539,f45544,f45545,f45588,f3303,f49037,f49038])).
fof(f49043,plain,(
  ( ! [X14,X13] : (s(bool,'T_0') = s(bool,X13) | s(bool,'F_0') = s(bool,X14)) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl176 | $spl188)),
  inference(global_subsumption,[],[f49042,f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f67,f64,f132,f133,f146,f156,f157,f200,f361,f363,f71,f54,f1085,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f4170,f4171,f3832,f3890,f3878,f2550,f2549,f3652,f3732,f3828,f3727,f2700,f763,f2698,f761,f5160,f5161,f2455,f1935,f1480,f82,f104,f2706,f11126,f11125,f11188,f11200,f11161,f11162,f11222,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13852,f13181,f13225,f13350,f13427,f13415,f15429,f17356,f22573,f18154,f18155,f22299,f22295,f22291,f22279,f22275,f22219,f22217,f18360,f18362,f18363,f22711,f18802,f19158,f21203,f22901,f22902,f22905,f19338,f19344,f19402,f19403,f19570,f22984,f20200,f23229,f23233,f23235,f24931,f24930,f24854,f24822,f24800,f23543,f23787,f25217,f25261,f25262,f25265,f25280,f25784,f25785,f24220,f24215,f25787,f25788,f23788,f23783,f25790,f23712,f20161,f20156,f25794,f19927,f19749,f19571,f19566,f25798,f19337,f19159,f18981,f18976,f18803,f18798,f25801,f25802,f13351,f13346,f25804,f13292,f13175,f13170,f13039,f11230,f9980,f9978,f9977,f9974,f6734,f6732,f6731,f6728,f24431,f25850,f25854,f28579,f28795,f28928,f28929,f28799,f28801,f28992,f29101,f29277,f29642,f30026,f30749,f30802,f30902,f31517,f31518,f31523,f32112,f32190,f32954,f32955,f32960,f33699,f33700,f33637,f33624,f34071,f34113,f34117,f35168,f35169,f35170,f35172,f36886,f36851,f36833,f37018,f37019,f37020,f37021,f36024,f36821,f37037,f37041,f36272,f36746,f37183,f37187,f37191,f36506,f36507,f36509,f36510,f37202,f37203,f36577,f36580,f37595,f13038,f3651,f17264,f17266,f15352,f22930,f22929,f22928,f786,f550,f475,f48,f47,f46,f25855,f19925,f19747,f19569,f18801,f762,f2762,f11302,f2704,f19406,f25856,f19914,f19736,f19558,f18790,f2699,f785,f2763,f22931,f37580,f37578,f37577,f37576,f37574,f37033,f36728,f34476,f34474,f34472,f34425,f24214,f23782,f22898,f19565,f19331,f19153,f18797,f16976,f15069,f13826,f13345,f13169,f13033,f4154,f3827,f3726,f3646,f2559,f368,f366,f37594,f37211,f37040,f37038,f37032,f36953,f36870,f36709,f36702,f36501,f36494,f36344,f36337,f36111,f36110,f36025,f36023,f36022,f34518,f34208,f34207,f34119,f34116,f34115,f33645,f33643,f33629,f33627,f33552,f33551,f32991,f32792,f32361,f32194,f32193,f32192,f32154,f32152,f32151,f32126,f32105,f32101,f31431,f31357,f31001,f30906,f30905,f30904,f30404,f30179,f30178,f30174,f30172,f30080,f30030,f30029,f30027,f30024,f29905,f29904,f29900,f29898,f29715,f29714,f29581,f29580,f29283,f29256,f29255,f29252,f29247,f29245,f29243,f29241,f25810,f25806,f25792,f25786,f22303,f22293,f22277,f21173,f19401,f19400,f17276,f17226,f15341,f13224,f13223,f11203,f3759,f3758,f2874,f1972,f1473,f370,f365,f39,f36343,f36874,f36872,f36708,f36696,f36575,f36504,f36500,f36488,f36331,f36184,f35841,f35840,f13854,f13853,f11293,f11201,f36263,f45254,f45391,f45392,f45396,f45400,f45453,f45454,f45538,f45539,f45544,f45545,f45588,f3303,f49037,f49038])).
fof(f49041,plain,(
  ( ! [X12,X11] : (s(bool,'F_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X11)) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl176 | $spl188)),
  inference(global_subsumption,[],[f49040,f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f67,f64,f132,f133,f146,f156,f157,f200,f361,f363,f71,f54,f1085,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f4170,f4171,f3832,f3890,f3878,f2550,f2549,f3652,f3732,f3828,f3727,f2700,f763,f2698,f761,f5160,f5161,f2455,f1935,f1480,f82,f104,f2706,f11126,f11125,f11188,f11200,f11161,f11162,f11222,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13852,f13181,f13225,f13350,f13427,f13415,f15429,f17356,f22573,f18154,f18155,f22299,f22295,f22291,f22279,f22275,f22219,f22217,f18360,f18362,f18363,f22711,f18802,f19158,f21203,f22901,f22902,f22905,f19338,f19344,f19402,f19403,f19570,f22984,f20200,f23229,f23233,f23235,f24931,f24930,f24854,f24822,f24800,f23543,f23787,f25217,f25261,f25262,f25265,f25280,f25784,f25785,f24220,f24215,f25787,f25788,f23788,f23783,f25790,f23712,f20161,f20156,f25794,f19927,f19749,f19571,f19566,f25798,f19337,f19159,f18981,f18976,f18803,f18798,f25801,f25802,f13351,f13346,f25804,f13292,f13175,f13170,f13039,f11230,f9980,f9978,f9977,f9974,f6734,f6732,f6731,f6728,f24431,f25850,f25854,f28579,f28795,f28928,f28929,f28799,f28801,f28992,f29101,f29277,f29642,f30026,f30749,f30802,f30902,f31517,f31518,f31523,f32112,f32190,f32954,f32955,f32960,f33699,f33700,f33637,f33624,f34071,f34113,f34117,f35168,f35169,f35170,f35172,f36886,f36851,f36833,f37018,f37019,f37020,f37021,f36024,f36821,f37037,f37041,f36272,f36746,f37183,f37187,f37191,f36506,f36507,f36509,f36510,f37202,f37203,f36577,f36580,f37595,f13038,f3651,f17264,f17266,f15352,f22930,f22929,f22928,f786,f550,f475,f48,f47,f46,f25855,f19925,f19747,f19569,f18801,f762,f2762,f11302,f2704,f19406,f25856,f19914,f19736,f19558,f18790,f2699,f785,f2763,f22931,f37580,f37578,f37577,f37576,f37574,f37033,f36728,f34476,f34474,f34472,f34425,f24214,f23782,f22898,f19565,f19331,f19153,f18797,f16976,f15069,f13826,f13345,f13169,f13033,f4154,f3827,f3726,f3646,f2559,f368,f366,f37594,f37211,f37040,f37038,f37032,f36953,f36870,f36709,f36702,f36501,f36494,f36344,f36337,f36111,f36110,f36025,f36023,f36022,f34518,f34208,f34207,f34119,f34116,f34115,f33645,f33643,f33629,f33627,f33552,f33551,f32991,f32792,f32361,f32194,f32193,f32192,f32154,f32152,f32151,f32126,f32105,f32101,f31431,f31357,f31001,f30906,f30905,f30904,f30404,f30179,f30178,f30174,f30172,f30080,f30030,f30029,f30027,f30024,f29905,f29904,f29900,f29898,f29715,f29714,f29581,f29580,f29283,f29256,f29255,f29252,f29247,f29245,f29243,f29241,f25810,f25806,f25792,f25786,f22303,f22293,f22277,f21173,f19401,f19400,f17276,f17226,f15341,f13224,f13223,f11203,f3759,f3758,f2874,f1972,f1473,f370,f365,f39,f36343,f36874,f36872,f36708,f36696,f36575,f36504,f36500,f36488,f36331,f36184,f35841,f35840,f13854,f13853,f11293,f11201,f36263,f45254,f45391,f45392,f45396,f45400,f45453,f45454,f45538,f45539,f45544,f45545,f45588,f3303,f49037,f49038])).
fof(f49039,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') = s(bool,X7) | s(bool,'T_0') = s(bool,X6)) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl176 | $spl188)),
  inference(global_subsumption,[],[f31623,f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f67,f64,f132,f133,f146,f156,f157,f200,f361,f363,f71,f54,f1085,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f4170,f4171,f3832,f3890,f3878,f2550,f2549,f3652,f3732,f3828,f3727,f2700,f763,f2698,f761,f5160,f5161,f2455,f1935,f1480,f82,f104,f2706,f11126,f11125,f11188,f11200,f11161,f11162,f11222,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13852,f13181,f13225,f13350,f13427,f13415,f15429,f17356,f22573,f18154,f18155,f22299,f22295,f22291,f22279,f22275,f22219,f22217,f18360,f18362,f18363,f22711,f18802,f19158,f21203,f22901,f22902,f22905,f19338,f19344,f19402,f19403,f19570,f22984,f20200,f23229,f23233,f23235,f24931,f24930,f24854,f24822,f24800,f23543,f23787,f25217,f25261,f25262,f25265,f25280,f25784,f25785,f24220,f24215,f25787,f25788,f23788,f23783,f25790,f23712,f20161,f20156,f25794,f19927,f19749,f19571,f19566,f25798,f19337,f19159,f18981,f18976,f18803,f18798,f25801,f25802,f13351,f13346,f25804,f13292,f13175,f13170,f13039,f11230,f9980,f9978,f9977,f9974,f6734,f6732,f6731,f6728,f24431,f25850,f25854,f28579,f28795,f28928,f28929,f28799,f28801,f28992,f29101,f29277,f29642,f30026,f30749,f30802,f30902,f31517,f31518,f31523,f32112,f32190,f32954,f32955,f32960,f33699,f33700,f33637,f33624,f34071,f34113,f34117,f35168,f35169,f35170,f35172,f36886,f36851,f36833,f37018,f37019,f37020,f37021,f36024,f36821,f37037,f37041,f36272,f36746,f37183,f37187,f37191,f36506,f36507,f36509,f36510,f37202,f37203,f36577,f36580,f37595,f13038,f3651,f17264,f17266,f15352,f22930,f22929,f22928,f786,f550,f475,f48,f47,f46,f25855,f19925,f19747,f19569,f18801,f762,f2762,f11302,f2704,f19406,f25856,f19914,f19736,f19558,f18790,f2699,f785,f2763,f22931,f37580,f37578,f37577,f37576,f37574,f37033,f36728,f34476,f34474,f34472,f34425,f24214,f23782,f22898,f19565,f19331,f19153,f18797,f16976,f15069,f13826,f13345,f13169,f13033,f4154,f3827,f3726,f3646,f2559,f368,f366,f37594,f37211,f37040,f37038,f37032,f36953,f36870,f36709,f36702,f36501,f36494,f36344,f36337,f36111,f36110,f36025,f36023,f36022,f34518,f34208,f34207,f34119,f34116,f34115,f33645,f33643,f33629,f33627,f33552,f33551,f32991,f32792,f32361,f32194,f32193,f32192,f32154,f32152,f32151,f32126,f32105,f32101,f31431,f31357,f31001,f30906,f30905,f30904,f30404,f30179,f30178,f30174,f30172,f30080,f30030,f30029,f30027,f30024,f29905,f29904,f29900,f29898,f29715,f29714,f29581,f29580,f29283,f29256,f29255,f29252,f29247,f29245,f29243,f29241,f25810,f25806,f25792,f25786,f22303,f22293,f22277,f21173,f19401,f19400,f17276,f17226,f15341,f13224,f13223,f11203,f3759,f3758,f2874,f1972,f1473,f370,f365,f39,f36343,f36874,f36872,f36708,f36696,f36575,f36504,f36500,f36488,f36331,f36184,f35841,f35840,f13854,f13853,f11293,f11201,f36263,f45254,f45391,f45392,f45396,f45400,f45453,f45454,f45538,f45539,f45544,f45545,f45588,f3303,f49037,f49038])).
fof(f49036,plain,(
  $false | ($spl4 | $spl9 | $spl36 | $spl46 | $spl54 | $spl188)),
  inference(subsumption_resolution,[],[f49035,f37])).
fof(f49035,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl54 | $spl188)),
  inference(forward_demodulation,[],[f39,f48954])).
fof(f48954,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl54 | $spl188)),
  inference(condensation,[],[f48953])).
fof(f48953,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,X0) = s(bool,X1)) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl54 | $spl188)),
  inference(global_subsumption,[],[f48952,f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f64,f132,f133,f146,f156,f157,f200,f361,f363,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f3832,f3890,f3878,f2550,f2549,f3652,f3732,f3828,f3727,f2700,f763,f2698,f761,f5160,f5161,f1935,f1480,f82,f104,f2706,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13181,f13225,f13350,f13427,f13415,f22573,f18154,f18155,f22299,f22291,f22279,f22275,f22219,f22217,f18360,f18362,f18363,f18802,f19158,f21203,f19338,f19344,f19402,f19403,f19570,f22984,f20200,f20161,f20156,f19927,f19749,f19571,f19566,f19337,f19159,f18981,f18976,f18803,f18798,f25801,f25802,f13351,f13346,f13175,f13170,f13039,f9980,f9978,f9977,f9974,f6734,f6732,f6731,f6728,f25850,f25854,f30902,f32190,f34071,f35168,f35169,f35170,f13038,f3651,f17264,f17266,f15352,f22930,f22929,f22928,f786,f550,f475,f48,f47,f46,f19925,f19747,f19569,f18801,f762,f2704,f19406,f19914,f19736,f19558,f18790,f2699,f785,f2763,f22931,f37580,f37578,f37577,f37576,f34476,f34474,f34472,f34425,f22898,f19565,f19331,f19153,f18797,f16976,f15069,f13826,f13345,f13169,f13033,f4154,f3827,f3726,f3646,f2559,f368,f366,f32792,f32194,f32193,f32192,f31357,f30906,f30905,f30904,f22293,f22277,f19401,f19400,f17276,f17226,f15341,f13224,f13223,f3759,f3758,f2874,f1972,f1473,f370,f365,f39,f3303])).
fof(f48952,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,g(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | ($spl4 | $spl9 | $spl36 | $spl54)),
  inference(condensation,[],[f48951])).
fof(f48951,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,g(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,X2) = s(bool,g(s(bool,X0)))) ) | ($spl4 | $spl9 | $spl36 | $spl54)),
  inference(condensation,[],[f48950])).
fof(f48950,plain,(
  ( ! [X35,X33,X36,X34] : (s(bool,'F_0') = s(bool,g(s(bool,X34))) | s(bool,'F_0') = s(bool,X33) | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35) | s(bool,X33) = s(bool,X36)) ) | ($spl4 | $spl9 | $spl36 | $spl54)),
  inference(subsumption_resolution,[],[f26357,f30026])).
fof(f26357,plain,(
  ( ! [X35,X33,X36,X34] : (s(bool,'F_0') = s(bool,g(s(bool,X34))) | s(bool,'F_0') = s(bool,X33) | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35) | s(bool,X33) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) ) | $spl54),
  inference(superposition,[],[f4883,f146])).
fof(f49034,plain,(
  $false | ($spl4 | $spl9 | $spl13 | $spl22 | $spl36 | $spl46 | $spl54 | $spl188)),
  inference(subsumption_resolution,[],[f6406,f48954])).
fof(f48986,plain,(
  $false | ($spl4 | $spl9 | $spl36 | $spl46 | $spl54 | $spl188 | $spl203)),
  inference(subsumption_resolution,[],[f37100,f48954])).
fof(f48984,plain,(
  $false | ($spl4 | $spl9 | $spl36 | $spl46 | $spl54 | $spl186 | $spl188)),
  inference(subsumption_resolution,[],[f41131,f48954])).
fof(f48983,plain,(
  $false | ($spl4 | $spl9 | $spl36 | $spl46 | $spl54 | $spl188 | $spl228)),
  inference(subsumption_resolution,[],[f40640,f48954])).
fof(f48982,plain,(
  $false | ($spl4 | $spl9 | $spl36 | $spl46 | $spl54 | $spl186 | $spl188)),
  inference(subsumption_resolution,[],[f33915,f48954])).
fof(f48942,plain,(
  $false | ($spl4 | $spl9 | $spl36 | $spl46 | $spl54)),
  inference(subsumption_resolution,[],[f48941,f71])).
fof(f48941,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl4 | $spl9 | $spl36 | $spl46 | $spl54)),
  inference(forward_demodulation,[],[f48940,f3303])).
fof(f48940,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | ($spl4 | $spl9 | $spl36 | $spl54)),
  inference(condensation,[],[f48939])).
fof(f48939,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | s(bool,X0) = s(bool,g(s(bool,'T_0')))) ) | ($spl4 | $spl9 | $spl36 | $spl54)),
  inference(condensation,[],[f48938])).
fof(f48938,plain,(
  ( ! [X8,X7] : (s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X7) | s(bool,X7) = s(bool,X8)) ) | ($spl4 | $spl9 | $spl36 | $spl54)),
  inference(subsumption_resolution,[],[f26350,f30026])).
fof(f26350,plain,(
  ( ! [X8,X7] : (s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X7) | s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | $spl54),
  inference(superposition,[],[f4883,f56])).
fof(f48937,plain,(
  $false | ($spl4 | $spl9 | $spl36 | $spl46 | $spl54)),
  inference(subsumption_resolution,[],[f48936,f71])).
fof(f48936,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl4 | $spl9 | $spl36 | $spl46 | $spl54)),
  inference(forward_demodulation,[],[f48935,f3303])).
fof(f48935,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | ($spl4 | $spl9 | $spl36 | $spl54)),
  inference(condensation,[],[f48934])).
fof(f48934,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | s(bool,X0) = s(bool,g(s(bool,'T_0')))) ) | ($spl4 | $spl9 | $spl36 | $spl54)),
  inference(condensation,[],[f48933])).
fof(f48933,plain,(
  ( ! [X2,X1] : (s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2)) ) | ($spl4 | $spl9 | $spl36 | $spl54)),
  inference(subsumption_resolution,[],[f26347,f30026])).
fof(f26347,plain,(
  ( ! [X2,X1] : (s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl54),
  inference(superposition,[],[f4883,f56])).
fof(f48932,plain,(
  $false | ($spl9 | $spl46 | $spl54)),
  inference(subsumption_resolution,[],[f48931,f71])).
fof(f48931,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl9 | $spl46 | $spl54)),
  inference(forward_demodulation,[],[f26345,f3303])).
fof(f26345,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | ($spl9 | $spl54)),
  inference(unit_resulting_resolution,[],[f71,f4883])).
fof(f48924,plain,(
  $false | ($spl4 | $spl9 | $spl36 | $spl42 | $spl46 | $spl76 | $spl186 | $spl228)),
  inference(subsumption_resolution,[],[f33915,f48548])).
fof(f48548,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl4 | $spl9 | $spl36 | $spl42 | $spl46 | $spl76 | $spl228)),
  inference(subsumption_resolution,[],[f48547,f3890])).
fof(f48547,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'T_0'),s(bool,X0)))))) ) | ($spl4 | $spl9 | $spl36 | $spl42 | $spl46 | $spl76 | $spl228)),
  inference(condensation,[],[f48477])).
fof(f48477,plain,(
  ( ! [X64,X63] : (s(bool,X64) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'T_0'),s(bool,X64))))) | s(bool,'T_0') = s(bool,X63) | s(bool,'T_0') = s(bool,X64)) ) | ($spl4 | $spl9 | $spl36 | $spl42 | $spl46 | $spl76 | $spl228)),
  inference(backward_demodulation,[],[f48476,f3170])).
fof(f48476,plain,(
  ( ! [X168] : (s(bool,'T_0') = s(bool,f(s(bool,X168)))) ) | ($spl4 | $spl9 | $spl36 | $spl46 | $spl76 | $spl228)),
  inference(global_subsumption,[],[f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f67,f64,f132,f133,f146,f156,f157,f200,f361,f363,f71,f54,f1085,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f4170,f4171,f3832,f3890,f3878,f2550,f2549,f3303,f3652,f3732,f3828,f3727,f2700,f763,f2698,f761,f5160,f5161,f2455,f1935,f1480,f82,f104,f2706,f11126,f11125,f11188,f11200,f11161,f11162,f11222,f11776,f11752,f11808,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13816,f13852,f13181,f13225,f13350,f13427,f13415,f15429,f17356,f22573,f18154,f18155,f22299,f22295,f22291,f22279,f22275,f22219,f22217,f18360,f18362,f18363,f22709,f22711,f18802,f22888,f19158,f21203,f21171,f22901,f22902,f22905,f19338,f19344,f19402,f19403,f19570,f22984,f20200,f23229,f23233,f23235,f24931,f24930,f24854,f24822,f24800,f23543,f23787,f25217,f24550,f25261,f25262,f25265,f25280,f25784,f25785,f24220,f24215,f25787,f25788,f23788,f23783,f25790,f23712,f20161,f20156,f25794,f19927,f19749,f19571,f19566,f25798,f19337,f19159,f18981,f18976,f18803,f18798,f25801,f25802,f13351,f13346,f25804,f13292,f13175,f13170,f13039,f11230,f9980,f9978,f9977,f9974,f6734,f6732,f6731,f6728,f24431,f25850,f25854,f28579,f28795,f28928,f28929,f28799,f28801,f28992,f29101,f29277,f29426,f29642,f30026,f30749,f30802,f30902,f31517,f31518,f31523,f32112,f32190,f32954,f32955,f32960,f33699,f33700,f33637,f33624,f34113,f34117,f35168,f35169,f35170,f35172,f36886,f36851,f36833,f37018,f37019,f37020,f37021,f36024,f36821,f37037,f37041,f36272,f36746,f37183,f37187,f37191,f36506,f36507,f36509,f36510,f37202,f37203,f37206,f36577,f36580,f36729,f37595,f37727,f37852,f13038,f3651,f17264,f17266,f15352,f22930,f22929,f22928,f786,f550,f475,f48,f47,f46,f25855,f19925,f19747,f19569,f18801,f762,f2762,f11302,f2704,f19406,f25856,f19914,f19736,f19558,f18790,f2699,f785,f2763,f22931,f37580,f37578,f37577,f37576,f37574,f37033,f36728,f34476,f34474,f34472,f34425,f24214,f23782,f22898,f19565,f19331,f19153,f18797,f16976,f15069,f13826,f13345,f13169,f13033,f4154,f3827,f3726,f3646,f2559,f368,f366,f37728,f37594,f37211,f37193,f37040,f37038,f37032,f36953,f36870,f36709,f36702,f36501,f36494,f36344,f36337,f36111,f36110,f36025,f36023,f36022,f34518,f34208,f34207,f34119,f34116,f34115,f33645,f33643,f33629,f33627,f33552,f33551,f32991,f32792,f32361,f32194,f32193,f32192,f32154,f32152,f32151,f32126,f32105,f32101,f31431,f31357,f31001,f30906,f30905,f30904,f30716,f30405,f30404,f30179,f30178,f30174,f30172,f30080,f30030,f30029,f30027,f30024,f29905,f29904,f29900,f29898,f29715,f29714,f29581,f29580,f29283,f29256,f29255,f29252,f29247,f29245,f29243,f29241,f25810,f25806,f25792,f25786,f22303,f22293,f22277,f21173,f19401,f19400,f17276,f17226,f15341,f13224,f13223,f11203,f3759,f3758,f2874,f1972,f1473,f370,f365,f39,f36343,f36874,f36872,f36708,f36696,f36575,f36504,f36500,f36488,f36331,f36184,f35841,f35840,f29424,f29422,f29397,f29395,f29389,f13854,f13853,f11293,f11201,f4984,f36263,f38764,f40698,f40699,f45254,f45391,f45392,f45396,f45400,f45453,f45454,f45538,f45539,f45544,f45545,f45588,f48458,f48459,f46149,f48467,f48475,f45501])).
fof(f45501,plain,(
  ( ! [X168] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X168)))) ) | ($spl4 | $spl9 | $spl36 | $spl76 | $spl228)),
  inference(forward_demodulation,[],[f45500,f11203])).
fof(f45500,plain,(
  ( ! [X168] : (s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,f(s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,f(s(bool,X168)))) ) | ($spl4 | $spl9 | $spl36 | $spl76 | $spl228)),
  inference(subsumption_resolution,[],[f45069,f37193])).
fof(f45069,plain,(
  ( ! [X168] : (s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,f(s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X168) | s(bool,'T_0') = s(bool,f(s(bool,X168)))) ) | ($spl4 | $spl9 | $spl36 | $spl228)),
  inference(superposition,[],[f38764,f37032])).
fof(f29389,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') != s(bool,f(s(bool,X1))) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X2)) ) | $spl76),
  inference(condensation,[],[f29388])).
fof(f29388,plain,(
  ( ! [X54,X56,X55,X53] : (s(bool,'F_0') != s(bool,f(s(bool,X54))) | s(bool,'T_0') = s(bool,X53) | s(bool,X54) = s(bool,X55) | s(bool,X55) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) ) | $spl76),
  inference(duplicate_literal_removal,[],[f29305])).
fof(f29305,plain,(
  ( ! [X54,X56,X55,X53] : (s(bool,'F_0') != s(bool,f(s(bool,X54))) | s(bool,'T_0') = s(bool,X53) | s(bool,X54) = s(bool,X55) | s(bool,'T_0') = s(bool,X53) | s(bool,X55) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) ) | $spl76),
  inference(superposition,[],[f4984,f146])).
fof(f29395,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,X2))))) | s(bool,'F_0') != s(bool,f(s(bool,X2)))) ) | $spl76),
  inference(condensation,[],[f29394])).
fof(f29394,plain,(
  ( ! [X28,X26,X27,X25] : (s(bool,'F_0') != s(bool,f(s(bool,X26))) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X27,bool),'const.pair.,_2'(s(X27,X28),s(bool,X26)))))) ) | $spl76),
  inference(duplicate_literal_removal,[],[f29298])).
fof(f29298,plain,(
  ( ! [X28,X26,X27,X25] : (s(bool,'F_0') != s(bool,f(s(bool,X26))) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X27,bool),'const.pair.,_2'(s(X27,X28),s(bool,X26)))))) ) | $spl76),
  inference(superposition,[],[f4984,f133])).
fof(f29397,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,X1),s(X0,X2))))) | s(bool,'F_0') != s(bool,f(s(bool,X1)))) ) | $spl76),
  inference(condensation,[],[f29396])).
fof(f29396,plain,(
  ( ! [X14,X15,X13,X16] : (s(bool,'F_0') != s(bool,f(s(bool,X14))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X15),'const.pair.,_2'(s(bool,X14),s(X15,X16)))))) ) | $spl76),
  inference(duplicate_literal_removal,[],[f29295])).
fof(f29295,plain,(
  ( ! [X14,X15,X13,X16] : (s(bool,'F_0') != s(bool,f(s(bool,X14))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X15),'const.pair.,_2'(s(bool,X14),s(X15,X16)))))) ) | $spl76),
  inference(superposition,[],[f4984,f132])).
fof(f29422,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,f(s(bool,X0))),s(X1,X2)))))) ) | $spl76),
  inference(condensation,[],[f29421])).
fof(f29421,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X11),'const.pair.,_2'(s(bool,f(s(bool,X9))),s(X11,X12)))))) ) | $spl76),
  inference(subsumption_resolution,[],[f29350,f40])).
fof(f29350,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,'F_0') != s(bool,X10) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X11),'const.pair.,_2'(s(bool,f(s(bool,X9))),s(X11,X12)))))) ) | $spl76),
  inference(superposition,[],[f4984,f132])).
fof(f29424,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,f(s(bool,X0)))))))) ) | $spl76),
  inference(condensation,[],[f29423])).
fof(f29423,plain,(
  ( ! [X24,X23,X21,X22] : (s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X23,bool),'const.pair.,_2'(s(X23,X24),s(bool,f(s(bool,X21)))))))) ) | $spl76),
  inference(subsumption_resolution,[],[f29353,f40])).
fof(f29353,plain,(
  ( ! [X24,X23,X21,X22] : (s(bool,'F_0') != s(bool,X22) | s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X23,bool),'const.pair.,_2'(s(X23,X24),s(bool,f(s(bool,X21)))))))) ) | $spl76),
  inference(superposition,[],[f4984,f133])).
fof(f30405,plain,(
  ( ! [X202,X201] : (s(bool,'T_0') = s(bool,X201) | s(bool,X202) = s(bool,f(s(bool,X201))) | s(bool,'F_0') = s(bool,X202)) ) | ($spl4 | $spl9 | $spl36 | $spl76)),
  inference(trivial_inequality_removal,[],[f30140])).
fof(f30140,plain,(
  ( ! [X202,X201] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X201) | s(bool,X202) = s(bool,f(s(bool,X201))) | s(bool,'F_0') = s(bool,X202)) ) | ($spl4 | $spl9 | $spl36 | $spl76)),
  inference(superposition,[],[f4984,f28992])).
fof(f30716,plain,(
  ( ! [X161,X160] : (s(bool,X160) != s(bool,f(s(bool,X160))) | s(bool,X160) = s(bool,X161) | s(bool,'F_0') = s(bool,X161)) ) | ($spl4 | $spl9 | $spl36 | $spl76)),
  inference(subsumption_resolution,[],[f30060,f71])).
fof(f30060,plain,(
  ( ! [X161,X160] : (s(bool,X160) != s(bool,f(s(bool,X160))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X160) = s(bool,X161) | s(bool,'F_0') = s(bool,X161)) ) | ($spl4 | $spl9 | $spl36 | $spl76)),
  inference(superposition,[],[f4984,f28992])).
fof(f37193,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | s(bool,'T_0') = s(bool,f(s(bool,X0)))) ) | ($spl4 | $spl9 | $spl36 | $spl76)),
  inference(condensation,[],[f37192])).
fof(f37192,plain,(
  ( ! [X337,X338] : (s(bool,'F_0') != s(bool,X338) | s(bool,'T_0') = s(bool,X337) | s(bool,'T_0') = s(bool,f(s(bool,X338)))) ) | ($spl4 | $spl9 | $spl36 | $spl76)),
  inference(subsumption_resolution,[],[f36444,f4984])).
fof(f36444,plain,(
  ( ! [X337,X338] : (s(bool,'F_0') != s(bool,X338) | s(bool,'T_0') = s(bool,X337) | s(bool,'T_0') = s(bool,f(s(bool,X338))) | s(bool,'F_0') = s(bool,f(s(bool,X337)))) ) | ($spl4 | $spl9 | $spl36 | $spl76)),
  inference(superposition,[],[f4984,f30029])).
fof(f37728,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) != s(bool,X4) | s(bool,'T_0') = s(bool,f(s(bool,X4))) | s(bool,X3) = s(bool,X5) | s(bool,'F_0') = s(bool,X5)) ) | ($spl4 | $spl9 | $spl36 | $spl76)),
  inference(superposition,[],[f37193,f28992])).
fof(f37852,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,X24) != s(bool,X25) | s(bool,'T_0') = s(bool,f(s(bool,X25))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X26),'const.pair.,_2'(s(bool,X24),s(X26,X27)))))) ) | ($spl4 | $spl9 | $spl36 | $spl76)),
  inference(subsumption_resolution,[],[f37737,f71])).
fof(f37737,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,X24) != s(bool,X25) | s(bool,'T_0') = s(bool,f(s(bool,X25))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X26),'const.pair.,_2'(s(bool,X24),s(X26,X27)))))) ) | ($spl4 | $spl9 | $spl36 | $spl76)),
  inference(superposition,[],[f37193,f132])).
fof(f37727,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,f(s(bool,X2))) | s(bool,X0) = s(bool,X2)) ) | ($spl4 | $spl9 | $spl36 | $spl76)),
  inference(superposition,[],[f37193,f30029])).
fof(f36729,plain,(
  ( ! [X267,X268] : (s(bool,'T_0') = s(bool,X267) | s(bool,'T_0') = s(bool,f(s(bool,X268))) | s(bool,X268) = s(bool,f(s(bool,X267)))) ) | ($spl4 | $spl9 | $spl36 | $spl76)),
  inference(trivial_inequality_removal,[],[f36657])).
fof(f36657,plain,(
  ( ! [X267,X268] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X267) | s(bool,'T_0') = s(bool,f(s(bool,X268))) | s(bool,X268) = s(bool,f(s(bool,X267)))) ) | ($spl4 | $spl9 | $spl36 | $spl76)),
  inference(superposition,[],[f4984,f30029])).
fof(f37206,plain,(
  ( ! [X200,X201] : (s(bool,X200) != s(bool,f(s(bool,X200))) | s(bool,'T_0') = s(bool,f(s(bool,X201))) | s(bool,X200) = s(bool,X201)) ) | ($spl4 | $spl9 | $spl36 | $spl76)),
  inference(subsumption_resolution,[],[f36548,f71])).
fof(f36548,plain,(
  ( ! [X200,X201] : (s(bool,X200) != s(bool,f(s(bool,X200))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(bool,X201))) | s(bool,X200) = s(bool,X201)) ) | ($spl4 | $spl9 | $spl36 | $spl76)),
  inference(superposition,[],[f4984,f30029])).
fof(f29426,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,f(s(bool,X0))) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl76),
  inference(condensation,[],[f29425])).
fof(f29425,plain,(
  ( ! [X52,X50,X51,X49] : (s(bool,'T_0') = s(bool,X49) | s(bool,X51) = s(bool,f(s(bool,X49))) | s(bool,'T_0') = s(bool,X50) | s(bool,X51) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) ) | $spl76),
  inference(subsumption_resolution,[],[f29360,f40])).
fof(f29360,plain,(
  ( ! [X52,X50,X51,X49] : (s(bool,'F_0') != s(bool,X50) | s(bool,'T_0') = s(bool,X49) | s(bool,X51) = s(bool,f(s(bool,X49))) | s(bool,'T_0') = s(bool,X50) | s(bool,X51) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) ) | $spl76),
  inference(superposition,[],[f4984,f146])).
fof(f21171,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') != s(bool,f(s(bool,X1))) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X2)) ) | $spl76),
  inference(condensation,[],[f21170])).
fof(f21170,plain,(
  ( ! [X171,X169,X170,X168] : (s(bool,'F_0') != s(bool,f(s(bool,X169))) | s(bool,'T_0') = s(bool,X168) | s(bool,X169) = s(bool,X170) | s(bool,X170) = s(bool,X171) | s(bool,'T_0') = s(bool,X171)) ) | $spl76),
  inference(duplicate_literal_removal,[],[f19188])).
fof(f19188,plain,(
  ( ! [X171,X169,X170,X168] : (s(bool,'F_0') != s(bool,f(s(bool,X169))) | s(bool,'T_0') = s(bool,X168) | s(bool,X169) = s(bool,X170) | s(bool,'T_0') = s(bool,X168) | s(bool,X170) = s(bool,X171) | s(bool,'T_0') = s(bool,X171)) ) | $spl76),
  inference(superposition,[],[f4984,f146])).
fof(f22888,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,f(s(bool,X0))) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl76),
  inference(condensation,[],[f22887])).
fof(f22887,plain,(
  ( ! [X505,X504,X506,X503] : (s(bool,'T_0') = s(bool,X503) | s(bool,X505) = s(bool,f(s(bool,X503))) | s(bool,'T_0') = s(bool,X504) | s(bool,X505) = s(bool,X506) | s(bool,'T_0') = s(bool,X506)) ) | $spl76),
  inference(subsumption_resolution,[],[f19069,f40])).
fof(f19069,plain,(
  ( ! [X505,X504,X506,X503] : (s(bool,'F_0') != s(bool,X504) | s(bool,'T_0') = s(bool,X503) | s(bool,X505) = s(bool,f(s(bool,X503))) | s(bool,'T_0') = s(bool,X504) | s(bool,X505) = s(bool,X506) | s(bool,'T_0') = s(bool,X506)) ) | $spl76),
  inference(superposition,[],[f4984,f146])).
fof(f22709,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') != s(bool,X1) | s(bool,X0) = s(bool,X1)) ) | ($spl36 | $spl76)),
  inference(condensation,[],[f22708])).
fof(f22708,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') != s(bool,X1) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X2)) ) | ($spl36 | $spl76)),
  inference(condensation,[],[f22707])).
fof(f22707,plain,(
  ( ! [X505,X504,X506,X503] : (s(bool,'F_0') != s(bool,X504) | s(bool,'T_0') = s(bool,X503) | s(bool,X504) = s(bool,X505) | s(bool,'T_0') = s(bool,X505) | s(bool,'T_0') = s(bool,X506)) ) | ($spl36 | $spl76)),
  inference(subsumption_resolution,[],[f18713,f11162])).
fof(f18713,plain,(
  ( ! [X505,X504,X506,X503] : (s(bool,'F_0') != s(bool,X504) | s(bool,'T_0') = s(bool,X503) | s(bool,X504) = s(bool,X505) | s(bool,'T_0') = s(bool,X505) | s(bool,X506) = s(bool,f(s(bool,X503))) | s(bool,'T_0') = s(bool,X506)) ) | $spl76),
  inference(superposition,[],[f4984,f146])).
fof(f13816,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,f(s(bool,X0)))))))) ) | $spl76),
  inference(condensation,[],[f13815])).
fof(f13815,plain,(
  ( ! [X360,X362,X361,X363] : (s(bool,'T_0') = s(bool,X360) | s(bool,'T_0') = s(bool,X361) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X362,bool),'const.pair.,_2'(s(X362,X363),s(bool,f(s(bool,X360)))))))) ) | $spl76),
  inference(subsumption_resolution,[],[f12970,f40])).
fof(f12970,plain,(
  ( ! [X360,X362,X361,X363] : (s(bool,'F_0') != s(bool,X361) | s(bool,'T_0') = s(bool,X360) | s(bool,'T_0') = s(bool,X361) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X362,bool),'const.pair.,_2'(s(X362,X363),s(bool,f(s(bool,X360)))))))) ) | $spl76),
  inference(superposition,[],[f4984,f133])).
fof(f11808,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,f(s(bool,X0))),s(X1,X2)))))) ) | $spl76),
  inference(condensation,[],[f11807])).
fof(f11807,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X11),'const.pair.,_2'(s(bool,f(s(bool,X9))),s(X11,X12)))))) ) | $spl76),
  inference(subsumption_resolution,[],[f11760,f40])).
fof(f11760,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,'F_0') != s(bool,X10) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X11),'const.pair.,_2'(s(bool,f(s(bool,X9))),s(X11,X12)))))) ) | $spl76),
  inference(superposition,[],[f4984,f132])).
fof(f11752,plain,(
  ( ! [X85,X83,X84] : (s(bool,'F_0') != s(bool,f(s(bool,X85))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X83,bool),'const.pair.,_2'(s(X83,X84),s(bool,X85)))))) ) | $spl76),
  inference(superposition,[],[f4984,f44])).
fof(f11776,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,X1),s(X0,X2))))) | s(bool,'F_0') != s(bool,f(s(bool,X1)))) ) | $spl76),
  inference(condensation,[],[f11775])).
fof(f11775,plain,(
  ( ! [X14,X15,X13,X16] : (s(bool,'F_0') != s(bool,f(s(bool,X14))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X15),'const.pair.,_2'(s(bool,X14),s(X15,X16)))))) ) | $spl76),
  inference(duplicate_literal_removal,[],[f11717])).
fof(f11717,plain,(
  ( ! [X14,X15,X13,X16] : (s(bool,'F_0') != s(bool,f(s(bool,X14))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X15),'const.pair.,_2'(s(bool,X14),s(X15,X16)))))) ) | $spl76),
  inference(superposition,[],[f4984,f132])).
fof(f48923,plain,(
  $false | ($spl4 | $spl9 | $spl36 | $spl42 | $spl46 | $spl76 | $spl228)),
  inference(subsumption_resolution,[],[f40640,f48548])).
fof(f48922,plain,(
  $false | ($spl4 | $spl9 | $spl36 | $spl42 | $spl46 | $spl76 | $spl186 | $spl228)),
  inference(subsumption_resolution,[],[f41131,f48548])).
fof(f48916,plain,(
  $false | ($spl4 | $spl9 | $spl36 | $spl42 | $spl46 | $spl76 | $spl203 | $spl228)),
  inference(subsumption_resolution,[],[f37100,f48548])).
fof(f48862,plain,(
  $false | ($spl4 | $spl9 | $spl13 | $spl22 | $spl36 | $spl42 | $spl46 | $spl76 | $spl228)),
  inference(subsumption_resolution,[],[f48577,f48548])).
fof(f48577,plain,(
  s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'T_0'),s(bool,'T_0'))))) | ($spl4 | $spl9 | $spl13 | $spl22 | $spl36 | $spl42 | $spl46 | $spl76 | $spl228)),
  inference(backward_demodulation,[],[f48548,f6406])).
fof(f48860,plain,(
  $false | ($spl4 | $spl9 | $spl36 | $spl42 | $spl46 | $spl76 | $spl228)),
  inference(subsumption_resolution,[],[f48549,f37])).
fof(f48549,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl9 | $spl36 | $spl42 | $spl46 | $spl76 | $spl228)),
  inference(backward_demodulation,[],[f48548,f39])).
fof(f48442,plain,(
  ( ! [X105,X102,X106] : (s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X105,bool),'const.pair.,_2'(s(X105,X106),s(bool,g(s(bool,f(s(bool,X102))))))))) | s(bool,'T_0') = s(bool,X102)) ) | $spl246),
  inference(cnf_transformation,[],[f48442_D])).
fof(f48442_D,plain,(
  ( ! [X105,X102,X106] : (s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X105,bool),'const.pair.,_2'(s(X105,X106),s(bool,g(s(bool,f(s(bool,X102))))))))) | s(bool,'T_0') = s(bool,X102)) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f48275,plain,(
  ( ! [X127,X125,X128,X126,X124] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X125,bool),'const.pair.,_2'(s(X125,X126),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X127,bool),'const.pair.,_2'(s(X127,X128),s(bool,g(s(bool,f(s(bool,X124))))))))) | s(bool,'T_0') = s(bool,X124)) ) | ($spl4 | $spl9 | $spl36 | $spl70)),
  inference(superposition,[],[f1972,f45445])).
fof(f45445,plain,(
  ( ! [X114] : (s(bool,'F_0') = s(bool,g(s(bool,f(s(bool,X114))))) | s(bool,'T_0') = s(bool,X114)) ) | ($spl4 | $spl9 | $spl36 | $spl70)),
  inference(forward_demodulation,[],[f45444,f11203])).
fof(f45444,plain,(
  ( ! [X114] : (s(bool,'F_0') = s(bool,g(s(bool,f(s(bool,X114))))) | s(bool,X114) = s(bool,f(s(bool,'F_0')))) ) | ($spl4 | $spl9 | $spl36 | $spl70)),
  inference(subsumption_resolution,[],[f44961,f37])).
fof(f44961,plain,(
  ( ! [X114] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,g(s(bool,f(s(bool,X114))))) | s(bool,X114) = s(bool,f(s(bool,'F_0')))) ) | ($spl4 | $spl9 | $spl36 | $spl70)),
  inference(superposition,[],[f4928,f37032])).
fof(f48443,plain,(
  ( ! [X123,X120] : (s(bool,X123) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,g(s(bool,f(s(bool,X120))))),s(bool,X123))))) | s(bool,'T_0') = s(bool,X120) | s(bool,'T_0') = s(bool,X123)) ) | $spl248),
  inference(cnf_transformation,[],[f48443_D])).
fof(f48443_D,plain,(
  ( ! [X123,X120] : (s(bool,X123) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,g(s(bool,f(s(bool,X120))))),s(bool,X123))))) | s(bool,'T_0') = s(bool,X120) | s(bool,'T_0') = s(bool,X123)) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f48274,plain,(
  ( ! [X123,X121,X122,X120] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X121),'const.pair.,_2'(s(bool,'F_0'),s(X121,X122))))) | s(bool,X123) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,g(s(bool,f(s(bool,X120))))),s(bool,X123))))) | s(bool,'T_0') = s(bool,X123) | s(bool,'T_0') = s(bool,X120)) ) | ($spl4 | $spl9 | $spl16 | $spl36 | $spl70)),
  inference(superposition,[],[f1557,f45445])).
fof(f48270,plain,(
  ( ! [X103,X105,X102,X106,X104] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X103),'const.pair.,_2'(s(bool,'F_0'),s(X103,X104))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X105,bool),'const.pair.,_2'(s(X105,X106),s(bool,g(s(bool,f(s(bool,X102))))))))) | s(bool,'T_0') = s(bool,X102)) ) | ($spl4 | $spl9 | $spl36 | $spl70)),
  inference(superposition,[],[f1473,f45445])).
fof(f48393,plain,(
  ( ! [X79] : (s(bool,'F_0') = s(bool,g(s(bool,f(s(bool,X79))))) | s(bool,X79) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | $spl244),
  inference(cnf_transformation,[],[f48393_D])).
fof(f48393_D,plain,(
  ( ! [X79] : (s(bool,'F_0') = s(bool,g(s(bool,f(s(bool,X79))))) | s(bool,X79) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f48352,plain,(
  ( ! [X78,X79] : (s(bool,'F_0') = s(bool,g(s(bool,f(s(bool,X79))))) | s(bool,'T_0') = s(bool,X78) | s(bool,X79) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl9 | $spl36 | $spl70)),
  inference(duplicate_literal_removal,[],[f48113])).
fof(f48113,plain,(
  ( ! [X78,X79] : (s(bool,'F_0') = s(bool,g(s(bool,f(s(bool,X79))))) | s(bool,'T_0') = s(bool,X78) | s(bool,'T_0') = s(bool,X78) | s(bool,X79) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl9 | $spl36 | $spl70)),
  inference(superposition,[],[f45445,f22295])).
fof(f47983,plain,(
  ( ! [X109,X105,X108] : (s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X108,bool),'const.pair.,_2'(s(X108,X109),s(bool,g(s(bool,X105))))))) | s(bool,'T_0') = s(bool,f(s(bool,X105)))) ) | $spl240),
  inference(cnf_transformation,[],[f47983_D])).
fof(f47983_D,plain,(
  ( ! [X109,X105,X108] : (s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X108,bool),'const.pair.,_2'(s(X108,X109),s(bool,g(s(bool,X105))))))) | s(bool,'T_0') = s(bool,f(s(bool,X105)))) ) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f47888,plain,(
  ( ! [X127,X130,X128,X131,X129] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X128,bool),'const.pair.,_2'(s(X128,X129),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X130,bool),'const.pair.,_2'(s(X130,X131),s(bool,g(s(bool,X127))))))) | s(bool,'T_0') = s(bool,f(s(bool,X127)))) ) | ($spl4 | $spl9 | $spl36 | $spl70 | $spl172)),
  inference(superposition,[],[f1972,f37094])).
fof(f37094,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,g(s(bool,X0))) | s(bool,'T_0') = s(bool,f(s(bool,X0)))) ) | ($spl4 | $spl9 | $spl36 | $spl70 | $spl172)),
  inference(condensation,[],[f37093])).
fof(f37093,plain,(
  ( ! [X167,X168] : (s(bool,'F_0') = s(bool,g(s(bool,X167))) | s(bool,'T_0') = s(bool,f(s(bool,X167))) | s(bool,'F_0') = s(bool,X168)) ) | ($spl4 | $spl9 | $spl36 | $spl70 | $spl172)),
  inference(subsumption_resolution,[],[f36227,f29278])).
fof(f36227,plain,(
  ( ! [X167,X168] : (~p(s(bool,X168)) | s(bool,'F_0') = s(bool,g(s(bool,X167))) | s(bool,'T_0') = s(bool,f(s(bool,X167))) | s(bool,'F_0') = s(bool,X168)) ) | ($spl4 | $spl9 | $spl36 | $spl70)),
  inference(superposition,[],[f4928,f30029])).
fof(f47984,plain,(
  ( ! [X123,X126] : (s(bool,X126) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,g(s(bool,X123))),s(bool,X126))))) | s(bool,'T_0') = s(bool,f(s(bool,X123))) | s(bool,'T_0') = s(bool,X126)) ) | $spl242),
  inference(cnf_transformation,[],[f47984_D])).
fof(f47984_D,plain,(
  ( ! [X123,X126] : (s(bool,X126) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,g(s(bool,X123))),s(bool,X126))))) | s(bool,'T_0') = s(bool,f(s(bool,X123))) | s(bool,'T_0') = s(bool,X126)) ) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f47887,plain,(
  ( ! [X125,X123,X126,X124] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X124),'const.pair.,_2'(s(bool,'F_0'),s(X124,X125))))) | s(bool,X126) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,g(s(bool,X123))),s(bool,X126))))) | s(bool,'T_0') = s(bool,X126) | s(bool,'T_0') = s(bool,f(s(bool,X123)))) ) | ($spl4 | $spl9 | $spl16 | $spl36 | $spl70 | $spl172)),
  inference(superposition,[],[f1557,f37094])).
fof(f47883,plain,(
  ( ! [X109,X107,X105,X108,X106] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X106),'const.pair.,_2'(s(bool,'F_0'),s(X106,X107))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X108,bool),'const.pair.,_2'(s(X108,X109),s(bool,g(s(bool,X105))))))) | s(bool,'T_0') = s(bool,f(s(bool,X105)))) ) | ($spl4 | $spl9 | $spl36 | $spl70 | $spl172)),
  inference(superposition,[],[f1473,f37094])).
fof(f44329,plain,(
  ( ! [X300,X301,X297] : (s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X300,bool),'const.pair.,_2'(s(X300,X301),s(bool,f(s(bool,X297))))))) | s(bool,'F_0') = s(bool,X297)) ) | $spl238),
  inference(cnf_transformation,[],[f44329_D])).
fof(f44329_D,plain,(
  ( ! [X300,X301,X297] : (s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X300,bool),'const.pair.,_2'(s(X300,X301),s(bool,f(s(bool,X297))))))) | s(bool,'F_0') = s(bool,X297)) ) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f43680,plain,(
  ( ! [X300,X298,X301,X297,X299] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X298),'const.pair.,_2'(s(bool,'F_0'),s(X298,X299))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X300,bool),'const.pair.,_2'(s(X300,X301),s(bool,f(s(bool,X297))))))) | s(bool,'F_0') = s(bool,X297)) ) | $spl204),
  inference(superposition,[],[f1473,f37102])).
fof(f42951,plain,(
  ( ! [X2,X1] : (s(bool,X1) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2) | s(bool,X1) = s(bool,X2)) ) | $spl236),
  inference(cnf_transformation,[],[f42951_D])).
fof(f42951_D,plain,(
  ( ! [X2,X1] : (s(bool,X1) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2) | s(bool,X1) = s(bool,X2)) ) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f42953,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,X1) != s(bool,f(s(bool,'T_0'))) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl4 | $spl9 | $spl36 | $spl40 | $spl47 | $spl186)),
  inference(condensation,[],[f42952])).
fof(f42952,plain,(
  ( ! [X39,X37,X38,X40] : (s(bool,X38) != s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X37) | s(bool,X37) = s(bool,X39) | s(bool,X38) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) ) | ($spl4 | $spl9 | $spl36 | $spl40 | $spl47 | $spl186)),
  inference(subsumption_resolution,[],[f42814,f30026])).
fof(f42814,plain,(
  ( ! [X39,X37,X38,X40] : (s(bool,X38) != s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X37) | s(bool,X37) = s(bool,X39) | s(bool,'T_0') = s(bool,X39) | s(bool,X38) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) ) | ($spl4 | $spl9 | $spl36 | $spl40 | $spl47 | $spl186)),
  inference(superposition,[],[f42136,f146])).
fof(f42136,plain,(
  ( ! [X164] : (s(bool,X164) != s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X164)) ) | ($spl4 | $spl9 | $spl36 | $spl40 | $spl47 | $spl186)),
  inference(subsumption_resolution,[],[f41809,f3304])).
fof(f41809,plain,(
  ( ! [X164] : (s(bool,X164) != s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X164) | s(bool,'T_0') = s(bool,g(s(bool,'T_0')))) ) | ($spl4 | $spl9 | $spl36 | $spl40 | $spl186)),
  inference(superposition,[],[f33724,f41609])).
fof(f42950,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,X1) != s(bool,f(s(bool,'T_0'))) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl4 | $spl9 | $spl36 | $spl40 | $spl47 | $spl186)),
  inference(condensation,[],[f42949])).
fof(f42949,plain,(
  ( ! [X35,X33,X36,X34] : (s(bool,X34) != s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X33) | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35) | s(bool,X33) = s(bool,X36)) ) | ($spl4 | $spl9 | $spl36 | $spl40 | $spl47 | $spl186)),
  inference(subsumption_resolution,[],[f42813,f30026])).
fof(f42813,plain,(
  ( ! [X35,X33,X36,X34] : (s(bool,X34) != s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X33) | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35) | s(bool,X33) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) ) | ($spl4 | $spl9 | $spl36 | $spl40 | $spl47 | $spl186)),
  inference(superposition,[],[f42136,f146])).
fof(f37100,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | $spl203),
  inference(cnf_transformation,[],[f37100_D])).
fof(f37100_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl203),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl203])])).
fof(f42133,plain,(
  ( ! [X153] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X153)) ) | ($spl2 | $spl4 | $spl9 | $spl36 | $spl40 | $spl194)),
  inference(forward_demodulation,[],[f41949,f1074])).
fof(f41949,plain,(
  ( ! [X153] : (s(bool,f(s(bool,'T_0'))) != s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X153)) ) | ($spl4 | $spl9 | $spl36 | $spl40 | $spl194)),
  inference(duplicate_literal_removal,[],[f41799])).
fof(f41799,plain,(
  ( ! [X153] : (s(bool,f(s(bool,'T_0'))) != s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X153) | s(bool,'T_0') = s(bool,X153)) ) | ($spl4 | $spl9 | $spl36 | $spl40 | $spl194)),
  inference(superposition,[],[f34914,f41609])).
fof(f42132,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') != s(bool,f(s(bool,'T_0')))) ) | ($spl4 | $spl9 | $spl36 | $spl40)),
  inference(condensation,[],[f42131])).
fof(f42131,plain,(
  ( ! [X149,X150] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(bool,X149) = s(bool,X150) | s(bool,'F_0') = s(bool,X150)) ) | ($spl4 | $spl9 | $spl36 | $spl40)),
  inference(subsumption_resolution,[],[f41796,f30026])).
fof(f41796,plain,(
  ( ! [X149,X150] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(bool,X149) = s(bool,X150) | s(bool,'F_0') = s(bool,X150) | s(bool,'T_0') = s(bool,X149)) ) | ($spl4 | $spl9 | $spl36 | $spl40)),
  inference(superposition,[],[f30026,f41609])).
fof(f42127,plain,(
  ( ! [X145] : (s(bool,X145) != s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X145)) ) | $spl234),
  inference(cnf_transformation,[],[f42127_D])).
fof(f42127_D,plain,(
  ( ! [X145] : (s(bool,X145) != s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X145)) ) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f41950,plain,(
  ( ! [X144,X145] : (s(bool,X145) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X144) | s(bool,'F_0') = s(bool,X145)) ) | ($spl4 | $spl9 | $spl36 | $spl40)),
  inference(duplicate_literal_removal,[],[f41793])).
fof(f41793,plain,(
  ( ! [X144,X145] : (s(bool,X145) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X144) | s(bool,'F_0') = s(bool,X145) | s(bool,'T_0') = s(bool,X144)) ) | ($spl4 | $spl9 | $spl36 | $spl40)),
  inference(superposition,[],[f29101,f41609])).
fof(f42119,plain,(
  ( ! [X132] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X132)) ) | ($spl4 | $spl9 | $spl36 | $spl40)),
  inference(forward_demodulation,[],[f41952,f11203])).
fof(f41952,plain,(
  ( ! [X132] : (s(bool,f(s(bool,'T_0'))) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X132)) ) | ($spl4 | $spl9 | $spl36 | $spl40)),
  inference(duplicate_literal_removal,[],[f41786])).
fof(f41786,plain,(
  ( ! [X132] : (s(bool,f(s(bool,'T_0'))) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X132) | s(bool,'T_0') = s(bool,X132)) ) | ($spl4 | $spl9 | $spl36 | $spl40)),
  inference(superposition,[],[f25786,f41609])).
fof(f42116,plain,(
  ( ! [X130] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X130)) ) | ($spl4 | $spl9 | $spl36 | $spl40)),
  inference(forward_demodulation,[],[f41953,f1085])).
fof(f41953,plain,(
  ( ! [X130] : (s(bool,f(s(bool,'T_0'))) != s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,'T_0') = s(bool,X130)) ) | ($spl4 | $spl9 | $spl36 | $spl40)),
  inference(duplicate_literal_removal,[],[f41784])).
fof(f41784,plain,(
  ( ! [X130] : (s(bool,f(s(bool,'T_0'))) != s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,'T_0') = s(bool,X130) | s(bool,'T_0') = s(bool,X130)) ) | ($spl4 | $spl9 | $spl36 | $spl40)),
  inference(superposition,[],[f24431,f41609])).
fof(f41955,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') != s(bool,f(s(bool,'T_0')))) ) | ($spl4 | $spl9 | $spl13 | $spl22 | $spl36 | $spl40)),
  inference(duplicate_literal_removal,[],[f41954])).
fof(f41954,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | ($spl4 | $spl9 | $spl13 | $spl22 | $spl36 | $spl40)),
  inference(condensation,[],[f41783])).
fof(f41783,plain,(
  ( ! [X127,X128,X129] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X128) | s(bool,X127) = s(bool,X129) | s(bool,X128) = s(bool,X129) | s(bool,'T_0') = s(bool,X127)) ) | ($spl4 | $spl9 | $spl13 | $spl22 | $spl36 | $spl40)),
  inference(superposition,[],[f22900,f41609])).
fof(f42114,plain,(
  ( ! [X116] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X116)) ) | ($spl4 | $spl9 | $spl13 | $spl22 | $spl36 | $spl40)),
  inference(subsumption_resolution,[],[f41778,f17356])).
fof(f41778,plain,(
  ( ! [X118,X116,X117] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X117,bool),'const.pair.,_2'(s(X117,X118),s(bool,X116))))) | s(bool,'T_0') = s(bool,X116)) ) | ($spl4 | $spl9 | $spl13 | $spl22 | $spl36 | $spl40)),
  inference(superposition,[],[f7282,f41609])).
fof(f40384,plain,(
  ( ! [X107,X110] : (s(bool,X110) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,f(s(bool,X107))),s(bool,X110))))) | s(bool,'F_0') = s(bool,X107) | s(bool,'T_0') = s(bool,X110)) ) | $spl232),
  inference(cnf_transformation,[],[f40384_D])).
fof(f40384_D,plain,(
  ( ! [X107,X110] : (s(bool,X110) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,f(s(bool,X107))),s(bool,X110))))) | s(bool,'F_0') = s(bool,X107) | s(bool,'T_0') = s(bool,X110)) ) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f40278,plain,(
  ( ! [X109,X107,X110,X108] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X108),'const.pair.,_2'(s(bool,'F_0'),s(X108,X109))))) | s(bool,X110) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,f(s(bool,X107))),s(bool,X110))))) | s(bool,'T_0') = s(bool,X110) | s(bool,'F_0') = s(bool,X107)) ) | ($spl16 | $spl204)),
  inference(superposition,[],[f1557,f37102])).
fof(f37101,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | $spl202),
  inference(cnf_transformation,[],[f37101_D])).
fof(f37101_D,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f38765,plain,(
  ( ! [X304,X305] : (s(bool,X305) != s(bool,g(s(bool,X304))) | s(bool,'F_0') = s(bool,X304) | s(bool,'F_0') = s(bool,X305)) ) | $spl230),
  inference(cnf_transformation,[],[f38765_D])).
fof(f38765_D,plain,(
  ( ! [X304,X305] : (s(bool,X305) != s(bool,g(s(bool,X304))) | s(bool,'F_0') = s(bool,X304) | s(bool,'F_0') = s(bool,X305)) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f36273,plain,(
  ( ! [X304,X305] : (s(bool,X305) != s(bool,g(s(bool,X304))) | s(bool,'F_0') = s(bool,X305) | s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X304)) ) | ($spl4 | $spl9 | $spl36 | $spl186)),
  inference(superposition,[],[f33724,f30029])).
fof(f38764,plain,(
  ( ! [X278] : (s(bool,X278) != s(bool,g(s(bool,X278))) | s(bool,'F_0') = s(bool,X278)) ) | $spl228),
  inference(cnf_transformation,[],[f38764_D])).
fof(f38764_D,plain,(
  ( ! [X278] : (s(bool,X278) != s(bool,g(s(bool,X278))) | s(bool,'F_0') = s(bool,X278)) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f36265,plain,(
  ( ! [X278] : (s(bool,X278) != s(bool,g(s(bool,X278))) | s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X278)) ) | ($spl4 | $spl9 | $spl36 | $spl47)),
  inference(superposition,[],[f3304,f30029])).
fof(f38763,plain,(
  ( ! [X271,X273,X275,X272,X274] : (s(X272,f(s(bool,X271))) != s(X272,g(s(bool,X271))) | s(bool,'F_0') = s(bool,X271) | s(bool,'T_0') = s(bool,X275) | s(bool,X275) = s(bool,sK6(X273,X274,X272,bool)) | s(fun(bool,X272),X273) = s(fun(bool,X272),X274)) ) | $spl226),
  inference(cnf_transformation,[],[f38763_D])).
fof(f38763_D,plain,(
  ( ! [X271,X273,X275,X272,X274] : (s(X272,f(s(bool,X271))) != s(X272,g(s(bool,X271))) | s(bool,'F_0') = s(bool,X271) | s(bool,'T_0') = s(bool,X275) | s(bool,X275) = s(bool,sK6(X273,X274,X272,bool)) | s(fun(bool,X272),X273) = s(fun(bool,X272),X274)) ) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f36263,plain,(
  ( ! [X271,X273,X275,X272,X274] : (s(X272,f(s(bool,X271))) != s(X272,g(s(bool,X271))) | s(fun(bool,X272),X273) = s(fun(bool,X272),X274) | s(bool,X275) = s(bool,sK6(X273,X274,X272,bool)) | s(bool,'T_0') = s(bool,X275) | s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X271)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f361,f30029])).
fof(f38717,plain,(
  $false | ($spl4 | $spl9 | $spl13 | $spl22 | $spl36 | $spl42 | $spl180)),
  inference(subsumption_resolution,[],[f38441,f38411])).
fof(f38411,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl4 | $spl9 | $spl36 | $spl42 | $spl180)),
  inference(subsumption_resolution,[],[f38410,f3890])).
fof(f38410,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'T_0'),s(bool,X0)))))) ) | ($spl4 | $spl9 | $spl36 | $spl42 | $spl180)),
  inference(condensation,[],[f38385])).
fof(f38385,plain,(
  ( ! [X64,X63] : (s(bool,X64) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'T_0'),s(bool,X64))))) | s(bool,'T_0') = s(bool,X63) | s(bool,'T_0') = s(bool,X64)) ) | ($spl4 | $spl9 | $spl36 | $spl42 | $spl180)),
  inference(backward_demodulation,[],[f38384,f3170])).
fof(f38384,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,f(s(bool,X0)))) ) | ($spl4 | $spl9 | $spl36 | $spl180)),
  inference(subsumption_resolution,[],[f38383,f71])).
fof(f38383,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,'F_0')) ) | ($spl4 | $spl9 | $spl36 | $spl180)),
  inference(condensation,[],[f38382])).
fof(f38382,plain,(
  ( ! [X107,X106] : (s(bool,'T_0') = s(bool,X107) | s(bool,'T_0') = s(bool,f(s(bool,X106))) | s(bool,'T_0') = s(bool,'F_0')) ) | ($spl4 | $spl9 | $spl36 | $spl180)),
  inference(subsumption_resolution,[],[f38271,f56])).
fof(f38271,plain,(
  ( ! [X107,X106] : (s(bool,X107) != s(bool,f(s(bool,X106))) | s(bool,'T_0') = s(bool,X107) | s(bool,'T_0') = s(bool,f(s(bool,X106))) | s(bool,'T_0') = s(bool,'F_0')) ) | ($spl4 | $spl9 | $spl36 | $spl180)),
  inference(superposition,[],[f33698,f30029])).
fof(f38441,plain,(
  s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'T_0'),s(bool,'T_0'))))) | ($spl4 | $spl9 | $spl13 | $spl22 | $spl36 | $spl42 | $spl180)),
  inference(backward_demodulation,[],[f38411,f6406])).
fof(f38715,plain,(
  $false | ($spl4 | $spl9 | $spl36 | $spl42 | $spl180)),
  inference(subsumption_resolution,[],[f38412,f37])).
fof(f38412,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl9 | $spl36 | $spl42 | $spl180)),
  inference(backward_demodulation,[],[f38411,f39])).
fof(f37596,plain,(
  ( ! [X196,X194,X192,X193] : (s(fun(X192,bool),X193) = s(fun(X192,bool),X194) | s(bool,'T_0') = s(bool,f(s(bool,f(s(X192,sK6(X193,X194,bool,X192)))))) | s(bool,'T_0') = s(bool,X196) | s(bool,X196) = s(bool,g(s(X192,sK6(X193,X194,bool,X192))))) ) | $spl224),
  inference(cnf_transformation,[],[f37596_D])).
fof(f37596_D,plain,(
  ( ! [X196,X194,X192,X193] : (s(fun(X192,bool),X193) = s(fun(X192,bool),X194) | s(bool,'T_0') = s(bool,f(s(bool,f(s(X192,sK6(X193,X194,bool,X192)))))) | s(bool,'T_0') = s(bool,X196) | s(bool,X196) = s(bool,g(s(X192,sK6(X193,X194,bool,X192))))) ) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f37595,plain,(
  ( ! [X196,X194,X192,X195,X193] : (s(fun(X192,bool),X193) = s(fun(X192,bool),X194) | s(bool,X196) = s(bool,g(s(X192,sK6(X193,X194,bool,X192)))) | s(bool,'T_0') = s(bool,X196) | s(bool,'T_0') = s(bool,f(s(bool,f(s(X192,sK6(X193,X194,bool,X192)))))) | s(bool,'F_0') = s(bool,X195)) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f37569,f40])).
fof(f37569,plain,(
  ( ! [X196,X194,X192,X195,X193] : (s(bool,'T_0') != s(bool,X195) | s(fun(X192,bool),X193) = s(fun(X192,bool),X194) | s(bool,X196) = s(bool,g(s(X192,sK6(X193,X194,bool,X192)))) | s(bool,'T_0') = s(bool,X196) | s(bool,'T_0') = s(bool,f(s(bool,f(s(X192,sK6(X193,X194,bool,X192)))))) | s(bool,'F_0') = s(bool,X195)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f366,f30029])).
fof(f37583,plain,(
  ( ! [X177,X180,X178] : (s(fun(fun(bool,bool),bool),X177) = s(fun(fun(bool,bool),bool),X178) | s(bool,'T_0') = s(bool,X180) | s(bool,X180) = s(bool,g(s(fun(bool,bool),sK6(X177,X178,bool,fun(bool,bool)))))) ) | $spl220),
  inference(cnf_transformation,[],[f37583_D])).
fof(f37583_D,plain,(
  ( ! [X177,X180,X178] : (s(fun(fun(bool,bool),bool),X177) = s(fun(fun(bool,bool),bool),X178) | s(bool,'T_0') = s(bool,X180) | s(bool,X180) = s(bool,g(s(fun(bool,bool),sK6(X177,X178,bool,fun(bool,bool)))))) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f37584,plain,(
  ( ! [X179] : (s(bool,'T_0') != s(bool,f(s(fun(bool,bool),X179)))) ) | $spl222),
  inference(cnf_transformation,[],[f37584_D])).
fof(f37584_D,plain,(
  ( ! [X179] : (s(bool,'T_0') != s(bool,f(s(fun(bool,bool),X179)))) ) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f37535,plain,(
  ( ! [X182,X184,X183,X181] : (s(bool,'T_0') != s(bool,f(s(fun(bool,bool),X183))) | s(fun(fun(bool,bool),bool),X181) = s(fun(fun(bool,bool),bool),X182) | s(bool,X184) = s(bool,g(s(fun(bool,bool),sK6(X181,X182,bool,fun(bool,bool))))) | s(bool,'T_0') = s(bool,X184)) ) | $spl44),
  inference(superposition,[],[f366,f3302])).
fof(f37534,plain,(
  ( ! [X177,X180,X178,X179] : (s(bool,'T_0') != s(bool,f(s(fun(bool,bool),X179))) | s(fun(fun(bool,bool),bool),X177) = s(fun(fun(bool,bool),bool),X178) | s(bool,X180) = s(bool,g(s(fun(bool,bool),sK6(X177,X178,bool,fun(bool,bool))))) | s(bool,'T_0') = s(bool,X180)) ) | $spl44),
  inference(superposition,[],[f366,f3302])).
fof(f37240,plain,(
  $false | ($spl9 | $spl40 | $spl202)),
  inference(subsumption_resolution,[],[f37239,f71])).
fof(f37239,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl40 | $spl202)),
  inference(backward_demodulation,[],[f37101,f2561])).
fof(f37182,plain,(
  ( ! [X240] : (s(bool,X240) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X240)))) ) | $spl218),
  inference(cnf_transformation,[],[f37182_D])).
fof(f37182_D,plain,(
  ( ! [X240] : (s(bool,X240) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X240)))) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f36742,plain,(
  ( ! [X239,X240] : (s(bool,X240) != s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X239) | s(bool,'T_0') = s(bool,f(s(bool,X240)))) ) | ($spl4 | $spl9 | $spl36 | $spl186)),
  inference(duplicate_literal_removal,[],[f36411])).
fof(f36411,plain,(
  ( ! [X239,X240] : (s(bool,X240) != s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X239) | s(bool,'T_0') = s(bool,f(s(bool,X240))) | s(bool,'F_0') = s(bool,X239)) ) | ($spl4 | $spl9 | $spl36 | $spl186)),
  inference(superposition,[],[f33724,f30029])).
fof(f37181,plain,(
  ( ! [X222,X223] : (s(bool,X222) != s(bool,X223) | s(bool,'T_0') = s(bool,f(s(bool,X222))) | s(bool,'T_0') = s(bool,X223)) ) | $spl216),
  inference(cnf_transformation,[],[f37181_D])).
fof(f37181_D,plain,(
  ( ! [X222,X223] : (s(bool,X222) != s(bool,X223) | s(bool,'T_0') = s(bool,f(s(bool,X222))) | s(bool,'T_0') = s(bool,X223)) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f36746,plain,(
  ( ! [X222,X223,X221] : (s(bool,X222) != s(bool,X223) | s(bool,'T_0') = s(bool,X223) | s(bool,'F_0') = s(bool,X221) | s(bool,'T_0') = s(bool,f(s(bool,X222)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(duplicate_literal_removal,[],[f36404])).
fof(f36404,plain,(
  ( ! [X222,X223,X221] : (s(bool,X222) != s(bool,X223) | s(bool,'T_0') = s(bool,X223) | s(bool,'F_0') = s(bool,X221) | s(bool,'T_0') = s(bool,f(s(bool,X222))) | s(bool,'F_0') = s(bool,X221)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f29101,f30029])).
fof(f37134,plain,(
  ( ! [X383,X377,X379,X382,X378,X381] : (~p(s(bool,happ(s(fun(X381,bool),X382),s(X381,X383)))) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(X377,bool),X378),s(X377,X379))))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X381,bool),X382),s(fun(X377,bool),X378),s('type.pair.prod'(X381,X377),'const.pair.,_2'(s(X381,X383),s(X377,X379))))))) ) | $spl214),
  inference(cnf_transformation,[],[f37134_D])).
fof(f37134_D,plain,(
  ( ! [X383,X377,X379,X382,X378,X381] : (~p(s(bool,happ(s(fun(X381,bool),X382),s(X381,X383)))) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(X377,bool),X378),s(X377,X379))))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X381,bool),X382),s(fun(X377,bool),X378),s('type.pair.prod'(X381,X377),'const.pair.,_2'(s(X381,X383),s(X377,X379))))))) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f37133,plain,(
  ( ! [X383,X377,X379,X380,X382,X378,X381] : (~p(s(bool,happ(s(fun(X381,bool),X382),s(X381,X383)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X381,bool),X382),s(fun(X377,bool),X378),s('type.pair.prod'(X381,X377),'const.pair.,_2'(s(X381,X383),s(X377,X379)))))) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(X377,bool),X378),s(X377,X379))))) | s(bool,'F_0') = s(bool,X380)) ) | ($spl4 | $spl9 | $spl36 | $spl172)),
  inference(subsumption_resolution,[],[f36301,f29278])).
fof(f36301,plain,(
  ( ! [X383,X377,X379,X380,X382,X378,X381] : (~p(s(bool,X380)) | ~p(s(bool,happ(s(fun(X381,bool),X382),s(X381,X383)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X381,bool),X382),s(fun(X377,bool),X378),s('type.pair.prod'(X381,X377),'const.pair.,_2'(s(X381,X383),s(X377,X379)))))) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(X377,bool),X378),s(X377,X379))))) | s(bool,'F_0') = s(bool,X380)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f785,f30029])).
fof(f37132,plain,(
  ( ! [X372,X374,X376,X373,X371] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X373,bool),X376),s(fun(X371,bool),X372),s('type.pair.prod'(X373,X371),X374)))) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(X371,bool),X372),s(X371,'const.pair.SND_1'(s('type.pair.prod'(X373,X371),X374))))))) | ~p(s(bool,happ(s(fun(X373,bool),X376),s(X373,'const.pair.FST_1'(s('type.pair.prod'(X373,X371),X374))))))) ) | $spl212),
  inference(cnf_transformation,[],[f37132_D])).
fof(f37132_D,plain,(
  ( ! [X372,X374,X376,X373,X371] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X373,bool),X376),s(fun(X371,bool),X372),s('type.pair.prod'(X373,X371),X374)))) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(X371,bool),X372),s(X371,'const.pair.SND_1'(s('type.pair.prod'(X373,X371),X374))))))) | ~p(s(bool,happ(s(fun(X373,bool),X376),s(X373,'const.pair.FST_1'(s('type.pair.prod'(X373,X371),X374))))))) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f37131,plain,(
  ( ! [X372,X374,X376,X373,X375,X371] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X373,bool),X376),s(fun(X371,bool),X372),s('type.pair.prod'(X373,X371),X374)))) | ~p(s(bool,happ(s(fun(X373,bool),X376),s(X373,'const.pair.FST_1'(s('type.pair.prod'(X373,X371),X374)))))) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(X371,bool),X372),s(X371,'const.pair.SND_1'(s('type.pair.prod'(X373,X371),X374))))))) | s(bool,'F_0') = s(bool,X375)) ) | ($spl4 | $spl9 | $spl36 | $spl172)),
  inference(subsumption_resolution,[],[f36300,f29278])).
fof(f36300,plain,(
  ( ! [X372,X374,X376,X373,X375,X371] : (~p(s(bool,X375)) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X373,bool),X376),s(fun(X371,bool),X372),s('type.pair.prod'(X373,X371),X374)))) | ~p(s(bool,happ(s(fun(X373,bool),X376),s(X373,'const.pair.FST_1'(s('type.pair.prod'(X373,X371),X374)))))) | s(bool,'T_0') = s(bool,f(s(bool,happ(s(fun(X371,bool),X372),s(X371,'const.pair.SND_1'(s('type.pair.prod'(X373,X371),X374))))))) | s(bool,'F_0') = s(bool,X375)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f48,f30029])).
fof(f37130,plain,(
  ( ! [X366,X368,X365,X367,X369] : (p(s(bool,happ(s(fun(X365,bool),X366),s(X365,'const.pair.FST_1'(s('type.pair.prod'(X365,X367),X369)))))) | s(bool,'T_0') = s(bool,f(s(bool,'const.pair.PROD_ALL_3'(s(fun(X365,bool),X366),s(fun(X367,bool),X368),s('type.pair.prod'(X365,X367),X369)))))) ) | $spl210),
  inference(cnf_transformation,[],[f37130_D])).
fof(f37130_D,plain,(
  ( ! [X366,X368,X365,X367,X369] : (p(s(bool,happ(s(fun(X365,bool),X366),s(X365,'const.pair.FST_1'(s('type.pair.prod'(X365,X367),X369)))))) | s(bool,'T_0') = s(bool,f(s(bool,'const.pair.PROD_ALL_3'(s(fun(X365,bool),X366),s(fun(X367,bool),X368),s('type.pair.prod'(X365,X367),X369)))))) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f37129,plain,(
  ( ! [X366,X368,X370,X365,X367,X369] : (p(s(bool,happ(s(fun(X365,bool),X366),s(X365,'const.pair.FST_1'(s('type.pair.prod'(X365,X367),X369)))))) | s(bool,'T_0') = s(bool,f(s(bool,'const.pair.PROD_ALL_3'(s(fun(X365,bool),X366),s(fun(X367,bool),X368),s('type.pair.prod'(X365,X367),X369))))) | s(bool,'F_0') = s(bool,X370)) ) | ($spl4 | $spl9 | $spl36 | $spl172)),
  inference(subsumption_resolution,[],[f36299,f29278])).
fof(f36299,plain,(
  ( ! [X366,X368,X370,X365,X367,X369] : (~p(s(bool,X370)) | p(s(bool,happ(s(fun(X365,bool),X366),s(X365,'const.pair.FST_1'(s('type.pair.prod'(X365,X367),X369)))))) | s(bool,'T_0') = s(bool,f(s(bool,'const.pair.PROD_ALL_3'(s(fun(X365,bool),X366),s(fun(X367,bool),X368),s('type.pair.prod'(X365,X367),X369))))) | s(bool,'F_0') = s(bool,X370)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f46,f30029])).
fof(f37128,plain,(
  ( ! [X360,X362,X359,X361,X363] : (p(s(bool,happ(s(fun(X361,bool),X362),s(X361,'const.pair.SND_1'(s('type.pair.prod'(X359,X361),X363)))))) | s(bool,'T_0') = s(bool,f(s(bool,'const.pair.PROD_ALL_3'(s(fun(X359,bool),X360),s(fun(X361,bool),X362),s('type.pair.prod'(X359,X361),X363)))))) ) | $spl208),
  inference(cnf_transformation,[],[f37128_D])).
fof(f37128_D,plain,(
  ( ! [X360,X362,X359,X361,X363] : (p(s(bool,happ(s(fun(X361,bool),X362),s(X361,'const.pair.SND_1'(s('type.pair.prod'(X359,X361),X363)))))) | s(bool,'T_0') = s(bool,f(s(bool,'const.pair.PROD_ALL_3'(s(fun(X359,bool),X360),s(fun(X361,bool),X362),s('type.pair.prod'(X359,X361),X363)))))) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f37127,plain,(
  ( ! [X364,X360,X362,X359,X361,X363] : (p(s(bool,happ(s(fun(X361,bool),X362),s(X361,'const.pair.SND_1'(s('type.pair.prod'(X359,X361),X363)))))) | s(bool,'T_0') = s(bool,f(s(bool,'const.pair.PROD_ALL_3'(s(fun(X359,bool),X360),s(fun(X361,bool),X362),s('type.pair.prod'(X359,X361),X363))))) | s(bool,'F_0') = s(bool,X364)) ) | ($spl4 | $spl9 | $spl36 | $spl172)),
  inference(subsumption_resolution,[],[f36298,f29278])).
fof(f36298,plain,(
  ( ! [X364,X360,X362,X359,X361,X363] : (~p(s(bool,X364)) | p(s(bool,happ(s(fun(X361,bool),X362),s(X361,'const.pair.SND_1'(s('type.pair.prod'(X359,X361),X363)))))) | s(bool,'T_0') = s(bool,f(s(bool,'const.pair.PROD_ALL_3'(s(fun(X359,bool),X360),s(fun(X361,bool),X362),s('type.pair.prod'(X359,X361),X363))))) | s(bool,'F_0') = s(bool,X364)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f47,f30029])).
fof(f37104,plain,(
  ( ! [X302,X301,X303] : (s(bool,X301) != s(bool,X302) | s(bool,'F_0') = s(bool,X301) | s(bool,'F_0') = s(bool,X303) | s(bool,X302) = s(bool,X303)) ) | $spl206),
  inference(cnf_transformation,[],[f37104_D])).
fof(f37104_D,plain,(
  ( ! [X302,X301,X303] : (s(bool,X301) != s(bool,X302) | s(bool,'F_0') = s(bool,X301) | s(bool,'F_0') = s(bool,X303) | s(bool,X302) = s(bool,X303)) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f36272,plain,(
  ( ! [X302,X301,X303] : (s(bool,X301) != s(bool,X302) | s(bool,X302) = s(bool,X303) | s(bool,'F_0') = s(bool,X303) | s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X301)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30026,f30029])).
fof(f37102,plain,(
  ( ! [X277] : (s(bool,'F_0') = s(bool,f(s(bool,X277))) | s(bool,'F_0') = s(bool,X277)) ) | $spl204),
  inference(cnf_transformation,[],[f37102_D])).
fof(f37102_D,plain,(
  ( ! [X277] : (s(bool,'F_0') = s(bool,f(s(bool,X277))) | s(bool,'F_0') = s(bool,X277)) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f36264,plain,(
  ( ! [X277] : (s(bool,'F_0') = s(bool,f(s(bool,X277))) | s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X277)) ) | ($spl4 | $spl9 | $spl36 | $spl40)),
  inference(superposition,[],[f2561,f30029])).
fof(f37017,plain,(
  ( ! [X136] : (s(bool,X136) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3))) | s(bool,'T_0') = s(bool,f(s(bool,X136)))) ) | $spl200),
  inference(cnf_transformation,[],[f37017_D])).
fof(f37017_D,plain,(
  ( ! [X136] : (s(bool,X136) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3))) | s(bool,'T_0') = s(bool,f(s(bool,X136)))) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f36819,plain,(
  ( ! [X156,X157] : (s(bool,'F_0') = s(bool,X156) | s(bool,'T_0') = s(bool,f(s(bool,X157))) | s(bool,X157) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) ) | ($spl2 | $spl4 | $spl9 | $spl36)),
  inference(duplicate_literal_removal,[],[f36096])).
fof(f36096,plain,(
  ( ! [X156,X157] : (s(bool,'F_0') = s(bool,X156) | s(bool,'T_0') = s(bool,f(s(bool,X157))) | s(bool,X157) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3))) | s(bool,'F_0') = s(bool,X156)) ) | ($spl2 | $spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30029,f32150])).
fof(f37016,plain,(
  ( ! [X134] : (s(bool,X134) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X134)))) ) | $spl198),
  inference(cnf_transformation,[],[f37016_D])).
fof(f37016_D,plain,(
  ( ! [X134] : (s(bool,X134) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X134)))) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f36820,plain,(
  ( ! [X142,X141] : (s(bool,'F_0') = s(bool,X141) | s(bool,'T_0') = s(bool,f(s(bool,X142))) | s(bool,X142) = s(bool,f(s(bool,'F_0')))) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(duplicate_literal_removal,[],[f36090])).
fof(f36090,plain,(
  ( ! [X142,X141] : (s(bool,'F_0') = s(bool,X141) | s(bool,'T_0') = s(bool,f(s(bool,X142))) | s(bool,X142) = s(bool,f(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X141)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(superposition,[],[f30029,f30750])).
fof(f36937,plain,(
  ( ! [X166,X168] : (s(bool,'T_0') = s(bool,f(s(bool,X166))) | s(bool,X166) = s(bool,X168) | s(bool,'T_0') = s(bool,f(s(bool,X168)))) ) | $spl196),
  inference(cnf_transformation,[],[f36937_D])).
fof(f36937_D,plain,(
  ( ! [X166,X168] : (s(bool,'T_0') = s(bool,f(s(bool,X166))) | s(bool,X166) = s(bool,X168) | s(bool,'T_0') = s(bool,f(s(bool,X168)))) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f36821,plain,(
  ( ! [X121,X120,X119] : (s(bool,'F_0') = s(bool,X120) | s(bool,'T_0') = s(bool,f(s(bool,X121))) | s(bool,X119) = s(bool,X121) | s(bool,'T_0') = s(bool,f(s(bool,X119)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(duplicate_literal_removal,[],[f36086])).
fof(f36086,plain,(
  ( ! [X121,X120,X119] : (s(bool,'F_0') = s(bool,X120) | s(bool,'T_0') = s(bool,f(s(bool,X121))) | s(bool,X119) = s(bool,X121) | s(bool,'T_0') = s(bool,f(s(bool,X119))) | s(bool,'F_0') = s(bool,X120)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30029,f30029])).
fof(f36831,plain,(
  ( ! [X136,X135] : (s(bool,X136) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3))) | s(bool,'T_0') = s(bool,f(s(bool,X136))) | s(bool,'F_0') = s(bool,X135)) ) | ($spl2 | $spl4 | $spl9 | $spl36)),
  inference(duplicate_literal_removal,[],[f35985])).
fof(f35985,plain,(
  ( ! [X136,X135] : (s(bool,X136) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3))) | s(bool,'T_0') = s(bool,f(s(bool,X136))) | s(bool,'F_0') = s(bool,X135) | s(bool,'F_0') = s(bool,X135)) ) | ($spl2 | $spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30029,f32150])).
fof(f36832,plain,(
  ( ! [X134,X133] : (s(bool,X134) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X134))) | s(bool,'F_0') = s(bool,X133)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(duplicate_literal_removal,[],[f35984])).
fof(f35984,plain,(
  ( ! [X134,X133] : (s(bool,X134) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X134))) | s(bool,'F_0') = s(bool,X133) | s(bool,'F_0') = s(bool,X133)) ) | ($spl4 | $spl9 | $spl36 | $spl176)),
  inference(superposition,[],[f30029,f30750])).
fof(f36833,plain,(
  ( ! [X127,X128,X129] : (s(bool,X128) = s(bool,X129) | s(bool,'T_0') = s(bool,f(s(bool,X129))) | s(bool,'F_0') = s(bool,X127) | s(bool,'T_0') = s(bool,f(s(bool,X128)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(duplicate_literal_removal,[],[f35982])).
fof(f35982,plain,(
  ( ! [X127,X128,X129] : (s(bool,X128) = s(bool,X129) | s(bool,'T_0') = s(bool,f(s(bool,X129))) | s(bool,'F_0') = s(bool,X127) | s(bool,'T_0') = s(bool,f(s(bool,X128))) | s(bool,'F_0') = s(bool,X127)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30029,f30029])).
fof(f36851,plain,(
  ( ! [X132,X130,X131] : (s(bool,'F_0') = s(bool,X131) | s(bool,'T_0') = s(bool,f(s(bool,X130))) | s(bool,'T_0') = s(bool,f(s(bool,X132))) | s(bool,X130) = s(bool,X132)) ) | ($spl4 | $spl9 | $spl36)),
  inference(duplicate_literal_removal,[],[f35908])).
fof(f35908,plain,(
  ( ! [X132,X130,X131] : (s(bool,'F_0') = s(bool,X131) | s(bool,'T_0') = s(bool,f(s(bool,X130))) | s(bool,'F_0') = s(bool,X131) | s(bool,'T_0') = s(bool,f(s(bool,X132))) | s(bool,X130) = s(bool,X132)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30029,f30029])).
fof(f36886,plain,(
  ( ! [X167,X166,X168] : (s(bool,'T_0') = s(bool,f(s(bool,X166))) | s(bool,'F_0') = s(bool,X167) | s(bool,'T_0') = s(bool,f(s(bool,X168))) | s(bool,X166) = s(bool,X168)) ) | ($spl4 | $spl9 | $spl36)),
  inference(duplicate_literal_removal,[],[f35811])).
fof(f35811,plain,(
  ( ! [X167,X166,X168] : (s(bool,'T_0') = s(bool,f(s(bool,X166))) | s(bool,'F_0') = s(bool,X167) | s(bool,'F_0') = s(bool,X167) | s(bool,'T_0') = s(bool,f(s(bool,X168))) | s(bool,X166) = s(bool,X168)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f30029,f30029])).
fof(f25243,plain,(
  ( ! [X243,X242] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X243),s(bool,X242))) | s(bool,X242) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | $spl136),
  inference(cnf_transformation,[],[f25243_D])).
fof(f25243_D,plain,(
  ( ! [X243,X242] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X243),s(bool,X242))) | s(bool,X242) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f35559,plain,(
  ( ! [X103,X102,X104] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X104),s(bool,X103))) | s(bool,'T_0') = s(bool,X102) | s(bool,X103) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl116)),
  inference(duplicate_literal_removal,[],[f35469])).
fof(f35469,plain,(
  ( ! [X103,X102,X104] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X104),s(bool,X103))) | s(bool,'T_0') = s(bool,X102) | s(bool,'T_0') = s(bool,X102) | s(bool,X103) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl116)),
  inference(superposition,[],[f9594,f22295])).
fof(f34914,plain,(
  ( ! [X161] : (s(bool,X161) != s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X161)) ) | $spl194),
  inference(cnf_transformation,[],[f34914_D])).
fof(f34914_D,plain,(
  ( ! [X161] : (s(bool,X161) != s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X161)) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f34802,plain,(
  ( ! [X161,X160] : (s(bool,X161) != s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X161) | s(bool,'F_0') = s(bool,X160)) ) | ($spl2 | $spl4 | $spl9 | $spl36)),
  inference(duplicate_literal_removal,[],[f34658])).
fof(f34658,plain,(
  ( ! [X161,X160] : (s(bool,X161) != s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X161) | s(bool,'F_0') = s(bool,X160) | s(bool,'F_0') = s(bool,X160)) ) | ($spl2 | $spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f29101,f32150])).
fof(f34481,plain,(
  ( ! [X127,X126] : (s(fun(fun(bool,bool),bool),X126) = s(fun(fun(bool,bool),bool),X127) | s(bool,'T_0') = s(bool,f(s(fun(bool,bool),sK6(X126,X127,bool,fun(bool,bool)))))) ) | $spl190),
  inference(cnf_transformation,[],[f34481_D])).
fof(f34481_D,plain,(
  ( ! [X127,X126] : (s(fun(fun(bool,bool),bool),X126) = s(fun(fun(bool,bool),bool),X127) | s(bool,'T_0') = s(bool,f(s(fun(bool,bool),sK6(X126,X127,bool,fun(bool,bool)))))) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f34482,plain,(
  ( ! [X128] : (s(bool,'T_0') = s(bool,g(s(fun(bool,bool),X128)))) ) | $spl192),
  inference(cnf_transformation,[],[f34482_D])).
fof(f34482_D,plain,(
  ( ! [X128] : (s(bool,'T_0') = s(bool,g(s(fun(bool,bool),X128)))) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f34331,plain,(
  ( ! [X130,X131,X129] : (s(bool,'T_0') = s(bool,g(s(fun(bool,bool),X131))) | s(fun(fun(bool,bool),bool),X129) = s(fun(fun(bool,bool),bool),X130) | s(bool,'T_0') = s(bool,f(s(fun(bool,bool),sK6(X129,X130,bool,fun(bool,bool)))))) ) | $spl44),
  inference(superposition,[],[f2559,f3302])).
fof(f34330,plain,(
  ( ! [X127,X128,X126] : (s(bool,'T_0') = s(bool,g(s(fun(bool,bool),X128))) | s(fun(fun(bool,bool),bool),X126) = s(fun(fun(bool,bool),bool),X127) | s(bool,'T_0') = s(bool,f(s(fun(bool,bool),sK6(X126,X127,bool,fun(bool,bool)))))) ) | $spl44),
  inference(superposition,[],[f2559,f3302])).
fof(f34071,plain,(
  ( ! [X2,X1] : (s(bool,X1) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2) | s(bool,X1) = s(bool,X2)) ) | $spl188),
  inference(cnf_transformation,[],[f34071_D])).
fof(f34071_D,plain,(
  ( ! [X2,X1] : (s(bool,X1) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2) | s(bool,X1) = s(bool,X2)) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f34073,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,X1) != s(bool,g(s(bool,'T_0'))) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl4 | $spl9 | $spl36 | $spl186)),
  inference(condensation,[],[f34072])).
fof(f34072,plain,(
  ( ! [X39,X37,X38,X40] : (s(bool,X38) != s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X37) | s(bool,X37) = s(bool,X39) | s(bool,X38) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) ) | ($spl4 | $spl9 | $spl36 | $spl186)),
  inference(subsumption_resolution,[],[f33931,f30026])).
fof(f33931,plain,(
  ( ! [X39,X37,X38,X40] : (s(bool,X38) != s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X37) | s(bool,X37) = s(bool,X39) | s(bool,'T_0') = s(bool,X39) | s(bool,X38) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) ) | $spl186),
  inference(superposition,[],[f33724,f146])).
fof(f34070,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,X1) != s(bool,g(s(bool,'T_0'))) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl4 | $spl9 | $spl36 | $spl186)),
  inference(condensation,[],[f34069])).
fof(f34069,plain,(
  ( ! [X35,X33,X36,X34] : (s(bool,X34) != s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X33) | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35) | s(bool,X33) = s(bool,X36)) ) | ($spl4 | $spl9 | $spl36 | $spl186)),
  inference(subsumption_resolution,[],[f33930,f30026])).
fof(f33930,plain,(
  ( ! [X35,X33,X36,X34] : (s(bool,X34) != s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X33) | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35) | s(bool,X33) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) ) | $spl186),
  inference(superposition,[],[f33724,f146])).
fof(f33724,plain,(
  ( ! [X111] : (s(bool,X111) != s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X111)) ) | $spl186),
  inference(cnf_transformation,[],[f33724_D])).
fof(f33724_D,plain,(
  ( ! [X111] : (s(bool,X111) != s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X111)) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f33623,plain,(
  ( ! [X111,X110] : (s(bool,X111) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X110) | s(bool,'F_0') = s(bool,X111)) ) | ($spl4 | $spl9 | $spl36 | $spl70)),
  inference(duplicate_literal_removal,[],[f33580])).
fof(f33580,plain,(
  ( ! [X111,X110] : (s(bool,X111) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X110) | s(bool,'F_0') = s(bool,X111) | s(bool,'T_0') = s(bool,X110)) ) | ($spl4 | $spl9 | $spl36 | $spl70)),
  inference(superposition,[],[f29101,f26087])).
fof(f26087,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl9 | $spl70)),
  inference(subsumption_resolution,[],[f26009,f71])).
fof(f26009,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl70),
  inference(superposition,[],[f25908,f56])).
fof(f25908,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | $spl70),
  inference(unit_resulting_resolution,[],[f37,f4928])).
fof(f33721,plain,(
  ( ! [X103,X104] : (s(bool,X103) != s(bool,X104) | s(bool,X103) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,'F_0') = s(bool,X104)) ) | $spl184),
  inference(cnf_transformation,[],[f33721_D])).
fof(f33721_D,plain,(
  ( ! [X103,X104] : (s(bool,X103) != s(bool,X104) | s(bool,X103) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,'F_0') = s(bool,X104)) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f33624,plain,(
  ( ! [X103,X102,X104] : (s(bool,X103) != s(bool,X104) | s(bool,'T_0') = s(bool,X102) | s(bool,'F_0') = s(bool,X104) | s(bool,X103) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(duplicate_literal_removal,[],[f33577])).
fof(f33577,plain,(
  ( ! [X103,X102,X104] : (s(bool,X103) != s(bool,X104) | s(bool,'T_0') = s(bool,X102) | s(bool,'F_0') = s(bool,X104) | s(bool,'T_0') = s(bool,X102) | s(bool,X103) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f29101,f22295])).
fof(f33720,plain,(
  ( ! [X95] : (s(bool,X95) != s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2))))) | s(bool,'F_0') = s(bool,X95)) ) | $spl182),
  inference(cnf_transformation,[],[f33720_D])).
fof(f33720_D,plain,(
  ( ! [X95] : (s(bool,X95) != s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2))))) | s(bool,'F_0') = s(bool,X95)) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f33625,plain,(
  ( ! [X94,X95] : (s(bool,X95) != s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2))))) | s(bool,'T_0') = s(bool,X94) | s(bool,'F_0') = s(bool,X95)) ) | ($spl1 | $spl4 | $spl9 | $spl36)),
  inference(duplicate_literal_removal,[],[f33574])).
fof(f33574,plain,(
  ( ! [X94,X95] : (s(bool,X95) != s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2))))) | s(bool,'T_0') = s(bool,X94) | s(bool,'F_0') = s(bool,X95) | s(bool,'T_0') = s(bool,X94)) ) | ($spl1 | $spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f29101,f1101])).
fof(f33698,plain,(
  ( ! [X127] : (s(bool,X127) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X127)) ) | $spl180),
  inference(cnf_transformation,[],[f33698_D])).
fof(f33698_D,plain,(
  ( ! [X127] : (s(bool,X127) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X127)) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f33651,plain,(
  ( ! [X127,X126] : (s(bool,X127) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X127) | s(bool,'F_0') = s(bool,X126)) ) | ($spl4 | $spl9 | $spl36 | $spl41)),
  inference(duplicate_literal_removal,[],[f33518])).
fof(f33518,plain,(
  ( ! [X127,X126] : (s(bool,X127) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X127) | s(bool,'F_0') = s(bool,X126) | s(bool,'F_0') = s(bool,X126)) ) | ($spl4 | $spl9 | $spl36 | $spl41)),
  inference(superposition,[],[f29101,f32149])).
fof(f32149,plain,(
  ( ! [X166] : (s(bool,X166) = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X166)) ) | ($spl4 | $spl9 | $spl36 | $spl41)),
  inference(subsumption_resolution,[],[f32061,f37])).
fof(f32061,plain,(
  ( ! [X166] : (~p(s(bool,'T_0')) | s(bool,X166) = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X166)) ) | ($spl4 | $spl9 | $spl36 | $spl41)),
  inference(superposition,[],[f30024,f27394])).
fof(f27394,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | ($spl9 | $spl41)),
  inference(unit_resulting_resolution,[],[f2562,f71,f2562,f2562,f146])).
fof(f31441,plain,(
  ( ! [X308,X304,X307] : (s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X307,bool),'const.pair.,_2'(s(X307,X308),s(bool,g(s(bool,X304))))))) | s(bool,'F_0') = s(bool,X304)) ) | $spl178),
  inference(cnf_transformation,[],[f31441_D])).
fof(f31441_D,plain,(
  ( ! [X308,X304,X307] : (s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X307,bool),'const.pair.,_2'(s(X307,X308),s(bool,g(s(bool,X304))))))) | s(bool,'F_0') = s(bool,X304)) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f30854,plain,(
  ( ! [X308,X304,X306,X305,X307] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X305),'const.pair.,_2'(s(bool,'F_0'),s(X305,X306))))) | s(bool,'F_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X307,bool),'const.pair.,_2'(s(X307,X308),s(bool,g(s(bool,X304))))))) | s(bool,'F_0') = s(bool,X304)) ) | $spl54),
  inference(superposition,[],[f1473,f4883])).
fof(f30750,plain,(
  ( ! [X199] : (s(bool,X199) = s(bool,f(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X199)) ) | $spl176),
  inference(cnf_transformation,[],[f30750_D])).
fof(f30750_D,plain,(
  ( ! [X199] : (s(bool,X199) = s(bool,f(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X199)) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f30802,plain,(
  ( ! [X198,X199] : (s(bool,'T_0') = s(bool,X198) | s(bool,X199) = s(bool,f(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X199)) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f30295,f40])).
fof(f30295,plain,(
  ( ! [X198,X199] : (s(bool,'F_0') != s(bool,X198) | s(bool,'T_0') = s(bool,X198) | s(bool,X199) = s(bool,f(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X199)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f25786,f28992])).
fof(f30715,plain,(
  ( ! [X151,X147,X148] : (s(bool,X148) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X147),s(bool,X148))))) | s(bool,'F_0') = s(bool,X151) | s(bool,X147) = s(bool,X151) | s(bool,'T_0') = s(bool,X148)) ) | $spl174),
  inference(cnf_transformation,[],[f30715_D])).
fof(f30715_D,plain,(
  ( ! [X151,X147,X148] : (s(bool,X148) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X147),s(bool,X148))))) | s(bool,'F_0') = s(bool,X151) | s(bool,X147) = s(bool,X151) | s(bool,'T_0') = s(bool,X148)) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f30268,plain,(
  ( ! [X134,X132,X130,X133,X131] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X131),'const.pair.,_2'(s(bool,'F_0'),s(X131,X132))))) | s(bool,X133) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X130),s(bool,X133))))) | s(bool,'T_0') = s(bool,X133) | s(bool,X130) = s(bool,X134) | s(bool,'F_0') = s(bool,X134)) ) | ($spl4 | $spl9 | $spl16 | $spl36)),
  inference(superposition,[],[f1557,f28992])).
fof(f30213,plain,(
  ( ! [X151,X149,X147,X150,X148] : (s(bool,X148) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X147),s(bool,X148))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X149),'const.pair.,_2'(s(bool,'F_0'),s(X149,X150))))) | s(bool,'T_0') = s(bool,X148) | s(bool,X147) = s(bool,X151) | s(bool,'F_0') = s(bool,X151)) ) | ($spl4 | $spl9 | $spl16 | $spl36)),
  inference(superposition,[],[f1557,f28992])).
fof(f30749,plain,(
  ( ! [X198,X199] : (s(bool,'T_0') = s(bool,X198) | s(bool,X199) = s(bool,f(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X199)) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f30138,f40])).
fof(f30138,plain,(
  ( ! [X198,X199] : (s(bool,'F_0') != s(bool,X198) | s(bool,'T_0') = s(bool,X198) | s(bool,X199) = s(bool,f(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X199)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f25786,f28992])).
fof(f30111,plain,(
  ( ! [X134,X132,X130,X133,X131] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X131),'const.pair.,_2'(s(bool,'F_0'),s(X131,X132))))) | s(bool,X133) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X130),s(bool,X133))))) | s(bool,'T_0') = s(bool,X133) | s(bool,X130) = s(bool,X134) | s(bool,'F_0') = s(bool,X134)) ) | ($spl4 | $spl9 | $spl16 | $spl36)),
  inference(superposition,[],[f1557,f28992])).
fof(f30056,plain,(
  ( ! [X151,X149,X147,X150,X148] : (s(bool,X148) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X147),s(bool,X148))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X149),'const.pair.,_2'(s(bool,'F_0'),s(X149,X150))))) | s(bool,'T_0') = s(bool,X148) | s(bool,X147) = s(bool,X151) | s(bool,'F_0') = s(bool,X151)) ) | ($spl4 | $spl9 | $spl16 | $spl36)),
  inference(superposition,[],[f1557,f28992])).
fof(f29278,plain,(
  ( ! [X100] : (p(s(bool,X100)) | s(bool,'F_0') = s(bool,X100)) ) | $spl172),
  inference(cnf_transformation,[],[f29278_D])).
fof(f29278_D,plain,(
  ( ! [X100] : (p(s(bool,X100)) | s(bool,'F_0') = s(bool,X100)) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f29277,plain,(
  ( ! [X101,X100] : (s(bool,'T_0') = s(bool,X101) | p(s(bool,X100)) | s(bool,'F_0') = s(bool,X100)) ) | ($spl4 | $spl9 | $spl36)),
  inference(forward_demodulation,[],[f29276,f1085])).
fof(f29276,plain,(
  ( ! [X101,X100] : (p(s(bool,X100)) | s(bool,X101) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,'F_0') = s(bool,X100)) ) | ($spl4 | $spl9 | $spl36)),
  inference(subsumption_resolution,[],[f29129,f25261])).
fof(f29129,plain,(
  ( ! [X101,X100] : (p(s(bool,X100)) | s(bool,X101) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,'T_0') = s(bool,X101) | s(bool,'F_0') = s(bool,X100)) ) | ($spl4 | $spl9 | $spl36)),
  inference(superposition,[],[f200,f25792])).
fof(f26480,plain,(
  ( ! [X101,X104] : (s(bool,X104) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,g(s(bool,X101))),s(bool,X104))))) | s(bool,'F_0') = s(bool,X101) | s(bool,'T_0') = s(bool,X104)) ) | $spl170),
  inference(cnf_transformation,[],[f26480_D])).
fof(f26480_D,plain,(
  ( ! [X101,X104] : (s(bool,X104) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,g(s(bool,X101))),s(bool,X104))))) | s(bool,'F_0') = s(bool,X101) | s(bool,'T_0') = s(bool,X104)) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f26448,plain,(
  ( ! [X103,X101,X102,X104] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X102),'const.pair.,_2'(s(bool,'F_0'),s(X102,X103))))) | s(bool,X104) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,g(s(bool,X101))),s(bool,X104))))) | s(bool,'T_0') = s(bool,X104) | s(bool,'F_0') = s(bool,X101)) ) | ($spl16 | $spl54)),
  inference(superposition,[],[f1557,f4883])).
fof(f25851,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') = s(bool,X4) | s(bool,X3) = s(bool,X4)) ) | $spl166),
  inference(cnf_transformation,[],[f25851_D])).
fof(f25851_D,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') = s(bool,X4) | s(bool,X3) = s(bool,X4)) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f25854,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK6(X1,X2,bool,X0)))) | s(bool,'F_0') != s(bool,X3) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(condensation,[],[f25853])).
fof(f25853,plain,(
  ( ! [X540,X542,X538,X541,X543,X539] : (s(bool,'F_0') != s(bool,X541) | s(fun(X538,bool),X539) = s(fun(X538,bool),X540) | s(bool,'T_0') = s(bool,g(s(X538,sK6(X539,X540,bool,X538)))) | s(bool,X541) = s(bool,X542) | s(bool,'T_0') = s(bool,X542) | s(bool,'T_0') = s(bool,X543)) )),
  inference(subsumption_resolution,[],[f18721,f368])).
fof(f18721,plain,(
  ( ! [X540,X542,X538,X541,X543,X539] : (s(bool,'F_0') != s(bool,X541) | s(fun(X538,bool),X539) = s(fun(X538,bool),X540) | s(bool,'T_0') = s(bool,g(s(X538,sK6(X539,X540,bool,X538)))) | s(bool,X541) = s(bool,X542) | s(bool,'T_0') = s(bool,X542) | s(bool,X543) = s(bool,f(s(X538,sK6(X539,X540,bool,X538)))) | s(bool,'T_0') = s(bool,X543)) )),
  inference(superposition,[],[f370,f146])).
fof(f25850,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK6(X1,X2,bool,X0)))) | s(bool,'F_0') != s(bool,X3) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(condensation,[],[f25849])).
fof(f25849,plain,(
  ( ! [X540,X542,X538,X541,X543,X539] : (s(bool,'F_0') != s(bool,X541) | s(fun(X538,bool),X539) = s(fun(X538,bool),X540) | s(bool,'T_0') = s(bool,g(s(X538,sK6(X539,X540,bool,X538)))) | s(bool,'T_0') = s(bool,X542) | s(bool,X541) = s(bool,X543) | s(bool,'T_0') = s(bool,X543)) )),
  inference(subsumption_resolution,[],[f18899,f368])).
fof(f18899,plain,(
  ( ! [X540,X542,X538,X541,X543,X539] : (s(bool,'F_0') != s(bool,X541) | s(fun(X538,bool),X539) = s(fun(X538,bool),X540) | s(bool,'T_0') = s(bool,g(s(X538,sK6(X539,X540,bool,X538)))) | s(bool,X542) = s(bool,f(s(X538,sK6(X539,X540,bool,X538)))) | s(bool,'T_0') = s(bool,X542) | s(bool,X541) = s(bool,X543) | s(bool,'T_0') = s(bool,X543)) )),
  inference(superposition,[],[f370,f146])).
fof(f2560,plain,(
  ( ! [X4,X2,X3] : (s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,sK6(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X2,X3,bool,bool))))) ) | $spl38),
  inference(cnf_transformation,[],[f2560_D])).
fof(f2560_D,plain,(
  ( ! [X4,X2,X3] : (s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,sK6(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X2,X3,bool,bool))))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f2562,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl41),
  inference(cnf_transformation,[],[f2562_D])).
fof(f2562_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f6728,plain,(
  ( ! [X4,X2,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X2,X3,bool,bool)))) | s(bool,X4) = s(bool,sK6(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f370,f56])).
fof(f6731,plain,(
  ( ! [X12,X13,X11] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X11,X12,bool,bool)))) | s(bool,X13) = s(bool,sK6(X11,X12,bool,bool)) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f370,f56])).
fof(f4987,plain,(
  ( ! [X445,X449,X446,X448] : (s(fun(bool,bool),X445) = s(fun(bool,bool),X446) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X448),'const.pair.,_2'(s(bool,sK6(X445,X446,bool,bool)),s(X448,X449))))) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X445,X446,bool,bool))))) ) | $spl78),
  inference(cnf_transformation,[],[f4987_D])).
fof(f4987_D,plain,(
  ( ! [X445,X449,X446,X448] : (s(fun(bool,bool),X445) = s(fun(bool,bool),X446) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X448),'const.pair.,_2'(s(bool,sK6(X445,X446,bool,bool)),s(X448,X449))))) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X445,X446,bool,bool))))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f4984,plain,(
  ( ! [X10] : (s(bool,'F_0') != s(bool,f(s(bool,X10))) | s(bool,'T_0') = s(bool,X10)) ) | $spl76),
  inference(cnf_transformation,[],[f4984_D])).
fof(f4984_D,plain,(
  ( ! [X10] : (s(bool,'F_0') != s(bool,f(s(bool,X10))) | s(bool,'T_0') = s(bool,X10)) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f6732,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,'F_0') != s(bool,f(s(bool,X16))) | s(fun(bool,bool),X14) = s(fun(bool,bool),X15) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X14,X15,bool,bool)))) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X17),'const.pair.,_2'(s(bool,sK6(X14,X15,bool,bool)),s(X17,X18)))))) )),
  inference(superposition,[],[f370,f132])).
fof(f25814,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(fun(bool,bool),X24) = s(fun(bool,bool),X25) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X27),'const.pair.,_2'(s(bool,X26),s(X27,X28))))) | s(bool,X26) = s(bool,sK6(X24,X25,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X24,X25,bool,bool))))) ) | $spl164),
  inference(cnf_transformation,[],[f25814_D])).
fof(f25814_D,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(fun(bool,bool),X24) = s(fun(bool,bool),X25) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X27),'const.pair.,_2'(s(bool,X26),s(X27,X28))))) | s(bool,X26) = s(bool,sK6(X24,X25,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X24,X25,bool,bool))))) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f6734,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X24) = s(fun(bool,bool),X25) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X24,X25,bool,bool)))) | s(bool,X26) = s(bool,sK6(X24,X25,bool,bool)) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X27),'const.pair.,_2'(s(bool,X26),s(X27,X28)))))) )),
  inference(superposition,[],[f370,f132])).
fof(f9974,plain,(
  ( ! [X4,X2,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X2,X3,bool,bool)))) | s(bool,X4) = s(bool,sK6(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f370,f56])).
fof(f9977,plain,(
  ( ! [X12,X13,X11] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X11,X12,bool,bool)))) | s(bool,X13) = s(bool,sK6(X11,X12,bool,bool)) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f370,f56])).
fof(f9978,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,'F_0') != s(bool,f(s(bool,X16))) | s(fun(bool,bool),X14) = s(fun(bool,bool),X15) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X14,X15,bool,bool)))) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X17),'const.pair.,_2'(s(bool,sK6(X14,X15,bool,bool)),s(X17,X18)))))) )),
  inference(superposition,[],[f370,f132])).
fof(f9980,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X24) = s(fun(bool,bool),X25) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X24,X25,bool,bool)))) | s(bool,X26) = s(bool,sK6(X24,X25,bool,bool)) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X27),'const.pair.,_2'(s(bool,X26),s(X27,X28)))))) )),
  inference(superposition,[],[f370,f132])).
fof(f25809,plain,(
  ( ! [X736,X737,X739,X740] : (s(fun(bool,bool),X736) = s(fun(bool,bool),X737) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X739,bool),'const.pair.,_2'(s(X739,X740),s(bool,sK6(X736,X737,bool,bool)))))) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X736,X737,bool,bool))))) ) | $spl162),
  inference(cnf_transformation,[],[f25809_D])).
fof(f25809_D,plain,(
  ( ! [X736,X737,X739,X740] : (s(fun(bool,bool),X736) = s(fun(bool,bool),X737) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X739,bool),'const.pair.,_2'(s(X739,X740),s(bool,sK6(X736,X737,bool,bool)))))) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X736,X737,bool,bool))))) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f13039,plain,(
  ( ! [X736,X738,X737,X739,X740] : (s(bool,'F_0') != s(bool,f(s(bool,X738))) | s(fun(bool,bool),X736) = s(fun(bool,bool),X737) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X736,X737,bool,bool)))) | s(bool,'T_0') = s(bool,X738) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X739,bool),'const.pair.,_2'(s(X739,X740),s(bool,sK6(X736,X737,bool,bool))))))) )),
  inference(superposition,[],[f370,f133])).
fof(f2448,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl26),
  inference(cnf_transformation,[],[f2448_D])).
fof(f2448_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f2449,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK6(X0,X1,bool,bool)) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) | $spl28),
  inference(cnf_transformation,[],[f2449_D])).
fof(f2449_D,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK6(X0,X1,bool,bool)) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f25808,plain,(
  ( ! [X695,X696,X692] : (s(bool,X692) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X695,bool),'const.pair.,_2'(s(X695,X696),s(bool,X692)))))) ) | $spl160),
  inference(cnf_transformation,[],[f25808_D])).
fof(f25808_D,plain,(
  ( ! [X695,X696,X692] : (s(bool,X692) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X695,bool),'const.pair.,_2'(s(X695,X696),s(bool,X692)))))) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f13170,plain,(
  ( ! [X693,X695,X696,X692,X694] : (s(bool,X692) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X693) = s(fun(bool,bool),X694) | s(bool,'T_0') = s(bool,sK6(X693,X694,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X695,bool),'const.pair.,_2'(s(X695,X696),s(bool,X692)))))) )),
  inference(superposition,[],[f365,f133])).
fof(f4983,plain,(
  ( ! [X8,X9] : (s(fun(bool,bool),X8) = s(fun(bool,bool),X9) | s(bool,'T_0') = s(bool,sK6(X8,X9,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X8,X9,bool,bool))))) ) | $spl74),
  inference(cnf_transformation,[],[f4983_D])).
fof(f4983_D,plain,(
  ( ! [X8,X9] : (s(fun(bool,bool),X8) = s(fun(bool,bool),X9) | s(bool,'T_0') = s(bool,sK6(X8,X9,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X8,X9,bool,bool))))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f25807,plain,(
  ( ! [X724,X726,X725] : (s(bool,'F_0') != s(bool,f(s(bool,X724))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X725,bool),'const.pair.,_2'(s(X725,X726),s(bool,X724)))))) ) | $spl158),
  inference(cnf_transformation,[],[f25807_D])).
fof(f25807_D,plain,(
  ( ! [X724,X726,X725] : (s(bool,'F_0') != s(bool,f(s(bool,X724))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X725,bool),'const.pair.,_2'(s(X725,X726),s(bool,X724)))))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f13175,plain,(
  ( ! [X724,X726,X722,X725,X723] : (s(bool,'F_0') != s(bool,f(s(bool,X724))) | s(fun(bool,bool),X722) = s(fun(bool,bool),X723) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X722,X723,bool,bool)))) | s(bool,'T_0') = s(bool,sK6(X722,X723,bool,bool)) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X725,bool),'const.pair.,_2'(s(X725,X726),s(bool,X724)))))) )),
  inference(superposition,[],[f370,f133])).
fof(f25805,plain,(
  ( ! [X644,X645,X643] : (s(bool,X643) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X644,bool),'const.pair.,_2'(s(X644,X645),s(bool,X643)))))) ) | $spl156),
  inference(cnf_transformation,[],[f25805_D])).
fof(f25805_D,plain,(
  ( ! [X644,X645,X643] : (s(bool,X643) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X644,bool),'const.pair.,_2'(s(X644,X645),s(bool,X643)))))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f25804,plain,(
  ( ! [X644,X642,X645,X641,X643] : (s(bool,X643) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK6(X641,X642,bool,bool)) | s(fun(bool,bool),X641) = s(fun(bool,bool),X642) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X644,bool),'const.pair.,_2'(s(X644,X645),s(bool,X643)))))) ) | $spl36),
  inference(global_subsumption,[],[f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f64,f132,f133,f146,f156,f157,f200,f361,f363,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f3832,f3890,f3878,f2550,f2549,f3652,f3732,f3828,f3727,f2700,f763,f2874,f2698,f761,f3651,f786,f550,f475,f48,f47,f46,f762,f2699,f2704,f785,f2763,f4154,f3827,f3726,f3646,f368,f366,f3759,f3758,f2455,f1935,f1480,f1972,f1473,f39,f82,f104,f365,f370,f2706,f2559,f11126,f11125,f11188,f11161,f11162,f11203,f11222,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13826,f13033,f13038,f13169,f13223,f13224,f13225,f13345,f13350,f13427,f13415,f15352,f15341,f15429,f15069,f17276,f17266,f17264,f17226,f17356,f16976,f22573,f18154,f18155,f22303,f22293,f22291,f22279,f22277,f22275,f22219,f22217,f18360,f18362,f18363,f22711,f18790,f18797,f18801,f18802,f22898,f19153,f19158,f21203,f21173,f19331,f19338,f19400,f19401,f19402,f19403,f19406,f22928,f22929,f22930,f22931,f19558,f19565,f19569,f19570,f22984,f20200,f20161,f20156,f25794,f19927,f19749,f19571,f19566,f25798,f19337,f19159,f18981,f18976,f18803,f18798,f25801,f25802,f13351,f13346])).
fof(f25803,plain,(
  ( ! [X673,X675,X671,X672,X674] : (s(fun(bool,bool),X671) = s(fun(bool,bool),X672) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X674,bool),'const.pair.,_2'(s(X674,X675),s(bool,X673))))) | s(bool,X673) = s(bool,sK6(X671,X672,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X671,X672,bool,bool))))) ) | $spl154),
  inference(cnf_transformation,[],[f25803_D])).
fof(f25803_D,plain,(
  ( ! [X673,X675,X671,X672,X674] : (s(fun(bool,bool),X671) = s(fun(bool,bool),X672) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X674,bool),'const.pair.,_2'(s(X674,X675),s(bool,X673))))) | s(bool,X673) = s(bool,sK6(X671,X672,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X671,X672,bool,bool))))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f13351,plain,(
  ( ! [X673,X675,X671,X672,X674] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X671) = s(fun(bool,bool),X672) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X671,X672,bool,bool)))) | s(bool,X673) = s(bool,sK6(X671,X672,bool,bool)) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X674,bool),'const.pair.,_2'(s(X674,X675),s(bool,X673)))))) )),
  inference(superposition,[],[f370,f133])).
fof(f25802,plain,(
  ( ! [X825,X824,X827] : (s(bool,'F_0') != s(bool,f(s(bool,X827))) | s(fun(bool,bool),X824) = s(fun(bool,bool),X825) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X824,X825,bool,bool)))) | s(bool,'T_0') = s(bool,X827) | s(bool,'T_0') = s(bool,sK6(X824,X825,bool,bool))) )),
  inference(forward_demodulation,[],[f15074,f43])).
fof(f15074,plain,(
  ( ! [X825,X826,X824,X828,X827] : (s(bool,'F_0') != s(bool,f(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X826),'const.pair.,_2'(s(bool,X827),s(X826,X828))))))) | s(fun(bool,bool),X824) = s(fun(bool,bool),X825) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X824,X825,bool,bool)))) | s(bool,'T_0') = s(bool,X827) | s(bool,'T_0') = s(bool,sK6(X824,X825,bool,bool))) )),
  inference(superposition,[],[f370,f156])).
fof(f25801,plain,(
  ( ! [X922,X918,X919] : (s(bool,'F_0') != s(bool,f(s(bool,X922))) | s(fun(bool,bool),X918) = s(fun(bool,bool),X919) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X918,X919,bool,bool)))) | s(bool,'T_0') = s(bool,X922) | s(bool,'T_0') = s(bool,sK6(X918,X919,bool,bool))) )),
  inference(forward_demodulation,[],[f16981,f44])).
fof(f16981,plain,(
  ( ! [X922,X920,X918,X921,X919] : (s(bool,'F_0') != s(bool,f(s(bool,'const.pair.SND_1'(s('type.pair.prod'(X920,bool),'const.pair.,_2'(s(X920,X921),s(bool,X922))))))) | s(fun(bool,bool),X918) = s(fun(bool,bool),X919) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X918,X919,bool,bool)))) | s(bool,'T_0') = s(bool,X922) | s(bool,'T_0') = s(bool,sK6(X918,X919,bool,bool))) )),
  inference(superposition,[],[f370,f157])).
fof(f2452,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) ) | $spl32),
  inference(cnf_transformation,[],[f2452_D])).
fof(f2452_D,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f25800,plain,(
  ( ! [X984,X980] : (s(bool,X980) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X984) | s(bool,X980) = s(bool,X984)) ) | $spl152),
  inference(cnf_transformation,[],[f25800_D])).
fof(f25800_D,plain,(
  ( ! [X984,X980] : (s(bool,X980) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X984) | s(bool,X980) = s(bool,X984)) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f18798,plain,(
  ( ! [X984,X982,X980,X983,X981] : (s(bool,X980) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X981) = s(fun(bool,bool),X982) | s(bool,'T_0') = s(bool,sK6(X981,X982,bool,bool)) | s(bool,X980) = s(bool,X983) | s(bool,'T_0') = s(bool,X983) | s(bool,X984) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X984)) )),
  inference(superposition,[],[f365,f146])).
fof(f25799,plain,(
  ( ! [X1014,X1012] : (s(bool,'F_0') != s(bool,f(s(bool,X1012))) | s(bool,'T_0') = s(bool,X1014) | s(bool,X1012) = s(bool,X1014)) ) | $spl150),
  inference(cnf_transformation,[],[f25799_D])).
fof(f25799_D,plain,(
  ( ! [X1014,X1012] : (s(bool,'F_0') != s(bool,f(s(bool,X1012))) | s(bool,'T_0') = s(bool,X1014) | s(bool,X1012) = s(bool,X1014)) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f18803,plain,(
  ( ! [X1011,X1013,X1010,X1014,X1012] : (s(bool,'F_0') != s(bool,f(s(bool,X1012))) | s(fun(bool,bool),X1010) = s(fun(bool,bool),X1011) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X1010,X1011,bool,bool)))) | s(bool,X1012) = s(bool,X1013) | s(bool,'T_0') = s(bool,X1013) | s(bool,X1014) = s(bool,sK6(X1010,X1011,bool,bool)) | s(bool,'T_0') = s(bool,X1014)) )),
  inference(superposition,[],[f370,f146])).
fof(f18976,plain,(
  ( ! [X984,X982,X980,X983,X981] : (s(bool,X980) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X981) = s(fun(bool,bool),X982) | s(bool,'T_0') = s(bool,sK6(X981,X982,bool,bool)) | s(bool,X983) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X983) | s(bool,X980) = s(bool,X984) | s(bool,'T_0') = s(bool,X984)) )),
  inference(superposition,[],[f365,f146])).
fof(f18981,plain,(
  ( ! [X1011,X1013,X1010,X1014,X1012] : (s(bool,'F_0') != s(bool,f(s(bool,X1012))) | s(fun(bool,bool),X1010) = s(fun(bool,bool),X1011) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X1010,X1011,bool,bool)))) | s(bool,X1013) = s(bool,sK6(X1010,X1011,bool,bool)) | s(bool,'T_0') = s(bool,X1013) | s(bool,X1012) = s(bool,X1014) | s(bool,'T_0') = s(bool,X1014)) )),
  inference(superposition,[],[f370,f146])).
fof(f25797,plain,(
  ( ! [X1011,X1013,X1010,X1014] : (s(fun(bool,bool),X1010) = s(fun(bool,bool),X1011) | s(bool,'T_0') = s(bool,X1014) | s(bool,X1013) = s(bool,X1014) | s(bool,X1013) = s(bool,sK6(X1010,X1011,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X1010,X1011,bool,bool))))) ) | $spl148),
  inference(cnf_transformation,[],[f25797_D])).
fof(f25797_D,plain,(
  ( ! [X1011,X1013,X1010,X1014] : (s(fun(bool,bool),X1010) = s(fun(bool,bool),X1011) | s(bool,'T_0') = s(bool,X1014) | s(bool,X1013) = s(bool,X1014) | s(bool,X1013) = s(bool,sK6(X1010,X1011,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X1010,X1011,bool,bool))))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f19159,plain,(
  ( ! [X1011,X1013,X1010,X1014,X1012] : (s(bool,'F_0') != s(bool,f(s(bool,X1012))) | s(fun(bool,bool),X1010) = s(fun(bool,bool),X1011) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X1010,X1011,bool,bool)))) | s(bool,X1013) = s(bool,sK6(X1010,X1011,bool,bool)) | s(bool,'T_0') = s(bool,X1012) | s(bool,X1013) = s(bool,X1014) | s(bool,'T_0') = s(bool,X1014)) )),
  inference(superposition,[],[f370,f146])).
fof(f25796,plain,(
  ( ! [X1013,X1014,X1012] : (s(bool,'F_0') != s(bool,f(s(bool,X1012))) | s(bool,'T_0') = s(bool,X1014) | s(bool,X1013) = s(bool,X1014) | s(bool,X1012) = s(bool,X1013)) ) | $spl146),
  inference(cnf_transformation,[],[f25796_D])).
fof(f25796_D,plain,(
  ( ! [X1013,X1014,X1012] : (s(bool,'F_0') != s(bool,f(s(bool,X1012))) | s(bool,'T_0') = s(bool,X1014) | s(bool,X1013) = s(bool,X1014) | s(bool,X1012) = s(bool,X1013)) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f19337,plain,(
  ( ! [X1002,X1006,X1004,X1003,X1005] : (s(bool,'F_0') != s(bool,f(s(bool,X1004))) | s(fun(bool,bool),X1002) = s(fun(bool,bool),X1003) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X1002,X1003,bool,bool)))) | s(bool,X1004) = s(bool,X1005) | s(bool,'T_0') = s(bool,sK6(X1002,X1003,bool,bool)) | s(bool,X1005) = s(bool,X1006) | s(bool,'T_0') = s(bool,X1006)) )),
  inference(superposition,[],[f370,f146])).
fof(f25795,plain,(
  ( ! [X895,X893,X894] : (s(bool,X894) = s(bool,X895) | s(bool,X893) = s(bool,g(s(bool,'F_0'))) | s(bool,X893) = s(bool,X894) | s(bool,'T_0') = s(bool,X895)) ) | $spl144),
  inference(cnf_transformation,[],[f25795_D])).
fof(f25795_D,plain,(
  ( ! [X895,X893,X894] : (s(bool,X894) = s(bool,X895) | s(bool,X893) = s(bool,g(s(bool,'F_0'))) | s(bool,X893) = s(bool,X894) | s(bool,'T_0') = s(bool,X895)) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f25798,plain,(
  ( ! [X895,X893,X894,X892,X891] : (s(fun(bool,bool),X891) = s(fun(bool,bool),X892) | s(bool,X893) = s(bool,X894) | s(bool,X894) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK6(X891,X892,bool,bool)) | s(bool,X893) = s(bool,X895) | s(bool,'T_0') = s(bool,X895)) ) | $spl36),
  inference(global_subsumption,[],[f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f64,f132,f133,f146,f156,f157,f200,f361,f363,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f3832,f3890,f3878,f2550,f2549,f3652,f3732,f3828,f3727,f2700,f763,f2874,f2698,f761,f3651,f786,f550,f475,f48,f47,f46,f762,f2699,f2704,f785,f2763,f4154,f3827,f3726,f3646,f368,f366,f3759,f3758,f2455,f1935,f1480,f1972,f1473,f39,f82,f104,f365,f370,f2706,f2559,f11126,f11125,f11188,f11161,f11162,f11203,f11222,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13826,f13033,f13038,f13169,f13223,f13224,f13225,f13345,f13350,f13427,f13415,f15352,f15341,f15429,f15069,f17276,f17266,f17264,f17226,f17356,f16976,f22573,f18154,f18155,f22303,f22293,f22291,f22279,f22277,f22275,f22219,f22217,f18360,f18362,f18363,f22711,f18790,f18797,f18801,f18802,f22898,f19153,f19158,f21203,f21173,f19331,f19338,f19400,f19401,f19402,f19403,f19406,f22928,f22929,f22930,f22931,f19558,f19565,f19569,f19570,f22984,f20200,f20161,f20156,f25794,f19927,f19749,f19571,f19566])).
fof(f25793,plain,(
  ( ! [X922,X924,X923,X921,X925] : (s(fun(bool,bool),X921) = s(fun(bool,bool),X922) | s(bool,X923) = s(bool,sK6(X921,X922,bool,bool)) | s(bool,X923) = s(bool,X924) | s(bool,'T_0') = s(bool,X925) | s(bool,X924) = s(bool,X925) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X921,X922,bool,bool))))) ) | $spl142),
  inference(cnf_transformation,[],[f25793_D])).
fof(f25793_D,plain,(
  ( ! [X922,X924,X923,X921,X925] : (s(fun(bool,bool),X921) = s(fun(bool,bool),X922) | s(bool,X923) = s(bool,sK6(X921,X922,bool,bool)) | s(bool,X923) = s(bool,X924) | s(bool,'T_0') = s(bool,X925) | s(bool,X924) = s(bool,X925) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X921,X922,bool,bool))))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f19571,plain,(
  ( ! [X922,X924,X923,X921,X925] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X921) = s(fun(bool,bool),X922) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X921,X922,bool,bool)))) | s(bool,X923) = s(bool,X924) | s(bool,X924) = s(bool,sK6(X921,X922,bool,bool)) | s(bool,X923) = s(bool,X925) | s(bool,'T_0') = s(bool,X925)) )),
  inference(superposition,[],[f370,f146])).
fof(f19749,plain,(
  ( ! [X1011,X1013,X1010,X1014,X1012] : (s(bool,'F_0') != s(bool,f(s(bool,X1012))) | s(fun(bool,bool),X1010) = s(fun(bool,bool),X1011) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X1010,X1011,bool,bool)))) | s(bool,X1013) = s(bool,sK6(X1010,X1011,bool,bool)) | s(bool,X1013) = s(bool,X1014) | s(bool,'T_0') = s(bool,X1014) | s(bool,'T_0') = s(bool,X1012)) )),
  inference(superposition,[],[f370,f146])).
fof(f19927,plain,(
  ( ! [X1011,X1013,X1010,X1014,X1012] : (s(bool,'F_0') != s(bool,f(s(bool,X1012))) | s(fun(bool,bool),X1010) = s(fun(bool,bool),X1011) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X1010,X1011,bool,bool)))) | s(bool,X1012) = s(bool,X1013) | s(bool,X1013) = s(bool,X1014) | s(bool,'T_0') = s(bool,X1014) | s(bool,'T_0') = s(bool,sK6(X1010,X1011,bool,bool))) )),
  inference(superposition,[],[f370,f146])).
fof(f25794,plain,(
  ( ! [X895,X893,X894,X892,X891] : (s(bool,X894) = s(bool,X895) | s(bool,'T_0') = s(bool,X895) | s(bool,'T_0') = s(bool,sK6(X891,X892,bool,bool)) | s(fun(bool,bool),X891) = s(fun(bool,bool),X892) | s(bool,X893) = s(bool,X894) | s(bool,X893) = s(bool,g(s(bool,'F_0')))) ) | $spl36),
  inference(global_subsumption,[],[f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f64,f132,f133,f146,f156,f157,f200,f361,f363,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f3832,f3890,f3878,f2550,f2549,f3652,f3732,f3828,f3727,f2700,f763,f2874,f2698,f761,f3651,f786,f550,f475,f48,f47,f46,f762,f2699,f2704,f785,f2763,f4154,f3827,f3726,f3646,f368,f366,f3759,f3758,f2455,f1935,f1480,f1972,f1473,f39,f82,f104,f365,f370,f2706,f2559,f11126,f11125,f11188,f11161,f11162,f11203,f11222,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13826,f13033,f13038,f13169,f13223,f13224,f13225,f13345,f13350,f13427,f13415,f15352,f15341,f15429,f15069,f17276,f17266,f17264,f17226,f17356,f16976,f22573,f18154,f18155,f22303,f22293,f22291,f22279,f22277,f22275,f22219,f22217,f18360,f18362,f18363,f22711,f18790,f18797,f18801,f18802,f22898,f19153,f19158,f21203,f21173,f19331,f19338,f19400,f19401,f19402,f19403,f19406,f22928,f22929,f22930,f22931,f19558,f19565,f19569,f19570,f22984,f20200,f20161,f20156])).
fof(f20161,plain,(
  ( ! [X922,X924,X923,X921,X925] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X921) = s(fun(bool,bool),X922) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X921,X922,bool,bool)))) | s(bool,X923) = s(bool,X924) | s(bool,X924) = s(bool,X925) | s(bool,'T_0') = s(bool,X925) | s(bool,X923) = s(bool,sK6(X921,X922,bool,bool))) )),
  inference(superposition,[],[f370,f146])).
fof(f25791,plain,(
  ( ! [X637] : (s(bool,X637) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X637) = s(bool,g(s(bool,'F_0')))) ) | $spl140),
  inference(cnf_transformation,[],[f25791_D])).
fof(f25791_D,plain,(
  ( ! [X637] : (s(bool,X637) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X637) = s(bool,g(s(bool,'F_0')))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f25790,plain,(
  ( ! [X635,X636,X637] : (s(bool,X637) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,'T_0') = s(bool,sK6(X635,X636,bool,bool)) | s(fun(bool,bool),X635) = s(fun(bool,bool),X636) | s(bool,X637) = s(bool,g(s(bool,'F_0')))) ) | ($spl4 | $spl36)),
  inference(global_subsumption,[],[f36,f35,f34,f37,f40,f43,f44,f45,f56,f58,f67,f64,f132,f133,f146,f156,f157,f200,f361,f363,f54,f1085,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f3832,f3890,f3878,f2550,f2549,f3652,f3732,f3828,f3727,f2700,f763,f2874,f2698,f761,f3651,f786,f550,f475,f48,f47,f46,f762,f2762,f2699,f2704,f785,f2763,f4154,f3827,f3726,f3646,f368,f366,f3759,f3758,f2455,f1935,f1480,f1972,f1473,f39,f82,f104,f365,f370,f2706,f2559,f11126,f11125,f11188,f11161,f11162,f11203,f11222,f13741,f12750,f12751,f12752,f12791,f12807,f12808,f12809,f12891,f13826,f13033,f13038,f13169,f13223,f13224,f13225,f13345,f13350,f13427,f13415,f15352,f15341,f15429,f15069,f17276,f17266,f17264,f17226,f17356,f16976,f22573,f18154,f18155,f22303,f22295,f22293,f22291,f22279,f22277,f22275,f22219,f22217,f18360,f18362,f18363,f22711,f18790,f18797,f18801,f18802,f22898,f19153,f19158,f21203,f21173,f19331,f19338,f19400,f19401,f19402,f19403,f19406,f22928,f22929,f22930,f22931,f19558,f19565,f19569,f19570,f22984,f20200,f23229,f23233,f23235,f24931,f24930,f24854,f24822,f24800,f23782,f23787,f25217,f24214,f25784,f25785,f25786,f24220,f24215,f25788,f23788,f23783])).
fof(f25789,plain,(
  ( ! [X656,X657,X655] : (s(fun(bool,bool),X655) = s(fun(bool,bool),X656) | s(bool,X657) = s(bool,sK6(X655,X656,bool,bool)) | s(bool,X657) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X655,X656,bool,bool))))) ) | $spl138),
  inference(cnf_transformation,[],[f25789_D])).
fof(f25789_D,plain,(
  ( ! [X656,X657,X655] : (s(fun(bool,bool),X655) = s(fun(bool,bool),X656) | s(bool,X657) = s(bool,sK6(X655,X656,bool,bool)) | s(bool,X657) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X655,X656,bool,bool))))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f23788,plain,(
  ( ! [X656,X657,X655] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X655) = s(fun(bool,bool),X656) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X655,X656,bool,bool)))) | s(bool,X657) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X657) = s(bool,sK6(X655,X656,bool,bool))) ) | $spl4),
  inference(superposition,[],[f370,f22295])).
fof(f25788,plain,(
  ( ! [X598,X597,X599] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X597) = s(fun(bool,bool),X598) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X597,X598,bool,bool)))) | s(bool,'T_0') = s(bool,X599) | s(bool,X599) = s(bool,sK6(X597,X598,bool,bool))) ) | $spl4),
  inference(forward_demodulation,[],[f23945,f1085])).
fof(f23945,plain,(
  ( ! [X598,X597,X599] : (s(bool,'F_0') != s(bool,f(s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))))) | s(fun(bool,bool),X597) = s(fun(bool,bool),X598) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X597,X598,bool,bool)))) | s(bool,'T_0') = s(bool,X599) | s(bool,X599) = s(bool,sK6(X597,X598,bool,bool))) ) | $spl4),
  inference(superposition,[],[f370,f22295])).
fof(f25219,plain,(
  ( ! [X153] : (s(bool,'F_0') != s(bool,f(s(bool,X153))) | s(bool,X153) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | $spl130),
  inference(cnf_transformation,[],[f25219_D])).
fof(f25219_D,plain,(
  ( ! [X153] : (s(bool,'F_0') != s(bool,f(s(bool,X153))) | s(bool,X153) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f24220,plain,(
  ( ! [X716,X714,X715] : (s(bool,'F_0') != s(bool,f(s(bool,X716))) | s(fun(bool,bool),X714) = s(fun(bool,bool),X715) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X714,X715,bool,bool)))) | s(bool,'T_0') = s(bool,sK6(X714,X715,bool,bool)) | s(bool,X716) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | $spl4),
  inference(superposition,[],[f370,f22295])).
fof(f25784,plain,(
  ( ! [X716,X714,X715] : (s(bool,'T_0') = s(bool,sK6(X714,X715,bool,bool)) | s(bool,'F_0') != s(bool,f(s(bool,X716))) | s(fun(bool,bool),X714) = s(fun(bool,bool),X715) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X714,X715,bool,bool)))) | s(bool,'T_0') = s(bool,X716)) ) | $spl4),
  inference(forward_demodulation,[],[f24410,f1085])).
fof(f24410,plain,(
  ( ! [X716,X714,X715] : (s(bool,'F_0') != s(bool,f(s(bool,X716))) | s(fun(bool,bool),X714) = s(fun(bool,bool),X715) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X714,X715,bool,bool)))) | s(bool,'T_0') = s(bool,X716) | s(bool,sK6(X714,X715,bool,bool)) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | $spl4),
  inference(superposition,[],[f370,f22295])).
fof(f25772,plain,(
  $false | ($spl13 | $spl22 | $spl48 | $spl92)),
  inference(subsumption_resolution,[],[f25485,f4877])).
fof(f25485,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0')))) ) | ($spl13 | $spl22 | $spl48 | $spl92)),
  inference(backward_demodulation,[],[f4877,f9836])).
fof(f9836,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) ) | ($spl13 | $spl22 | $spl92)),
  inference(forward_demodulation,[],[f9763,f43])).
fof(f9763,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X0),s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,'F_0')))))))) ) | ($spl13 | $spl22 | $spl92)),
  inference(unit_resulting_resolution,[],[f6406,f9595])).
fof(f9595,plain,(
  ( ! [X173,X174] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X173),s(bool,X174))) | s(bool,'T_0') = s(bool,X174)) ) | ($spl13 | $spl22 | $spl92)),
  inference(subsumption_resolution,[],[f9401,f439])).
fof(f9401,plain,(
  ( ! [X173,X174] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X173),s(bool,X174))) | s(bool,'T_0') = s(bool,X174) | s(bool,'T_0') = s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2)))))) ) | ($spl13 | $spl22 | $spl92)),
  inference(superposition,[],[f439,f8188])).
fof(f8188,plain,(
  ( ! [X6,X7,X5] : (s(bool,X7) = s(bool,happ(s(fun(bool,bool),X5),s(bool,X6))) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X7)) ) | ($spl13 | $spl22 | $spl92)),
  inference(duplicate_literal_removal,[],[f8150])).
fof(f8150,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X6) | s(bool,X7) = s(bool,happ(s(fun(bool,bool),X5),s(bool,X6))) | s(bool,'T_0') = s(bool,X7)) ) | ($spl13 | $spl22 | $spl92)),
  inference(superposition,[],[f7469,f56])).
fof(f7469,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),X1),s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl13 | $spl22 | $spl92)),
  inference(superposition,[],[f7444,f40])).
fof(f7444,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) ) | ($spl13 | $spl22 | $spl92)),
  inference(forward_demodulation,[],[f7419,f43])).
fof(f7419,plain,(
  ( ! [X0] : (s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0'))) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,'F_0')))))) ) | ($spl13 | $spl22 | $spl92)),
  inference(unit_resulting_resolution,[],[f6406,f7405,f200])).
fof(f25771,plain,(
  $false | ($spl13 | $spl22 | $spl48)),
  inference(subsumption_resolution,[],[f25470,f4877])).
fof(f25470,plain,(
  s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'T_0'),s(bool,'T_0'))))) | ($spl13 | $spl22 | $spl48)),
  inference(backward_demodulation,[],[f4877,f6406])).
fof(f25768,plain,(
  $false | $spl48),
  inference(subsumption_resolution,[],[f25438,f37])).
fof(f25438,plain,(
  ~p(s(bool,'T_0')) | $spl48),
  inference(backward_demodulation,[],[f4877,f39])).
fof(f25237,plain,(
  ( ! [X218,X219] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X219),s(bool,X218))) | s(bool,X218) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | $spl134),
  inference(cnf_transformation,[],[f25237_D])).
fof(f25237_D,plain,(
  ( ! [X218,X219] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X219),s(bool,X218))) | s(bool,X218) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f25378,plain,(
  ( ! [X222,X223,X226] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X226),s(bool,X222))) | s(bool,'T_0') = s(bool,X223) | s(bool,X222) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl13 | $spl22 | $spl36 | $spl92)),
  inference(subsumption_resolution,[],[f24277,f17356])).
fof(f24277,plain,(
  ( ! [X222,X225,X223,X226,X224] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X224,bool),'const.pair.,_2'(s(X224,X225),s(bool,X223))))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X226),s(bool,X222))) | s(bool,'T_0') = s(bool,X223) | s(bool,X222) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl13 | $spl22 | $spl92)),
  inference(superposition,[],[f9272,f22295])).
fof(f9272,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X78,bool),'const.pair.,_2'(s(X78,X79),s(bool,X77))))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X76),s(bool,X77)))) ) | ($spl13 | $spl22 | $spl92)),
  inference(subsumption_resolution,[],[f9188,f7282])).
fof(f9188,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X78,bool),'const.pair.,_2'(s(X78,X79),s(bool,X77))))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X76),s(bool,X77))) | s(bool,'T_0') = s(bool,X77)) ) | ($spl13 | $spl22 | $spl92)),
  inference(superposition,[],[f104,f7469])).
fof(f25377,plain,(
  ( ! [X212,X216,X213] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X216),s(bool,X212))) | s(bool,'T_0') = s(bool,X213) | s(bool,X212) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl9 | $spl13 | $spl22 | $spl36 | $spl92)),
  inference(subsumption_resolution,[],[f24275,f15429])).
fof(f24275,plain,(
  ( ! [X212,X216,X215,X213,X214] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X214),'const.pair.,_2'(s(bool,X213),s(X214,X215))))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X216),s(bool,X212))) | s(bool,'T_0') = s(bool,X213) | s(bool,X212) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl9 | $spl13 | $spl22 | $spl92)),
  inference(superposition,[],[f8988,f22295])).
fof(f8988,plain,(
  ( ! [X76,X74,X77,X75] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X76),'const.pair.,_2'(s(bool,X75),s(X76,X77))))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X74),s(bool,X75)))) ) | ($spl9 | $spl13 | $spl22 | $spl92)),
  inference(subsumption_resolution,[],[f8854,f4170])).
fof(f8854,plain,(
  ( ! [X76,X74,X77,X75] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X76),'const.pair.,_2'(s(bool,X75),s(X76,X77))))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X74),s(bool,X75))) | s(bool,'T_0') = s(bool,X75)) ) | ($spl13 | $spl22 | $spl92)),
  inference(superposition,[],[f82,f7469])).
fof(f25316,plain,(
  ( ! [X499,X500] : (s(bool,'F_0') != s(bool,f(s(bool,X500))) | s(bool,'T_0') = s(bool,X499) | s(bool,X500) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl13 | $spl22 | $spl76 | $spl92)),
  inference(subsumption_resolution,[],[f24163,f9595])).
fof(f24163,plain,(
  ( ! [X499,X500,X498] : (s(bool,'F_0') != s(bool,f(s(bool,X500))) | s(bool,'T_0') = s(bool,X499) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X498),s(bool,X499))) | s(bool,X500) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl13 | $spl22 | $spl76 | $spl92)),
  inference(superposition,[],[f11770,f22295])).
fof(f11770,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,happ(s(fun(bool,bool),X1),s(bool,X0))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl13 | $spl22 | $spl76 | $spl92)),
  inference(condensation,[],[f11769])).
fof(f11769,plain,(
  ( ! [X26,X27,X25] : (s(bool,'F_0') != s(bool,f(s(bool,happ(s(fun(bool,bool),X26),s(bool,X27))))) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,X27)) ) | ($spl13 | $spl22 | $spl76 | $spl92)),
  inference(duplicate_literal_removal,[],[f11722])).
fof(f11722,plain,(
  ( ! [X26,X27,X25] : (s(bool,'F_0') != s(bool,f(s(bool,happ(s(fun(bool,bool),X26),s(bool,X27))))) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,X25)) ) | ($spl13 | $spl22 | $spl76 | $spl92)),
  inference(superposition,[],[f4984,f8188])).
fof(f25313,plain,(
  ( ! [X488,X487,X489] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X489),s(bool,X488))) | s(bool,'T_0') = s(bool,X487) | s(bool,X488) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl13 | $spl22 | $spl92)),
  inference(subsumption_resolution,[],[f24160,f9595])).
fof(f24160,plain,(
  ( ! [X486,X488,X487,X489] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X489),s(bool,X488))) | s(bool,'T_0') = s(bool,X487) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X486),s(bool,X487))) | s(bool,X488) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl13 | $spl22 | $spl92)),
  inference(superposition,[],[f9720,f22295])).
fof(f9720,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X1),s(bool,happ(s(fun(bool,bool),X2),s(bool,X0))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl13 | $spl22 | $spl92)),
  inference(condensation,[],[f9719])).
fof(f9719,plain,(
  ( ! [X35,X33,X36,X34] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X36),s(bool,happ(s(fun(bool,bool),X34),s(bool,X35))))) | s(bool,'T_0') = s(bool,X33) | s(bool,'T_0') = s(bool,X35)) ) | ($spl13 | $spl22 | $spl92)),
  inference(duplicate_literal_removal,[],[f9630])).
fof(f9630,plain,(
  ( ! [X35,X33,X36,X34] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X36),s(bool,happ(s(fun(bool,bool),X34),s(bool,X35))))) | s(bool,'T_0') = s(bool,X33) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,X33)) ) | ($spl13 | $spl22 | $spl92)),
  inference(superposition,[],[f9573,f8188])).
fof(f9573,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X1),s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl13 | $spl22 | $spl92)),
  inference(condensation,[],[f9572])).
fof(f9572,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X1),s(bool,X2))) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) ) | ($spl13 | $spl22 | $spl92)),
  inference(duplicate_literal_removal,[],[f9292])).
fof(f9292,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X1),s(bool,X2))) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | ($spl13 | $spl22 | $spl92)),
  inference(superposition,[],[f8188,f40])).
fof(f25218,plain,(
  ( ! [X151] : (s(bool,'T_0') != s(bool,X151) | s(bool,X151) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | $spl128),
  inference(cnf_transformation,[],[f25218_D])).
fof(f25218_D,plain,(
  ( ! [X151] : (s(bool,'T_0') != s(bool,X151) | s(bool,X151) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f24515,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') != s(bool,X3) | s(bool,X3) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl13 | $spl22 | $spl36)),
  inference(duplicate_literal_removal,[],[f24514])).
fof(f24514,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') != s(bool,X3) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X3) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl13 | $spl22 | $spl36)),
  inference(condensation,[],[f24102])).
fof(f24102,plain,(
  ( ! [X280,X277,X279,X278] : (s(bool,'T_0') != s(bool,X278) | s(bool,'T_0') = s(bool,X279) | s(bool,X277) = s(bool,X280) | s(bool,X279) = s(bool,X280) | s(bool,'T_0') = s(bool,X277) | s(bool,X278) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl13 | $spl22 | $spl36)),
  inference(superposition,[],[f22900,f22295])).
fof(f24528,plain,(
  ( ! [X243,X241,X242] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X243),s(bool,X242))) | s(bool,'T_0') = s(bool,X241) | s(bool,X242) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl13 | $spl22 | $spl92)),
  inference(duplicate_literal_removal,[],[f24092])).
fof(f24092,plain,(
  ( ! [X243,X241,X242] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X243),s(bool,X242))) | s(bool,'T_0') = s(bool,X241) | s(bool,'T_0') = s(bool,X241) | s(bool,X242) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl13 | $spl22 | $spl92)),
  inference(superposition,[],[f9595,f22295])).
fof(f24529,plain,(
  ( ! [X239,X240,X238] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X240),s(bool,X239))) | s(bool,'T_0') = s(bool,X238) | s(bool,X239) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl13 | $spl22 | $spl92)),
  inference(duplicate_literal_removal,[],[f24091])).
fof(f24091,plain,(
  ( ! [X239,X240,X238] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X240),s(bool,X239))) | s(bool,'T_0') = s(bool,X238) | s(bool,'T_0') = s(bool,X238) | s(bool,X239) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl13 | $spl22 | $spl92)),
  inference(superposition,[],[f9573,f22295])).
fof(f25238,plain,(
  ( ! [X229,X227,X228] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X229),s(bool,X228))) | s(bool,'T_0') = s(bool,X227) | s(bool,X228) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl13 | $spl22 | $spl36 | $spl92)),
  inference(subsumption_resolution,[],[f24088,f17356])).
fof(f24088,plain,(
  ( ! [X231,X229,X227,X230,X228] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X229),s(bool,X228))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X230,bool),'const.pair.,_2'(s(X230,X231),s(bool,X227))))) | s(bool,'T_0') = s(bool,X227) | s(bool,X228) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl13 | $spl22 | $spl92)),
  inference(superposition,[],[f9272,f22295])).
fof(f25236,plain,(
  ( ! [X218,X219,X217] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X219),s(bool,X218))) | s(bool,'T_0') = s(bool,X217) | s(bool,X218) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl9 | $spl13 | $spl22 | $spl36 | $spl92)),
  inference(subsumption_resolution,[],[f24086,f15429])).
fof(f24086,plain,(
  ( ! [X218,X220,X219,X217,X221] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X219),s(bool,X218))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X220),'const.pair.,_2'(s(bool,X217),s(X220,X221))))) | s(bool,'T_0') = s(bool,X217) | s(bool,X218) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl9 | $spl13 | $spl22 | $spl92)),
  inference(superposition,[],[f8988,f22295])).
fof(f25231,plain,(
  ( ! [X180,X179] : (s(bool,X179) = s(bool,happ(s(fun(bool,bool),X180),s(bool,X179))) | s(bool,X179) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | $spl132),
  inference(cnf_transformation,[],[f25231_D])).
fof(f25231_D,plain,(
  ( ! [X180,X179] : (s(bool,X179) = s(bool,happ(s(fun(bool,bool),X180),s(bool,X179))) | s(bool,X179) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f25235,plain,(
  ( ! [X200,X199,X201] : (s(bool,X200) = s(bool,happ(s(fun(bool,bool),X201),s(bool,X200))) | s(bool,'T_0') = s(bool,X199) | s(bool,X200) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl13 | $spl22 | $spl36 | $spl92)),
  inference(subsumption_resolution,[],[f24082,f17356])).
fof(f24082,plain,(
  ( ! [X202,X200,X199,X203,X201] : (s(bool,X200) = s(bool,happ(s(fun(bool,bool),X201),s(bool,X200))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X202,bool),'const.pair.,_2'(s(X202,X203),s(bool,X199))))) | s(bool,'T_0') = s(bool,X199) | s(bool,X200) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl13 | $spl22 | $spl92)),
  inference(superposition,[],[f8148,f22295])).
fof(f8148,plain,(
  ( ! [X70,X72,X71,X69] : (s(bool,X71) = s(bool,happ(s(fun(bool,bool),X72),s(bool,X71))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X69,bool),'const.pair.,_2'(s(X69,X70),s(bool,X71)))))) ) | ($spl13 | $spl22 | $spl92)),
  inference(superposition,[],[f7469,f44])).
fof(f25232,plain,(
  ( ! [X187,X188,X186] : (s(bool,X187) = s(bool,happ(s(fun(bool,bool),X188),s(bool,X187))) | s(bool,'T_0') = s(bool,X186) | s(bool,X187) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl9 | $spl13 | $spl22 | $spl36 | $spl92)),
  inference(subsumption_resolution,[],[f24079,f15429])).
fof(f24079,plain,(
  ( ! [X189,X187,X190,X188,X186] : (s(bool,X187) = s(bool,happ(s(fun(bool,bool),X188),s(bool,X187))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X189),'const.pair.,_2'(s(bool,X186),s(X189,X190))))) | s(bool,'T_0') = s(bool,X186) | s(bool,X187) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl9 | $spl13 | $spl22 | $spl92)),
  inference(superposition,[],[f7526,f22295])).
fof(f7526,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,X15) = s(bool,happ(s(fun(bool,bool),X16),s(bool,X15))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X17),'const.pair.,_2'(s(bool,X15),s(X17,X18)))))) ) | ($spl9 | $spl13 | $spl22 | $spl92)),
  inference(subsumption_resolution,[],[f7476,f71])).
fof(f7476,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,X15) = s(bool,happ(s(fun(bool,bool),X16),s(bool,X15))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X17),'const.pair.,_2'(s(bool,X15),s(X17,X18)))))) ) | ($spl13 | $spl22 | $spl92)),
  inference(superposition,[],[f7444,f132])).
fof(f24537,plain,(
  ( ! [X180,X178,X179] : (s(bool,X179) = s(bool,happ(s(fun(bool,bool),X180),s(bool,X179))) | s(bool,'T_0') = s(bool,X178) | s(bool,X179) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl13 | $spl22 | $spl92)),
  inference(duplicate_literal_removal,[],[f24077])).
fof(f24077,plain,(
  ( ! [X180,X178,X179] : (s(bool,X179) = s(bool,happ(s(fun(bool,bool),X180),s(bool,X179))) | s(bool,'T_0') = s(bool,X178) | s(bool,'T_0') = s(bool,X178) | s(bool,X179) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl13 | $spl22 | $spl92)),
  inference(superposition,[],[f7469,f22295])).
fof(f25230,plain,(
  ( ! [X175,X174] : (s(bool,'T_0') != s(bool,X175) | s(bool,'T_0') = s(bool,X174) | s(bool,X175) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl13 | $spl22 | $spl36)),
  inference(subsumption_resolution,[],[f24076,f17356])).
fof(f24076,plain,(
  ( ! [X177,X175,X176,X174] : (s(bool,'T_0') != s(bool,X175) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X176,bool),'const.pair.,_2'(s(X176,X177),s(bool,X174))))) | s(bool,'T_0') = s(bool,X174) | s(bool,X175) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl13 | $spl22)),
  inference(superposition,[],[f7282,f22295])).
fof(f24550,plain,(
  ( ! [X152,X153] : (s(bool,'F_0') != s(bool,f(s(bool,X153))) | s(bool,'T_0') = s(bool,X152) | s(bool,X153) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl76)),
  inference(duplicate_literal_removal,[],[f24067])).
fof(f24067,plain,(
  ( ! [X152,X153] : (s(bool,'F_0') != s(bool,f(s(bool,X153))) | s(bool,'T_0') = s(bool,X152) | s(bool,'T_0') = s(bool,X152) | s(bool,X153) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl76)),
  inference(superposition,[],[f4984,f22295])).
fof(f25217,plain,(
  ( ! [X151,X150] : (s(bool,'T_0') != s(bool,X151) | s(bool,'T_0') = s(bool,X150) | s(bool,X151) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl36)),
  inference(forward_demodulation,[],[f24551,f11203])).
fof(f24551,plain,(
  ( ! [X151,X150] : (s(bool,X151) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X150) | s(bool,X151) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl36)),
  inference(duplicate_literal_removal,[],[f24066])).
fof(f24066,plain,(
  ( ! [X151,X150] : (s(bool,X151) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X150) | s(bool,'T_0') = s(bool,X150) | s(bool,X151) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | ($spl4 | $spl36)),
  inference(superposition,[],[f2455,f22295])).
fof(f24970,plain,(
  ( ! [X94,X92] : (s(bool,X94) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X92) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X92) = s(bool,X94)) ) | $spl124),
  inference(cnf_transformation,[],[f24970_D])).
fof(f24970_D,plain,(
  ( ! [X94,X92] : (s(bool,X94) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X92) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X92) = s(bool,X94)) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f24800,plain,(
  ( ! [X109,X110,X108] : (s(bool,'T_0') = s(bool,X109) | s(bool,X108) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X110) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X108) = s(bool,X110)) ) | $spl4),
  inference(duplicate_literal_removal,[],[f23487])).
fof(f23487,plain,(
  ( ! [X109,X110,X108] : (s(bool,'T_0') = s(bool,X109) | s(bool,'T_0') = s(bool,X109) | s(bool,X108) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X110) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X108) = s(bool,X110)) ) | $spl4),
  inference(superposition,[],[f22295,f22295])).
fof(f24822,plain,(
  ( ! [X111,X112,X113] : (s(bool,X112) = s(bool,X113) | s(bool,'T_0') = s(bool,X111) | s(bool,X113) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X112) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | $spl4),
  inference(duplicate_literal_removal,[],[f23414])).
fof(f23414,plain,(
  ( ! [X111,X112,X113] : (s(bool,X112) = s(bool,X113) | s(bool,'T_0') = s(bool,X111) | s(bool,X113) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,'T_0') = s(bool,X111) | s(bool,X112) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | $spl4),
  inference(superposition,[],[f22295,f22295])).
fof(f24971,plain,(
  ( ! [X103] : (s(bool,X103) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,'F_0') = s(bool,X103)) ) | $spl126),
  inference(cnf_transformation,[],[f24971_D])).
fof(f24971_D,plain,(
  ( ! [X103] : (s(bool,X103) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,'F_0') = s(bool,X103)) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f24854,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | $spl4),
  inference(duplicate_literal_removal,[],[f23387])).
fof(f23387,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,'T_0') = s(bool,X0)) ) | $spl4),
  inference(superposition,[],[f22295,f40])).
fof(f24930,plain,(
  ( ! [X103,X102] : (s(bool,'T_0') = s(bool,X102) | s(bool,X103) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,'F_0') = s(bool,X103)) ) | $spl4),
  inference(duplicate_literal_removal,[],[f23279])).
fof(f23279,plain,(
  ( ! [X103,X102] : (s(bool,'T_0') = s(bool,X102) | s(bool,X103) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,'F_0') = s(bool,X103) | s(bool,'T_0') = s(bool,X102)) ) | $spl4),
  inference(superposition,[],[f22295,f40])).
fof(f24931,plain,(
  ( ! [X94,X92,X93] : (s(bool,'T_0') = s(bool,X93) | s(bool,X94) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X92) = s(bool,X94) | s(bool,X92) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | $spl4),
  inference(duplicate_literal_removal,[],[f23278])).
fof(f23278,plain,(
  ( ! [X94,X92,X93] : (s(bool,'T_0') = s(bool,X93) | s(bool,X94) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2))) | s(bool,X92) = s(bool,X94) | s(bool,'T_0') = s(bool,X93) | s(bool,X92) = s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) ) | $spl4),
  inference(superposition,[],[f22295,f22295])).
fof(f16614,plain,(
  ( ! [X243,X246,X242,X248,X247] : (s(bool,X246) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,happ(s(fun(bool,bool),X242),s(bool,X243))),s(bool,X246))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X247,bool),'const.pair.,_2'(s(X247,X248),s(bool,X243))))) | s(bool,'T_0') = s(bool,X246)) ) | $spl122),
  inference(cnf_transformation,[],[f16614_D])).
fof(f16614_D,plain,(
  ( ! [X243,X246,X242,X248,X247] : (s(bool,X246) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,happ(s(fun(bool,bool),X242),s(bool,X243))),s(bool,X246))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X247,bool),'const.pair.,_2'(s(X247,X248),s(bool,X243))))) | s(bool,'T_0') = s(bool,X246)) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f16423,plain,(
  ( ! [X243,X246,X244,X242,X248,X247,X245] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X244),'const.pair.,_2'(s(bool,'F_0'),s(X244,X245))))) | s(bool,X246) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,happ(s(fun(bool,bool),X242),s(bool,X243))),s(bool,X246))))) | s(bool,'T_0') = s(bool,X246) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X247,bool),'const.pair.,_2'(s(X247,X248),s(bool,X243)))))) ) | ($spl13 | $spl16 | $spl22 | $spl92)),
  inference(superposition,[],[f1557,f9272])).
fof(f11670,plain,(
  ( ! [X142,X140,X145,X141] : (s(bool,happ(s(fun(bool,bool),X140),s(bool,happ(s(fun(bool,bool),X141),s(bool,X142))))) = s(bool,happ(s(fun(bool,bool),X145),s(bool,happ(s(fun(bool,bool),X140),s(bool,happ(s(fun(bool,bool),X141),s(bool,X142))))))) | s(bool,'T_0') = s(bool,X142)) ) | $spl120),
  inference(cnf_transformation,[],[f11670_D])).
fof(f11670_D,plain,(
  ( ! [X142,X140,X145,X141] : (s(bool,happ(s(fun(bool,bool),X140),s(bool,happ(s(fun(bool,bool),X141),s(bool,X142))))) = s(bool,happ(s(fun(bool,bool),X145),s(bool,happ(s(fun(bool,bool),X140),s(bool,happ(s(fun(bool,bool),X141),s(bool,X142))))))) | s(bool,'T_0') = s(bool,X142)) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f11395,plain,(
  ( ! [X144,X142,X140,X145,X143,X141] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X143),'const.pair.,_2'(s(bool,'F_0'),s(X143,X144))))) | s(bool,happ(s(fun(bool,bool),X140),s(bool,happ(s(fun(bool,bool),X141),s(bool,X142))))) = s(bool,happ(s(fun(bool,bool),X145),s(bool,happ(s(fun(bool,bool),X140),s(bool,happ(s(fun(bool,bool),X141),s(bool,X142))))))) | s(bool,'T_0') = s(bool,X142)) ) | ($spl9 | $spl13 | $spl22 | $spl92)),
  inference(superposition,[],[f7526,f9720])).
fof(f10183,plain,(
  ( ! [X94,X95,X93,X98] : (s(bool,X98) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,happ(s(fun(bool,bool),X93),s(bool,happ(s(fun(bool,bool),X94),s(bool,X95))))),s(bool,X98))))) | s(bool,'T_0') = s(bool,X95) | s(bool,'T_0') = s(bool,X98)) ) | $spl118),
  inference(cnf_transformation,[],[f10183_D])).
fof(f10183_D,plain,(
  ( ! [X94,X95,X93,X98] : (s(bool,X98) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,happ(s(fun(bool,bool),X93),s(bool,happ(s(fun(bool,bool),X94),s(bool,X95))))),s(bool,X98))))) | s(bool,'T_0') = s(bool,X95) | s(bool,'T_0') = s(bool,X98)) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f10103,plain,(
  ( ! [X94,X97,X95,X93,X98,X96] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X96),'const.pair.,_2'(s(bool,'F_0'),s(X96,X97))))) | s(bool,X98) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,happ(s(fun(bool,bool),X93),s(bool,happ(s(fun(bool,bool),X94),s(bool,X95))))),s(bool,X98))))) | s(bool,'T_0') = s(bool,X98) | s(bool,'T_0') = s(bool,X95)) ) | ($spl13 | $spl16 | $spl22 | $spl92)),
  inference(superposition,[],[f1557,f9720])).
fof(f9590,plain,(
  ( ! [X76,X77,X75] : (s(bool,X77) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,happ(s(fun(bool,bool),X75),s(bool,X76))),s(bool,X77))))) | s(bool,'T_0') = s(bool,X76) | s(bool,'T_0') = s(bool,X77)) ) | $spl114),
  inference(cnf_transformation,[],[f9590_D])).
fof(f9590_D,plain,(
  ( ! [X76,X77,X75] : (s(bool,X77) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,happ(s(fun(bool,bool),X75),s(bool,X76))),s(bool,X77))))) | s(bool,'T_0') = s(bool,X76) | s(bool,'T_0') = s(bool,X77)) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f960,plain,(
  ( ! [X15,X11,X16] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X15),'const.pair.,_2'(s(bool,X11),s(X15,X16)))))) ) | $spl14),
  inference(cnf_transformation,[],[f960_D])).
fof(f960_D,plain,(
  ( ! [X15,X11,X16] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X15),'const.pair.,_2'(s(bool,X11),s(X15,X16)))))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f9612,plain,(
  ( ! [X94,X92,X90,X95,X93,X91] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X93),'const.pair.,_2'(s(bool,X92),s(X93,X94))))) | s(bool,X95) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,happ(s(fun(bool,bool),X90),s(bool,X91))),s(bool,X95))))) | s(bool,'T_0') = s(bool,X95) | s(bool,'T_0') = s(bool,X91)) ) | ($spl9 | $spl13 | $spl16 | $spl22 | $spl92)),
  inference(subsumption_resolution,[],[f9456,f4170])).
fof(f9456,plain,(
  ( ! [X94,X92,X90,X95,X93,X91] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X93),'const.pair.,_2'(s(bool,X92),s(X93,X94))))) | s(bool,X95) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,happ(s(fun(bool,bool),X90),s(bool,X91))),s(bool,X95))))) | s(bool,'T_0') = s(bool,X95) | s(bool,'T_0') = s(bool,X91) | s(bool,'T_0') = s(bool,X92)) ) | ($spl13 | $spl16 | $spl22 | $spl92)),
  inference(superposition,[],[f1557,f8188])).
fof(f9594,plain,(
  ( ! [X111,X110] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X110),s(bool,X111))) | s(bool,'T_0') = s(bool,X111)) ) | $spl116),
  inference(cnf_transformation,[],[f9594_D])).
fof(f9594_D,plain,(
  ( ! [X111,X110] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X110),s(bool,X111))) | s(bool,'T_0') = s(bool,X111)) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f9593,plain,(
  ( ! [X111,X109,X112,X110,X113] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X110),s(bool,X111))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X112,bool),'const.pair.,_2'(s(X112,X113),s(bool,X109))))) | s(bool,'T_0') = s(bool,X111)) ) | ($spl13 | $spl22 | $spl92)),
  inference(subsumption_resolution,[],[f9397,f7282])).
fof(f9397,plain,(
  ( ! [X111,X109,X112,X110,X113] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X110),s(bool,X111))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X112,bool),'const.pair.,_2'(s(X112,X113),s(bool,X109))))) | s(bool,'T_0') = s(bool,X111) | s(bool,'T_0') = s(bool,X109)) ) | ($spl13 | $spl22 | $spl92)),
  inference(superposition,[],[f7282,f8188])).
fof(f9589,plain,(
  ( ! [X78,X76,X74,X79,X77,X75] : (s(bool,X77) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,happ(s(fun(bool,bool),X75),s(bool,X76))),s(bool,X77))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X78),'const.pair.,_2'(s(bool,X74),s(X78,X79))))) | s(bool,'T_0') = s(bool,X77) | s(bool,'T_0') = s(bool,X76)) ) | ($spl9 | $spl13 | $spl16 | $spl22 | $spl92)),
  inference(subsumption_resolution,[],[f9388,f4170])).
fof(f9388,plain,(
  ( ! [X78,X76,X74,X79,X77,X75] : (s(bool,X77) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,happ(s(fun(bool,bool),X75),s(bool,X76))),s(bool,X77))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X78),'const.pair.,_2'(s(bool,X74),s(X78,X79))))) | s(bool,'T_0') = s(bool,X77) | s(bool,'T_0') = s(bool,X76) | s(bool,'T_0') = s(bool,X74)) ) | ($spl13 | $spl16 | $spl22 | $spl92)),
  inference(superposition,[],[f1557,f8188])).
fof(f7913,plain,(
  ( ! [X70,X74,X72,X71,X73] : (p(s(bool,happ(s(fun(X72,bool),X73),s(X72,'const.pair.SND_1'(s('type.pair.prod'(X70,X72),X74)))))) | s(bool,'T_0') = s(bool,'const.pair.PROD_ALL_3'(s(fun(X70,bool),X71),s(fun(X72,bool),X73),s('type.pair.prod'(X70,X72),X74)))) ) | $spl112),
  inference(cnf_transformation,[],[f7913_D])).
fof(f7913_D,plain,(
  ( ! [X70,X74,X72,X71,X73] : (p(s(bool,happ(s(fun(X72,bool),X73),s(X72,'const.pair.SND_1'(s('type.pair.prod'(X70,X72),X74)))))) | s(bool,'T_0') = s(bool,'const.pair.PROD_ALL_3'(s(fun(X70,bool),X71),s(fun(X72,bool),X73),s('type.pair.prod'(X70,X72),X74)))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f7910,plain,(
  ~p(s(bool,g(s(bool,'F_0')))) | $spl107),
  inference(cnf_transformation,[],[f7910_D])).
fof(f7910_D,plain,(
  ~p(s(bool,g(s(bool,'F_0')))) <=> ~$spl107),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl107])])).
fof(f7613,plain,(
  ( ! [X70,X74,X72,X71,X73] : (~p(s(bool,g(s(bool,'F_0')))) | p(s(bool,happ(s(fun(X72,bool),X73),s(X72,'const.pair.SND_1'(s('type.pair.prod'(X70,X72),X74)))))) | s(bool,'T_0') = s(bool,'const.pair.PROD_ALL_3'(s(fun(X70,bool),X71),s(fun(X72,bool),X73),s('type.pair.prod'(X70,X72),X74)))) ) | $spl32),
  inference(superposition,[],[f47,f2452])).
fof(f7912,plain,(
  ( ! [X78,X76,X79,X77,X75] : (p(s(bool,happ(s(fun(X75,bool),X76),s(X75,'const.pair.FST_1'(s('type.pair.prod'(X75,X77),X79)))))) | s(bool,'T_0') = s(bool,'const.pair.PROD_ALL_3'(s(fun(X75,bool),X76),s(fun(X77,bool),X78),s('type.pair.prod'(X75,X77),X79)))) ) | $spl110),
  inference(cnf_transformation,[],[f7912_D])).
fof(f7912_D,plain,(
  ( ! [X78,X76,X79,X77,X75] : (p(s(bool,happ(s(fun(X75,bool),X76),s(X75,'const.pair.FST_1'(s('type.pair.prod'(X75,X77),X79)))))) | s(bool,'T_0') = s(bool,'const.pair.PROD_ALL_3'(s(fun(X75,bool),X76),s(fun(X77,bool),X78),s('type.pair.prod'(X75,X77),X79)))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f7614,plain,(
  ( ! [X78,X76,X79,X77,X75] : (~p(s(bool,g(s(bool,'F_0')))) | p(s(bool,happ(s(fun(X75,bool),X76),s(X75,'const.pair.FST_1'(s('type.pair.prod'(X75,X77),X79)))))) | s(bool,'T_0') = s(bool,'const.pair.PROD_ALL_3'(s(fun(X75,bool),X76),s(fun(X77,bool),X78),s('type.pair.prod'(X75,X77),X79)))) ) | $spl32),
  inference(superposition,[],[f46,f2452])).
fof(f7911,plain,(
  ( ! [X80,X83,X81,X84,X82] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X82,bool),X84),s(fun(X80,bool),X81),s('type.pair.prod'(X82,X80),X83)))) | s(bool,'T_0') = s(bool,happ(s(fun(X80,bool),X81),s(X80,'const.pair.SND_1'(s('type.pair.prod'(X82,X80),X83))))) | ~p(s(bool,happ(s(fun(X82,bool),X84),s(X82,'const.pair.FST_1'(s('type.pair.prod'(X82,X80),X83))))))) ) | $spl108),
  inference(cnf_transformation,[],[f7911_D])).
fof(f7911_D,plain,(
  ( ! [X80,X83,X81,X84,X82] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X82,bool),X84),s(fun(X80,bool),X81),s('type.pair.prod'(X82,X80),X83)))) | s(bool,'T_0') = s(bool,happ(s(fun(X80,bool),X81),s(X80,'const.pair.SND_1'(s('type.pair.prod'(X82,X80),X83))))) | ~p(s(bool,happ(s(fun(X82,bool),X84),s(X82,'const.pair.FST_1'(s('type.pair.prod'(X82,X80),X83))))))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f7615,plain,(
  ( ! [X80,X83,X81,X84,X82] : (~p(s(bool,g(s(bool,'F_0')))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X82,bool),X84),s(fun(X80,bool),X81),s('type.pair.prod'(X82,X80),X83)))) | ~p(s(bool,happ(s(fun(X82,bool),X84),s(X82,'const.pair.FST_1'(s('type.pair.prod'(X82,X80),X83)))))) | s(bool,'T_0') = s(bool,happ(s(fun(X80,bool),X81),s(X80,'const.pair.SND_1'(s('type.pair.prod'(X82,X80),X83)))))) ) | $spl32),
  inference(superposition,[],[f48,f2452])).
fof(f7908,plain,(
  ( ! [X90,X88,X87,X85,X89,X86] : (~p(s(bool,happ(s(fun(X88,bool),X89),s(X88,X90)))) | s(bool,'T_0') = s(bool,happ(s(fun(X85,bool),X86),s(X85,X87))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X88,bool),X89),s(fun(X85,bool),X86),s('type.pair.prod'(X88,X85),'const.pair.,_2'(s(X88,X90),s(X85,X87))))))) ) | $spl104),
  inference(cnf_transformation,[],[f7908_D])).
fof(f7908_D,plain,(
  ( ! [X90,X88,X87,X85,X89,X86] : (~p(s(bool,happ(s(fun(X88,bool),X89),s(X88,X90)))) | s(bool,'T_0') = s(bool,happ(s(fun(X85,bool),X86),s(X85,X87))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X88,bool),X89),s(fun(X85,bool),X86),s('type.pair.prod'(X88,X85),'const.pair.,_2'(s(X88,X90),s(X85,X87))))))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f7616,plain,(
  ( ! [X90,X88,X87,X85,X89,X86] : (~p(s(bool,g(s(bool,'F_0')))) | ~p(s(bool,happ(s(fun(X88,bool),X89),s(X88,X90)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X88,bool),X89),s(fun(X85,bool),X86),s('type.pair.prod'(X88,X85),'const.pair.,_2'(s(X88,X90),s(X85,X87)))))) | s(bool,'T_0') = s(bool,happ(s(fun(X85,bool),X86),s(X85,X87)))) ) | $spl32),
  inference(superposition,[],[f785,f2452])).
fof(f7907,plain,(
  ( ! [X91] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X91),s(bool,'F_0')))) ) | $spl102),
  inference(cnf_transformation,[],[f7907_D])).
fof(f7907_D,plain,(
  ( ! [X91] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X91),s(bool,'F_0')))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f7900,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | $spl98),
  inference(cnf_transformation,[],[f7900_D])).
fof(f7900_D,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f7617,plain,(
  ( ! [X91] : (s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X91),s(bool,'F_0')))) ) | ($spl13 | $spl22 | $spl32 | $spl92)),
  inference(superposition,[],[f7444,f2452])).
fof(f7905,plain,(
  ( ! [X101,X99,X100] : (s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X99),'const.pair.,_2'(s(bool,X100),s(X99,X101)))))) ) | ($spl9 | $spl16 | $spl32)),
  inference(subsumption_resolution,[],[f7622,f4170])).
fof(f7622,plain,(
  ( ! [X101,X99,X100] : (s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X100) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X99),'const.pair.,_2'(s(bool,X100),s(X99,X101)))))) ) | ($spl16 | $spl32)),
  inference(superposition,[],[f961,f2452])).
fof(f7902,plain,(
  ( ! [X12,X13,X11] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X12),'const.pair.,_2'(s(bool,X11),s(X12,X13))))) | s(bool,'F_0') = s(bool,g(s(bool,'F_0')))) ) | ($spl9 | $spl32)),
  inference(subsumption_resolution,[],[f7658,f4170])).
fof(f7658,plain,(
  ( ! [X12,X13,X11] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X12),'const.pair.,_2'(s(bool,X11),s(X12,X13))))) | s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X11)) ) | $spl32),
  inference(superposition,[],[f82,f2452])).
fof(f7898,plain,(
  ( ! [X14,X15,X16] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X15,bool),'const.pair.,_2'(s(X15,X16),s(bool,X14))))) | s(bool,'F_0') = s(bool,g(s(bool,'F_0')))) ) | ($spl13 | $spl22 | $spl32)),
  inference(subsumption_resolution,[],[f7659,f7282])).
fof(f7659,plain,(
  ( ! [X14,X15,X16] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X15,bool),'const.pair.,_2'(s(X15,X16),s(bool,X14))))) | s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X14)) ) | $spl32),
  inference(superposition,[],[f104,f2452])).
fof(f7894,plain,(
  ( ! [X50] : (s(bool,X50) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,g(s(bool,'F_0'))),s(bool,X50))))) | s(bool,'T_0') = s(bool,X50)) ) | $spl96),
  inference(cnf_transformation,[],[f7894_D])).
fof(f7894_D,plain,(
  ( ! [X50] : (s(bool,X50) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,g(s(bool,'F_0'))),s(bool,X50))))) | s(bool,'T_0') = s(bool,X50)) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f7893,plain,(
  ( ! [X47,X50,X48,X49] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X48),'const.pair.,_2'(s(bool,X47),s(X48,X49))))) | s(bool,X50) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,g(s(bool,'F_0'))),s(bool,X50))))) | s(bool,'T_0') = s(bool,X50)) ) | ($spl9 | $spl16 | $spl32)),
  inference(subsumption_resolution,[],[f7668,f4170])).
fof(f7668,plain,(
  ( ! [X47,X50,X48,X49] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X48),'const.pair.,_2'(s(bool,X47),s(X48,X49))))) | s(bool,X50) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,g(s(bool,'F_0'))),s(bool,X50))))) | s(bool,'T_0') = s(bool,X50) | s(bool,'T_0') = s(bool,X47)) ) | ($spl16 | $spl32)),
  inference(superposition,[],[f1557,f2452])).
fof(f7887,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'F_0')))) ) | $spl94),
  inference(cnf_transformation,[],[f7887_D])).
fof(f7887_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'F_0')))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f2447,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'F_0'))) | $spl27),
  inference(cnf_transformation,[],[f2447_D])).
fof(f2447_D,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'F_0'))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f7688,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,g(s(bool,'F_0'))) | s(bool,X0) = s(bool,g(s(bool,'F_0')))) ) | $spl32),
  inference(equality_factoring,[],[f2452])).
fof(f7876,plain,(
  $false | ($spl13 | $spl22 | $spl26 | $spl32)),
  inference(subsumption_resolution,[],[f7821,f7735])).
fof(f7735,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2)) ) | ($spl26 | $spl32)),
  inference(duplicate_literal_removal,[],[f7570])).
fof(f7570,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl26 | $spl32)),
  inference(superposition,[],[f2452,f2448])).
fof(f7821,plain,(
  s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'T_0'),s(bool,'T_0'))))) | ($spl13 | $spl22 | $spl26 | $spl32)),
  inference(backward_demodulation,[],[f7735,f6406])).
fof(f7875,plain,(
  $false | ($spl26 | $spl32)),
  inference(subsumption_resolution,[],[f7795,f37])).
fof(f7795,plain,(
  ~p(s(bool,'T_0')) | ($spl26 | $spl32)),
  inference(backward_demodulation,[],[f7735,f39])).
fof(f7794,plain,(
  $false | ($spl9 | $spl13 | $spl22 | $spl26 | $spl32 | $spl92)),
  inference(subsumption_resolution,[],[f7793,f71])).
fof(f7793,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl9 | $spl13 | $spl22 | $spl26 | $spl32 | $spl92)),
  inference(backward_demodulation,[],[f7792,f7444])).
fof(f7792,plain,(
  ( ! [X32] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X32),s(bool,'F_0')))) ) | ($spl9 | $spl13 | $spl22 | $spl26 | $spl32 | $spl92)),
  inference(subsumption_resolution,[],[f7791,f71])).
fof(f7791,plain,(
  ( ! [X32] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X32),s(bool,'F_0')))) ) | ($spl13 | $spl22 | $spl26 | $spl32 | $spl92)),
  inference(forward_demodulation,[],[f7548,f2448])).
fof(f7548,plain,(
  ( ! [X32] : (s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X32),s(bool,'F_0')))) ) | ($spl13 | $spl22 | $spl32 | $spl92)),
  inference(superposition,[],[f2452,f7444])).
fof(f7790,plain,(
  $false | ($spl1 | $spl9 | $spl13 | $spl26 | $spl32)),
  inference(subsumption_resolution,[],[f7789,f71])).
fof(f7789,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl9 | $spl13 | $spl26 | $spl32)),
  inference(forward_demodulation,[],[f7788,f2448])).
fof(f7788,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | ($spl1 | $spl9 | $spl13 | $spl32)),
  inference(subsumption_resolution,[],[f7544,f439])).
fof(f7544,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2))))) | ($spl1 | $spl9 | $spl32)),
  inference(superposition,[],[f2452,f1103])).
fof(f7786,plain,(
  $false | ($spl9 | $spl22 | $spl26 | $spl32)),
  inference(subsumption_resolution,[],[f7785,f71])).
fof(f7785,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl22 | $spl26 | $spl32)),
  inference(condensation,[],[f7784])).
fof(f7784,plain,(
  ( ! [X25] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X25)) ) | ($spl22 | $spl26 | $spl32)),
  inference(forward_demodulation,[],[f7783,f2448])).
fof(f7783,plain,(
  ( ! [X25] : (s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X25)) ) | ($spl22 | $spl32)),
  inference(forward_demodulation,[],[f7737,f43])).
fof(f7737,plain,(
  ( ! [X25] : (s(bool,g(s(bool,'F_0'))) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X25)) ) | ($spl22 | $spl32)),
  inference(duplicate_literal_removal,[],[f7542])).
fof(f7542,plain,(
  ( ! [X25] : (s(bool,g(s(bool,'F_0'))) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) ) | ($spl22 | $spl32)),
  inference(superposition,[],[f2452,f5991])).
fof(f7782,plain,(
  $false | ($spl9 | $spl22 | $spl26 | $spl32)),
  inference(subsumption_resolution,[],[f7781,f71])).
fof(f7781,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl22 | $spl26 | $spl32)),
  inference(condensation,[],[f7780])).
fof(f7780,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X0)) ) | ($spl22 | $spl26 | $spl32)),
  inference(forward_demodulation,[],[f7779,f2448])).
fof(f7779,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl22 | $spl32)),
  inference(forward_demodulation,[],[f7739,f43])).
fof(f7739,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,g(s(bool,'F_0'))) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,X0)))))) ) | ($spl22 | $spl32)),
  inference(condensation,[],[f7738])).
fof(f7738,plain,(
  ( ! [X24,X23] : (s(bool,g(s(bool,'F_0'))) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,X24))))) | s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,X24)) ) | ($spl22 | $spl32)),
  inference(duplicate_literal_removal,[],[f7541])).
fof(f7541,plain,(
  ( ! [X24,X23] : (s(bool,g(s(bool,'F_0'))) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,X24))))) | s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X23)) ) | ($spl22 | $spl32)),
  inference(superposition,[],[f2452,f5520])).
fof(f7778,plain,(
  $false | ($spl1 | $spl9 | $spl26 | $spl32)),
  inference(subsumption_resolution,[],[f7777,f71])).
fof(f7777,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl9 | $spl26 | $spl32)),
  inference(condensation,[],[f7776])).
fof(f7776,plain,(
  ( ! [X21] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X21)) ) | ($spl1 | $spl9 | $spl26 | $spl32)),
  inference(forward_demodulation,[],[f7775,f2448])).
fof(f7775,plain,(
  ( ! [X21] : (s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X21)) ) | ($spl1 | $spl9 | $spl32)),
  inference(forward_demodulation,[],[f7740,f1103])).
fof(f7740,plain,(
  ( ! [X21] : (s(bool,g(s(bool,'F_0'))) = s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2))))) | s(bool,'T_0') = s(bool,X21)) ) | ($spl1 | $spl32)),
  inference(duplicate_literal_removal,[],[f7540])).
fof(f7540,plain,(
  ( ! [X21] : (s(bool,g(s(bool,'F_0'))) = s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2))))) | s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) ) | ($spl1 | $spl32)),
  inference(superposition,[],[f2452,f1101])).
fof(f7769,plain,(
  $false | ($spl9 | $spl16 | $spl26 | $spl32)),
  inference(subsumption_resolution,[],[f7768,f71])).
fof(f7768,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl9 | $spl16 | $spl26 | $spl32)),
  inference(condensation,[],[f7761])).
fof(f7761,plain,(
  ( ! [X12] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X12)) ) | ($spl9 | $spl16 | $spl26 | $spl32)),
  inference(backward_demodulation,[],[f7760,f961])).
fof(f7760,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,X1),s(X0,X2)))))) ) | ($spl9 | $spl26 | $spl32)),
  inference(subsumption_resolution,[],[f7759,f4170])).
fof(f7759,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,X1),s(X0,X2)))))) ) | ($spl26 | $spl32)),
  inference(forward_demodulation,[],[f7742,f2448])).
fof(f7742,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,X1),s(X0,X2))))) | s(bool,X1) = s(bool,g(s(bool,'F_0')))) ) | $spl32),
  inference(condensation,[],[f7741])).
fof(f7741,plain,(
  ( ! [X14,X15,X13,X16] : (s(bool,X14) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X15),'const.pair.,_2'(s(bool,X14),s(X15,X16)))))) ) | $spl32),
  inference(duplicate_literal_removal,[],[f7538])).
fof(f7538,plain,(
  ( ! [X14,X15,X13,X16] : (s(bool,X14) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X15),'const.pair.,_2'(s(bool,X14),s(X15,X16)))))) ) | $spl32),
  inference(superposition,[],[f2452,f132])).
fof(f7758,plain,(
  $false | ($spl9 | $spl26 | $spl32)),
  inference(subsumption_resolution,[],[f7757,f71])).
fof(f7757,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl26 | $spl32)),
  inference(condensation,[],[f7756])).
fof(f7756,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X0)) ) | ($spl26 | $spl32)),
  inference(forward_demodulation,[],[f7747,f2448])).
fof(f7747,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl32),
  inference(duplicate_literal_removal,[],[f7532])).
fof(f7532,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl32),
  inference(superposition,[],[f2452,f40])).
fof(f7755,plain,(
  $false | ($spl9 | $spl26 | $spl32)),
  inference(subsumption_resolution,[],[f7754,f71])).
fof(f7754,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl9 | $spl26 | $spl32)),
  inference(forward_demodulation,[],[f7529,f2448])).
fof(f7529,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | ($spl9 | $spl32)),
  inference(unit_resulting_resolution,[],[f71,f2452])).
fof(f7753,plain,(
  $false | ($spl1 | $spl9 | $spl13 | $spl26 | $spl32)),
  inference(subsumption_resolution,[],[f7752,f71])).
fof(f7752,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl9 | $spl13 | $spl26 | $spl32)),
  inference(forward_demodulation,[],[f7751,f2448])).
fof(f7751,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | ($spl1 | $spl9 | $spl13 | $spl32)),
  inference(forward_demodulation,[],[f7530,f1103])).
fof(f7530,plain,(
  s(bool,g(s(bool,'F_0'))) = s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2))))) | ($spl13 | $spl32)),
  inference(unit_resulting_resolution,[],[f439,f2452])).
fof(f7750,plain,(
  $false | ($spl9 | $spl13 | $spl22 | $spl26 | $spl32)),
  inference(subsumption_resolution,[],[f7749,f71])).
fof(f7749,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl13 | $spl22 | $spl26 | $spl32)),
  inference(forward_demodulation,[],[f7748,f2448])).
fof(f7748,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | ($spl13 | $spl22 | $spl32)),
  inference(forward_demodulation,[],[f7531,f43])).
fof(f7531,plain,(
  s(bool,g(s(bool,'F_0'))) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,'F_0'))))) | ($spl13 | $spl22 | $spl32)),
  inference(unit_resulting_resolution,[],[f6406,f2452])).
fof(f7404,plain,(
  ( ! [X4,X2,X0,X3] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X2,bool),X3),s(fun(bool,bool),X0),s('type.pair.prod'(X2,bool),X4)))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X2,bool),X4))) | ~p(s(bool,happ(s(fun(X2,bool),X3),s(X2,'const.pair.FST_1'(s('type.pair.prod'(X2,bool),X4))))))) ) | $spl90),
  inference(cnf_transformation,[],[f7404_D])).
fof(f7404_D,plain,(
  ( ! [X4,X2,X0,X3] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X2,bool),X3),s(fun(bool,bool),X0),s('type.pair.prod'(X2,bool),X4)))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X2,bool),X4))) | ~p(s(bool,happ(s(fun(X2,bool),X3),s(X2,'const.pair.FST_1'(s('type.pair.prod'(X2,bool),X4))))))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f7381,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,happ(s(fun(bool,bool),X6),s(bool,'F_0')))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X7,bool),X8),s(fun(bool,bool),X5),s('type.pair.prod'(X7,bool),X9)))) | ~p(s(bool,happ(s(fun(X7,bool),X8),s(X7,'const.pair.FST_1'(s('type.pair.prod'(X7,bool),X9)))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X7,bool),X9)))) ) | ($spl9 | $spl28 | $spl41 | $spl46 | $spl74)),
  inference(superposition,[],[f761,f5071])).
fof(f5071,plain,(
  ( ! [X8,X9] : (s(fun(bool,bool),X8) = s(fun(bool,bool),X9)) ) | ($spl9 | $spl28 | $spl41 | $spl46 | $spl74)),
  inference(global_subsumption,[],[f4983,f36,f35,f34,f37,f39,f40,f43,f44,f45,f56,f58,f64,f82,f104,f132,f133,f146,f156,f157,f200,f361,f363,f365,f761,f762,f763,f785,f48,f47,f46,f475,f550,f786,f366,f368,f370,f71,f1480,f1935,f2442,f2443,f2444,f2445,f2446,f2548,f2551,f2559,f2562,f2698,f2699,f2700,f2704,f2706,f2763,f2771,f2772,f2797,f2874,f2449,f3301,f3278,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f4154,f3646,f3651,f4170,f4171,f3726,f3736,f3758,f3759,f3827,f3832,f3890,f3878,f2550,f2549,f3303,f3652,f3732,f1972,f1473,f3828,f3727])).
fof(f3736,plain,(
  ( ! [X21,X19,X20,X18] : (s(bool,X18) = s(bool,f(s(bool,X18))) | s(bool,X18) = s(bool,X19) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X20),'const.pair.,_2'(s(bool,X19),s(X20,X21)))))) ) | ($spl9 | $spl41)),
  inference(superposition,[],[f2771,f132])).
fof(f3278,plain,(
  ( ! [X103,X101,X99,X102,X100,X98] : (~p(s(bool,happ(s(fun(bool,bool),X100),s(bool,'T_0')))) | ~p(s(bool,happ(s(fun(X101,bool),X102),s(X101,X103)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X101,bool),X102),s(fun(bool,bool),X100),s('type.pair.prod'(X101,bool),'const.pair.,_2'(s(X101,X103),s(bool,sK6(X98,X99,bool,bool))))))) | s(fun(bool,bool),X98) = s(fun(bool,bool),X99)) ) | $spl28),
  inference(superposition,[],[f785,f2449])).
fof(f2797,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,f(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | ($spl9 | $spl41)),
  inference(superposition,[],[f2771,f56])).
fof(f2772,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl41),
  inference(superposition,[],[f2562,f56])).
fof(f2771,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | ($spl9 | $spl41)),
  inference(unit_resulting_resolution,[],[f71,f2562,f56])).
fof(f7380,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X2,bool),X3),s(fun(bool,bool),X0),s('type.pair.prod'(X2,bool),X4)))) | ~p(s(bool,happ(s(fun(X2,bool),X3),s(X2,'const.pair.FST_1'(s('type.pair.prod'(X2,bool),X4)))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X2,bool),X4)))) ) | ($spl9 | $spl28 | $spl41 | $spl46 | $spl74)),
  inference(superposition,[],[f761,f5071])).
fof(f6754,plain,(
  ( ! [X37,X36] : (s(fun(fun(bool,bool),bool),X36) = s(fun(fun(bool,bool),bool),X37) | s(bool,'T_0') = s(bool,g(s(fun(bool,bool),sK6(X36,X37,bool,fun(bool,bool)))))) ) | $spl86),
  inference(cnf_transformation,[],[f6754_D])).
fof(f6754_D,plain,(
  ( ! [X37,X36] : (s(fun(fun(bool,bool),bool),X36) = s(fun(fun(bool,bool),bool),X37) | s(bool,'T_0') = s(bool,g(s(fun(bool,bool),sK6(X36,X37,bool,fun(bool,bool)))))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f6755,plain,(
  ( ! [X38] : (s(bool,'F_0') != s(bool,f(s(fun(bool,bool),X38)))) ) | $spl88),
  inference(cnf_transformation,[],[f6755_D])).
fof(f6755_D,plain,(
  ( ! [X38] : (s(bool,'F_0') != s(bool,f(s(fun(bool,bool),X38)))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f6739,plain,(
  ( ! [X39,X41,X40] : (s(bool,'F_0') != s(bool,f(s(fun(bool,bool),X41))) | s(fun(fun(bool,bool),bool),X39) = s(fun(fun(bool,bool),bool),X40) | s(bool,'T_0') = s(bool,g(s(fun(bool,bool),sK6(X39,X40,bool,fun(bool,bool)))))) ) | ($spl9 | $spl28 | $spl41 | $spl46 | $spl74)),
  inference(superposition,[],[f370,f5071])).
fof(f6738,plain,(
  ( ! [X37,X38,X36] : (s(bool,'F_0') != s(bool,f(s(fun(bool,bool),X38))) | s(fun(fun(bool,bool),bool),X36) = s(fun(fun(bool,bool),bool),X37) | s(bool,'T_0') = s(bool,g(s(fun(bool,bool),sK6(X36,X37,bool,fun(bool,bool)))))) ) | ($spl9 | $spl28 | $spl41 | $spl46 | $spl74)),
  inference(superposition,[],[f370,f5071])).
fof(f6722,plain,(
  $false | ($spl9 | $spl13 | $spl16 | $spl22 | $spl24)),
  inference(subsumption_resolution,[],[f6721,f71])).
fof(f6721,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl9 | $spl13 | $spl16 | $spl22 | $spl24)),
  inference(condensation,[],[f6714])).
fof(f6714,plain,(
  ( ! [X12] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X12)) ) | ($spl9 | $spl13 | $spl16 | $spl22 | $spl24)),
  inference(backward_demodulation,[],[f6713,f961])).
fof(f6713,plain,(
  ( ! [X61,X59,X60] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X60),'const.pair.,_2'(s(bool,X59),s(X60,X61)))))) ) | ($spl9 | $spl13 | $spl22 | $spl24)),
  inference(subsumption_resolution,[],[f6711,f71])).
fof(f6711,plain,(
  ( ! [X61,X59,X60] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X60),'const.pair.,_2'(s(bool,X59),s(X60,X61)))))) ) | ($spl13 | $spl22 | $spl24)),
  inference(backward_demodulation,[],[f6710,f1473])).
fof(f6710,plain,(
  ( ! [X59,X60,X58] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X58,bool),'const.pair.,_2'(s(X58,X59),s(bool,X60)))))) ) | ($spl13 | $spl22 | $spl24)),
  inference(subsumption_resolution,[],[f6662,f2300])).
fof(f6662,plain,(
  ( ! [X59,X60,X58] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X60),s(bool,X60))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X58,bool),'const.pair.,_2'(s(X58,X59),s(bool,X60)))))) ) | ($spl13 | $spl22)),
  inference(superposition,[],[f6455,f44])).
fof(f6455,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X0),s(bool,X0))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl13 | $spl22)),
  inference(superposition,[],[f6406,f40])).
fof(f5066,plain,(
  $false | ($spl35 | $spl36 | $spl42)),
  inference(subsumption_resolution,[],[f5028,f4999])).
fof(f4999,plain,(
  ( ! [X63] : (s(bool,'T_0') = s(bool,X63)) ) | ($spl35 | $spl36 | $spl42)),
  inference(global_subsumption,[],[f3170,f36,f35,f34,f37,f39,f40,f43,f44,f45,f56,f58,f64,f82,f104,f132,f133,f146,f156,f157,f200,f361,f363,f365,f761,f762,f763,f785,f48,f47,f46,f475,f550,f786,f366,f368,f370,f1480,f1935,f2442,f2443,f2444,f2445,f2446,f2454,f2528,f2540,f2548,f2551,f2559,f2698,f2699,f2700,f2704,f2706,f2763,f2874,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f4154,f3646,f3651,f3726,f3758,f3759,f3827,f3832,f3890,f3878,f2550,f2549,f2455,f3652,f3732,f1972,f1473,f3828,f4994,f3727])).
fof(f4994,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl35),
  inference(global_subsumption,[],[f36,f35,f34,f37,f39,f40,f43,f44,f45,f56,f58,f64,f82,f104,f132,f133,f146,f156,f157,f200,f361,f363,f365,f761,f762,f763,f785,f48,f47,f46,f475,f550,f786,f366,f368,f370,f1480,f1935,f2442,f2443,f2444,f2445,f2446,f2454,f2528,f2540,f2548,f2551,f2559,f2698,f2699,f2700,f2704,f2706,f2763,f2874,f4122,f3463,f3464,f3465,f3490,f3525,f3535,f3536,f4154,f3646,f3651,f3726,f3758,f3759,f3827,f3832,f3890,f3878,f2550,f2549,f3652,f3732,f1972,f1473,f3828])).
fof(f2540,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl35),
  inference(trivial_inequality_removal,[],[f2535])).
fof(f2535,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl35),
  inference(superposition,[],[f2454,f56])).
fof(f2528,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl35),
  inference(superposition,[],[f2454,f40])).
fof(f5028,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl35 | $spl36 | $spl42)),
  inference(backward_demodulation,[],[f4999,f4994])).
fof(f5065,plain,(
  $false | ($spl35 | $spl36 | $spl42)),
  inference(subsumption_resolution,[],[f5000,f37])).
fof(f5000,plain,(
  ~p(s(bool,'T_0')) | ($spl35 | $spl36 | $spl42)),
  inference(backward_demodulation,[],[f4999,f39])).
fof(f4997,plain,(
  $false | ($spl22 | $spl31 | $spl35 | $spl42)),
  inference(global_subsumption,[],[f3170,f36,f35,f34,f37,f39,f40,f43,f44,f45,f56,f58,f64,f82,f104,f132,f133,f146,f156,f157,f200,f361,f363,f365,f761,f762,f763,f785,f48,f47,f46,f475,f550,f786,f366,f368,f370,f1480,f1935,f2360,f2418,f2442,f2443,f2444,f2445,f2446,f2454,f2541,f2528,f2540,f2548,f2551,f2559,f2567,f2639,f2640,f2683,f2655,f2656,f2698,f2699,f2700,f2704,f2706,f2708,f2763,f2874,f3141,f3133,f2912,f2913,f3121,f2943,f2944,f3104,f3094,f3008,f3015,f3019,f3020,f3053,f4122,f3463,f3464,f3465,f4060,f3490,f3525,f3535,f3536,f4154,f3646,f3651,f3726,f3758,f3759,f3827,f3832,f3890,f3878,f2550,f2549,f2293,f3652,f3732,f1972,f1473,f3828,f4994,f2451])).
fof(f4060,plain,(
  ( ! [X94,X95,X93,X96] : (s(bool,'T_0') = s(bool,X93) | s(bool,X94) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,X93))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X95),'const.pair.,_2'(s(bool,X94),s(X95,X96)))))) ) | $spl22),
  inference(duplicate_literal_removal,[],[f3488])).
fof(f3488,plain,(
  ( ! [X94,X95,X93,X96] : (s(bool,'T_0') = s(bool,X93) | s(bool,X94) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,X93))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X95),'const.pair.,_2'(s(bool,X94),s(X95,X96))))) | s(bool,'T_0') = s(bool,X93)) ) | $spl22),
  inference(superposition,[],[f132,f2293])).
fof(f3053,plain,(
  ( ! [X111,X114,X112,X110,X115,X113] : (~p(s(bool,happ(s(fun(bool,bool),X112),s(bool,X111)))) | ~p(s(bool,happ(s(fun(X113,bool),X114),s(X113,X115)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X113,bool),X114),s(fun(bool,bool),X112),s('type.pair.prod'(X113,bool),'const.pair.,_2'(s(X113,X115),s(bool,f(s(bool,X110)))))))) | s(bool,'T_0') = s(bool,X110) | s(bool,'T_0') = s(bool,X111)) ) | ($spl22 | $spl35)),
  inference(superposition,[],[f785,f2683])).
fof(f3020,plain,(
  ( ! [X269,X267,X268,X270] : (s(X269,f(s(bool,f(s(bool,X270))))) != s(X269,g(s(bool,f(s(bool,X270))))) | s(fun(bool,X269),X267) = s(fun(bool,X269),X268) | s(bool,'T_0') = s(bool,X270) | s(bool,'T_0') = s(bool,sK6(X267,X268,X269,bool))) ) | ($spl22 | $spl35)),
  inference(superposition,[],[f45,f2683])).
fof(f3019,plain,(
  ( ! [X263,X265,X262,X264,X266] : (~p(s(bool,happ(s(fun(bool,bool),X265),s(bool,f(s(bool,X264)))))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X262,bool),X266),s(fun(bool,bool),X265),s('type.pair.prod'(X262,bool),X263)))) | ~p(s(bool,happ(s(fun(X262,bool),X266),s(X262,'const.pair.FST_1'(s('type.pair.prod'(X262,bool),X263)))))) | s(bool,'T_0') = s(bool,X264) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X262,bool),X263)))) ) | ($spl22 | $spl35)),
  inference(superposition,[],[f48,f2683])).
fof(f3015,plain,(
  ( ! [X249,X250,X248,X247] : (s(bool,f(s(bool,X250))) != s(bool,f(s(X247,sK6(X248,X249,bool,X247)))) | s(fun(X247,bool),X248) = s(fun(X247,bool),X249) | s(bool,'T_0') = s(bool,X250) | s(bool,'T_0') = s(bool,g(s(X247,sK6(X248,X249,bool,X247))))) ) | ($spl22 | $spl35)),
  inference(superposition,[],[f45,f2683])).
fof(f3008,plain,(
  ( ! [X212,X210,X213,X211,X209,X214] : (~p(s(bool,happ(s(fun(bool,bool),X211),s(bool,f(s(bool,X210)))))) | ~p(s(bool,happ(s(fun(X212,bool),X213),s(X212,X214)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X212,bool),X213),s(fun(bool,bool),X211),s('type.pair.prod'(X212,bool),'const.pair.,_2'(s(X212,X214),s(bool,X209)))))) | s(bool,'T_0') = s(bool,X210) | s(bool,'T_0') = s(bool,X209)) ) | ($spl22 | $spl35)),
  inference(superposition,[],[f785,f2683])).
fof(f3094,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,f(s(bool,X0))) = s(bool,f(s(bool,f(s(bool,X0)))))) ) | ($spl22 | $spl35)),
  inference(condensation,[],[f3093])).
fof(f3093,plain,(
  ( ! [X74,X75] : (s(bool,f(s(bool,X75))) = s(bool,f(s(bool,f(s(bool,X75))))) | s(bool,'T_0') = s(bool,X74) | s(bool,'T_0') = s(bool,X75)) ) | ($spl22 | $spl35)),
  inference(duplicate_literal_removal,[],[f2975])).
fof(f2975,plain,(
  ( ! [X74,X75] : (s(bool,f(s(bool,X75))) = s(bool,f(s(bool,f(s(bool,X75))))) | s(bool,'T_0') = s(bool,X74) | s(bool,'T_0') = s(bool,X75) | s(bool,'T_0') = s(bool,X74)) ) | ($spl22 | $spl35)),
  inference(superposition,[],[f2567,f2683])).
fof(f3104,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,f(s(bool,X0))))) ) | ($spl22 | $spl35)),
  inference(condensation,[],[f3103])).
fof(f3103,plain,(
  ( ! [X14,X15] : (~p(s(bool,f(s(bool,X15)))) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,X15)) ) | ($spl22 | $spl35)),
  inference(duplicate_literal_removal,[],[f2961])).
fof(f2961,plain,(
  ( ! [X14,X15] : (~p(s(bool,f(s(bool,X15)))) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X14)) ) | ($spl22 | $spl35)),
  inference(superposition,[],[f58,f2683])).
fof(f2944,plain,(
  ( ! [X92,X90,X93,X91] : (s(bool,X93) = s(bool,f(s(bool,X92))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X90,bool),'const.pair.,_2'(s(X90,X91),s(bool,X92))))) | s(bool,'T_0') = s(bool,X93)) ) | ($spl22 | $spl35)),
  inference(superposition,[],[f2683,f44])).
fof(f2943,plain,(
  ( ! [X88,X87,X89,X86] : (s(bool,X89) = s(bool,f(s(bool,X87))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X86),'const.pair.,_2'(s(bool,X87),s(X86,X88))))) | s(bool,'T_0') = s(bool,X89)) ) | ($spl22 | $spl35)),
  inference(superposition,[],[f2683,f43])).
fof(f3121,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,f(s(bool,f(s(bool,X1))))) | s(bool,'T_0') = s(bool,X1)) ) | ($spl22 | $spl35)),
  inference(condensation,[],[f3120])).
fof(f3120,plain,(
  ( ! [X17,X18,X16] : (s(bool,X18) = s(bool,f(s(bool,f(s(bool,X17))))) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X17)) ) | ($spl22 | $spl35)),
  inference(duplicate_literal_removal,[],[f2920])).
fof(f2920,plain,(
  ( ! [X17,X18,X16] : (s(bool,X18) = s(bool,f(s(bool,f(s(bool,X17))))) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X17) | s(bool,'T_0') = s(bool,X16)) ) | ($spl22 | $spl35)),
  inference(superposition,[],[f2683,f2683])).
fof(f2913,plain,(
  ( ! [X92,X90,X93,X91] : (s(bool,X92) = s(bool,f(s(bool,X93))) | s(bool,'T_0') = s(bool,X93) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X90,bool),'const.pair.,_2'(s(X90,X91),s(bool,X92)))))) ) | ($spl22 | $spl35)),
  inference(superposition,[],[f2683,f44])).
fof(f2912,plain,(
  ( ! [X88,X87,X89,X86] : (s(bool,X87) = s(bool,f(s(bool,X89))) | s(bool,'T_0') = s(bool,X89) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X86),'const.pair.,_2'(s(bool,X87),s(X86,X88)))))) ) | ($spl22 | $spl35)),
  inference(superposition,[],[f2683,f43])).
fof(f3133,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,f(s(bool,X0))) = s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,X1)) ) | ($spl22 | $spl35)),
  inference(condensation,[],[f3132])).
fof(f3132,plain,(
  ( ! [X17,X18,X16] : (s(bool,f(s(bool,X17))) = s(bool,f(s(bool,X18))) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,X17)) ) | ($spl22 | $spl35)),
  inference(duplicate_literal_removal,[],[f2892])).
fof(f2892,plain,(
  ( ! [X17,X18,X16] : (s(bool,f(s(bool,X17))) = s(bool,f(s(bool,X18))) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,X17) | s(bool,'T_0') = s(bool,X16)) ) | ($spl22 | $spl35)),
  inference(superposition,[],[f2683,f2683])).
fof(f3141,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,f(s(bool,X0)))) ) | ($spl22 | $spl35)),
  inference(condensation,[],[f3140])).
fof(f3140,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | ($spl22 | $spl35)),
  inference(duplicate_literal_removal,[],[f2886])).
fof(f2886,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | ($spl22 | $spl35)),
  inference(superposition,[],[f2683,f40])).
fof(f2708,plain,(
  ( ! [X59,X57,X58,X56,X55] : (~p(s(bool,happ(s(fun(bool,bool),X56),s(bool,X55)))) | ~p(s(bool,happ(s(fun(X57,bool),X58),s(X57,X59)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X57,bool),X58),s(fun(bool,bool),X56),s('type.pair.prod'(X57,bool),'const.pair.,_2'(s(X57,X59),s(bool,f(s(bool,X55)))))))) | s(bool,'T_0') = s(bool,X55)) ) | ($spl22 | $spl35)),
  inference(superposition,[],[f785,f2567])).
fof(f2656,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X16,bool),'const.pair.,_2'(s(X16,X17),s(bool,X15))))) | s(bool,'F_0') = s(bool,f(s(bool,X15))) | s(bool,'T_0') = s(bool,X15)) ) | ($spl22 | $spl35)),
  inference(superposition,[],[f104,f2567])).
fof(f2655,plain,(
  ( ! [X14,X12,X13] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X13),'const.pair.,_2'(s(bool,X12),s(X13,X14))))) | s(bool,'F_0') = s(bool,f(s(bool,X12))) | s(bool,'T_0') = s(bool,X12)) ) | ($spl22 | $spl35)),
  inference(superposition,[],[f82,f2567])).
fof(f2683,plain,(
  ( ! [X2,X3] : (s(bool,X3) = s(bool,f(s(bool,X2))) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X3)) ) | ($spl22 | $spl35)),
  inference(duplicate_literal_removal,[],[f2642])).
fof(f2642,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X3) = s(bool,f(s(bool,X2))) | s(bool,'T_0') = s(bool,X3)) ) | ($spl22 | $spl35)),
  inference(superposition,[],[f2567,f56])).
fof(f2640,plain,(
  ( ! [X57,X58,X56] : (s(bool,X58) = s(bool,f(s(bool,X58))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X56,bool),'const.pair.,_2'(s(X56,X57),s(bool,X58)))))) ) | ($spl22 | $spl35)),
  inference(superposition,[],[f2567,f44])).
fof(f2639,plain,(
  ( ! [X54,X55,X53] : (s(bool,X54) = s(bool,f(s(bool,X54))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X53),'const.pair.,_2'(s(bool,X54),s(X53,X55)))))) ) | ($spl22 | $spl35)),
  inference(superposition,[],[f2567,f43])).
fof(f2567,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl22 | $spl35)),
  inference(superposition,[],[f2541,f40])).
fof(f2541,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | ($spl22 | $spl35)),
  inference(forward_demodulation,[],[f2524,f43])).
fof(f2524,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,f(s(bool,'F_0'))))))) | ($spl22 | $spl35)),
  inference(unit_resulting_resolution,[],[f2454,f2293])).
fof(f2418,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') = s(bool,X16) | s(bool,X17) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,X16))))) | s(bool,'T_0') = s(bool,X17)) ) | $spl22),
  inference(duplicate_literal_removal,[],[f2369])).
fof(f2369,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,X16) | s(bool,X17) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,X16))))) | s(bool,'T_0') = s(bool,X17)) ) | $spl22),
  inference(superposition,[],[f2293,f56])).
fof(f2360,plain,(
  ( ! [X54,X52,X53] : (s(bool,X53) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,X53))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X52),'const.pair.,_2'(s(bool,X53),s(X52,X54)))))) ) | $spl22),
  inference(superposition,[],[f2293,f43])).
fof(f4996,plain,(
  $false | ($spl22 | $spl31 | $spl35 | $spl76)),
  inference(global_subsumption,[],[f4984,f36,f35,f34,f37,f39,f40,f43,f44,f45,f56,f58,f64,f82,f104,f132,f133,f146,f156,f157,f200,f361,f363,f365,f761,f762,f763,f785,f48,f47,f46,f475,f550,f786,f366,f368,f370,f1480,f1935,f2360,f2418,f2442,f2443,f2444,f2445,f2446,f2454,f2541,f2528,f2540,f2548,f2551,f2559,f2567,f2639,f2640,f2683,f2655,f2656,f2698,f2699,f2700,f2704,f2706,f2708,f2763,f2874,f3141,f3133,f2912,f2913,f3121,f2943,f2944,f3104,f3094,f3008,f3015,f3019,f3020,f3053,f4122,f3463,f3464,f3465,f4060,f3490,f3525,f3535,f3536,f4154,f3646,f3651,f3726,f3758,f3759,f3827,f3832,f3890,f3878,f2550,f2549,f2293,f3652,f3732,f1972,f1473,f3828,f4994,f2451])).
fof(f4995,plain,(
  $false | ($spl22 | $spl31 | $spl35)),
  inference(global_subsumption,[],[f36,f35,f34,f37,f39,f40,f43,f44,f45,f56,f58,f64,f82,f104,f132,f133,f146,f156,f157,f200,f361,f363,f365,f761,f762,f763,f785,f48,f47,f46,f475,f550,f786,f366,f368,f370,f1480,f1935,f2360,f2418,f2442,f2443,f2444,f2445,f2446,f2454,f2541,f2528,f2540,f2548,f2551,f2559,f2567,f2639,f2640,f2683,f2655,f2656,f2698,f2699,f2700,f2704,f2706,f2708,f2763,f2874,f3141,f3133,f2912,f2913,f3121,f2943,f2944,f3104,f3094,f3008,f3015,f3019,f3020,f3053,f4122,f3463,f3464,f3465,f4060,f3490,f3525,f3535,f3536,f4154,f3646,f3651,f3726,f3758,f3759,f3827,f3832,f3890,f3878,f2550,f2549,f2293,f3652,f3732,f1972,f1473,f3828,f4994,f2451])).
fof(f4993,plain,(
  ( ! [X407,X403,X406] : (s(bool,X403) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X406),'const.pair.,_2'(s(bool,X403),s(X406,X407)))))) ) | $spl84),
  inference(cnf_transformation,[],[f4993_D])).
fof(f4993_D,plain,(
  ( ! [X407,X403,X406] : (s(bool,X403) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X406),'const.pair.,_2'(s(bool,X403),s(X406,X407)))))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f4990,plain,(
  ( ! [X405,X407,X403,X404,X406] : (s(bool,X403) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X404) = s(fun(bool,bool),X405) | s(bool,'T_0') = s(bool,sK6(X404,X405,bool,bool)) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X406),'const.pair.,_2'(s(bool,X403),s(X406,X407)))))) ) | $spl27),
  inference(subsumption_resolution,[],[f3727,f2447])).
fof(f4992,plain,(
  ( ! [X436,X437,X435] : (s(bool,'F_0') != s(bool,f(s(bool,X435))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X436),'const.pair.,_2'(s(bool,X435),s(X436,X437)))))) ) | $spl82),
  inference(cnf_transformation,[],[f4992_D])).
fof(f4992_D,plain,(
  ( ! [X436,X437,X435] : (s(bool,'F_0') != s(bool,f(s(bool,X435))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X436),'const.pair.,_2'(s(bool,X435),s(X436,X437)))))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f3732,plain,(
  ( ! [X436,X434,X437,X433,X435] : (s(bool,'F_0') != s(bool,f(s(bool,X435))) | s(fun(bool,bool),X433) = s(fun(bool,bool),X434) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X433,X434,bool,bool)))) | s(bool,'T_0') = s(bool,sK6(X433,X434,bool,bool)) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X436),'const.pair.,_2'(s(bool,X435),s(X436,X437)))))) )),
  inference(superposition,[],[f370,f132])).
fof(f4989,plain,(
  ( ! [X366,X368,X367] : (s(bool,X366) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X367),'const.pair.,_2'(s(bool,X366),s(X367,X368)))))) ) | $spl80),
  inference(cnf_transformation,[],[f4989_D])).
fof(f4989_D,plain,(
  ( ! [X366,X368,X367] : (s(bool,X366) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X367),'const.pair.,_2'(s(bool,X366),s(X367,X368)))))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f4988,plain,(
  ( ! [X364,X366,X368,X365,X367] : (s(fun(bool,bool),X364) = s(fun(bool,bool),X365) | s(bool,'T_0') = s(bool,sK6(X364,X365,bool,bool)) | s(bool,X366) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X367),'const.pair.,_2'(s(bool,X366),s(X367,X368)))))) ) | $spl34),
  inference(subsumption_resolution,[],[f3828,f2453])).
fof(f2453,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl34),
  inference(cnf_transformation,[],[f2453_D])).
fof(f2453_D,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f3652,plain,(
  ( ! [X445,X447,X449,X446,X448] : (s(bool,'F_0') != s(bool,f(s(bool,X447))) | s(fun(bool,bool),X445) = s(fun(bool,bool),X446) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X445,X446,bool,bool)))) | s(bool,'T_0') = s(bool,X447) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X448),'const.pair.,_2'(s(bool,sK6(X445,X446,bool,bool)),s(X448,X449)))))) )),
  inference(superposition,[],[f370,f132])).
fof(f2549,plain,(
  ( ! [X6,X7,X5] : (s(bool,'F_0') != s(bool,f(s(bool,X7))) | s(fun(bool,bool),X5) = s(fun(bool,bool),X6) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X5,X6,bool,bool)))) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,sK6(X5,X6,bool,bool))) )),
  inference(superposition,[],[f370,f56])).
fof(f2550,plain,(
  ( ! [X10,X8,X9] : (s(bool,'F_0') != s(bool,f(s(bool,X10))) | s(fun(bool,bool),X8) = s(fun(bool,bool),X9) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X8,X9,bool,bool)))) | s(bool,'T_0') = s(bool,sK6(X8,X9,bool,bool)) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f370,f56])).
fof(f4980,plain,(
  $false | ($spl9 | $spl14 | $spl16)),
  inference(subsumption_resolution,[],[f4979,f71])).
fof(f4979,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl14 | $spl16)),
  inference(condensation,[],[f4978])).
fof(f4978,plain,(
  ( ! [X12] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X12)) ) | ($spl14 | $spl16)),
  inference(forward_demodulation,[],[f961,f960])).
fof(f4965,plain,(
  $false | ($spl9 | $spl16 | $spl20)),
  inference(subsumption_resolution,[],[f4964,f71])).
fof(f4964,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl9 | $spl16 | $spl20)),
  inference(condensation,[],[f4957])).
fof(f4957,plain,(
  ( ! [X12] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X12)) ) | ($spl9 | $spl16 | $spl20)),
  inference(backward_demodulation,[],[f4956,f961])).
fof(f4956,plain,(
  ( ! [X2,X0,X3] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),'const.pair.,_2'(s(bool,X0),s(X2,X3)))))) ) | ($spl9 | $spl20)),
  inference(subsumption_resolution,[],[f2292,f4170])).
fof(f4930,plain,(
  ( ! [X93] : (s(bool,X93) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,X93))))) | s(bool,'F_0') = s(bool,g(s(bool,X93)))) ) | $spl72),
  inference(cnf_transformation,[],[f4930_D])).
fof(f4930_D,plain,(
  ( ! [X93] : (s(bool,X93) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,X93))))) | s(bool,'F_0') = s(bool,g(s(bool,X93)))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f4745,plain,(
  ( ! [X92,X93] : (s(bool,X93) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,X93))))) | s(bool,'T_0') = s(bool,X92) | s(bool,'F_0') = s(bool,g(s(bool,X93)))) ) | ($spl22 | $spl47)),
  inference(duplicate_literal_removal,[],[f4528])).
fof(f4528,plain,(
  ( ! [X92,X93] : (s(bool,X93) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,X93))))) | s(bool,'T_0') = s(bool,X92) | s(bool,'F_0') = s(bool,g(s(bool,X93))) | s(bool,'T_0') = s(bool,X92)) ) | ($spl22 | $spl47)),
  inference(superposition,[],[f2293,f3344])).
fof(f3344,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,g(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | ($spl22 | $spl47)),
  inference(superposition,[],[f3338,f56])).
fof(f3338,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | ($spl22 | $spl47)),
  inference(forward_demodulation,[],[f3319,f43])).
fof(f3319,plain,(
  s(bool,g(s(bool,'T_0'))) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,g(s(bool,'T_0'))))))) | ($spl22 | $spl47)),
  inference(unit_resulting_resolution,[],[f3304,f2293])).
fof(f4912,plain,(
  ( ! [X74,X72,X73] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X73,bool),'const.pair.,_2'(s(X73,X74),s(bool,X72))))) | s(bool,'T_0') = s(bool,X72)) ) | $spl68),
  inference(cnf_transformation,[],[f4912_D])).
fof(f4912_D,plain,(
  ( ! [X74,X72,X73] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X73,bool),'const.pair.,_2'(s(X73,X74),s(bool,X72))))) | s(bool,'T_0') = s(bool,X72)) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f4893,plain,(
  ( ! [X85,X86,X82] : (s(bool,'F_0') = s(bool,g(s(bool,X82))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X85),'const.pair.,_2'(s(bool,X82),s(X85,X86)))))) ) | $spl56),
  inference(cnf_transformation,[],[f4893_D])).
fof(f4893_D,plain,(
  ( ! [X85,X86,X82] : (s(bool,'F_0') = s(bool,g(s(bool,X82))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X85),'const.pair.,_2'(s(bool,X82),s(X85,X86)))))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f4523,plain,(
  ( ! [X70,X68,X66,X69,X67,X65] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X67),'const.pair.,_2'(s(bool,X66),s(X67,X68))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X69,bool),'const.pair.,_2'(s(X69,X70),s(bool,X65))))) | s(bool,'F_0') = s(bool,g(s(bool,X66))) | s(bool,'T_0') = s(bool,X65)) ) | ($spl16 | $spl22 | $spl47)),
  inference(superposition,[],[f1258,f3344])).
fof(f1258,plain,(
  ( ! [X45,X43,X46,X44,X42] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X45),'const.pair.,_2'(s(bool,X44),s(X45,X46))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X42,bool),'const.pair.,_2'(s(X42,X43),s(bool,X44)))))) ) | $spl16),
  inference(superposition,[],[f961,f44])).
fof(f4929,plain,(
  ( ! [X61,X59,X64,X62,X60,X63] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X61),'const.pair.,_2'(s(bool,X60),s(X61,X62))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X63),'const.pair.,_2'(s(bool,X59),s(X63,X64))))) | s(bool,'F_0') = s(bool,g(s(bool,X60)))) ) | ($spl9 | $spl16 | $spl22 | $spl47)),
  inference(subsumption_resolution,[],[f4522,f4170])).
fof(f4522,plain,(
  ( ! [X61,X59,X64,X62,X60,X63] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X61),'const.pair.,_2'(s(bool,X60),s(X61,X62))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X63),'const.pair.,_2'(s(bool,X59),s(X63,X64))))) | s(bool,'F_0') = s(bool,g(s(bool,X60))) | s(bool,'T_0') = s(bool,X59)) ) | ($spl16 | $spl22 | $spl47)),
  inference(superposition,[],[f1257,f3344])).
fof(f4749,plain,(
  ( ! [X47,X50,X48,X49] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X49),'const.pair.,_2'(s(bool,X48),s(X49,X50))))) | s(bool,'T_0') = s(bool,X47) | s(bool,'F_0') = s(bool,g(s(bool,X48)))) ) | ($spl16 | $spl22 | $spl47)),
  inference(duplicate_literal_removal,[],[f4519])).
fof(f4519,plain,(
  ( ! [X47,X50,X48,X49] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X49),'const.pair.,_2'(s(bool,X48),s(X49,X50))))) | s(bool,'T_0') = s(bool,X47) | s(bool,'F_0') = s(bool,g(s(bool,X48))) | s(bool,'T_0') = s(bool,X47)) ) | ($spl16 | $spl22 | $spl47)),
  inference(superposition,[],[f961,f3344])).
fof(f4928,plain,(
  ( ! [X15] : (~p(s(bool,X15)) | s(bool,'F_0') = s(bool,g(s(bool,X15)))) ) | $spl70),
  inference(cnf_transformation,[],[f4928_D])).
fof(f4928_D,plain,(
  ( ! [X15] : (~p(s(bool,X15)) | s(bool,'F_0') = s(bool,g(s(bool,X15)))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f4752,plain,(
  ( ! [X14,X15] : (~p(s(bool,X15)) | s(bool,'T_0') = s(bool,X14) | s(bool,'F_0') = s(bool,g(s(bool,X15)))) ) | ($spl22 | $spl47)),
  inference(duplicate_literal_removal,[],[f4511])).
fof(f4511,plain,(
  ( ! [X14,X15] : (~p(s(bool,X15)) | s(bool,'T_0') = s(bool,X14) | s(bool,'F_0') = s(bool,g(s(bool,X15))) | s(bool,'T_0') = s(bool,X14)) ) | ($spl22 | $spl47)),
  inference(superposition,[],[f58,f3344])).
fof(f4897,plain,(
  ( ! [X121,X122] : (s(bool,'F_0') = s(bool,g(s(bool,sK6(X121,X122,bool,bool)))) | s(fun(bool,bool),X121) = s(fun(bool,bool),X122)) ) | $spl58),
  inference(cnf_transformation,[],[f4897_D])).
fof(f4897_D,plain,(
  ( ! [X121,X122] : (s(bool,'F_0') = s(bool,g(s(bool,sK6(X121,X122,bool,bool)))) | s(fun(bool,bool),X121) = s(fun(bool,bool),X122)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f4758,plain,(
  ( ! [X231,X230,X232] : (s(bool,'T_0') = s(bool,X232) | s(fun(bool,bool),X230) = s(fun(bool,bool),X231) | s(bool,'F_0') = s(bool,g(s(bool,sK6(X230,X231,bool,bool))))) ) | ($spl22 | $spl28 | $spl47)),
  inference(duplicate_literal_removal,[],[f4484])).
fof(f4484,plain,(
  ( ! [X231,X230,X232] : (s(bool,'T_0') = s(bool,X232) | s(fun(bool,bool),X230) = s(fun(bool,bool),X231) | s(bool,'F_0') = s(bool,g(s(bool,sK6(X230,X231,bool,bool)))) | s(bool,'T_0') = s(bool,X232)) ) | ($spl22 | $spl28 | $spl47)),
  inference(superposition,[],[f2449,f3344])).
fof(f4925,plain,(
  ( ! [X229,X227,X228,X226] : (s(bool,'F_0') = s(bool,g(s(bool,X226))) | s(bool,'T_0') = s(bool,X227) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X228),'const.pair.,_2'(s(bool,X226),s(X228,X229)))))) ) | ($spl16 | $spl22 | $spl47)),
  inference(forward_demodulation,[],[f4759,f44])).
fof(f4759,plain,(
  ( ! [X229,X227,X225,X228,X226,X224] : (s(bool,'T_0') = s(bool,X227) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X228),'const.pair.,_2'(s(bool,X226),s(X228,X229))))) | s(bool,'F_0') = s(bool,g(s(bool,'const.pair.SND_1'(s('type.pair.prod'(X224,bool),'const.pair.,_2'(s(X224,X225),s(bool,X226)))))))) ) | ($spl16 | $spl22 | $spl47)),
  inference(duplicate_literal_removal,[],[f4483])).
fof(f4483,plain,(
  ( ! [X229,X227,X225,X228,X226,X224] : (s(bool,'T_0') = s(bool,X227) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X228),'const.pair.,_2'(s(bool,X226),s(X228,X229))))) | s(bool,'F_0') = s(bool,g(s(bool,'const.pair.SND_1'(s('type.pair.prod'(X224,bool),'const.pair.,_2'(s(X224,X225),s(bool,X226))))))) | s(bool,'T_0') = s(bool,X227)) ) | ($spl16 | $spl22 | $spl47)),
  inference(superposition,[],[f1258,f3344])).
fof(f4922,plain,(
  ( ! [X191,X194,X195,X193] : (s(bool,'F_0') = s(bool,g(s(bool,X191))) | s(bool,'T_0') = s(bool,X193) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X194),'const.pair.,_2'(s(bool,X191),s(X194,X195)))))) ) | ($spl16 | $spl22 | $spl47)),
  inference(forward_demodulation,[],[f4763,f43])).
fof(f4763,plain,(
  ( ! [X191,X194,X192,X190,X195,X193] : (s(bool,'T_0') = s(bool,X193) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X194),'const.pair.,_2'(s(bool,X191),s(X194,X195))))) | s(bool,'F_0') = s(bool,g(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X190),'const.pair.,_2'(s(bool,X191),s(X190,X192)))))))) ) | ($spl16 | $spl22 | $spl47)),
  inference(duplicate_literal_removal,[],[f4476])).
fof(f4476,plain,(
  ( ! [X191,X194,X192,X190,X195,X193] : (s(bool,'T_0') = s(bool,X193) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X194),'const.pair.,_2'(s(bool,X191),s(X194,X195))))) | s(bool,'F_0') = s(bool,g(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X190),'const.pair.,_2'(s(bool,X191),s(X190,X192))))))) | s(bool,'T_0') = s(bool,X193)) ) | ($spl16 | $spl22 | $spl47)),
  inference(superposition,[],[f1257,f3344])).
fof(f4437,plain,(
  ( ! [X76,X74,X72,X71,X75,X73] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X73,bool),'const.pair.,_2'(s(X73,X74),s(bool,X72))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X75),'const.pair.,_2'(s(bool,X71),s(X75,X76))))) | s(bool,'F_0') = s(bool,g(s(bool,X71))) | s(bool,'T_0') = s(bool,X72)) ) | ($spl16 | $spl22 | $spl47)),
  inference(superposition,[],[f1258,f3344])).
fof(f4911,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X55),'const.pair.,_2'(s(bool,X54),s(X55,X56))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X57),'const.pair.,_2'(s(bool,X53),s(X57,X58))))) | s(bool,'F_0') = s(bool,g(s(bool,X53)))) ) | ($spl9 | $spl16 | $spl22 | $spl47)),
  inference(subsumption_resolution,[],[f4434,f4170])).
fof(f4434,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X55),'const.pair.,_2'(s(bool,X54),s(X55,X56))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X57),'const.pair.,_2'(s(bool,X53),s(X57,X58))))) | s(bool,'F_0') = s(bool,g(s(bool,X53))) | s(bool,'T_0') = s(bool,X54)) ) | ($spl16 | $spl22 | $spl47)),
  inference(superposition,[],[f1257,f3344])).
fof(f4909,plain,(
  ( ! [X92,X88,X91] : (s(bool,X91) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,g(s(bool,X88))),s(bool,X91))))) | s(bool,'T_0') = s(bool,X92) | s(bool,X88) = s(bool,X92) | s(bool,'T_0') = s(bool,X91)) ) | $spl66),
  inference(cnf_transformation,[],[f4909_D])).
fof(f4909_D,plain,(
  ( ! [X92,X88,X91] : (s(bool,X91) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,g(s(bool,X88))),s(bool,X91))))) | s(bool,'T_0') = s(bool,X92) | s(bool,X88) = s(bool,X92) | s(bool,'T_0') = s(bool,X91)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f4401,plain,(
  ( ! [X92,X90,X88,X91,X89] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X89),'const.pair.,_2'(s(bool,'F_0'),s(X89,X90))))) | s(bool,X91) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,g(s(bool,X88))),s(bool,X91))))) | s(bool,'T_0') = s(bool,X91) | s(bool,X88) = s(bool,X92) | s(bool,'T_0') = s(bool,X92)) ) | ($spl16 | $spl22 | $spl47)),
  inference(superposition,[],[f1557,f3344])).
fof(f4907,plain,(
  ( ! [X59,X57,X54,X58] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X57),'const.pair.,_2'(s(bool,g(s(bool,X54))),s(X57,X58))))) | s(bool,'T_0') = s(bool,X59) | s(bool,X54) = s(bool,X59)) ) | $spl62),
  inference(cnf_transformation,[],[f4907_D])).
fof(f4907_D,plain,(
  ( ! [X59,X57,X54,X58] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X57),'const.pair.,_2'(s(bool,g(s(bool,X54))),s(X57,X58))))) | s(bool,'T_0') = s(bool,X59) | s(bool,X54) = s(bool,X59)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f4398,plain,(
  ( ! [X76,X74,X72,X77,X75,X73] : (s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X73,bool),'const.pair.,_2'(s(X73,X74),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X75),'const.pair.,_2'(s(bool,g(s(bool,X72))),s(X75,X76))))) | s(bool,X72) = s(bool,X77) | s(bool,'T_0') = s(bool,X77)) ) | ($spl16 | $spl22 | $spl47)),
  inference(superposition,[],[f1258,f3344])).
fof(f4395,plain,(
  ( ! [X59,X57,X54,X58,X56,X55] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X55),'const.pair.,_2'(s(bool,'F_0'),s(X55,X56))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X57),'const.pair.,_2'(s(bool,g(s(bool,X54))),s(X57,X58))))) | s(bool,X54) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) ) | ($spl16 | $spl22 | $spl47)),
  inference(superposition,[],[f1257,f3344])).
fof(f4900,plain,(
  ( ! [X30] : (s(bool,X30) = s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2))))) | s(bool,'F_0') = s(bool,g(s(bool,X30)))) ) | $spl60),
  inference(cnf_transformation,[],[f4900_D])).
fof(f4900_D,plain,(
  ( ! [X30] : (s(bool,X30) = s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2))))) | s(bool,'F_0') = s(bool,g(s(bool,X30)))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f4809,plain,(
  ( ! [X41,X40] : (s(bool,'T_0') = s(bool,X40) | s(bool,'F_0') = s(bool,g(s(bool,X41))) | s(bool,X41) = s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2)))))) ) | ($spl1 | $spl22 | $spl47)),
  inference(duplicate_literal_removal,[],[f4370])).
fof(f4370,plain,(
  ( ! [X41,X40] : (s(bool,'T_0') = s(bool,X40) | s(bool,'F_0') = s(bool,g(s(bool,X41))) | s(bool,X41) = s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2))))) | s(bool,'T_0') = s(bool,X40)) ) | ($spl1 | $spl22 | $spl47)),
  inference(superposition,[],[f3344,f1101])).
fof(f4878,plain,(
  ( ! [X42,X40] : (s(bool,'F_0') = s(bool,g(s(bool,X40))) | s(bool,X40) = s(bool,X42) | s(bool,'F_0') = s(bool,g(s(bool,X42)))) ) | $spl50),
  inference(cnf_transformation,[],[f4878_D])).
fof(f4878_D,plain,(
  ( ! [X42,X40] : (s(bool,'F_0') = s(bool,g(s(bool,X40))) | s(bool,X40) = s(bool,X42) | s(bool,'F_0') = s(bool,g(s(bool,X42)))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f4811,plain,(
  ( ! [X24,X23,X25] : (s(bool,'T_0') = s(bool,X24) | s(bool,'F_0') = s(bool,g(s(bool,X25))) | s(bool,X23) = s(bool,X25) | s(bool,'F_0') = s(bool,g(s(bool,X23)))) ) | ($spl22 | $spl47)),
  inference(duplicate_literal_removal,[],[f4365])).
fof(f4365,plain,(
  ( ! [X24,X23,X25] : (s(bool,'T_0') = s(bool,X24) | s(bool,'F_0') = s(bool,g(s(bool,X25))) | s(bool,X23) = s(bool,X25) | s(bool,'F_0') = s(bool,g(s(bool,X23))) | s(bool,'T_0') = s(bool,X24)) ) | ($spl22 | $spl47)),
  inference(superposition,[],[f3344,f3344])).
fof(f4824,plain,(
  ( ! [X35,X36,X34] : (s(bool,X35) = s(bool,X36) | s(bool,'F_0') = s(bool,g(s(bool,X36))) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,g(s(bool,X35)))) ) | ($spl22 | $spl47)),
  inference(duplicate_literal_removal,[],[f4315])).
fof(f4315,plain,(
  ( ! [X35,X36,X34] : (s(bool,X35) = s(bool,X36) | s(bool,'F_0') = s(bool,g(s(bool,X36))) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,g(s(bool,X35))) | s(bool,'T_0') = s(bool,X34)) ) | ($spl22 | $spl47)),
  inference(superposition,[],[f3344,f3344])).
fof(f4825,plain,(
  ( ! [X30,X29] : (s(bool,X30) = s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2))))) | s(bool,'F_0') = s(bool,g(s(bool,X30))) | s(bool,'T_0') = s(bool,X29)) ) | ($spl1 | $spl22 | $spl47)),
  inference(duplicate_literal_removal,[],[f4313])).
fof(f4313,plain,(
  ( ! [X30,X29] : (s(bool,X30) = s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2))))) | s(bool,'F_0') = s(bool,g(s(bool,X30))) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) ) | ($spl1 | $spl22 | $spl47)),
  inference(superposition,[],[f3344,f1101])).
fof(f4883,plain,(
  ( ! [X36] : (s(bool,'F_0') = s(bool,g(s(bool,X36))) | s(bool,'F_0') = s(bool,X36)) ) | $spl54),
  inference(cnf_transformation,[],[f4883_D])).
fof(f4883_D,plain,(
  ( ! [X36] : (s(bool,'F_0') = s(bool,g(s(bool,X36))) | s(bool,'F_0') = s(bool,X36)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f4832,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,g(s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl22 | $spl47)),
  inference(duplicate_literal_removal,[],[f4305])).
fof(f4305,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,g(s(bool,X1))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | ($spl22 | $spl47)),
  inference(superposition,[],[f3344,f40])).
fof(f4833,plain,(
  ( ! [X123,X121,X122] : (s(bool,'T_0') = s(bool,X123) | s(bool,'F_0') = s(bool,g(s(bool,sK6(X121,X122,bool,bool)))) | s(fun(bool,bool),X121) = s(fun(bool,bool),X122)) ) | ($spl22 | $spl28 | $spl47)),
  inference(duplicate_literal_removal,[],[f4302])).
fof(f4302,plain,(
  ( ! [X123,X121,X122] : (s(bool,'T_0') = s(bool,X123) | s(bool,'F_0') = s(bool,g(s(bool,sK6(X121,X122,bool,bool)))) | s(bool,'T_0') = s(bool,X123) | s(fun(bool,bool),X121) = s(fun(bool,bool),X122)) ) | ($spl22 | $spl28 | $spl47)),
  inference(superposition,[],[f3344,f2449])).
fof(f4896,plain,(
  ( ! [X118,X120,X119,X117] : (s(bool,'F_0') = s(bool,g(s(bool,X117))) | s(bool,'T_0') = s(bool,X118) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X119),'const.pair.,_2'(s(bool,X117),s(X119,X120)))))) ) | ($spl16 | $spl22 | $spl47)),
  inference(forward_demodulation,[],[f4834,f44])).
fof(f4834,plain,(
  ( ! [X118,X116,X120,X119,X117,X115] : (s(bool,'T_0') = s(bool,X118) | s(bool,'F_0') = s(bool,g(s(bool,'const.pair.SND_1'(s('type.pair.prod'(X115,bool),'const.pair.,_2'(s(X115,X116),s(bool,X117))))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X119),'const.pair.,_2'(s(bool,X117),s(X119,X120)))))) ) | ($spl16 | $spl22 | $spl47)),
  inference(duplicate_literal_removal,[],[f4301])).
fof(f4301,plain,(
  ( ! [X118,X116,X120,X119,X117,X115] : (s(bool,'T_0') = s(bool,X118) | s(bool,'F_0') = s(bool,g(s(bool,'const.pair.SND_1'(s('type.pair.prod'(X115,bool),'const.pair.,_2'(s(X115,X116),s(bool,X117))))))) | s(bool,'T_0') = s(bool,X118) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X119),'const.pair.,_2'(s(bool,X117),s(X119,X120)))))) ) | ($spl16 | $spl22 | $spl47)),
  inference(superposition,[],[f3344,f1258])).
fof(f4895,plain,(
  ( ! [X114,X113] : (s(bool,'F_0') = s(bool,g(s(bool,X113))) | s(bool,'T_0') = s(bool,X114) | s(bool,'F_0') = s(bool,X113)) ) | ($spl22 | $spl47)),
  inference(forward_demodulation,[],[f4835,f44])).
fof(f4835,plain,(
  ( ! [X111,X114,X112,X113] : (s(bool,'T_0') = s(bool,X114) | s(bool,'F_0') = s(bool,g(s(bool,'const.pair.SND_1'(s('type.pair.prod'(X111,bool),'const.pair.,_2'(s(X111,X112),s(bool,X113))))))) | s(bool,'F_0') = s(bool,X113)) ) | ($spl22 | $spl47)),
  inference(duplicate_literal_removal,[],[f4300])).
fof(f4300,plain,(
  ( ! [X111,X114,X112,X113] : (s(bool,'T_0') = s(bool,X114) | s(bool,'F_0') = s(bool,g(s(bool,'const.pair.SND_1'(s('type.pair.prod'(X111,bool),'const.pair.,_2'(s(X111,X112),s(bool,X113))))))) | s(bool,'T_0') = s(bool,X114) | s(bool,'F_0') = s(bool,X113)) ) | ($spl22 | $spl47)),
  inference(superposition,[],[f3344,f104])).
fof(f4892,plain,(
  ( ! [X85,X86,X84,X82] : (s(bool,'F_0') = s(bool,g(s(bool,X82))) | s(bool,'T_0') = s(bool,X84) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X85),'const.pair.,_2'(s(bool,X82),s(X85,X86)))))) ) | ($spl16 | $spl22 | $spl47)),
  inference(forward_demodulation,[],[f4838,f43])).
fof(f4838,plain,(
  ( ! [X85,X83,X81,X86,X84,X82] : (s(bool,'T_0') = s(bool,X84) | s(bool,'F_0') = s(bool,g(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X81),'const.pair.,_2'(s(bool,X82),s(X81,X83))))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X85),'const.pair.,_2'(s(bool,X82),s(X85,X86)))))) ) | ($spl16 | $spl22 | $spl47)),
  inference(duplicate_literal_removal,[],[f4294])).
fof(f4294,plain,(
  ( ! [X85,X83,X81,X86,X84,X82] : (s(bool,'T_0') = s(bool,X84) | s(bool,'F_0') = s(bool,g(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X81),'const.pair.,_2'(s(bool,X82),s(X81,X83))))))) | s(bool,'T_0') = s(bool,X84) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X85),'const.pair.,_2'(s(bool,X82),s(X85,X86)))))) ) | ($spl16 | $spl22 | $spl47)),
  inference(superposition,[],[f3344,f1257])).
fof(f4890,plain,(
  ( ! [X71,X69] : (s(bool,'F_0') = s(bool,g(s(bool,X69))) | s(bool,'T_0') = s(bool,X71) | s(bool,'F_0') = s(bool,X69)) ) | ($spl22 | $spl47)),
  inference(forward_demodulation,[],[f4841,f43])).
fof(f4841,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,'T_0') = s(bool,X71) | s(bool,'F_0') = s(bool,g(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X68),'const.pair.,_2'(s(bool,X69),s(X68,X70))))))) | s(bool,'F_0') = s(bool,X69)) ) | ($spl22 | $spl47)),
  inference(duplicate_literal_removal,[],[f4291])).
fof(f4291,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,'T_0') = s(bool,X71) | s(bool,'F_0') = s(bool,g(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X68),'const.pair.,_2'(s(bool,X69),s(X68,X70))))))) | s(bool,'T_0') = s(bool,X71) | s(bool,'F_0') = s(bool,X69)) ) | ($spl22 | $spl47)),
  inference(superposition,[],[f3344,f82])).
fof(f4853,plain,(
  ( ! [X39,X37,X38] : (s(bool,'T_0') = s(bool,X38) | s(bool,'F_0') = s(bool,g(s(bool,X37))) | s(bool,'F_0') = s(bool,g(s(bool,X39))) | s(bool,X37) = s(bool,X39)) ) | ($spl22 | $spl47)),
  inference(duplicate_literal_removal,[],[f4276])).
fof(f4276,plain,(
  ( ! [X39,X37,X38] : (s(bool,'T_0') = s(bool,X38) | s(bool,'F_0') = s(bool,g(s(bool,X37))) | s(bool,'T_0') = s(bool,X38) | s(bool,'F_0') = s(bool,g(s(bool,X39))) | s(bool,X37) = s(bool,X39)) ) | ($spl22 | $spl47)),
  inference(superposition,[],[f3344,f3344])).
fof(f4882,plain,(
  ( ! [X37,X35] : (s(bool,X35) = s(bool,X37) | s(bool,'T_0') = s(bool,g(s(bool,X35))) | s(bool,'T_0') = s(bool,X37)) ) | $spl52),
  inference(cnf_transformation,[],[f4882_D])).
fof(f4882_D,plain,(
  ( ! [X37,X35] : (s(bool,X35) = s(bool,X37) | s(bool,'T_0') = s(bool,g(s(bool,X35))) | s(bool,'T_0') = s(bool,X37)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f4263,plain,(
  ( ! [X37,X35,X36] : (s(bool,'F_0') = s(bool,X36) | s(bool,X35) = s(bool,X37) | s(bool,'T_0') = s(bool,X37) | s(bool,'F_0') = s(bool,g(s(bool,X36))) | s(bool,'T_0') = s(bool,g(s(bool,X35)))) ) | ($spl22 | $spl47)),
  inference(superposition,[],[f3344,f3344])).
fof(f4869,plain,(
  ( ! [X41,X42,X40] : (s(bool,'F_0') = s(bool,g(s(bool,X40))) | s(bool,'T_0') = s(bool,X41) | s(bool,'F_0') = s(bool,g(s(bool,X42))) | s(bool,X40) = s(bool,X42)) ) | ($spl22 | $spl47)),
  inference(duplicate_literal_removal,[],[f4225])).
fof(f4225,plain,(
  ( ! [X41,X42,X40] : (s(bool,'F_0') = s(bool,g(s(bool,X40))) | s(bool,'T_0') = s(bool,X41) | s(bool,'T_0') = s(bool,X41) | s(bool,'F_0') = s(bool,g(s(bool,X42))) | s(bool,X40) = s(bool,X42)) ) | ($spl22 | $spl47)),
  inference(superposition,[],[f3344,f3344])).
fof(f3302,plain,(
  ( ! [X74,X75] : (s(fun(bool,bool),X74) = s(fun(bool,bool),X75)) ) | $spl44),
  inference(cnf_transformation,[],[f3302_D])).
fof(f3302_D,plain,(
  ( ! [X74,X75] : (s(fun(bool,bool),X74) = s(fun(bool,bool),X75)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f3304,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | $spl47),
  inference(cnf_transformation,[],[f3304_D])).
fof(f3304_D,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f3301,plain,(
  ( ! [X74,X75] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X74) = s(fun(bool,bool),X75)) ) | ($spl9 | $spl28 | $spl41)),
  inference(forward_demodulation,[],[f3285,f2771])).
fof(f3285,plain,(
  ( ! [X74,X75] : (s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X74) = s(fun(bool,bool),X75)) ) | $spl28),
  inference(duplicate_literal_removal,[],[f3271])).
fof(f3271,plain,(
  ( ! [X74,X75] : (s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X74) = s(fun(bool,bool),X75) | s(fun(bool,bool),X74) = s(fun(bool,bool),X75)) ) | $spl28),
  inference(superposition,[],[f45,f2449])).
fof(f2454,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl35),
  inference(cnf_transformation,[],[f2454_D])).
fof(f2454_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f2481,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X34) = s(fun(bool,bool),X35) | s(bool,'T_0') = s(bool,sK6(X34,X35,bool,bool))) ) | $spl26),
  inference(superposition,[],[f365,f2448])).
fof(f3170,plain,(
  ( ! [X64,X63] : (s(bool,X64) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,f(s(bool,X63))),s(bool,X64))))) | s(bool,'T_0') = s(bool,X63) | s(bool,'T_0') = s(bool,X64)) ) | $spl42),
  inference(cnf_transformation,[],[f3170_D])).
fof(f3170_D,plain,(
  ( ! [X64,X63] : (s(bool,X64) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,f(s(bool,X63))),s(bool,X64))))) | s(bool,'T_0') = s(bool,X63) | s(bool,'T_0') = s(bool,X64)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f2292,plain,(
  ( ! [X2,X0,X3] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),'const.pair.,_2'(s(bool,X0),s(X2,X3))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl20),
  inference(cnf_transformation,[],[f2292_D])).
fof(f2292_D,plain,(
  ( ! [X2,X0,X3] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),'const.pair.,_2'(s(bool,X0),s(X2,X3))))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f3044,plain,(
  ( ! [X78,X76,X79,X77,X75] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X77),'const.pair.,_2'(s(bool,X76),s(X77,X78))))) | s(bool,X79) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,f(s(bool,X75))),s(bool,X79))))) | s(bool,'T_0') = s(bool,X79) | s(bool,'T_0') = s(bool,X75) | s(bool,'T_0') = s(bool,X76)) ) | ($spl16 | $spl22 | $spl35)),
  inference(superposition,[],[f1557,f2683])).
fof(f2972,plain,(
  ( ! [X66,X64,X62,X65,X63] : (s(bool,X64) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,f(s(bool,X63))),s(bool,X64))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X65),'const.pair.,_2'(s(bool,X62),s(X65,X66))))) | s(bool,'T_0') = s(bool,X64) | s(bool,'T_0') = s(bool,X63) | s(bool,'T_0') = s(bool,X62)) ) | ($spl16 | $spl22 | $spl35)),
  inference(superposition,[],[f1557,f2683])).
fof(f2551,plain,(
  ( ! [X12,X13,X11] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X11,X12,bool,bool)))) | s(bool,X13) = s(bool,sK6(X11,X12,bool,bool)) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f370,f56])).
fof(f2548,plain,(
  ( ! [X4,X2,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,g(s(bool,sK6(X2,X3,bool,bool)))) | s(bool,X4) = s(bool,sK6(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f370,f56])).
fof(f2446,plain,(
  ( ! [X12,X13,X11] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,'T_0') = s(bool,sK6(X11,X12,bool,bool)) | s(bool,X13) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f365,f56])).
fof(f2455,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) ) | $spl36),
  inference(cnf_transformation,[],[f2455_D])).
fof(f2455_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f2445,plain,(
  ( ! [X10,X8,X9] : (s(bool,X8) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X9) = s(fun(bool,bool),X10) | s(bool,'T_0') = s(bool,sK6(X9,X10,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f365,f56])).
fof(f2444,plain,(
  ( ! [X6,X7,X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7) | s(bool,'T_0') = s(bool,sK6(X6,X7,bool,bool)) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f365,f56])).
fof(f2443,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,sK6(X2,X3,bool,bool)) | s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f365,f56])).
fof(f2451,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl31),
  inference(cnf_transformation,[],[f2451_D])).
fof(f2451_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f2442,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK6(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f365,f40])).
fof(f2293,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,X1))))) | s(bool,'T_0') = s(bool,X1)) ) | $spl22),
  inference(cnf_transformation,[],[f2293_D])).
fof(f2293_D,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,X1))))) | s(bool,'T_0') = s(bool,X1)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f2311,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,X45) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,X45))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X43),'const.pair.,_2'(s(bool,X42),s(X43,X44))))) | s(bool,'T_0') = s(bool,X45) | s(bool,'T_0') = s(bool,X42)) ) | ($spl1 | $spl9 | $spl16)),
  inference(forward_demodulation,[],[f2106,f1103])).
fof(f2106,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X43),'const.pair.,_2'(s(bool,X42),s(X43,X44))))) | s(bool,X45) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2))))),s(bool,X45))))) | s(bool,'T_0') = s(bool,X45) | s(bool,'T_0') = s(bool,X42)) ) | ($spl1 | $spl16)),
  inference(superposition,[],[f1557,f1101])).
fof(f2105,plain,(
  ( ! [X39,X41,X38,X40] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X39),'const.pair.,_2'(s(bool,X38),s(X39,X40))))) | s(bool,X41) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,X41))))) | s(bool,'T_0') = s(bool,X41) | s(bool,'T_0') = s(bool,X38)) ) | $spl16),
  inference(superposition,[],[f1557,f40])).
fof(f2305,plain,(
  ( ! [X97,X95,X93,X96] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X96),'const.pair.,_2'(s(bool,X93),s(X96,X97))))) | s(bool,X95) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,X95))))) | s(bool,'T_0') = s(bool,X95) | s(bool,'T_0') = s(bool,X93)) ) | $spl16),
  inference(forward_demodulation,[],[f2054,f43])).
fof(f2054,plain,(
  ( ! [X94,X92,X97,X95,X93,X96] : (s(bool,X95) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,X95))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X96),'const.pair.,_2'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X92),'const.pair.,_2'(s(bool,X93),s(X92,X94))))),s(X96,X97))))) | s(bool,'T_0') = s(bool,X95) | s(bool,'T_0') = s(bool,X93)) ) | $spl16),
  inference(superposition,[],[f1557,f961])).
fof(f2303,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,X1),s(X0,X2))))) | s(bool,X3) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,X3))))) | s(bool,'T_0') = s(bool,X3)) ) | $spl16),
  inference(condensation,[],[f2302])).
fof(f2302,plain,(
  ( ! [X85,X83,X81,X79,X84,X82] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X82),'const.pair.,_2'(s(bool,X79),s(X82,X83))))) | s(bool,X81) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,X81))))) | s(bool,'T_0') = s(bool,X81) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X84),'const.pair.,_2'(s(bool,X79),s(X84,X85)))))) ) | $spl16),
  inference(forward_demodulation,[],[f2052,f43])).
fof(f2052,plain,(
  ( ! [X80,X78,X85,X83,X81,X79,X84,X82] : (s(bool,X81) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,X81))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X82),'const.pair.,_2'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X78),'const.pair.,_2'(s(bool,X79),s(X78,X80))))),s(X82,X83))))) | s(bool,'T_0') = s(bool,X81) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X84),'const.pair.,_2'(s(bool,X79),s(X84,X85)))))) ) | $spl16),
  inference(superposition,[],[f1557,f1257])).
fof(f2300,plain,(
  ( ! [X68,X66,X69,X67,X63] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X66),'const.pair.,_2'(s(bool,X63),s(X66,X67))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X68,bool),'const.pair.,_2'(s(X68,X69),s(bool,X63)))))) ) | $spl24),
  inference(cnf_transformation,[],[f2300_D])).
fof(f2300_D,plain,(
  ( ! [X68,X66,X69,X67,X63] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X66),'const.pair.,_2'(s(bool,X63),s(X66,X67))))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X68,bool),'const.pair.,_2'(s(X68,X69),s(bool,X63)))))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f2299,plain,(
  ( ! [X68,X66,X69,X67,X65,X63] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X66),'const.pair.,_2'(s(bool,X63),s(X66,X67))))) | s(bool,X65) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,X65))))) | s(bool,'T_0') = s(bool,X65) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X68,bool),'const.pair.,_2'(s(X68,X69),s(bool,X63)))))) ) | $spl16),
  inference(forward_demodulation,[],[f2050,f43])).
fof(f2050,plain,(
  ( ! [X68,X66,X64,X62,X69,X67,X65,X63] : (s(bool,X65) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,X65))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X66),'const.pair.,_2'(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X62),'const.pair.,_2'(s(bool,X63),s(X62,X64))))),s(X66,X67))))) | s(bool,'T_0') = s(bool,X65) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X68,bool),'const.pair.,_2'(s(X68,X69),s(bool,X63)))))) ) | $spl16),
  inference(superposition,[],[f1557,f1258])).
fof(f2294,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,X25) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,X25))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X26),'const.pair.,_2'(s(bool,X24),s(X26,X27))))) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,X24)) ) | ($spl1 | $spl9 | $spl16)),
  inference(forward_demodulation,[],[f2041,f1103])).
fof(f2041,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,X25) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2))))),s(bool,X25))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X26),'const.pair.,_2'(s(bool,X24),s(X26,X27))))) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,X24)) ) | ($spl1 | $spl16)),
  inference(superposition,[],[f1557,f1101])).
fof(f2036,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,'F_0'),s(bool,X1))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X2),'const.pair.,_2'(s(bool,X0),s(X2,X3))))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl16),
  inference(superposition,[],[f1557,f40])).
fof(f1229,plain,(
  $false | ($spl9 | $spl18)),
  inference(subsumption_resolution,[],[f1208,f71])).
fof(f1208,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl18),
  inference(superposition,[],[f43,f967])).
fof(f1228,plain,(
  $false | ($spl9 | $spl18)),
  inference(subsumption_resolution,[],[f1202,f71])).
fof(f1202,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl18),
  inference(superposition,[],[f967,f43])).
fof(f1073,plain,(
  $false | ($spl10 | $spl13)),
  inference(subsumption_resolution,[],[f439,f437])).
fof(f1072,plain,(
  $false | ($spl1 | $spl12)),
  inference(subsumption_resolution,[],[f1071,f37])).
fof(f1071,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl12)),
  inference(forward_demodulation,[],[f51,f438])).
fof(f438,plain,(
  s(bool,'T_0') = s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2))))) | $spl12),
  inference(cnf_transformation,[],[f438_D])).
fof(f438_D,plain,(
  s(bool,'T_0') = s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2))))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1070,plain,(
  $false | ($spl0 | $spl5)),
  inference(subsumption_resolution,[],[f1023,f55])).
fof(f1023,plain,(
  p(s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) | $spl0),
  inference(forward_demodulation,[],[f1011,f44])).
fof(f1011,plain,(
  p(s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,'const.pair.SND_1'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2)))))))) | $spl0),
  inference(unit_resulting_resolution,[],[f50,f47])).
fof(f1069,plain,(
  $false | ($spl2 | $spl6)),
  inference(subsumption_resolution,[],[f1068,f39])).
fof(f1068,plain,(
  p(s(bool,'F_0')) | ($spl2 | $spl6)),
  inference(forward_demodulation,[],[f52,f69])).
fof(f1066,plain,(
  $false | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f1065,f39])).
fof(f1065,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f1023,f69])).
fof(f1064,plain,(
  $false | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f1063,f39])).
fof(f1063,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f1020,f69])).
fof(f1020,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) | $spl0),
  inference(forward_demodulation,[],[f1012,f43])).
fof(f1012,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,'const.pair.FST_1'(s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2)))))))) | $spl0),
  inference(unit_resulting_resolution,[],[f50,f46])).
fof(f1031,plain,(
  $false | ($spl0 | $spl3 | $spl6)),
  inference(global_subsumption,[],[f69,f36,f35,f34,f37,f39,f40,f43,f44,f45,f48,f47,f46,f56,f58,f64,f82,f104,f132,f133,f146,f156,f157,f200,f361,f363,f365,f366,f368,f370,f475,f550,f761,f762,f763,f785,f786,f50,f1010,f1013,f1020,f53])).
fof(f1013,plain,(
  s(bool,'T_0') = s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2))))) | $spl0),
  inference(unit_resulting_resolution,[],[f50,f58])).
fof(f1010,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(subsumption_resolution,[],[f1005,f37])).
fof(f1005,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f53,f56])).
fof(f1030,plain,(
  $false | ($spl0 | $spl3)),
  inference(global_subsumption,[],[f36,f35,f34,f37,f39,f40,f43,f44,f45,f48,f47,f46,f56,f58,f64,f82,f104,f132,f133,f146,f156,f157,f200,f361,f363,f365,f366,f368,f370,f475,f550,f761,f762,f763,f785,f786,f50,f1010,f1013,f1020,f53])).
fof(f1029,plain,(
  $false | ($spl0 | $spl3)),
  inference(global_subsumption,[],[f36,f35,f34,f37,f39,f40,f43,f44,f45,f48,f47,f46,f56,f58,f64,f82,f104,f132,f133,f146,f156,f157,f200,f361,f363,f365,f366,f368,f370,f475,f550,f761,f762,f763,f785,f786,f53,f50,f1010,f1013,f1020])).
fof(f1022,plain,(
  $false | ($spl0 | $spl3 | $spl9)),
  inference(subsumption_resolution,[],[f1021,f39])).
fof(f1021,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl3 | $spl9)),
  inference(forward_demodulation,[],[f1020,f1003])).
fof(f1003,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3))) | ($spl3 | $spl9)),
  inference(unit_resulting_resolution,[],[f71,f53,f200])).
fof(f1002,plain,(
  $false | ($spl0 | $spl13)),
  inference(subsumption_resolution,[],[f1001,f39])).
fof(f1001,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl13)),
  inference(forward_demodulation,[],[f50,f502])).
fof(f961,plain,(
  ( ! [X14,X12,X13] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X13),'const.pair.,_2'(s(bool,X12),s(X13,X14))))) | s(bool,'T_0') = s(bool,X12)) ) | $spl16),
  inference(cnf_transformation,[],[f961_D])).
fof(f961_D,plain,(
  ( ! [X14,X12,X13] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X13),'const.pair.,_2'(s(bool,X12),s(X13,X14))))) | s(bool,'T_0') = s(bool,X12)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f985,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X30),'const.pair.,_2'(s(bool,'F_0'),s(X30,X31))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X32),'const.pair.,_2'(s(bool,X29),s(X32,X33))))) | s(bool,'T_0') = s(bool,X29)) ) | ($spl1 | $spl9)),
  inference(forward_demodulation,[],[f843,f371])).
fof(f371,plain,(
  s(bool,'F_0') = s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2))))) | ($spl1 | $spl9)),
  inference(unit_resulting_resolution,[],[f71,f360])).
fof(f360,plain,(
  ( ! [X9] : (s(bool,X9) = s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2))))) | s(bool,'T_0') = s(bool,X9)) ) | $spl1),
  inference(subsumption_resolution,[],[f212,f37])).
fof(f212,plain,(
  ( ! [X9] : (~p(s(bool,'T_0')) | s(bool,X9) = s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2))))) | s(bool,'T_0') = s(bool,X9)) ) | $spl1),
  inference(superposition,[],[f51,f56])).
fof(f843,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X30),'const.pair.,_2'(s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2))))),s(X30,X31))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X32),'const.pair.,_2'(s(bool,X29),s(X32,X33))))) | s(bool,'T_0') = s(bool,X29)) ) | ($spl1 | $spl9)),
  inference(superposition,[],[f683,f360])).
fof(f683,plain,(
  ( ! [X59,X57,X58,X56,X55] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X55),'const.pair.,_2'(s(bool,X56),s(X55,X57))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X58),'const.pair.,_2'(s(bool,X56),s(X58,X59)))))) ) | ($spl1 | $spl9)),
  inference(forward_demodulation,[],[f586,f371])).
fof(f586,plain,(
  ( ! [X59,X57,X58,X56,X55] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X58),'const.pair.,_2'(s(bool,X56),s(X58,X59))))) | s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X55),'const.pair.,_2'(s(bool,X56),s(X55,X57))))) = s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2)))))) ) | $spl1),
  inference(superposition,[],[f377,f43])).
fof(f377,plain,(
  ( ! [X14,X15,X16] : (s(bool,X15) = s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X14),'const.pair.,_2'(s(bool,X15),s(X14,X16)))))) ) | $spl1),
  inference(superposition,[],[f360,f43])).
fof(f984,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X19),'const.pair.,_2'(s(bool,X18),s(X19,X20))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X21),'const.pair.,_2'(s(bool,X17),s(X21,X22))))) | s(bool,'T_0') = s(bool,X17)) ) | ($spl1 | $spl9 | $spl13)),
  inference(subsumption_resolution,[],[f841,f632])).
fof(f632,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X0) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,X0),s(X1,X2)))))) ) | ($spl1 | $spl13)),
  inference(superposition,[],[f439,f377])).
fof(f841,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X19),'const.pair.,_2'(s(bool,X18),s(X19,X20))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X21),'const.pair.,_2'(s(bool,X17),s(X21,X22))))) | s(bool,'T_0') = s(bool,X17) | s(bool,'T_0') = s(bool,X18)) ) | ($spl1 | $spl9)),
  inference(superposition,[],[f683,f56])).
fof(f983,plain,(
  ( ! [X14,X12,X15,X13,X11,X16] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X13),'const.pair.,_2'(s(bool,X12),s(X13,X14))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X15),'const.pair.,_2'(s(bool,X11),s(X15,X16))))) | s(bool,'T_0') = s(bool,X11)) ) | ($spl1 | $spl9 | $spl13)),
  inference(subsumption_resolution,[],[f840,f632])).
fof(f840,plain,(
  ( ! [X14,X12,X15,X13,X11,X16] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X13),'const.pair.,_2'(s(bool,X12),s(X13,X14))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X15),'const.pair.,_2'(s(bool,X11),s(X15,X16))))) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X11)) ) | ($spl1 | $spl9)),
  inference(superposition,[],[f683,f56])).
fof(f838,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'F_0'),s(X1,X2))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X3),'const.pair.,_2'(s(bool,X0),s(X3,X4))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl9)),
  inference(superposition,[],[f683,f40])).
fof(f966,plain,(
  ( ! [X70,X72,X71,X69,X73] : (s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X72),'const.pair.,_2'(s(bool,'F_0'),s(X72,X73))))) | s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X70),'const.pair.,_2'(s(bool,X69),s(X70,X71))))) | s(bool,'T_0') = s(bool,X69)) ) | ($spl1 | $spl9)),
  inference(forward_demodulation,[],[f801,f371])).
fof(f801,plain,(
  ( ! [X70,X72,X71,X69,X73] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X70),'const.pair.,_2'(s(bool,X69),s(X70,X71))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X72),'const.pair.,_2'(s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2))))),s(X72,X73))))) | s(bool,'T_0') = s(bool,X69)) ) | ($spl1 | $spl9)),
  inference(superposition,[],[f683,f360])).
fof(f962,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X19),'const.pair.,_2'(s(bool,X18),s(X19,X20))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X21),'const.pair.,_2'(s(bool,X17),s(X21,X22))))) | s(bool,'T_0') = s(bool,X18)) ) | ($spl1 | $spl9 | $spl13)),
  inference(subsumption_resolution,[],[f792,f632])).
fof(f792,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X19),'const.pair.,_2'(s(bool,X18),s(X19,X20))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X21),'const.pair.,_2'(s(bool,X17),s(X21,X22))))) | s(bool,'T_0') = s(bool,X17) | s(bool,'T_0') = s(bool,X18)) ) | ($spl1 | $spl9)),
  inference(superposition,[],[f683,f56])).
fof(f959,plain,(
  ( ! [X14,X12,X15,X13,X11,X16] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X13),'const.pair.,_2'(s(bool,X12),s(X13,X14))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X15),'const.pair.,_2'(s(bool,X11),s(X15,X16))))) | s(bool,'T_0') = s(bool,X12)) ) | ($spl1 | $spl9 | $spl13)),
  inference(subsumption_resolution,[],[f791,f632])).
fof(f791,plain,(
  ( ! [X14,X12,X15,X13,X11,X16] : (s(bool,'F_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X13),'const.pair.,_2'(s(bool,X12),s(X13,X14))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X15),'const.pair.,_2'(s(bool,X11),s(X15,X16))))) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X11)) ) | ($spl1 | $spl9)),
  inference(superposition,[],[f683,f56])).
fof(f437,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2)))))) ) | $spl10),
  inference(cnf_transformation,[],[f437_D])).
fof(f437_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2)))))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f439,plain,(
  s(bool,'T_0') != s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2))))) | $spl13),
  inference(cnf_transformation,[],[f439_D])).
fof(f439_D,plain,(
  s(bool,'T_0') != s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2))))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f412,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2))))) | s(bool,X0) = s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2)))))) ) | $spl1),
  inference(equality_factoring,[],[f360])).
fof(f71,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl9),
  inference(cnf_transformation,[],[f71_D])).
fof(f71_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f64,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f40])).
fof(f50,plain,(
  p(s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2)))))) | $spl0),
  inference(cnf_transformation,[],[f50_D])).
fof(f50_D,plain,(
  p(s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2)))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f52,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) | $spl2),
  inference(cnf_transformation,[],[f52_D])).
fof(f52_D,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f34,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2))))))),
  inference(cnf_transformation,[],[f29])).
fof(f29,plain,(
  (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2)))))) | (p(s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) & p(s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))))) & (~p(s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2)))))) | ~p(s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) | ~p(s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5])],[f28])).
fof(f28,plain,(
  ? [X0,X1,X2,X3,X4,X5] : ((p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X0,bool),X5),s(fun(X1,bool),X4),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X2)))))) | (p(s(bool,happ(s(fun(X0,bool),X5),s(X0,X3)))) & p(s(bool,happ(s(fun(X1,bool),X4),s(X1,X2)))))) & (~p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X0,bool),X5),s(fun(X1,bool),X4),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X2)))))) | ~p(s(bool,happ(s(fun(X0,bool),X5),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X1,bool),X4),s(X1,X2))))))),
  inference(flattening,[],[f27])).
fof(f27,plain,(
  ? [X0,X1,X2,X3,X4,X5] : ((p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X0,bool),X5),s(fun(X1,bool),X4),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X2)))))) | (p(s(bool,happ(s(fun(X0,bool),X5),s(X0,X3)))) & p(s(bool,happ(s(fun(X1,bool),X4),s(X1,X2)))))) & (~p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X0,bool),X5),s(fun(X1,bool),X4),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X2)))))) | (~p(s(bool,happ(s(fun(X0,bool),X5),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X1,bool),X4),s(X1,X2)))))))),
  inference(nnf_transformation,[],[f25])).
fof(f25,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X0,bool),X5),s(fun(X1,bool),X4),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X2)))))) <~> (p(s(bool,happ(s(fun(X0,bool),X5),s(X0,X3)))) & p(s(bool,happ(s(fun(X1,bool),X4),s(X1,X2))))))),
  inference(ennf_transformation,[],[f13])).
fof(f13,plain,(
  ~! [X0,X1,X2,X3,X4,X5] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X0,bool),X5),s(fun(X1,bool),X4),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X2)))))) <=> (p(s(bool,happ(s(fun(X0,bool),X5),s(X0,X3)))) & p(s(bool,happ(s(fun(X1,bool),X4),s(X1,X2))))))),
  inference(flattening,[],[f12])).
fof(f12,plain,(
  ~! [X0,X1] : ! [X2,X3,X4,X5] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X0,bool),X5),s(fun(X1,bool),X4),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X3),s(X1,X2)))))) <=> (p(s(bool,happ(s(fun(X0,bool),X5),s(X0,X3)))) & p(s(bool,happ(s(fun(X1,bool),X4),s(X1,X2))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10,X11] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X6,bool),X11),s(fun(X7,bool),X10),s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,X9),s(X7,X8)))))) <=> (p(s(bool,happ(s(fun(X6,bool),X11),s(X6,X9)))) & p(s(bool,happ(s(fun(X7,bool),X10),s(X7,X8))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10,X11] : (p(s(bool,'const.pair.PROD_ALL_3'(s(fun(X6,bool),X11),s(fun(X7,bool),X10),s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,X9),s(X7,X8)))))) <=> (p(s(bool,happ(s(fun(X6,bool),X11),s(X6,X9)))) & p(s(bool,happ(s(fun(X7,bool),X10),s(X7,X8))))))),
  file('thm.pair.PROD_ALL_THM','thm.pair.PROD_ALL_THM')).
fof(f54,plain,(
  p(s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) | $spl4),
  inference(cnf_transformation,[],[f54_D])).
fof(f54_D,plain,(
  p(s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f35,plain,(
  p(s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) | p(s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2))))))),
  inference(cnf_transformation,[],[f29])).
fof(f51,plain,(
  ~p(s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2)))))) | $spl1),
  inference(cnf_transformation,[],[f51_D])).
fof(f51_D,plain,(
  ~p(s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2)))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f53,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) | $spl3),
  inference(cnf_transformation,[],[f53_D])).
fof(f53_D,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f55,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) | $spl5),
  inference(cnf_transformation,[],[f55_D])).
fof(f55_D,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f36,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),sK4),s(sK1,sK2)))) | ~p(s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK3)))) | ~p(s(bool,'const.pair.PROD_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK4),s('type.pair.prod'(sK0,sK1),'const.pair.,_2'(s(sK0,sK3),s(sK1,sK2))))))),
  inference(cnf_transformation,[],[f29])).
% SZS output end Proof for thm.pair
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 353
% Passive clauses: 3310
% Generated clauses: 96059
% Final active clauses: 47
% Final passive clauses: 258
% Input formulas: 10
% Initial clauses: 15
% 
% Duplicate literals: 6459
% Trivial inequalities: 858
% Fw subsumption resolutions: 6853
% Fw demodulations: 9282
% Bw demodulations: 2311
% Condensations: 13470
% Global subsumptions: 84
% 
% Simple tautologies: 963
% Equational tautologies: 1171
% Forward subsumptions: 53363
% Fw demodulations to eq. taut.: 1
% Bw demodulations to eq. taut.: 185
% 
% Unit resulting resolution: 4045
% Forward superposition: 27914
% Backward superposition: 18579
% Self superposition: 630
% Equality factoring: 300
% Equality resolution: 35
% 
% Split clauses: 315
% Split components: 639
% SAT solver clauses: 4102
% SAT solver unit clauses: 13
% SAT solver binary clauses: 181
% 
% Sat splits: 315
% Sat splitting refutations: 76
% 
% Memory used [KB]: 46054
% Time elapsed: 6.451 s
% ------------------------------
% ------------------------------
% Success in time 7.753 s
