Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.real
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.real
% SZS output start Proof for thm.real
fof(f25072,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f25071,f439,f761,f759,f1380,f440,f762,f760])).
fof(f760,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15))) | $spl1),
  inference(cnf_transformation,[],[f760_D])).
fof(f760_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f762,plain,(
  s('type.num.num',sK14) != s('type.num.num',sK15) | $spl3),
  inference(cnf_transformation,[],[f762_D])).
fof(f762_D,plain,(
  s('type.num.num',sK14) != s('type.num.num',sK15) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f440,plain,(
  s('type.num.num',sK14) != s('type.num.num',sK15) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15)))),
  inference(cnf_transformation,[],[f316])).
fof(f316,plain,(
  (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15))) | s('type.num.num',sK14) = s('type.num.num',sK15)) & (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15))) | s('type.num.num',sK14) != s('type.num.num',sK15))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14,sK15])],[f315])).
fof(f315,plain,(
  ? [X0,X1] : ((s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))) | s('type.num.num',X0) != s('type.num.num',X1)))),
  inference(nnf_transformation,[],[f256])).
fof(f256,plain,(
  ? [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))) <~> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(ennf_transformation,[],[f119])).
fof(f119,plain,(
  ~! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X6))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X7))) <=> s('type.num.num',X6) = s('type.num.num',X7))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X6))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X7))) <=> s('type.num.num',X6) = s('type.num.num',X7))),
  file('thm.real.eq_ints|split|0','thm.real.eq_ints|split|0')).
fof(f1380,plain,(
  $false | ($spl1 | $spl2)),
  inference(trivial_inequality_removal,[],[f1379])).
fof(f1379,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f760,f761])).
fof(f759,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15))) | $spl0),
  inference(cnf_transformation,[],[f759_D])).
fof(f759_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f761,plain,(
  s('type.num.num',sK14) = s('type.num.num',sK15) | $spl2),
  inference(cnf_transformation,[],[f761_D])).
fof(f761_D,plain,(
  s('type.num.num',sK14) = s('type.num.num',sK15) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f439,plain,(
  s('type.num.num',sK14) = s('type.num.num',sK15) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15)))),
  inference(cnf_transformation,[],[f316])).
fof(f25071,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f25054,f762])).
fof(f25054,plain,(
  s('type.num.num',sK14) = s('type.num.num',sK15) | $spl0),
  inference(equality_resolution,[],[f11187])).
fof(f11187,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))) | s('type.num.num',sK15) = s('type.num.num',X0)) ) | $spl0),
  inference(superposition,[],[f589,f759])).
fof(f589,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f366])).
fof(f366,plain,(
  ! [X0,X1] : ((s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.num.num',X0) != s('type.num.num',X1) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f102])).
fof(f102,axiom,(
  ! [X7,X6] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X6))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X7))) <=> s('type.num.num',X6) = s('type.num.num',X7))),
  file('thm.real.eq_ints|split|0','thm.real.REAL_INJ')).
% SZS output end Proof for thm.real
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 999
% Passive clauses: 5143
% Generated clauses: 24595
% Final active clauses: 921
% Final passive clauses: 3702
% Input formulas: 117
% Initial clauses: 310
% 
% Duplicate literals: 233
% Trivial inequalities: 63
% Fw subsumption resolutions: 1980
% Fw demodulations: 2369
% Bw demodulations: 70
% 
% Simple tautologies: 240
% Equational tautologies: 31
% Forward subsumptions: 14446
% Fw demodulations to eq. taut.: 3
% 
% Binary resolution: 5818
% Factoring: 1271
% Forward superposition: 5286
% Backward superposition: 6995
% Self superposition: 31
% Equality factoring: 2
% Equality resolution: 16
% 
% Split clauses: 280
% Split components: 562
% SAT solver clauses: 482
% SAT solver unit clauses: 14
% SAT solver binary clauses: 67
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 1
% 
% Sat splits: 280
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 480
% TWLsolver calls for satisfiability: 428
% 
% Memory used [KB]: 13176
% Time elapsed: 0.433 s
% ------------------------------
% ------------------------------
% Success in time 0.438 s
