Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.realax
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.realax
% SZS output start Proof for thm.realax
fof(f1944,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f1931,f44,f61,f63,f48,f64])).
fof(f64,plain,(
  ~p(s(bool,'F_0')) | $spl3),
  inference(cnf_transformation,[],[f64_D])).
fof(f64_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f48,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f20])).
fof(f20,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.realax.TREAL_|49||48|',aHL_FALSITY)).
fof(f63,plain,(
  p(s(bool,'F_0')) | $spl2),
  inference(cnf_transformation,[],[f63_D])).
fof(f63_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f61,plain,(
  p(s(bool,'T_0')) | $spl0),
  inference(cnf_transformation,[],[f61_D])).
fof(f61_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f44,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f34])).
fof(f34,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f18])).
fof(f18,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f12])).
fof(f12,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.realax.TREAL_|49||48|','thm.bool.NOT_CLAUSES|split|2')).
fof(f1931,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1930,f61])).
fof(f1930,plain,(
  ~p(s(bool,'T_0')) | $spl3),
  inference(forward_demodulation,[],[f1929,f73])).
fof(f73,plain,(
  s(bool,'T_0') = s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|49|_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0'))) | $spl3),
  inference(subsumption_resolution,[],[f70,f64])).
fof(f70,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|49|_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0')))),
  inference(superposition,[],[f39,f49])).
fof(f49,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.realax.TREAL_|49||48|',aHL_BOOL_CASES)).
fof(f39,plain,(
  p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|49|_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0'))))),
  inference(cnf_transformation,[],[f17])).
fof(f17,plain,(
  p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|49|_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0'))))),
  inference(flattening,[],[f6])).
fof(f6,negated_conjecture,(
  ~~p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|49|_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0'))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ~p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|49|_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0'))))),
  file('thm.realax.TREAL_|49||48|','thm.realax.TREAL_|49||48|')).
fof(f1929,plain,(
  ~p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|49|_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0'))))),
  inference(subsumption_resolution,[],[f1915,f57])).
fof(f57,plain,(
  ( ! [X0,X1] : (s('type.hreal.hreal',X0) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1)))) )),
  inference(cnf_transformation,[],[f26])).
fof(f26,plain,(
  ! [X0,X1] : s('type.hreal.hreal',X0) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1)))),
  inference(flattening,[],[f25])).
fof(f25,plain,(
  ! [X0,X1] : ~s('type.hreal.hreal',X0) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X9,X10] : ~s('type.hreal.hreal',X9) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X9),s('type.hreal.hreal',X10)))),
  file('thm.realax.TREAL_|49||48|','thm.hreal.HREAL_NOZERO')).
fof(f1915,plain,(
  ~p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|49|_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0')))) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_|49|_0'),s('type.hreal.hreal','const.hreal.hreal_|49|_0'))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_|49|_0'),s('type.hreal.hreal','const.hreal.hreal_|49|_0'))),s('type.hreal.hreal','const.hreal.hreal_|49|_0')))),
  inference(superposition,[],[f705,f43])).
fof(f43,plain,(
  s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|49|_0') = s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_|49|_0'),s('type.hreal.hreal','const.hreal.hreal_|49|_0'))),s('type.hreal.hreal','const.hreal.hreal_|49|_0')))),
  inference(cnf_transformation,[],[f15])).
fof(f15,axiom,(
  s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|49|_0') = s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_|49|_0'),s('type.hreal.hreal','const.hreal.hreal_|49|_0'))),s('type.hreal.hreal','const.hreal.hreal_|49|_0')))),
  file('thm.realax.TREAL_|49||48|','thm.realax.treal_|49|')).
fof(f705,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0')))) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal','const.hreal.hreal_|49|_0'))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_|49|_0'),s('type.hreal.hreal',X1)))) )),
  inference(superposition,[],[f59,f42])).
fof(f42,plain,(
  s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0') = s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal','const.hreal.hreal_|49|_0'),s('type.hreal.hreal','const.hreal.hreal_|49|_0')))),
  inference(cnf_transformation,[],[f14])).
fof(f14,axiom,(
  s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0') = s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal','const.hreal.hreal_|49|_0'),s('type.hreal.hreal','const.hreal.hreal_|49|_0')))),
  file('thm.realax.TREAL_|49||48|','thm.realax.treal_|48|')).
fof(f59,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal',X2),s('type.hreal.hreal',X3)))))) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X3))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X2),s('type.hreal.hreal',X1)))) )),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal',X2),s('type.hreal.hreal',X3)))))) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X3))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X2),s('type.hreal.hreal',X1)))) & (s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X3))) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X2),s('type.hreal.hreal',X1))) | p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal',X2),s('type.hreal.hreal',X3))))))))),
  inference(nnf_transformation,[],[f29])).
fof(f29,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal',X2),s('type.hreal.hreal',X3)))))) <=> s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X3))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X2),s('type.hreal.hreal',X1))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X11,X12,X13,X14] : (p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal',X11),s('type.hreal.hreal',X12))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal',X13),s('type.hreal.hreal',X14)))))) <=> s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X11),s('type.hreal.hreal',X14))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X13),s('type.hreal.hreal',X12))))),
  file('thm.realax.TREAL_|49||48|','thm.realax.treal_eq')).
% SZS output end Proof for thm.realax
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 49
% Passive clauses: 315
% Generated clauses: 1838
% Final active clauses: 40
% Final passive clauses: 220
% Input formulas: 15
% Initial clauses: 22
% 
% Duplicate literals: 179
% Trivial inequalities: 25
% Fw subsumption resolutions: 251
% Fw demodulations: 107
% Bw demodulations: 7
% 
% Simple tautologies: 55
% Equational tautologies: 33
% Forward subsumptions: 812
% 
% Binary resolution: 66
% Factoring: 1
% Forward superposition: 558
% Backward superposition: 478
% Self superposition: 54
% Equality factoring: 58
% Equality resolution: 2
% 
% Split clauses: 54
% Split components: 109
% SAT solver clauses: 103
% SAT solver unit clauses: 3
% SAT solver binary clauses: 29
% SAT solver learnt clauses: 1
% 
% Sat splits: 54
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 97
% TWLsolver calls for satisfiability: 38
% 
% Memory used [KB]: 1407
% Time elapsed: 0.051 s
% ------------------------------
% ------------------------------
% Success in time 0.053 s
