Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.integer
% (40544)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1069
% Passive clauses: 7510
% Generated clauses: 30954
% Final active clauses: 941
% Final passive clauses: 5827
% Input formulas: 261
% Initial clauses: 512
% 
% Duplicate literals: 297
% Trivial inequalities: 94
% Fw subsumption resolutions: 1584
% Fw demodulations: 4622
% Bw demodulations: 125
% 
% Simple tautologies: 255
% Equational tautologies: 31
% Forward subsumptions: 16337
% Fw demodulations to eq. taut.: 4
% 
% Binary resolution: 8333
% Factoring: 1676
% Forward superposition: 7100
% Backward superposition: 6235
% Self superposition: 49
% Equality factoring: 2
% Equality resolution: 22
% 
% Split clauses: 463
% Split components: 928
% SAT solver clauses: 679
% SAT solver unit clauses: 15
% SAT solver binary clauses: 116
% 
% Sat splits: 463
% Sat splitting refutations: 8
% 
% TWLsolver clauses: 675
% TWLsolver calls for satisfiability: 445
% 
% Memory used [KB]: 13560
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.integer
% (40561)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 590
% Passive clauses: 3604
% Generated clauses: 21053
% Final active clauses: 465
% Final passive clauses: 2841
% Input formulas: 261
% Initial clauses: 512
% 
% Duplicate literals: 86
% Trivial inequalities: 18
% Fw subsumption resolutions: 486
% Fw demodulations: 3885
% Bw demodulations: 296
% Condensations: 60
% Global subsumptions: 23
% 
% Simple tautologies: 228
% Equational tautologies: 13
% Forward subsumptions: 8306
% Fw demodulations to eq. taut.: 16
% 
% Binary resolution: 3339
% Factoring: 99
% Forward superposition: 5268
% Backward superposition: 3203
% Self superposition: 72
% Equality factoring: 2
% Equality resolution: 9
% 
% Split clauses: 44
% Split components: 88
% SAT solver clauses: 4672
% SAT solver unit clauses: 607
% SAT solver binary clauses: 187
% SAT solver learnt clauses: 46
% SAT solver learnt literals: 12
% 
% Sat splits: 44
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 4580
% TWLsolver calls for satisfiability: 3647
% 
% Memory used [KB]: 21748
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.integer
% (40590)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 2708
% Passive clauses: 20093
% Generated clauses: 130457
% Final active clauses: 2356
% Final passive clauses: 5263
% Input formulas: 261
% Initial clauses: 512
% Discarded non-redundant clauses: 31305
% 
% Duplicate literals: 9
% Trivial inequalities: 56
% Fw subsumption resolutions: 1311
% Fw demodulations: 25600
% Bw demodulations: 1817
% Condensations: 33
% Global subsumptions: 5
% 
% Simple tautologies: 92
% Equational tautologies: 134
% Forward subsumptions: 66904
% Fw demodulations to eq. taut.: 346
% Bw demodulations to eq. taut.: 8
% 
% Unit resulting resolution: 57369
% Factoring: 7
% Forward superposition: 24918
% Backward superposition: 17193
% Self superposition: 306
% Equality factoring: 2
% Equality resolution: 19
% 
% Split clauses: 54
% Split components: 120
% SAT solver clauses: 21380
% SAT solver unit clauses: 5087
% SAT solver binary clauses: 10499
% 
% Sat splits: 54
% Sat splitting refutations: 1172
% 
% Memory used [KB]: 116416
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.integer
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.integer
% SZS output start Proof for thm.integer
fof(f70715,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f753,f754,f756,f757,f758,f759,f760,f761,f762,f764,f765,f835,f893,f894,f896,f897,f898,f899,f900,f901,f902,f904,f905,f975,f1133,f1197,f1198,f1199,f1200,f1203,f1204,f1205,f1206,f1207,f1208,f1209,f1210,f1211,f1212,f2285,f2286,f2287,f2288,f2289,f2290,f2292,f2366,f2424,f2454,f2455,f2456,f2457,f2458,f2460,f2461,f2531,f2589,f2590,f2591,f2592,f2593,f2594,f2596,f2670,f2738,f2768,f2769,f2770,f2771,f2772,f2774,f2775,f2845,f2899,f2900,f2901,f2902,f2903,f2904,f2906,f2980,f3039,f3040,f3041,f3042,f3043,f3044,f3046,f3120,f3200,f3462,f3463,f3464,f3465,f3466,f3467,f3469,f3543,f2896,f2895,f2889,f2888,f2767,f2766,f2764,f2763,f4118,f4133,f4134,f4135,f4136,f4137,f4138,f4140,f4214,f4273,f4274,f4275,f4276,f4277,f2726,f4278,f4280,f4354,f4420,f4678,f4679,f4680,f4681,f4682,f4683,f4685,f4759,f4130,f4129,f4123,f4122,f2453,f2452,f2450,f2449,f5333,f5352,f5353,f5354,f5355,f5356,f2412,f5357,f5359,f5433,f2411,f5489,f5823,f5802,f5887,f5888,f5889,f5890,f5893,f5894,f5895,f5896,f5897,f5898,f5899,f5900,f5901,f5902,f6057,f6143,f6144,f6145,f6146,f6147,f6148,f6149,f6150,f6151,f6154,f6155,f6156,f6157,f6158,f7507,f2426,f7618,f7650,f7651,f7652,f7653,f7654,f7656,f880,f7657,f7727,f7781,f7782,f7783,f7784,f7785,f3023,f7786,f7788,f7862,f7949,f7950,f7951,f7952,f7953,f7954,f7956,f8030,f8098,f8128,f8129,f8130,f8131,f8132,f8134,f1019,f8135,f8205,f8260,f8261,f8262,f8263,f8264,f3188,f8265,f8267,f8341,f8419,f8647,f8648,f8649,f8650,f8651,f3582,f8652,f8654,f8728,f8787,f8788,f8789,f8790,f8791,f8793,f8903,f8794,f8864,f8917,f8918,f8919,f8920,f8921,f4257,f8922,f8924,f8998,f9298,f9299,f9300,f9301,f9302,f4798,f9303,f9305,f9379,f9438,f9439,f9440,f9441,f9442,f9444,f9554,f9445,f9515,f8907,f8906,f8127,f8126,f9826,f10105,f10106,f10107,f10108,f10109,f881,f8086,f10110,f10112,f10186,f10252,f11146,f12752,f17239,f17241,f17242,f17243,f17246,f17316,f17383,f17384,f2428,f17390,f17391,f17521,f17523,f17394,f17396,f17397,f17399,f17401,f17475,f17524,f17599,f17601,f17602,f17603,f17606,f17676,f17361,f18943,f18945,f18946,f18947,f18950,f19020,f19243,f19244,f19245,f19248,f19249,f19102,f19101,f19260,f19095,f19094,f19262,f8914,f8913,f8124,f8123,f19186,f19112,f19110,f19108,f19107,f19105,f19263,f19289,f19291,f19292,f19293,f19296,f19366,f19420,f19422,f19423,f19425,f19427,f19501,f19587,f19589,f19590,f19592,f19594,f19668,f3187,f19745,f19973,f19975,f19976,f19978,f19980,f20054,f20112,f20114,f20115,f20116,f20119,f20189,f21210,f3204,f3202,f23374,f23376,f23377,f23378,f23381,f23451,f23599,f23600,f23606,f23607,f23747,f23749,f23610,f23612,f23613,f23615,f23617,f23691,f23750,f23897,f23899,f23900,f23901,f23904,f23974,f24057,f24059,f24060,f24061,f24064,f24134,f24201,f24448,f24461,f24462,f24464,f24466,f24468,f24475,f24476,f24477,f19417,f19416,f19410,f19409,f19288,f19287,f19285,f19284,f24483,f24485,f24380,f24306,f24304,f24302,f24301,f24299,f24488,f25125,f25127,f25128,f25130,f25132,f25206,f25271,f29157,f29159,f29160,f29161,f29164,f29234,f29455,f29456,f29457,f29458,f29459,f29460,f29461,f29490,f29492,f29493,f29495,f29497,f29571,f29641,f29642,f29643,f29644,f29646,f29647,f874,f29648,f29655,f29656,f29645,f29657,f29649,f7771,f7770,f7649,f7648,f24286,f876,f1016,f29663,f29665,f877,f30007,f30296,f30298,f30299,f30300,f30303,f30373,f1018,f30430,f30432,f30433,f30435,f30437,f30511,f4256,f30567,f30569,f30570,f30572,f30574,f30648,f2725,f30713,f30971,f30973,f30974,f30976,f30978,f31052,f4797,f31110,f31112,f31113,f31114,f31117,f31187,f9553,f32566,f33206,f2742,f2740,f36925,f36926,f36962,f36964,f36970,f37765,f37771,f37772,f37774,f37776,f37777,f37815,f37818,f37819,f37821,f37823,f37824,f39225,f39227,f39228,f39229,f39232,f39302,f23506,f40097,f40693,f43479,f43481,f43482,f43484,f43486,f43560,f43615,f43638,f43640,f43641,f43642,f43645,f43715,f43769,f43832,f43835,f43857,f43918,f43920,f43978,f43980,f44037,f44039,f44097,f44099,f44157,f44159,f44217,f44219,f44280,f44282,f44343,f44345,f44480,f44483,f44485,f44486,f44488,f44490,f44564,f44636,f44715,f44717,f44843,f44845,f44902,f44904,f44960,f44962,f45018,f45020,f45077,f45079,f45136,f45138,f45196,f43833,f45198,f43836,f46670,f46672,f46673,f46675,f46677,f46751,f47040,f47037,f47021,f47020,f47017,f47107,f47108,f47109,f47110,f47111,f47112,f47113,f47114,f47115,f47118,f47119,f47120,f47121,f47122,f47203,f47205,f47206,f47207,f47210,f47280,f47357,f47665,f47667,f47668,f47669,f47672,f47742,f47819,f47909,f47911,f47912,f47914,f47916,f47990,f48064,f48272,f48273,f48274,f48275,f48404,f48401,f48385,f48384,f48381,f48471,f48472,f48473,f48474,f48475,f48476,f48477,f48478,f48479,f48482,f48483,f48484,f48485,f48486,f48498,f48499,f48500,f48501,f48279,f48577,f48578,f48581,f48582,f48592,f48593,f48277,f48596,f48597,f48584,f48659,f48661,f48662,f48663,f48666,f48736,f48952,f48951,f48948,f48929,f48928,f49016,f49017,f49018,f49019,f49022,f49023,f49024,f49025,f49026,f49027,f49028,f49029,f49030,f49031,f49587,f49903,f49905,f872,f7778,f7777,f7646,f7645,f875,f873,f47510,f47509,f49915,f49916,f49917,f49926,f49927,f49934,f49935,f24463,f49936,f49945,f49946,f24465,f30427,f30426,f30420,f30419,f30295,f30294,f30292,f30291,f1013,f49948,f49949,f49955,f49957,f49965,f49967,f1017,f49972,f49973,f49974,f49975,f49976,f49977,f49979,f49981,f49947,f49939,f49937,f49928,f49920,f24469,f1015,f49918,f24467,f1014,f50604,f50606,f50607,f50608,f50611,f50681,f878,f50737,f50739,f50740,f50742,f50744,f50818,f8085,f50883,f50897,f50899,f50900,f50902,f50904,f50978,f3022,f51277,f51279,f51280,f51282,f51284,f51358,f3581,f51416,f51418,f51419,f51420,f51423,f51493,f8902,f52866,f8102,f8100,f54818,f57232,f57234,f57270,f57272,f57278,f58075,f58081,f58082,f58084,f58086,f58087,f58125,f58128,f58129,f58131,f58133,f58134,f61118,f61120,f61121,f61122,f61125,f61195,f19075,f61278,f61280,f61281,f61283,f61285,f61359,f29626,f63181,f63183,f63184,f63185,f63188,f63258,f63312,f63375,f63377,f43834,f63435,f63436,f43919,f63494,f63495,f43979,f63552,f63553,f44038,f63610,f63611,f44098,f63668,f63669,f44158,f63726,f63727,f44218,f63788,f63789,f44281,f63850,f63851,f44344,f63982,f63983,f44716,f64040,f64041,f44844,f64098,f64099,f44903,f64107,f64109,f64110,f64112,f64114,f64188,f64260,f64338,f64339,f44961,f64395,f64396,f45019,f64453,f64454,f45078,f64511,f64512,f45137,f64570,f63376,f64571,f45197,f63378,f64765,f64767,f64768,f64770,f64772,f64846,f46823,f66925,f66927,f66928,f66929,f66932,f67002,f67079,f67443,f67445,f67446,f67447,f67450,f67520,f48814,f67740,f67751,f67757,f67807,f67815,f67822,f67830,f67843,f67844,f67847,f67848,f67906,f67907,f67910,f67911,f67850,f68002,f68004,f68005,f68006,f68009,f68079,f68156,f68254,f68256,f68257,f879,f68259,f68261,f882,f68335,f68410,f884,f68540,f68541,f68542,f68543,f68555,f68556,f67744,f68557,f68558,f68545,f883,f70414,f70415,f70416])).
fof(f70416,plain,(
  $false | $spl22),
  inference(subsumption_resolution,[],[f70021,f704])).
fof(f704,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f337])).
fof(f337,plain,(
  ! [X0] : p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X0))))),
  inference(rectify,[],[f85])).
fof(f85,axiom,(
  ! [X6] : p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X6),s('type.integer.int',X6))))),
  file('thm.integer.INT_LE_MUL','thm.integer.INT_LE_REFL')).
fof(f70021,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) | $spl22),
  inference(backward_demodulation,[],[f69282,f50722])).
fof(f50722,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | $spl22),
  inference(forward_demodulation,[],[f742,f1015])).
fof(f742,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))))),
  inference(forward_demodulation,[],[f505,f721])).
fof(f721,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),
  inference(cnf_transformation,[],[f89])).
fof(f89,axiom,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),
  file('thm.integer.INT_LE_MUL','thm.integer.int_of_num|split|0')).
fof(f505,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))))),
  inference(cnf_transformation,[],[f412])).
fof(f412,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) & ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13])],[f363])).
fof(f363,plain,(
  ? [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) & ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))))),
  inference(flattening,[],[f362])).
fof(f362,plain,(
  ? [X0,X1] : ((p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1))))) & ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))))),
  inference(ennf_transformation,[],[f263])).
fof(f263,plain,(
  ~! [X0,X1] : ((p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1))))) => p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ((p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X6)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X7))))) => p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X6),s('type.integer.int',X7)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ((p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X6)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X7))))) => p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X6),s('type.integer.int',X7)))))))),
  file('thm.integer.INT_LE_MUL','thm.integer.INT_LE_MUL')).
fof(f69282,plain,(
  ( ! [X52] : (s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',X52)))) ) | $spl22),
  inference(forward_demodulation,[],[f69243,f721])).
fof(f69243,plain,(
  ( ! [X52] : (s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X52)))) ) | $spl22),
  inference(trivial_inequality_removal,[],[f69190])).
fof(f69190,plain,(
  ( ! [X52] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X52)))) ) | $spl22),
  inference(superposition,[],[f55016,f518])).
fof(f518,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f276])).
fof(f276,plain,(
  ! [X0] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X6)))),
  file('thm.integer.INT_LE_MUL','thm.integer.INT_MUL_LZERO')).
fof(f55016,plain,(
  ( ! [X4] : (s('type.integer.int',X4) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',X4)) ) | $spl22),
  inference(superposition,[],[f50529,f712])).
fof(f712,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) | s('type.integer.int',X1) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f498])).
fof(f498,plain,(
  ! [X0,X1] : ((s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.integer.int',X1) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))),
  inference(nnf_transformation,[],[f341])).
fof(f341,plain,(
  ! [X0,X1] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) <=> s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f88])).
fof(f88,axiom,(
  ! [X6,X7] : (s('type.integer.int',X6) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X6),s('type.integer.int',X7))) <=> s('type.integer.int',X7) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.integer.INT_LE_MUL','thm.integer.INT_ADD_RID_UNIQ')).
fof(f50529,plain,(
  ( ! [X203] : (s('type.integer.int',X203) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',X203)))) ) | $spl22),
  inference(subsumption_resolution,[],[f50528,f1015])).
fof(f50528,plain,(
  ( ! [X203] : (s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int',sK12) | s('type.integer.int',X203) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',X203)))) ) | $spl22),
  inference(forward_demodulation,[],[f50259,f721])).
fof(f50259,plain,(
  ( ! [X203] : (s('type.integer.int',X203) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',X203))) | s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl22),
  inference(superposition,[],[f710,f1015])).
fof(f710,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X1) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f497])).
fof(f497,plain,(
  ! [X0,X1] : ((s('type.integer.int',X1) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))),
  inference(nnf_transformation,[],[f340])).
fof(f340,plain,(
  ! [X0,X1] : (s('type.integer.int',X1) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) <=> s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X6,X7] : (s('type.integer.int',X7) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X6),s('type.integer.int',X7))) <=> s('type.integer.int',X6) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.integer.INT_LE_MUL','thm.integer.INT_ADD_LID_UNIQ')).
fof(f70415,plain,(
  $false | ($spl22 | $spl74)),
  inference(subsumption_resolution,[],[f70002,f667])).
fof(f667,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f308])).
fof(f308,plain,(
  ! [X0] : ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0))))),
  inference(flattening,[],[f307])).
fof(f307,plain,(
  ! [X0] : ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0))))),
  inference(rectify,[],[f216])).
fof(f216,axiom,(
  ! [X6] : ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int',X6))))),
  file('thm.integer.INT_LE_MUL','thm.integer.INT_LT_REFL')).
fof(f70002,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl22 | $spl74)),
  inference(backward_demodulation,[],[f69282,f8100])).
fof(f70414,plain,(
  $false | $spl22),
  inference(trivial_inequality_removal,[],[f70139])).
fof(f70139,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int','const.integer.int_|48|_0') | $spl22),
  inference(superposition,[],[f50726,f69282])).
fof(f50726,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | $spl22),
  inference(resolution,[],[f50722,f687])).
fof(f687,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) != s('type.integer.int',X1)) )),
  inference(cnf_transformation,[],[f489])).
fof(f489,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | s('type.integer.int',X0) = s('type.integer.int',X1)) & (s('type.integer.int',X0) != s('type.integer.int',X1) | (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))))),
  inference(flattening,[],[f488])).
fof(f488,plain,(
  ! [X0,X1] : (((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) | s('type.integer.int',X0) = s('type.integer.int',X1)) & (s('type.integer.int',X0) != s('type.integer.int',X1) | (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))))),
  inference(nnf_transformation,[],[f324])).
fof(f324,plain,(
  ! [X0,X1] : ((p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) <=> s('type.integer.int',X0) = s('type.integer.int',X1))),
  inference(rectify,[],[f99])).
fof(f99,axiom,(
  ! [X6,X7] : ((p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X7),s('type.integer.int',X6))))) <=> s('type.integer.int',X6) = s('type.integer.int',X7))),
  file('thm.integer.INT_LE_MUL','thm.integer.INT_LE_ANTISYM')).
fof(f883,plain,(
  ( ! [X17] : (~p(s(bool,X17))) ) | $spl16),
  inference(cnf_transformation,[],[f883_D])).
fof(f883_D,plain,(
  ( ! [X17] : (~p(s(bool,X17))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f68545,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | $spl212),
  inference(cnf_transformation,[],[f68545_D])).
fof(f68545_D,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f68558,plain,(
  ( ! [X16] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | ~p(s(bool,X16))) ) | $spl4),
  inference(resolution,[],[f64175,f51475])).
fof(f51475,plain,(
  ( ! [X134,X133] : (sP9(X133,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)),X134) | ~p(s(bool,X134))) ) | $spl4),
  inference(resolution,[],[f877,f610])).
fof(f610,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X2)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f454])).
fof(f454,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | (p(s(bool,X1)) & p(s(bool,X2))) | (p(s(bool,X1)) & p(s(bool,X0)))) & (((~p(s(bool,X1)) | ~p(s(bool,X2))) & (~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f453])).
fof(f453,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | (p(s(bool,X0)) & p(s(bool,X1))) | (p(s(bool,X0)) & p(s(bool,X2)))) & (((~p(s(bool,X0)) | ~p(s(bool,X1))) & (~p(s(bool,X0)) | ~p(s(bool,X2)))) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f452])).
fof(f452,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X0)) & p(s(bool,X1))) | (p(s(bool,X0)) & p(s(bool,X2))))) & (((~p(s(bool,X0)) | ~p(s(bool,X1))) & (~p(s(bool,X0)) | ~p(s(bool,X2)))) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f406])).
fof(f406,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X0)) & p(s(bool,X1))) | (p(s(bool,X0)) & p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f64175,plain,(
  ( ! [X212,X213,X214] : (~sP9(X213,X212,X214) | s(bool,X212) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f833,f614])).
fof(f614,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f456])).
fof(f456,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,X0)) | (~p(s(bool,X1)) & ~p(s(bool,X2))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | (p(s(bool,X0)) & (p(s(bool,X1)) | p(s(bool,X2))))))),
  inference(flattening,[],[f455])).
fof(f455,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X0)) | (~p(s(bool,X1)) & ~p(s(bool,X2)))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | (p(s(bool,X0)) & (p(s(bool,X1)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f407])).
fof(f407,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X0)) & (p(s(bool,X1)) | p(s(bool,X2)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f287,f406])).
fof(f287,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X0)) & (p(s(bool,X1)) | p(s(bool,X2)))) <=> ((p(s(bool,X0)) & p(s(bool,X1))) | (p(s(bool,X0)) & p(s(bool,X2)))))),
  inference(rectify,[],[f211])).
fof(f211,axiom,(
  ! [X8,X15,X14] : ((p(s(bool,X8)) & (p(s(bool,X15)) | p(s(bool,X14)))) <=> ((p(s(bool,X8)) & p(s(bool,X15))) | (p(s(bool,X8)) & p(s(bool,X14)))))),
  file('thm.integer.INT_LE_MUL','thm.bool.LEFT_AND_OVER_OR')).
fof(f833,plain,(
  ( ! [X154] : (~p(s(bool,X154)) | s(bool,X154) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f743,f647])).
fof(f647,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f377])).
fof(f377,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f376])).
fof(f376,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f298])).
fof(f298,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f84])).
fof(f84,axiom,(
  ! [X22,X23] : ((p(s(bool,X22)) => p(s(bool,X23))) => ((p(s(bool,X23)) => p(s(bool,X22))) => s(bool,X22) = s(bool,X23)))),
  file('thm.integer.INT_LE_MUL','thm.bool.IMP_ANTISYM_AX')).
fof(f743,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))))),
  inference(forward_demodulation,[],[f504,f721])).
fof(f504,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13))))),
  inference(cnf_transformation,[],[f412])).
fof(f68557,plain,(
  ( ! [X15] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | ~p(s(bool,X15))) ) | $spl4),
  inference(resolution,[],[f64175,f51478])).
fof(f51478,plain,(
  ( ! [X140,X139] : (sP9(X139,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)),X140) | ~p(s(bool,X139))) ) | $spl4),
  inference(resolution,[],[f877,f611])).
fof(f611,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f454])).
fof(f67744,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | $spl204),
  inference(cnf_transformation,[],[f67744_D])).
fof(f67744_D,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f68556,plain,(
  ( ! [X14] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | ~p(s(bool,X14))) ) | $spl22),
  inference(resolution,[],[f64175,f50022])).
fof(f50022,plain,(
  ( ! [X130,X131] : (sP9(X130,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X131) | ~p(s(bool,X131))) ) | $spl22),
  inference(backward_demodulation,[],[f1015,f957])).
fof(f957,plain,(
  ( ! [X130,X131] : (sP9(X130,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)),X131) | ~p(s(bool,X131))) )),
  inference(resolution,[],[f744,f610])).
fof(f744,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))))),
  inference(forward_demodulation,[],[f503,f721])).
fof(f503,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12))))),
  inference(cnf_transformation,[],[f412])).
fof(f68555,plain,(
  ( ! [X13] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | ~p(s(bool,X13))) ) | $spl22),
  inference(resolution,[],[f64175,f50025])).
fof(f50025,plain,(
  ( ! [X136,X137] : (sP9(X136,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X137) | ~p(s(bool,X136))) ) | $spl22),
  inference(backward_demodulation,[],[f1015,f960])).
fof(f960,plain,(
  ( ! [X136,X137] : (sP9(X136,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)),X137) | ~p(s(bool,X136))) )),
  inference(resolution,[],[f744,f611])).
fof(f68543,plain,(
  ( ! [X19] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | ~p(s(bool,X19))) ) | $spl4),
  inference(resolution,[],[f64168,f51467])).
fof(f51467,plain,(
  ( ! [X118,X117] : (sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)),X117,X118) | ~p(s(bool,X118))) ) | $spl4),
  inference(resolution,[],[f877,f600])).
fof(f600,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f449])).
fof(f449,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | (p(s(bool,X2)) & p(s(bool,X0))) | (p(s(bool,X1)) & p(s(bool,X0)))) & (((~p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f448])).
fof(f448,plain,(
  ! [X0,X2,X1] : ((~sP8(X0,X2,X1) | (p(s(bool,X1)) & p(s(bool,X0))) | (p(s(bool,X2)) & p(s(bool,X0)))) & (((~p(s(bool,X1)) | ~p(s(bool,X0))) & (~p(s(bool,X2)) | ~p(s(bool,X0)))) | sP8(X0,X2,X1)))),
  inference(flattening,[],[f447])).
fof(f447,plain,(
  ! [X0,X2,X1] : ((~sP8(X0,X2,X1) | ((p(s(bool,X1)) & p(s(bool,X0))) | (p(s(bool,X2)) & p(s(bool,X0))))) & (((~p(s(bool,X1)) | ~p(s(bool,X0))) & (~p(s(bool,X2)) | ~p(s(bool,X0)))) | sP8(X0,X2,X1)))),
  inference(nnf_transformation,[],[f404])).
fof(f404,plain,(
  ! [X0,X2,X1] : (sP8(X0,X2,X1) <=> ((p(s(bool,X1)) & p(s(bool,X0))) | (p(s(bool,X2)) & p(s(bool,X0)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f64168,plain,(
  ( ! [X191,X192,X193] : (~sP8(X191,X192,X193) | s(bool,X191) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f833,f605])).
fof(f605,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~sP8(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f451])).
fof(f451,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ~p(s(bool,X2))) | ~p(s(bool,X0)) | sP8(X0,X2,X1)) & (~sP8(X0,X2,X1) | ((p(s(bool,X1)) | p(s(bool,X2))) & p(s(bool,X0)))))),
  inference(flattening,[],[f450])).
fof(f450,plain,(
  ! [X0,X1,X2] : ((((~p(s(bool,X1)) & ~p(s(bool,X2))) | ~p(s(bool,X0))) | sP8(X0,X2,X1)) & (~sP8(X0,X2,X1) | ((p(s(bool,X1)) | p(s(bool,X2))) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f405])).
fof(f405,plain,(
  ! [X0,X1,X2] : (((p(s(bool,X1)) | p(s(bool,X2))) & p(s(bool,X0))) <=> sP8(X0,X2,X1))),
  inference(definition_folding,[],[f286,f404])).
fof(f286,plain,(
  ! [X0,X1,X2] : (((p(s(bool,X1)) | p(s(bool,X2))) & p(s(bool,X0))) <=> ((p(s(bool,X1)) & p(s(bool,X0))) | (p(s(bool,X2)) & p(s(bool,X0)))))),
  inference(rectify,[],[f190])).
fof(f190,axiom,(
  ! [X8,X15,X14] : (((p(s(bool,X15)) | p(s(bool,X14))) & p(s(bool,X8))) <=> ((p(s(bool,X15)) & p(s(bool,X8))) | (p(s(bool,X14)) & p(s(bool,X8)))))),
  file('thm.integer.INT_LE_MUL','thm.bool.RIGHT_AND_OVER_OR')).
fof(f68542,plain,(
  ( ! [X18] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | ~p(s(bool,X18))) ) | $spl4),
  inference(resolution,[],[f64168,f51469])).
fof(f51469,plain,(
  ( ! [X121,X122] : (sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)),X121,X122) | ~p(s(bool,X121))) ) | $spl4),
  inference(resolution,[],[f877,f601])).
fof(f601,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f449])).
fof(f68541,plain,(
  ( ! [X17] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | ~p(s(bool,X17))) ) | $spl22),
  inference(resolution,[],[f64168,f50018])).
fof(f50018,plain,(
  ( ! [X114,X115] : (sP8('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X114,X115) | ~p(s(bool,X115))) ) | $spl22),
  inference(backward_demodulation,[],[f1015,f949])).
fof(f949,plain,(
  ( ! [X114,X115] : (sP8('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)),X114,X115) | ~p(s(bool,X115))) )),
  inference(resolution,[],[f744,f600])).
fof(f68540,plain,(
  ( ! [X16] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | ~p(s(bool,X16))) ) | $spl22),
  inference(resolution,[],[f64168,f50020])).
fof(f50020,plain,(
  ( ! [X118,X119] : (sP8('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X118,X119) | ~p(s(bool,X118))) ) | $spl22),
  inference(backward_demodulation,[],[f1015,f951])).
fof(f951,plain,(
  ( ! [X118,X119] : (sP8('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)),X118,X119) | ~p(s(bool,X118))) )),
  inference(resolution,[],[f744,f601])).
fof(f884,plain,(
  ( ! [X16] : (p(s(bool,X16))) ) | $spl18),
  inference(cnf_transformation,[],[f884_D])).
fof(f884_D,plain,(
  ( ! [X16] : (p(s(bool,X16))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f68410,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl210),
  inference(cnf_transformation,[],[f68410_D])).
fof(f68410_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f68335,plain,(
  ( ! [X243,X241,X242] : (s(bool,X241) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | p(s(bool,X242)) | ~p(s(bool,X243)) | s(bool,X241) != s(bool,X243)) ) | $spl4),
  inference(resolution,[],[f51491,f735])).
fof(f735,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f523])).
fof(f523,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f414])).
fof(f414,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f413])).
fof(f413,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f367])).
fof(f367,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f366])).
fof(f366,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f277])).
fof(f277,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f251])).
fof(f251,axiom,(
  ! [X6,X63,X7,X64] : ((s(bool,X6) = s(bool,X63) & (p(s(bool,X63)) => s(bool,X7) = s(bool,X64))) => ((p(s(bool,X6)) => p(s(bool,X7))) <=> (p(s(bool,X63)) => p(s(bool,X64)))))),
  file('thm.integer.INT_LE_MUL','thm.bool.IMP_CONG')).
fof(f51491,plain,(
  ( ! [X157] : (~p(s(bool,X157)) | s(bool,X157) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl4),
  inference(resolution,[],[f877,f647])).
fof(f882,plain,(
  ( ! [X12,X11] : (s(bool,X11) != s(bool,X12) | ~p(s(bool,X11))) ) | $spl14),
  inference(cnf_transformation,[],[f882_D])).
fof(f882_D,plain,(
  ( ! [X12,X11] : (s(bool,X11) != s(bool,X12) | ~p(s(bool,X11))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f68261,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl4),
  inference(resolution,[],[f51491,f526])).
fof(f526,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f414])).
fof(f68259,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) ) | $spl4),
  inference(resolution,[],[f51491,f525])).
fof(f525,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f414])).
fof(f879,plain,(
  ( ! [X6,X4] : (s(bool,X4) != s(bool,X6) | p(s(bool,X4))) ) | $spl8),
  inference(cnf_transformation,[],[f879_D])).
fof(f879_D,plain,(
  ( ! [X6,X4] : (s(bool,X4) != s(bool,X6) | p(s(bool,X4))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f68257,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl4),
  inference(resolution,[],[f51491,f524])).
fof(f524,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f414])).
fof(f68256,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl4),
  inference(resolution,[],[f51491,f522])).
fof(f522,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f414])).
fof(f68254,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl4),
  inference(resolution,[],[f51491,f520])).
fof(f520,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f414])).
fof(f68156,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl208),
  inference(cnf_transformation,[],[f68156_D])).
fof(f68156_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f68079,plain,(
  ( ! [X231,X229,X230] : (s(bool,X229) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X230)) | p(s(bool,X231)) | s(bool,X229) != s(bool,X231)) ) | $spl1),
  inference(resolution,[],[f51356,f735])).
fof(f51356,plain,(
  ( ! [X159] : (p(s(bool,X159)) | s(bool,X159) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl1),
  inference(resolution,[],[f873,f644])).
fof(f644,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f377])).
fof(f68009,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl1),
  inference(resolution,[],[f51356,f526])).
fof(f68006,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X16) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl1),
  inference(resolution,[],[f51356,f525])).
fof(f68005,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl1),
  inference(resolution,[],[f51356,f524])).
fof(f68004,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl1),
  inference(resolution,[],[f51356,f522])).
fof(f68002,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl1),
  inference(resolution,[],[f51356,f520])).
fof(f67850,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | $spl206),
  inference(cnf_transformation,[],[f67850_D])).
fof(f67850_D,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f67911,plain,(
  ( ! [X16] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | ~p(s(bool,X16))) ) | ($spl4 | $spl22)),
  inference(resolution,[],[f50070,f51475])).
fof(f50070,plain,(
  ( ! [X212,X213,X214] : (~sP9(X213,X212,X214) | s(bool,X212) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl22),
  inference(backward_demodulation,[],[f1015,f44551])).
fof(f44551,plain,(
  ( ! [X212,X213,X214] : (~sP9(X213,X212,X214) | s(bool,X212) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f973,f614])).
fof(f973,plain,(
  ( ! [X154] : (~p(s(bool,X154)) | s(bool,X154) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f744,f647])).
fof(f67910,plain,(
  ( ! [X15] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | ~p(s(bool,X15))) ) | ($spl4 | $spl22)),
  inference(resolution,[],[f50070,f51478])).
fof(f67907,plain,(
  ( ! [X12] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | ~p(s(bool,X12))) ) | $spl22),
  inference(resolution,[],[f50070,f817])).
fof(f817,plain,(
  ( ! [X130,X131] : (sP9(X130,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)),X131) | ~p(s(bool,X131))) )),
  inference(resolution,[],[f743,f610])).
fof(f67906,plain,(
  ( ! [X11] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | ~p(s(bool,X11))) ) | $spl22),
  inference(resolution,[],[f50070,f820])).
fof(f820,plain,(
  ( ! [X136,X137] : (sP9(X136,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)),X137) | ~p(s(bool,X136))) )),
  inference(resolution,[],[f743,f611])).
fof(f67848,plain,(
  ( ! [X19] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | ~p(s(bool,X19))) ) | ($spl4 | $spl22)),
  inference(resolution,[],[f50069,f51467])).
fof(f50069,plain,(
  ( ! [X191,X192,X193] : (~sP8(X191,X192,X193) | s(bool,X191) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl22),
  inference(backward_demodulation,[],[f1015,f44544])).
fof(f44544,plain,(
  ( ! [X191,X192,X193] : (~sP8(X191,X192,X193) | s(bool,X191) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f973,f605])).
fof(f67847,plain,(
  ( ! [X18] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | ~p(s(bool,X18))) ) | ($spl4 | $spl22)),
  inference(resolution,[],[f50069,f51469])).
fof(f67844,plain,(
  ( ! [X15] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | ~p(s(bool,X15))) ) | $spl22),
  inference(resolution,[],[f50069,f809])).
fof(f809,plain,(
  ( ! [X114,X115] : (sP8('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)),X114,X115) | ~p(s(bool,X115))) )),
  inference(resolution,[],[f743,f600])).
fof(f67843,plain,(
  ( ! [X14] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | ~p(s(bool,X14))) ) | $spl22),
  inference(resolution,[],[f50069,f811])).
fof(f811,plain,(
  ( ! [X118,X119] : (sP8('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)),X118,X119) | ~p(s(bool,X118))) )),
  inference(resolution,[],[f743,f601])).
fof(f67830,plain,(
  ( ! [X14,X15] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | p(s(bool,X14)) | ~p(s(bool,X15))) ) | $spl22),
  inference(resolution,[],[f50068,f801])).
fof(f801,plain,(
  ( ! [X99,X98] : (~sP5('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)),X98,X99) | p(s(bool,X99)) | ~p(s(bool,X98))) )),
  inference(resolution,[],[f743,f572])).
fof(f572,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f439])).
fof(f439,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f438])).
fof(f438,plain,(
  ! [X1,X0,X2] : ((~sP5(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP5(X1,X0,X2)))),
  inference(flattening,[],[f437])).
fof(f437,plain,(
  ! [X1,X0,X2] : ((~sP5(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP5(X1,X0,X2)))),
  inference(nnf_transformation,[],[f400])).
fof(f400,plain,(
  ! [X1,X0,X2] : (sP5(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f50068,plain,(
  ( ! [X151,X149,X150] : (sP5(X149,X150,X151) | s(bool,X149) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl22),
  inference(backward_demodulation,[],[f1015,f44530])).
fof(f44530,plain,(
  ( ! [X151,X149,X150] : (sP5(X149,X150,X151) | s(bool,X149) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f973,f575])).
fof(f575,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f439])).
fof(f67822,plain,(
  ( ! [X12,X13] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | p(s(bool,X12)) | ~p(s(bool,X13))) ) | $spl22),
  inference(resolution,[],[f50067,f802])).
fof(f802,plain,(
  ( ! [X101,X100] : (~sP5(X100,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)),X101) | p(s(bool,X101)) | ~p(s(bool,X100))) )),
  inference(resolution,[],[f743,f572])).
fof(f50067,plain,(
  ( ! [X146,X147,X148] : (sP5(X147,X146,X148) | s(bool,X146) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl22),
  inference(backward_demodulation,[],[f1015,f44529])).
fof(f44529,plain,(
  ( ! [X146,X147,X148] : (sP5(X147,X146,X148) | s(bool,X146) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f973,f574])).
fof(f574,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f439])).
fof(f67815,plain,(
  ( ! [X14,X15] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | p(s(bool,X14)) | ~p(s(bool,X15))) ) | $spl22),
  inference(resolution,[],[f50066,f798])).
fof(f798,plain,(
  ( ! [X92,X93] : (~sP6('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)),X92,X93) | p(s(bool,X93)) | ~p(s(bool,X92))) )),
  inference(resolution,[],[f743,f568])).
fof(f568,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f436])).
fof(f436,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(rectify,[],[f435])).
fof(f435,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(flattening,[],[f434])).
fof(f434,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(nnf_transformation,[],[f401])).
fof(f401,plain,(
  ! [X2,X0,X1] : (sP6(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f50066,plain,(
  ( ! [X142,X140,X141] : (sP6(X140,X141,X142) | s(bool,X140) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl22),
  inference(backward_demodulation,[],[f1015,f44527])).
fof(f44527,plain,(
  ( ! [X142,X140,X141] : (sP6(X140,X141,X142) | s(bool,X140) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f973,f571])).
fof(f571,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f436])).
fof(f67807,plain,(
  ( ! [X12,X13] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | p(s(bool,X12)) | ~p(s(bool,X13))) ) | $spl22),
  inference(resolution,[],[f50065,f799])).
fof(f799,plain,(
  ( ! [X94,X95] : (~sP6(X94,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)),X95) | p(s(bool,X95)) | ~p(s(bool,X94))) )),
  inference(resolution,[],[f743,f568])).
fof(f50065,plain,(
  ( ! [X138,X139,X137] : (sP6(X138,X137,X139) | s(bool,X137) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl22),
  inference(backward_demodulation,[],[f1015,f44526])).
fof(f44526,plain,(
  ( ! [X138,X139,X137] : (sP6(X138,X137,X139) | s(bool,X137) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f973,f570])).
fof(f570,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f436])).
fof(f67757,plain,(
  ( ! [X15,X16] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | p(s(bool,X15)) | p(s(bool,X16))) ) | $spl22),
  inference(resolution,[],[f50064,f785])).
fof(f785,plain,(
  ( ! [X66,X67] : (~sP2('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)),X66,X67) | p(s(bool,X67)) | p(s(bool,X66))) )),
  inference(resolution,[],[f743,f551])).
fof(f551,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f428])).
fof(f428,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f427])).
fof(f427,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f426])).
fof(f426,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f396])).
fof(f396,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f50064,plain,(
  ( ! [X103,X101,X102] : (sP2(X101,X102,X103) | s(bool,X101) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl22),
  inference(backward_demodulation,[],[f1015,f44514])).
fof(f44514,plain,(
  ( ! [X103,X101,X102] : (sP2(X101,X102,X103) | s(bool,X101) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f973,f554])).
fof(f554,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f428])).
fof(f67751,plain,(
  ( ! [X17,X18] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | ~p(s(bool,X17)) | p(s(bool,X18))) ) | $spl22),
  inference(resolution,[],[f50063,f771])).
fof(f771,plain,(
  ( ! [X39,X38] : (~sP0('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)),X38,X39) | ~p(s(bool,X39)) | p(s(bool,X38))) )),
  inference(resolution,[],[f743,f534])).
fof(f534,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f420])).
fof(f420,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(rectify,[],[f419])).
fof(f419,plain,(
  ! [X2,X0,X1] : ((~sP0(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP0(X2,X0,X1)))),
  inference(flattening,[],[f418])).
fof(f418,plain,(
  ! [X2,X0,X1] : ((~sP0(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP0(X2,X0,X1)))),
  inference(nnf_transformation,[],[f393])).
fof(f393,plain,(
  ! [X2,X0,X1] : (sP0(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f50063,plain,(
  ( ! [X54,X55,X53] : (sP0(X53,X54,X55) | s(bool,X53) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl22),
  inference(backward_demodulation,[],[f1015,f44498])).
fof(f44498,plain,(
  ( ! [X54,X55,X53] : (sP0(X53,X54,X55) | s(bool,X53) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f973,f537])).
fof(f537,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f420])).
fof(f67740,plain,(
  ( ! [X15,X16] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | p(s(bool,X15)) | ~p(s(bool,X16))) ) | $spl22),
  inference(resolution,[],[f50062,f772])).
fof(f772,plain,(
  ( ! [X41,X40] : (~sP0(X40,X41,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | p(s(bool,X41)) | ~p(s(bool,X40))) )),
  inference(resolution,[],[f743,f534])).
fof(f50062,plain,(
  ( ! [X52,X50,X51] : (sP0(X51,X52,X50) | s(bool,X50) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl22),
  inference(backward_demodulation,[],[f1015,f44497])).
fof(f44497,plain,(
  ( ! [X52,X50,X51] : (sP0(X51,X52,X50) | s(bool,X50) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f973,f535])).
fof(f535,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f420])).
fof(f48814,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl190),
  inference(cnf_transformation,[],[f48814_D])).
fof(f48814_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f67520,plain,(
  ( ! [X231,X229,X230] : (s(bool,X229) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X230)) | p(s(bool,X231)) | s(bool,X229) != s(bool,X231)) ) | $spl22),
  inference(resolution,[],[f50058,f735])).
fof(f50058,plain,(
  ( ! [X161] : (p(s(bool,X161)) | s(bool,X161) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl22),
  inference(backward_demodulation,[],[f1015,f4212])).
fof(f4212,plain,(
  ( ! [X161] : (p(s(bool,X161)) | s(bool,X161) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) )),
  inference(resolution,[],[f889,f644])).
fof(f889,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))))),
  inference(resolution,[],[f744,f680])).
fof(f680,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f388])).
fof(f388,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(ennf_transformation,[],[f318])).
fof(f318,plain,(
  ! [X0,X1] : ~(p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f97])).
fof(f97,axiom,(
  ! [X6,X7] : ~(p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int',X6)))))),
  file('thm.integer.INT_LE_MUL','thm.integer.INT_LTE_ANTSYM')).
fof(f67450,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl22),
  inference(resolution,[],[f50058,f526])).
fof(f67447,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X16) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl22),
  inference(resolution,[],[f50058,f525])).
fof(f67446,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl22),
  inference(resolution,[],[f50058,f524])).
fof(f67445,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl22),
  inference(resolution,[],[f50058,f522])).
fof(f67443,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl22),
  inference(resolution,[],[f50058,f520])).
fof(f67079,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl202),
  inference(cnf_transformation,[],[f67079_D])).
fof(f67079_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f67002,plain,(
  ( ! [X231,X229,X230] : (s(bool,X229) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X230)) | p(s(bool,X231)) | s(bool,X229) != s(bool,X231)) )),
  inference(resolution,[],[f2978,f735])).
fof(f2978,plain,(
  ( ! [X161] : (p(s(bool,X161)) | s(bool,X161) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) )),
  inference(resolution,[],[f749,f644])).
fof(f749,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))))),
  inference(resolution,[],[f743,f680])).
fof(f66932,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) )),
  inference(resolution,[],[f2978,f526])).
fof(f66929,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X16) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) )),
  inference(resolution,[],[f2978,f525])).
fof(f66928,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f2978,f524])).
fof(f66927,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) )),
  inference(resolution,[],[f2978,f522])).
fof(f66925,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f2978,f520])).
fof(f46823,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl176),
  inference(cnf_transformation,[],[f46823_D])).
fof(f46823_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f64846,plain,(
  ( ! [X243,X241,X242] : (s(bool,X241) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X242)) | ~p(s(bool,X243)) | s(bool,X241) != s(bool,X243)) ) | $spl22),
  inference(resolution,[],[f50030,f735])).
fof(f50030,plain,(
  ( ! [X154] : (~p(s(bool,X154)) | s(bool,X154) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl22),
  inference(backward_demodulation,[],[f1015,f973])).
fof(f64772,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl22),
  inference(resolution,[],[f50030,f526])).
fof(f64770,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) ) | $spl22),
  inference(resolution,[],[f50030,f525])).
fof(f64768,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl22),
  inference(resolution,[],[f50030,f524])).
fof(f64767,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl22),
  inference(resolution,[],[f50030,f522])).
fof(f64765,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl22),
  inference(resolution,[],[f50030,f520])).
fof(f63378,plain,(
  ( ! [X35] : (s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl198),
  inference(cnf_transformation,[],[f63378_D])).
fof(f63378_D,plain,(
  ( ! [X35] : (s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f45197,plain,(
  ( ! [X30,X31] : (sP4('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X30,X31)) ) | $spl174),
  inference(cnf_transformation,[],[f45197_D])).
fof(f45197_D,plain,(
  ( ! [X30,X31] : (sP4('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X30,X31)) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f64571,plain,(
  ( ! [X35,X33,X34] : (sP4('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X33,X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl74),
  inference(forward_demodulation,[],[f64547,f721])).
fof(f64547,plain,(
  ( ! [X35,X33,X34] : (sP4('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')),X33,X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl74),
  inference(superposition,[],[f18991,f711])).
fof(f711,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) )),
  inference(cnf_transformation,[],[f498])).
fof(f18991,plain,(
  ( ! [X109,X110] : (sP4('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')),X109,X110)) ) | $spl74),
  inference(resolution,[],[f8100,f579])).
fof(f579,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f441])).
fof(f441,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(flattening,[],[f440])).
fof(f440,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(nnf_transformation,[],[f399])).
fof(f399,plain,(
  ! [X0,X1,X2] : (sP4(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f63376,plain,(
  ( ! [X32] : (s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int',X32)))) ) | $spl196),
  inference(cnf_transformation,[],[f63376_D])).
fof(f63376_D,plain,(
  ( ! [X32] : (s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int',X32)))) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f64570,plain,(
  ( ! [X30,X31,X32] : (sP4('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X30,X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int',X32)))) ) | $spl74),
  inference(forward_demodulation,[],[f64546,f721])).
fof(f64546,plain,(
  ( ! [X30,X31,X32] : (sP4('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')),X30,X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int',X32)))) ) | $spl74),
  inference(superposition,[],[f18991,f709])).
fof(f709,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X1) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) )),
  inference(cnf_transformation,[],[f497])).
fof(f45137,plain,(
  ( ! [X30,X31] : (sP4(X30,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X31)) ) | $spl172),
  inference(cnf_transformation,[],[f45137_D])).
fof(f45137_D,plain,(
  ( ! [X30,X31] : (sP4(X30,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X31)) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f64512,plain,(
  ( ! [X35,X33,X34] : (sP4(X33,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl74),
  inference(forward_demodulation,[],[f64489,f721])).
fof(f64489,plain,(
  ( ! [X35,X33,X34] : (sP4(X33,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')),X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl74),
  inference(superposition,[],[f18990,f711])).
fof(f18990,plain,(
  ( ! [X107,X108] : (sP4(X107,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')),X108)) ) | $spl74),
  inference(resolution,[],[f8100,f578])).
fof(f578,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f441])).
fof(f64511,plain,(
  ( ! [X30,X31,X32] : (sP4(X30,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int',X32)))) ) | $spl74),
  inference(forward_demodulation,[],[f64488,f721])).
fof(f64488,plain,(
  ( ! [X30,X31,X32] : (sP4(X30,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')),X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int',X32)))) ) | $spl74),
  inference(superposition,[],[f18990,f709])).
fof(f45078,plain,(
  ( ! [X30,X31] : (sP4(X30,X31,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl170),
  inference(cnf_transformation,[],[f45078_D])).
fof(f45078_D,plain,(
  ( ! [X30,X31] : (sP4(X30,X31,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f64454,plain,(
  ( ! [X35,X33,X34] : (sP4(X33,X34,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl74),
  inference(forward_demodulation,[],[f64431,f721])).
fof(f64431,plain,(
  ( ! [X35,X33,X34] : (sP4(X33,X34,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0'))) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl74),
  inference(superposition,[],[f18989,f711])).
fof(f18989,plain,(
  ( ! [X105,X106] : (sP4(X105,X106,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl74),
  inference(resolution,[],[f8100,f577])).
fof(f577,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f441])).
fof(f64453,plain,(
  ( ! [X30,X31,X32] : (sP4(X30,X31,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int',X32)))) ) | $spl74),
  inference(forward_demodulation,[],[f64430,f721])).
fof(f64430,plain,(
  ( ! [X30,X31,X32] : (sP4(X30,X31,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0'))) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int',X32)))) ) | $spl74),
  inference(superposition,[],[f18989,f709])).
fof(f45019,plain,(
  ( ! [X30,X31] : (sP5(X30,X31,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl168),
  inference(cnf_transformation,[],[f45019_D])).
fof(f45019_D,plain,(
  ( ! [X30,X31] : (sP5(X30,X31,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f64396,plain,(
  ( ! [X35,X33,X34] : (sP5(X33,X34,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl74),
  inference(forward_demodulation,[],[f64374,f721])).
fof(f64374,plain,(
  ( ! [X35,X33,X34] : (sP5(X33,X34,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0'))) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl74),
  inference(superposition,[],[f18988,f711])).
fof(f18988,plain,(
  ( ! [X103,X104] : (sP5(X103,X104,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl74),
  inference(resolution,[],[f8100,f573])).
fof(f573,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f439])).
fof(f64395,plain,(
  ( ! [X30,X31,X32] : (sP5(X30,X31,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int',X32)))) ) | $spl74),
  inference(forward_demodulation,[],[f64373,f721])).
fof(f64373,plain,(
  ( ! [X30,X31,X32] : (sP5(X30,X31,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0'))) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int',X32)))) ) | $spl74),
  inference(superposition,[],[f18988,f709])).
fof(f44961,plain,(
  ( ! [X30,X31] : (sP6(X30,X31,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl166),
  inference(cnf_transformation,[],[f44961_D])).
fof(f44961_D,plain,(
  ( ! [X30,X31] : (sP6(X30,X31,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f64339,plain,(
  ( ! [X35,X33,X34] : (sP6(X33,X34,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl74),
  inference(forward_demodulation,[],[f64317,f721])).
fof(f64317,plain,(
  ( ! [X35,X33,X34] : (sP6(X33,X34,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0'))) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl74),
  inference(superposition,[],[f18985,f711])).
fof(f18985,plain,(
  ( ! [X97,X98] : (sP6(X97,X98,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl74),
  inference(resolution,[],[f8100,f569])).
fof(f569,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f436])).
fof(f64338,plain,(
  ( ! [X30,X31,X32] : (sP6(X30,X31,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int',X32)))) ) | $spl74),
  inference(forward_demodulation,[],[f64316,f721])).
fof(f64316,plain,(
  ( ! [X30,X31,X32] : (sP6(X30,X31,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0'))) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int',X32)))) ) | $spl74),
  inference(superposition,[],[f18985,f709])).
fof(f64260,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl200),
  inference(cnf_transformation,[],[f64260_D])).
fof(f64260_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f64188,plain,(
  ( ! [X243,X241,X242] : (s(bool,X241) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | p(s(bool,X242)) | ~p(s(bool,X243)) | s(bool,X241) != s(bool,X243)) )),
  inference(resolution,[],[f833,f735])).
fof(f64114,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) )),
  inference(resolution,[],[f833,f526])).
fof(f64112,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f833,f525])).
fof(f64110,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f833,f524])).
fof(f64109,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f833,f522])).
fof(f64107,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f833,f520])).
fof(f44903,plain,(
  ( ! [X30,X31] : (sP2(X30,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X31)) ) | $spl164),
  inference(cnf_transformation,[],[f44903_D])).
fof(f44903_D,plain,(
  ( ! [X30,X31] : (sP2(X30,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X31)) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f64099,plain,(
  ( ! [X35,X33,X34] : (sP2(X33,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl74),
  inference(forward_demodulation,[],[f64077,f721])).
fof(f64077,plain,(
  ( ! [X35,X33,X34] : (sP2(X33,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')),X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl74),
  inference(superposition,[],[f18972,f711])).
fof(f18972,plain,(
  ( ! [X72,X71] : (sP2(X71,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')),X72)) ) | $spl74),
  inference(resolution,[],[f8100,f553])).
fof(f553,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f428])).
fof(f64098,plain,(
  ( ! [X30,X31,X32] : (sP2(X30,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int',X32)))) ) | $spl74),
  inference(forward_demodulation,[],[f64076,f721])).
fof(f64076,plain,(
  ( ! [X30,X31,X32] : (sP2(X30,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')),X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int',X32)))) ) | $spl74),
  inference(superposition,[],[f18972,f709])).
fof(f44844,plain,(
  ( ! [X30,X31] : (sP2(X30,X31,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl162),
  inference(cnf_transformation,[],[f44844_D])).
fof(f44844_D,plain,(
  ( ! [X30,X31] : (sP2(X30,X31,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f64041,plain,(
  ( ! [X35,X33,X34] : (sP2(X33,X34,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl74),
  inference(forward_demodulation,[],[f64019,f721])).
fof(f64019,plain,(
  ( ! [X35,X33,X34] : (sP2(X33,X34,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0'))) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl74),
  inference(superposition,[],[f18971,f711])).
fof(f18971,plain,(
  ( ! [X70,X69] : (sP2(X69,X70,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl74),
  inference(resolution,[],[f8100,f552])).
fof(f552,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f428])).
fof(f64040,plain,(
  ( ! [X30,X31,X32] : (sP2(X30,X31,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int',X32)))) ) | $spl74),
  inference(forward_demodulation,[],[f64018,f721])).
fof(f64018,plain,(
  ( ! [X30,X31,X32] : (sP2(X30,X31,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0'))) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int',X32)))) ) | $spl74),
  inference(superposition,[],[f18971,f709])).
fof(f44716,plain,(
  ( ! [X30,X31] : (sP0(X30,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X31)) ) | $spl160),
  inference(cnf_transformation,[],[f44716_D])).
fof(f44716_D,plain,(
  ( ! [X30,X31] : (sP0(X30,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X31)) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f63983,plain,(
  ( ! [X35,X33,X34] : (sP0(X33,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl74),
  inference(forward_demodulation,[],[f63961,f721])).
fof(f63961,plain,(
  ( ! [X35,X33,X34] : (sP0(X33,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')),X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl74),
  inference(superposition,[],[f18958,f711])).
fof(f18958,plain,(
  ( ! [X43,X44] : (sP0(X43,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')),X44)) ) | $spl74),
  inference(resolution,[],[f8100,f536])).
fof(f536,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f420])).
fof(f63982,plain,(
  ( ! [X30,X31,X32] : (sP0(X30,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int',X32)))) ) | $spl74),
  inference(forward_demodulation,[],[f63960,f721])).
fof(f63960,plain,(
  ( ! [X30,X31,X32] : (sP0(X30,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')),X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int',X32)))) ) | $spl74),
  inference(superposition,[],[f18958,f709])).
fof(f44344,plain,(
  ( ! [X30,X31] : (~sP9(X30,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X31)) ) | $spl156),
  inference(cnf_transformation,[],[f44344_D])).
fof(f44344_D,plain,(
  ( ! [X30,X31] : (~sP9(X30,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X31)) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f63851,plain,(
  ( ! [X35,X33,X34] : (~sP9(X33,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl22),
  inference(forward_demodulation,[],[f63838,f721])).
fof(f63838,plain,(
  ( ! [X35,X33,X34] : (~sP9(X33,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl22),
  inference(superposition,[],[f50805,f711])).
fof(f50805,plain,(
  ( ! [X146,X145] : (~sP9(X145,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))),X146)) ) | $spl22),
  inference(resolution,[],[f50722,f614])).
fof(f63850,plain,(
  ( ! [X30,X31,X32] : (~sP9(X30,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int',X32)))) ) | $spl22),
  inference(forward_demodulation,[],[f63837,f721])).
fof(f63837,plain,(
  ( ! [X30,X31,X32] : (~sP9(X30,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int',X32)))) ) | $spl22),
  inference(superposition,[],[f50805,f709])).
fof(f44281,plain,(
  ( ! [X30,X31] : (~sP8('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X30,X31)) ) | $spl154),
  inference(cnf_transformation,[],[f44281_D])).
fof(f44281_D,plain,(
  ( ! [X30,X31] : (~sP8('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X30,X31)) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f63789,plain,(
  ( ! [X35,X33,X34] : (~sP8('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X33,X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl22),
  inference(forward_demodulation,[],[f63776,f721])).
fof(f63776,plain,(
  ( ! [X35,X33,X34] : (~sP8('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X33,X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl22),
  inference(superposition,[],[f50798,f711])).
fof(f50798,plain,(
  ( ! [X132,X131] : (~sP8('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))),X131,X132)) ) | $spl22),
  inference(resolution,[],[f50722,f605])).
fof(f63788,plain,(
  ( ! [X30,X31,X32] : (~sP8('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X30,X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int',X32)))) ) | $spl22),
  inference(forward_demodulation,[],[f63775,f721])).
fof(f63775,plain,(
  ( ! [X30,X31,X32] : (~sP8('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X30,X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int',X32)))) ) | $spl22),
  inference(superposition,[],[f50798,f709])).
fof(f44218,plain,(
  ( ! [X30,X31] : (sP5('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X30,X31)) ) | $spl152),
  inference(cnf_transformation,[],[f44218_D])).
fof(f44218_D,plain,(
  ( ! [X30,X31] : (sP5('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X30,X31)) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f63727,plain,(
  ( ! [X35,X33,X34] : (sP5('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X33,X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl22),
  inference(forward_demodulation,[],[f63712,f721])).
fof(f63712,plain,(
  ( ! [X35,X33,X34] : (sP5('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X33,X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl22),
  inference(superposition,[],[f50784,f711])).
fof(f50784,plain,(
  ( ! [X103,X104] : (sP5('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))),X103,X104)) ) | $spl22),
  inference(resolution,[],[f50722,f575])).
fof(f63726,plain,(
  ( ! [X30,X31,X32] : (sP5('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X30,X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int',X32)))) ) | $spl22),
  inference(forward_demodulation,[],[f63711,f721])).
fof(f63711,plain,(
  ( ! [X30,X31,X32] : (sP5('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X30,X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int',X32)))) ) | $spl22),
  inference(superposition,[],[f50784,f709])).
fof(f44158,plain,(
  ( ! [X30,X31] : (sP5(X30,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X31)) ) | $spl150),
  inference(cnf_transformation,[],[f44158_D])).
fof(f44158_D,plain,(
  ( ! [X30,X31] : (sP5(X30,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X31)) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f63669,plain,(
  ( ! [X35,X33,X34] : (sP5(X33,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl22),
  inference(forward_demodulation,[],[f63654,f721])).
fof(f63654,plain,(
  ( ! [X35,X33,X34] : (sP5(X33,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl22),
  inference(superposition,[],[f50783,f711])).
fof(f50783,plain,(
  ( ! [X101,X102] : (sP5(X101,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))),X102)) ) | $spl22),
  inference(resolution,[],[f50722,f574])).
fof(f63668,plain,(
  ( ! [X30,X31,X32] : (sP5(X30,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int',X32)))) ) | $spl22),
  inference(forward_demodulation,[],[f63653,f721])).
fof(f63653,plain,(
  ( ! [X30,X31,X32] : (sP5(X30,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int',X32)))) ) | $spl22),
  inference(superposition,[],[f50783,f709])).
fof(f44098,plain,(
  ( ! [X30,X31] : (sP6('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X30,X31)) ) | $spl148),
  inference(cnf_transformation,[],[f44098_D])).
fof(f44098_D,plain,(
  ( ! [X30,X31] : (sP6('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X30,X31)) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f63611,plain,(
  ( ! [X35,X33,X34] : (sP6('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X33,X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl22),
  inference(forward_demodulation,[],[f63596,f721])).
fof(f63596,plain,(
  ( ! [X35,X33,X34] : (sP6('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X33,X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl22),
  inference(superposition,[],[f50781,f711])).
fof(f50781,plain,(
  ( ! [X97,X98] : (sP6('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))),X97,X98)) ) | $spl22),
  inference(resolution,[],[f50722,f571])).
fof(f63610,plain,(
  ( ! [X30,X31,X32] : (sP6('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X30,X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int',X32)))) ) | $spl22),
  inference(forward_demodulation,[],[f63595,f721])).
fof(f63595,plain,(
  ( ! [X30,X31,X32] : (sP6('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X30,X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int',X32)))) ) | $spl22),
  inference(superposition,[],[f50781,f709])).
fof(f44038,plain,(
  ( ! [X30,X31] : (sP6(X30,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X31)) ) | $spl146),
  inference(cnf_transformation,[],[f44038_D])).
fof(f44038_D,plain,(
  ( ! [X30,X31] : (sP6(X30,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X31)) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f63553,plain,(
  ( ! [X35,X33,X34] : (sP6(X33,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl22),
  inference(forward_demodulation,[],[f63538,f721])).
fof(f63538,plain,(
  ( ! [X35,X33,X34] : (sP6(X33,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl22),
  inference(superposition,[],[f50780,f711])).
fof(f50780,plain,(
  ( ! [X95,X96] : (sP6(X95,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))),X96)) ) | $spl22),
  inference(resolution,[],[f50722,f570])).
fof(f63552,plain,(
  ( ! [X30,X31,X32] : (sP6(X30,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int',X32)))) ) | $spl22),
  inference(forward_demodulation,[],[f63537,f721])).
fof(f63537,plain,(
  ( ! [X30,X31,X32] : (sP6(X30,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int',X32)))) ) | $spl22),
  inference(superposition,[],[f50780,f709])).
fof(f43979,plain,(
  ( ! [X30,X31] : (sP2('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X30,X31)) ) | $spl144),
  inference(cnf_transformation,[],[f43979_D])).
fof(f43979_D,plain,(
  ( ! [X30,X31] : (sP2('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X30,X31)) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f63495,plain,(
  ( ! [X35,X33,X34] : (sP2('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X33,X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl22),
  inference(forward_demodulation,[],[f63480,f721])).
fof(f63480,plain,(
  ( ! [X35,X33,X34] : (sP2('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X33,X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl22),
  inference(superposition,[],[f50768,f711])).
fof(f50768,plain,(
  ( ! [X72,X71] : (sP2('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))),X71,X72)) ) | $spl22),
  inference(resolution,[],[f50722,f554])).
fof(f63494,plain,(
  ( ! [X30,X31,X32] : (sP2('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X30,X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int',X32)))) ) | $spl22),
  inference(forward_demodulation,[],[f63479,f721])).
fof(f63479,plain,(
  ( ! [X30,X31,X32] : (sP2('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X30,X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int',X32)))) ) | $spl22),
  inference(superposition,[],[f50768,f709])).
fof(f43919,plain,(
  ( ! [X30,X31] : (sP0('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X30,X31)) ) | $spl142),
  inference(cnf_transformation,[],[f43919_D])).
fof(f43919_D,plain,(
  ( ! [X30,X31] : (sP0('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X30,X31)) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f63436,plain,(
  ( ! [X35,X33,X34] : (sP0('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X33,X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl22),
  inference(forward_demodulation,[],[f63422,f721])).
fof(f63422,plain,(
  ( ! [X35,X33,X34] : (sP0('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X33,X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl22),
  inference(superposition,[],[f50752,f711])).
fof(f50752,plain,(
  ( ! [X39,X40] : (sP0('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))),X39,X40)) ) | $spl22),
  inference(resolution,[],[f50722,f537])).
fof(f63435,plain,(
  ( ! [X30,X31,X32] : (sP0('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X30,X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int',X32)))) ) | $spl22),
  inference(forward_demodulation,[],[f63421,f721])).
fof(f63421,plain,(
  ( ! [X30,X31,X32] : (sP0('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X30,X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int',X32)))) ) | $spl22),
  inference(superposition,[],[f50752,f709])).
fof(f43834,plain,(
  ( ! [X30,X31] : (sP0(X30,X31,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl138),
  inference(cnf_transformation,[],[f43834_D])).
fof(f43834_D,plain,(
  ( ! [X30,X31] : (sP0(X30,X31,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f63377,plain,(
  ( ! [X35,X33,X34] : (sP0(X33,X34,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl22),
  inference(forward_demodulation,[],[f63362,f721])).
fof(f63362,plain,(
  ( ! [X35,X33,X34] : (sP0(X33,X34,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl22),
  inference(superposition,[],[f50751,f711])).
fof(f50751,plain,(
  ( ! [X37,X38] : (sP0(X37,X38,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl22),
  inference(resolution,[],[f50722,f535])).
fof(f63375,plain,(
  ( ! [X30,X31,X32] : (sP0(X30,X31,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int',X32)))) ) | $spl22),
  inference(forward_demodulation,[],[f63361,f721])).
fof(f63361,plain,(
  ( ! [X30,X31,X32] : (sP0(X30,X31,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int',X32)))) ) | $spl22),
  inference(superposition,[],[f50751,f709])).
fof(f63312,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl194),
  inference(cnf_transformation,[],[f63312_D])).
fof(f63312_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f63258,plain,(
  ( ! [X156,X157] : (p(s(bool,X156)) | p(s(bool,X157)) | s(bool,X157) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl22),
  inference(resolution,[],[f50735,f735])).
fof(f50735,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl22),
  inference(resolution,[],[f50722,f700])).
fof(f700,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f333])).
fof(f333,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f93])).
fof(f93,axiom,(
  ! [X6,X7] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X7),s('type.integer.int',X6)))))),
  file('thm.integer.INT_LE_MUL','thm.integer.INT_LE_TOTAL')).
fof(f63188,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X26) != s(bool,X27)) ) | $spl22),
  inference(resolution,[],[f50735,f526])).
fof(f63185,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl22),
  inference(resolution,[],[f50735,f525])).
fof(f63184,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl22),
  inference(resolution,[],[f50735,f524])).
fof(f63183,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X11) != s(bool,X12)) ) | $spl22),
  inference(resolution,[],[f50735,f522])).
fof(f63181,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl22),
  inference(resolution,[],[f50735,f520])).
fof(f29626,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl122),
  inference(cnf_transformation,[],[f29626_D])).
fof(f29626_D,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f61359,plain,(
  ( ! [X163,X164] : (p(s(bool,X163)) | ~p(s(bool,X164)) | s(bool,X164) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl77),
  inference(resolution,[],[f8101,f735])).
fof(f8101,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | $spl77),
  inference(cnf_transformation,[],[f8101_D])).
fof(f8101_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) <=> ~$spl77),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl77])])).
fof(f61285,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))))) | s(bool,X24) != s(bool,X26)) ) | $spl77),
  inference(resolution,[],[f8101,f526])).
fof(f61283,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl77),
  inference(resolution,[],[f8101,f525])).
fof(f61281,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl77),
  inference(resolution,[],[f8101,f524])).
fof(f61280,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))))) | s(bool,X9) != s(bool,X11)) ) | $spl77),
  inference(resolution,[],[f8101,f522])).
fof(f61278,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl77),
  inference(resolution,[],[f8101,f520])).
fof(f19075,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl94),
  inference(cnf_transformation,[],[f19075_D])).
fof(f19075_D,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f61195,plain,(
  ( ! [X158,X157] : (p(s(bool,X157)) | p(s(bool,X158)) | s(bool,X158) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl74),
  inference(resolution,[],[f8100,f735])).
fof(f61125,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X27) != s(bool,X28)) ) | $spl74),
  inference(resolution,[],[f8100,f526])).
fof(f61122,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X19) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl74),
  inference(resolution,[],[f8100,f525])).
fof(f61121,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl74),
  inference(resolution,[],[f8100,f524])).
fof(f61120,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X12) != s(bool,X13)) ) | $spl74),
  inference(resolution,[],[f8100,f522])).
fof(f61118,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl74),
  inference(resolution,[],[f8100,f520])).
fof(f58134,plain,(
  $false | ($spl3 | $spl22 | $spl128)),
  inference(subsumption_resolution,[],[f57814,f874])).
fof(f57814,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | ($spl22 | $spl128)),
  inference(superposition,[],[f50532,f29647])).
fof(f50532,plain,(
  ( ! [X206] : (s('type.integer.int',X206) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X206),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl22),
  inference(subsumption_resolution,[],[f50531,f1015])).
fof(f50531,plain,(
  ( ! [X206] : (s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int',sK12) | s('type.integer.int',X206) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X206),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl22),
  inference(forward_demodulation,[],[f50262,f721])).
fof(f50262,plain,(
  ( ! [X206] : (s('type.integer.int',X206) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X206),s('type.integer.int','const.integer.int_|48|_0'))) | s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl22),
  inference(superposition,[],[f712,f1015])).
fof(f58133,plain,(
  $false | ($spl126 | $spl128)),
  inference(subsumption_resolution,[],[f58132,f57633])).
fof(f57633,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) | $spl126),
  inference(forward_demodulation,[],[f57601,f721])).
fof(f57601,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) | $spl126),
  inference(superposition,[],[f29645,f717])).
fof(f717,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f347])).
fof(f347,plain,(
  ! [X0] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X0)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X6))),s('type.integer.int',X6)))),
  file('thm.integer.INT_LE_MUL','thm.integer.INT_ADD_LINV')).
fof(f58132,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) | $spl128),
  inference(forward_demodulation,[],[f57813,f721])).
fof(f57813,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) | $spl128),
  inference(superposition,[],[f718,f29647])).
fof(f718,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))) )),
  inference(cnf_transformation,[],[f348])).
fof(f348,plain,(
  ! [X0] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))),
  inference(rectify,[],[f83])).
fof(f83,axiom,(
  ! [X6] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X6),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X6)))))),
  file('thm.integer.INT_LE_MUL','thm.integer.INT_ADD_RINV')).
fof(f58131,plain,(
  $false | ($spl3 | $spl128)),
  inference(subsumption_resolution,[],[f58130,f874])).
fof(f58130,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | $spl128),
  inference(forward_demodulation,[],[f57812,f721])).
fof(f57812,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl128),
  inference(superposition,[],[f719,f29647])).
fof(f719,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(cnf_transformation,[],[f349])).
fof(f349,plain,(
  ! [X0] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X6] : s('type.integer.int',X6) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X6),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.integer.INT_LE_MUL','thm.integer.INT_ADD_RID')).
fof(f58129,plain,(
  $false | ($spl126 | $spl128)),
  inference(subsumption_resolution,[],[f57811,f29645])).
fof(f57811,plain,(
  ( ! [X39] : (s('type.integer.int',X39) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X39),s('type.integer.int',sK13)))) ) | $spl128),
  inference(superposition,[],[f734,f29647])).
fof(f734,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f361])).
fof(f361,plain,(
  ! [X0,X1] : s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X0)))),
  inference(rectify,[],[f175])).
fof(f175,axiom,(
  ! [X7,X6] : s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X6),s('type.integer.int',X7))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X7),s('type.integer.int',X6)))),
  file('thm.integer.INT_LE_MUL','thm.integer.INT_ADD_SYM')).
fof(f58128,plain,(
  $false | ($spl126 | $spl128)),
  inference(subsumption_resolution,[],[f57810,f29645])).
fof(f57810,plain,(
  ( ! [X38] : (s('type.integer.int',X38) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X38),s('type.integer.int',sK13)))) ) | $spl128),
  inference(superposition,[],[f734,f29647])).
fof(f58125,plain,(
  $false | ($spl3 | $spl128)),
  inference(subsumption_resolution,[],[f58124,f874])).
fof(f58124,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | $spl128),
  inference(forward_demodulation,[],[f58076,f721])).
fof(f58076,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl128),
  inference(trivial_inequality_removal,[],[f57792])).
fof(f57792,plain,(
  ( ! [X4] : (s('type.integer.int',X4) != s('type.integer.int',X4) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl128),
  inference(superposition,[],[f709,f29647])).
fof(f58087,plain,(
  $false | ($spl3 | $spl22 | $spl128)),
  inference(subsumption_resolution,[],[f57770,f874])).
fof(f57770,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | ($spl22 | $spl128)),
  inference(superposition,[],[f29647,f50532])).
fof(f58086,plain,(
  $false | ($spl126 | $spl128)),
  inference(subsumption_resolution,[],[f58085,f57633])).
fof(f58085,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) | $spl128),
  inference(forward_demodulation,[],[f57769,f721])).
fof(f57769,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13))) | $spl128),
  inference(superposition,[],[f29647,f718])).
fof(f58084,plain,(
  $false | ($spl3 | $spl128)),
  inference(subsumption_resolution,[],[f58083,f874])).
fof(f58083,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | $spl128),
  inference(forward_demodulation,[],[f57768,f721])).
fof(f57768,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl128),
  inference(superposition,[],[f29647,f719])).
fof(f58082,plain,(
  $false | ($spl126 | $spl128)),
  inference(subsumption_resolution,[],[f57767,f29645])).
fof(f57767,plain,(
  ( ! [X14] : (s('type.integer.int',X14) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X14),s('type.integer.int',sK13)))) ) | $spl128),
  inference(superposition,[],[f29647,f734])).
fof(f58081,plain,(
  $false | ($spl126 | $spl128)),
  inference(subsumption_resolution,[],[f57766,f29645])).
fof(f57766,plain,(
  ( ! [X13] : (s('type.integer.int',X13) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X13),s('type.integer.int',sK13)))) ) | $spl128),
  inference(superposition,[],[f29647,f734])).
fof(f58075,plain,(
  $false | ($spl126 | $spl128)),
  inference(trivial_inequality_removal,[],[f57815])).
fof(f57815,plain,(
  s('type.integer.int',sK13) != s('type.integer.int',sK13) | ($spl126 | $spl128)),
  inference(superposition,[],[f29645,f29647])).
fof(f57278,plain,(
  $false | ($spl3 | $spl124)),
  inference(subsumption_resolution,[],[f57277,f874])).
fof(f57277,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | $spl124),
  inference(forward_demodulation,[],[f57228,f721])).
fof(f57228,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl124),
  inference(trivial_inequality_removal,[],[f56951])).
fof(f56951,plain,(
  ( ! [X6] : (s('type.integer.int',X6) != s('type.integer.int',X6) | s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl124),
  inference(superposition,[],[f711,f29643])).
fof(f57272,plain,(
  $false | ($spl3 | $spl22 | $spl124)),
  inference(subsumption_resolution,[],[f56944,f874])).
fof(f56944,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | ($spl22 | $spl124)),
  inference(superposition,[],[f50529,f29643])).
fof(f57270,plain,(
  $false | ($spl3 | $spl124)),
  inference(subsumption_resolution,[],[f57269,f874])).
fof(f57269,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | $spl124),
  inference(forward_demodulation,[],[f56942,f721])).
fof(f56942,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl124),
  inference(superposition,[],[f720,f29643])).
fof(f720,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f350])).
fof(f350,plain,(
  ! [X0] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X6] : s('type.integer.int',X6) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X6)))),
  file('thm.integer.INT_LE_MUL','thm.integer.INT_ADD_LID')).
fof(f57234,plain,(
  $false | ($spl3 | $spl22 | $spl124)),
  inference(subsumption_resolution,[],[f56914,f874])).
fof(f56914,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | ($spl22 | $spl124)),
  inference(superposition,[],[f29643,f50529])).
fof(f57232,plain,(
  $false | ($spl3 | $spl124)),
  inference(subsumption_resolution,[],[f57231,f874])).
fof(f57231,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | $spl124),
  inference(forward_demodulation,[],[f56912,f721])).
fof(f56912,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl124),
  inference(superposition,[],[f29643,f720])).
fof(f54818,plain,(
  $false | ($spl22 | $spl98)),
  inference(subsumption_resolution,[],[f19262,f50726])).
fof(f8100,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl74),
  inference(cnf_transformation,[],[f8100_D])).
fof(f8100_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f8102,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | $spl76),
  inference(cnf_transformation,[],[f8102_D])).
fof(f8102_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f52866,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl22),
  inference(forward_demodulation,[],[f52865,f721])).
fof(f52865,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | $spl22),
  inference(forward_demodulation,[],[f52864,f721])).
fof(f52864,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | $spl22),
  inference(forward_demodulation,[],[f52863,f716])).
fof(f716,plain,(
  ( ! [X0] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))) )),
  inference(cnf_transformation,[],[f346])).
fof(f346,plain,(
  ! [X0] : s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6] : s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X6)))))),
  file('thm.integer.INT_LE_MUL','thm.integer.INT_NEG_GT|48|')).
fof(f52863,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | $spl22),
  inference(subsumption_resolution,[],[f52854,f721])).
fof(f52854,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | $spl22),
  inference(superposition,[],[f50726,f714])).
fof(f714,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f344])).
fof(f344,plain,(
  ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : (s('type.integer.int',X6) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X6)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X6)))))))),
  file('thm.integer.INT_LE_MUL','thm.integer.INT_LT_NEGTOTAL')).
fof(f8902,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | p(s(bool,X8))) ) | $spl78),
  inference(cnf_transformation,[],[f8902_D])).
fof(f8902_D,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | p(s(bool,X8))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f51493,plain,(
  ( ! [X159,X160] : (p(s(bool,X159)) | p(s(bool,X160)) | s(bool,X160) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl4),
  inference(resolution,[],[f877,f735])).
fof(f51423,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X29) != s(bool,X30)) ) | $spl4),
  inference(resolution,[],[f877,f526])).
fof(f51420,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X21) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl4),
  inference(resolution,[],[f877,f525])).
fof(f51419,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl4),
  inference(resolution,[],[f877,f524])).
fof(f51418,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X14) != s(bool,X15)) ) | $spl4),
  inference(resolution,[],[f877,f522])).
fof(f51416,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl4),
  inference(resolution,[],[f877,f520])).
fof(f3581,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X2))) ) | $spl58),
  inference(cnf_transformation,[],[f3581_D])).
fof(f3581_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X2))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f51358,plain,(
  ( ! [X161,X162] : (p(s(bool,X161)) | ~p(s(bool,X162)) | s(bool,X162) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl1),
  inference(resolution,[],[f873,f735])).
fof(f51284,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X22) != s(bool,X24)) ) | $spl1),
  inference(resolution,[],[f873,f526])).
fof(f51282,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl1),
  inference(resolution,[],[f873,f525])).
fof(f51280,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl1),
  inference(resolution,[],[f873,f524])).
fof(f51279,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X7) != s(bool,X9)) ) | $spl1),
  inference(resolution,[],[f873,f522])).
fof(f51277,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl1),
  inference(resolution,[],[f873,f520])).
fof(f3022,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X4))) ) | $spl46),
  inference(cnf_transformation,[],[f3022_D])).
fof(f3022_D,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X4))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f50978,plain,(
  ( ! [X163,X164] : (p(s(bool,X163)) | ~p(s(bool,X164)) | s(bool,X164) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) )),
  inference(resolution,[],[f749,f735])).
fof(f50904,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X24) != s(bool,X26)) )),
  inference(resolution,[],[f749,f526])).
fof(f50902,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) )),
  inference(resolution,[],[f749,f525])).
fof(f50900,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) )),
  inference(resolution,[],[f749,f524])).
fof(f50899,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f749,f522])).
fof(f50897,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) )),
  inference(resolution,[],[f749,f520])).
fof(f50883,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl22),
  inference(forward_demodulation,[],[f50882,f721])).
fof(f50882,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | $spl22),
  inference(forward_demodulation,[],[f50881,f721])).
fof(f50881,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | $spl22),
  inference(forward_demodulation,[],[f50880,f716])).
fof(f50880,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | $spl22),
  inference(subsumption_resolution,[],[f50879,f704])).
fof(f50879,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | $spl22),
  inference(forward_demodulation,[],[f50858,f721])).
fof(f50858,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | $spl22),
  inference(superposition,[],[f50722,f714])).
fof(f8085,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl70),
  inference(cnf_transformation,[],[f8085_D])).
fof(f8085_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f50818,plain,(
  ( ! [X161,X162] : (p(s(bool,X161)) | ~p(s(bool,X162)) | s(bool,X162) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl22),
  inference(resolution,[],[f50722,f735])).
fof(f50744,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))))) | s(bool,X22) != s(bool,X24)) ) | $spl22),
  inference(resolution,[],[f50722,f526])).
fof(f50742,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl22),
  inference(resolution,[],[f50722,f525])).
fof(f50740,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl22),
  inference(resolution,[],[f50722,f524])).
fof(f50739,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))))) | s(bool,X7) != s(bool,X9)) ) | $spl22),
  inference(resolution,[],[f50722,f522])).
fof(f50737,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl22),
  inference(resolution,[],[f50722,f520])).
fof(f878,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | p(s(bool,X5))) ) | $spl6),
  inference(cnf_transformation,[],[f878_D])).
fof(f878_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | p(s(bool,X5))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f50681,plain,(
  ( ! [X156,X157] : (p(s(bool,X156)) | p(s(bool,X157)) | s(bool,X157) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f743,f735])).
fof(f50611,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X26) != s(bool,X27)) )),
  inference(resolution,[],[f743,f526])).
fof(f50608,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f743,f525])).
fof(f50607,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f743,f524])).
fof(f50606,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X11) != s(bool,X12)) )),
  inference(resolution,[],[f743,f522])).
fof(f50604,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f743,f520])).
fof(f1014,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int',sK12) | $spl23),
  inference(cnf_transformation,[],[f1014_D])).
fof(f1014_D,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int',sK12) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f24467,plain,(
  ( ! [X205] : (s('type.integer.int',X205) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK12),s('type.integer.int',X205)))) ) | $spl116),
  inference(cnf_transformation,[],[f24467_D])).
fof(f24467_D,plain,(
  ( ! [X205] : (s('type.integer.int',X205) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK12),s('type.integer.int',X205)))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f49918,plain,(
  ( ! [X205] : (s('type.integer.int',X205) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK12),s('type.integer.int',X205))) | s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int',sK12)) )),
  inference(forward_demodulation,[],[f42792,f21723])).
fof(f21723,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))),
  inference(forward_demodulation,[],[f21722,f4120])).
fof(f4120,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))),
  inference(resolution,[],[f889,f713])).
fof(f713,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_|48|_0'))))) )),
  inference(cnf_transformation,[],[f499])).
fof(f499,plain,(
  ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK24])],[f391])).
fof(f391,plain,(
  ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_|48|_0')))) | ? [X1] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))),
  inference(ennf_transformation,[],[f343])).
fof(f343,plain,(
  ! [X0] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_|48|_0')))) => ? [X1] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))),
  inference(flattening,[],[f342])).
fof(f342,plain,(
  ! [X0] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_|48|_0')))) => ? [X1] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))),
  inference(rectify,[],[f113])).
fof(f113,axiom,(
  ! [X10] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_|48|_0')))) => ? [X51] : s('type.integer.int',X10) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X51))))),
  file('thm.integer.INT_LE_MUL','thm.integer.NUM_POSINT_EX')).
fof(f21722,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))),
  inference(subsumption_resolution,[],[f21417,f889])).
fof(f21417,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))),
  inference(superposition,[],[f713,f4120])).
fof(f42792,plain,(
  ( ! [X205] : (s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int',sK12) | s('type.integer.int',X205) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))),s('type.integer.int',X205)))) )),
  inference(forward_demodulation,[],[f42484,f721])).
fof(f42484,plain,(
  ( ! [X205] : (s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X205) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))),s('type.integer.int',X205)))) )),
  inference(superposition,[],[f710,f21723])).
fof(f1015,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | $spl22),
  inference(cnf_transformation,[],[f1015_D])).
fof(f1015_D,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f24469,plain,(
  ( ! [X202] : (s('type.integer.int',X202) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK12),s('type.integer.int',X202)))) ) | $spl118),
  inference(cnf_transformation,[],[f24469_D])).
fof(f24469_D,plain,(
  ( ! [X202] : (s('type.integer.int',X202) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK12),s('type.integer.int',X202)))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f49920,plain,(
  ( ! [X202] : (s('type.integer.int',X202) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK12),s('type.integer.int',X202))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12)) )),
  inference(forward_demodulation,[],[f42786,f21723])).
fof(f42786,plain,(
  ( ! [X202] : (s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | s('type.integer.int',X202) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))),s('type.integer.int',X202)))) )),
  inference(forward_demodulation,[],[f42481,f721])).
fof(f42481,plain,(
  ( ! [X202] : (s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X202) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))),s('type.integer.int',X202)))) )),
  inference(superposition,[],[f709,f21723])).
fof(f49928,plain,(
  ( ! [X10] : (s('type.integer.int',X10) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK12),s('type.integer.int',X10))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12)) )),
  inference(forward_demodulation,[],[f42580,f21723])).
fof(f42580,plain,(
  ( ! [X10] : (s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | s('type.integer.int',X10) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))),s('type.integer.int',X10)))) )),
  inference(forward_demodulation,[],[f42294,f721])).
fof(f42294,plain,(
  ( ! [X10] : (s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X10) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))),s('type.integer.int',X10)))) )),
  inference(superposition,[],[f21723,f709])).
fof(f49937,plain,(
  ( ! [X205] : (s('type.integer.int',X205) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK12),s('type.integer.int',X205))) | s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int',sK12)) )),
  inference(forward_demodulation,[],[f33069,f4120])).
fof(f33069,plain,(
  ( ! [X205] : (s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int',sK12) | s('type.integer.int',X205) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))),s('type.integer.int',X205)))) )),
  inference(forward_demodulation,[],[f32766,f721])).
fof(f32766,plain,(
  ( ! [X205] : (s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X205) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))),s('type.integer.int',X205)))) )),
  inference(superposition,[],[f710,f4120])).
fof(f49939,plain,(
  ( ! [X202] : (s('type.integer.int',X202) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK12),s('type.integer.int',X202))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12)) )),
  inference(forward_demodulation,[],[f33065,f4120])).
fof(f33065,plain,(
  ( ! [X202] : (s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | s('type.integer.int',X202) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))),s('type.integer.int',X202)))) )),
  inference(forward_demodulation,[],[f32763,f721])).
fof(f32763,plain,(
  ( ! [X202] : (s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X202) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))),s('type.integer.int',X202)))) )),
  inference(superposition,[],[f709,f4120])).
fof(f49947,plain,(
  ( ! [X10] : (s('type.integer.int',X10) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK12),s('type.integer.int',X10))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12)) )),
  inference(forward_demodulation,[],[f32858,f4120])).
fof(f32858,plain,(
  ( ! [X10] : (s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | s('type.integer.int',X10) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))),s('type.integer.int',X10)))) )),
  inference(forward_demodulation,[],[f32577,f721])).
fof(f32577,plain,(
  ( ! [X10] : (s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X10) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))),s('type.integer.int',X10)))) )),
  inference(superposition,[],[f4120,f709])).
fof(f49981,plain,(
  $false | ($spl22 | $spl42 | $spl88)),
  inference(subsumption_resolution,[],[f49980,f667])).
fof(f49980,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl22 | $spl42 | $spl88)),
  inference(forward_demodulation,[],[f17740,f17521])).
fof(f17740,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl22 | $spl42)),
  inference(forward_demodulation,[],[f2740,f1015])).
fof(f49979,plain,(
  $false | ($spl22 | $spl42 | $spl88)),
  inference(subsumption_resolution,[],[f49978,f667])).
fof(f49978,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl22 | $spl42 | $spl88)),
  inference(forward_demodulation,[],[f24522,f17521])).
fof(f24522,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl22 | $spl42)),
  inference(backward_demodulation,[],[f1015,f2740])).
fof(f49977,plain,(
  $false | ($spl22 | $spl42 | $spl88)),
  inference(subsumption_resolution,[],[f24550,f17521])).
fof(f24550,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | ($spl22 | $spl42)),
  inference(backward_demodulation,[],[f1015,f23371])).
fof(f23371,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | $spl42),
  inference(resolution,[],[f2740,f695])).
fof(f695,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) != s('type.integer.int',X1) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f494])).
fof(f494,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & s('type.integer.int',X0) != s('type.integer.int',X1))) & (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(flattening,[],[f493])).
fof(f493,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & s('type.integer.int',X0) != s('type.integer.int',X1))) & ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f329])).
fof(f329,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & s('type.integer.int',X0) != s('type.integer.int',X1)))),
  inference(flattening,[],[f328])).
fof(f328,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & ~s('type.integer.int',X0) = s('type.integer.int',X1)))),
  inference(rectify,[],[f96])).
fof(f96,axiom,(
  ! [X6,X7] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) <=> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) & ~s('type.integer.int',X6) = s('type.integer.int',X7)))),
  file('thm.integer.INT_LE_MUL','thm.integer.INT_LT_LE')).
fof(f49976,plain,(
  $false | ($spl112 | $spl118)),
  inference(subsumption_resolution,[],[f36614,f24469])).
fof(f36614,plain,(
  ( ! [X13] : (s('type.integer.int',X13) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK12),s('type.integer.int',X13)))) ) | $spl112),
  inference(superposition,[],[f24463,f734])).
fof(f49975,plain,(
  $false | ($spl112 | $spl118)),
  inference(subsumption_resolution,[],[f36615,f24469])).
fof(f36615,plain,(
  ( ! [X14] : (s('type.integer.int',X14) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK12),s('type.integer.int',X14)))) ) | $spl112),
  inference(superposition,[],[f24463,f734])).
fof(f49974,plain,(
  $false | ($spl112 | $spl118)),
  inference(subsumption_resolution,[],[f36659,f24469])).
fof(f36659,plain,(
  ( ! [X38] : (s('type.integer.int',X38) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK12),s('type.integer.int',X38)))) ) | $spl112),
  inference(superposition,[],[f734,f24463])).
fof(f49973,plain,(
  $false | ($spl112 | $spl118)),
  inference(subsumption_resolution,[],[f36660,f24469])).
fof(f36660,plain,(
  ( ! [X39] : (s('type.integer.int',X39) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK12),s('type.integer.int',X39)))) ) | $spl112),
  inference(superposition,[],[f734,f24463])).
fof(f49972,plain,(
  $false | ($spl116 | $spl118)),
  inference(subsumption_resolution,[],[f24467,f24469])).
fof(f1017,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | $spl24),
  inference(cnf_transformation,[],[f1017_D])).
fof(f1017_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f49967,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))))),
  inference(forward_demodulation,[],[f49966,f721])).
fof(f49966,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))))),
  inference(forward_demodulation,[],[f33082,f4120])).
fof(f33082,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))),
  inference(forward_demodulation,[],[f33081,f721])).
fof(f33081,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))),
  inference(forward_demodulation,[],[f33080,f4120])).
fof(f33080,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))),
  inference(subsumption_resolution,[],[f33079,f889])).
fof(f33079,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))),
  inference(forward_demodulation,[],[f33078,f721])).
fof(f33078,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))),
  inference(forward_demodulation,[],[f32774,f716])).
fof(f32774,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))),
  inference(superposition,[],[f714,f4120])).
fof(f49965,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))))),
  inference(forward_demodulation,[],[f49964,f721])).
fof(f49964,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))))),
  inference(forward_demodulation,[],[f49963,f4120])).
fof(f49963,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))),
  inference(subsumption_resolution,[],[f49962,f889])).
fof(f49962,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))),
  inference(forward_demodulation,[],[f49961,f4120])).
fof(f49961,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))),
  inference(forward_demodulation,[],[f49960,f721])).
fof(f49960,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))),
  inference(forward_demodulation,[],[f33083,f716])).
fof(f33083,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))),
  inference(forward_demodulation,[],[f32775,f721])).
fof(f32775,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))),
  inference(superposition,[],[f714,f4120])).
fof(f49957,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))))),
  inference(forward_demodulation,[],[f49956,f721])).
fof(f49956,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))))),
  inference(forward_demodulation,[],[f42806,f21723])).
fof(f42806,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))),
  inference(forward_demodulation,[],[f42805,f721])).
fof(f42805,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))),
  inference(forward_demodulation,[],[f42804,f21723])).
fof(f42804,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))),
  inference(subsumption_resolution,[],[f42803,f889])).
fof(f42803,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))),
  inference(forward_demodulation,[],[f42802,f721])).
fof(f42802,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))),
  inference(forward_demodulation,[],[f42492,f716])).
fof(f42492,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))),
  inference(superposition,[],[f714,f21723])).
fof(f49955,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))))),
  inference(forward_demodulation,[],[f49954,f721])).
fof(f49954,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))))),
  inference(forward_demodulation,[],[f49953,f21723])).
fof(f49953,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))),
  inference(subsumption_resolution,[],[f49952,f889])).
fof(f49952,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))),
  inference(forward_demodulation,[],[f49951,f21723])).
fof(f49951,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))),
  inference(forward_demodulation,[],[f49950,f721])).
fof(f49950,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))),
  inference(forward_demodulation,[],[f42807,f716])).
fof(f42807,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))),
  inference(forward_demodulation,[],[f42493,f721])).
fof(f42493,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))),
  inference(superposition,[],[f714,f21723])).
fof(f49949,plain,(
  ( ! [X12,X13] : (p(s(bool,X13)) | p(s(bool,X12)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))))) ) | $spl82),
  inference(subsumption_resolution,[],[f47560,f644])).
fof(f47560,plain,(
  ( ! [X12,X13] : (s(bool,X12) != s(bool,X13) | p(s(bool,X13)) | p(s(bool,X12)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))))) ) | $spl82),
  inference(superposition,[],[f9553,f644])).
fof(f49948,plain,(
  ( ! [X14,X15] : (p(s(bool,X15)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | p(s(bool,X14))) ) | $spl82),
  inference(subsumption_resolution,[],[f47561,f644])).
fof(f47561,plain,(
  ( ! [X14,X15] : (s(bool,X14) != s(bool,X15) | p(s(bool,X15)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | p(s(bool,X14))) ) | $spl82),
  inference(superposition,[],[f9553,f644])).
fof(f1013,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | $spl21),
  inference(cnf_transformation,[],[f1013_D])).
fof(f1013_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f30291,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))))),
  inference(resolution,[],[f744,f686])).
fof(f686,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) = s('type.integer.int',X1) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f489])).
fof(f30292,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))))),
  inference(resolution,[],[f744,f686])).
fof(f30294,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))))),
  inference(resolution,[],[f744,f691])).
fof(f691,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) = s('type.integer.int',X1) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f492])).
fof(f492,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1)) & ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & s('type.integer.int',X0) != s('type.integer.int',X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(flattening,[],[f491])).
fof(f491,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1))) & ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & s('type.integer.int',X0) != s('type.integer.int',X1)) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f327])).
fof(f327,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6,X7] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) <=> (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) | s('type.integer.int',X6) = s('type.integer.int',X7)))),
  file('thm.integer.INT_LE_MUL','thm.integer.INT_LE_LT')).
fof(f30295,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12)),
  inference(resolution,[],[f744,f696])).
fof(f696,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f494])).
fof(f30419,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12)),
  inference(resolution,[],[f889,f649])).
fof(f649,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1)) )),
  inference(cnf_transformation,[],[f301])).
fof(f301,plain,(
  ! [X0,X1] : (s('type.integer.int',X0) = s('type.integer.int',X1) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f214])).
fof(f214,axiom,(
  ! [X6,X7] : (s('type.integer.int',X6) = s('type.integer.int',X7) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int',X6)))))),
  file('thm.integer.INT_LE_MUL','thm.integer.INT_LT_TOTAL')).
fof(f30420,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12)),
  inference(resolution,[],[f889,f649])).
fof(f30426,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))))),
  inference(resolution,[],[f889,f691])).
fof(f30427,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))))),
  inference(resolution,[],[f889,f696])).
fof(f24465,plain,(
  ( ! [X208] : (s('type.integer.int',X208) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X208),s('type.integer.int',sK12)))) ) | $spl114),
  inference(cnf_transformation,[],[f24465_D])).
fof(f24465_D,plain,(
  ( ! [X208] : (s('type.integer.int',X208) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X208),s('type.integer.int',sK12)))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f49946,plain,(
  ( ! [X11] : (s('type.integer.int',X11) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X11),s('type.integer.int',sK12))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12)) )),
  inference(forward_demodulation,[],[f32861,f4120])).
fof(f32861,plain,(
  ( ! [X11] : (s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | s('type.integer.int',X11) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X11),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))) )),
  inference(forward_demodulation,[],[f32578,f721])).
fof(f32578,plain,(
  ( ! [X11] : (s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X11) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X11),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))) )),
  inference(superposition,[],[f4120,f711])).
fof(f49945,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12)),
  inference(forward_demodulation,[],[f49944,f721])).
fof(f49944,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12)),
  inference(forward_demodulation,[],[f49943,f4120])).
fof(f49943,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))),
  inference(subsumption_resolution,[],[f49942,f889])).
fof(f49942,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))),
  inference(forward_demodulation,[],[f49941,f4120])).
fof(f49941,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))),
  inference(forward_demodulation,[],[f49940,f721])).
fof(f49940,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))),
  inference(forward_demodulation,[],[f32864,f716])).
fof(f32864,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))),
  inference(forward_demodulation,[],[f32579,f721])).
fof(f32579,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))),
  inference(superposition,[],[f4120,f714])).
fof(f49936,plain,(
  ( ! [X208] : (s('type.integer.int',X208) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X208),s('type.integer.int',sK12))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12)) )),
  inference(forward_demodulation,[],[f33071,f4120])).
fof(f33071,plain,(
  ( ! [X208] : (s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | s('type.integer.int',X208) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X208),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))) )),
  inference(forward_demodulation,[],[f32769,f721])).
fof(f32769,plain,(
  ( ! [X208] : (s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X208) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X208),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))) )),
  inference(superposition,[],[f711,f4120])).
fof(f24463,plain,(
  ( ! [X211] : (s('type.integer.int',X211) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X211),s('type.integer.int',sK12)))) ) | $spl112),
  inference(cnf_transformation,[],[f24463_D])).
fof(f24463_D,plain,(
  ( ! [X211] : (s('type.integer.int',X211) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X211),s('type.integer.int',sK12)))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f49935,plain,(
  ( ! [X211] : (s('type.integer.int',X211) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X211),s('type.integer.int',sK12))) | s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int',sK12)) )),
  inference(forward_demodulation,[],[f33075,f4120])).
fof(f33075,plain,(
  ( ! [X211] : (s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int',sK12) | s('type.integer.int',X211) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X211),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))) )),
  inference(forward_demodulation,[],[f32772,f721])).
fof(f32772,plain,(
  ( ! [X211] : (s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X211) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X211),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))) )),
  inference(superposition,[],[f712,f4120])).
fof(f49934,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12)),
  inference(forward_demodulation,[],[f49933,f721])).
fof(f49933,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12)),
  inference(forward_demodulation,[],[f49932,f4120])).
fof(f49932,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))),
  inference(subsumption_resolution,[],[f49931,f889])).
fof(f49931,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))),
  inference(forward_demodulation,[],[f49930,f4120])).
fof(f49930,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))),
  inference(forward_demodulation,[],[f49929,f721])).
fof(f49929,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))),
  inference(forward_demodulation,[],[f33084,f716])).
fof(f33084,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))),
  inference(forward_demodulation,[],[f32776,f721])).
fof(f32776,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))),
  inference(superposition,[],[f714,f4120])).
fof(f49927,plain,(
  ( ! [X11] : (s('type.integer.int',X11) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X11),s('type.integer.int',sK12))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12)) )),
  inference(forward_demodulation,[],[f42583,f21723])).
fof(f42583,plain,(
  ( ! [X11] : (s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | s('type.integer.int',X11) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X11),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))))) )),
  inference(forward_demodulation,[],[f42295,f721])).
fof(f42295,plain,(
  ( ! [X11] : (s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X11) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X11),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))))) )),
  inference(superposition,[],[f21723,f711])).
fof(f49926,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12)),
  inference(forward_demodulation,[],[f49925,f721])).
fof(f49925,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12)),
  inference(forward_demodulation,[],[f49924,f21723])).
fof(f49924,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))))))),
  inference(subsumption_resolution,[],[f49923,f889])).
fof(f49923,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))))))),
  inference(forward_demodulation,[],[f49922,f21723])).
fof(f49922,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))))))),
  inference(forward_demodulation,[],[f49921,f721])).
fof(f49921,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))))))),
  inference(forward_demodulation,[],[f42586,f716])).
fof(f42586,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))))))),
  inference(forward_demodulation,[],[f42296,f721])).
fof(f42296,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))))))),
  inference(superposition,[],[f21723,f714])).
fof(f49917,plain,(
  ( ! [X208] : (s('type.integer.int',X208) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X208),s('type.integer.int',sK12))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12)) )),
  inference(forward_demodulation,[],[f42793,f21723])).
fof(f42793,plain,(
  ( ! [X208] : (s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | s('type.integer.int',X208) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X208),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))))) )),
  inference(forward_demodulation,[],[f42487,f721])).
fof(f42487,plain,(
  ( ! [X208] : (s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X208) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X208),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))))) )),
  inference(superposition,[],[f711,f21723])).
fof(f49916,plain,(
  ( ! [X211] : (s('type.integer.int',X211) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X211),s('type.integer.int',sK12))) | s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int',sK12)) )),
  inference(forward_demodulation,[],[f42799,f21723])).
fof(f42799,plain,(
  ( ! [X211] : (s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int',sK12) | s('type.integer.int',X211) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X211),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))))) )),
  inference(forward_demodulation,[],[f42490,f721])).
fof(f42490,plain,(
  ( ! [X211] : (s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X211) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X211),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))))) )),
  inference(superposition,[],[f712,f21723])).
fof(f49915,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12)),
  inference(forward_demodulation,[],[f49914,f721])).
fof(f49914,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12)),
  inference(forward_demodulation,[],[f49913,f21723])).
fof(f49913,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))))))),
  inference(subsumption_resolution,[],[f49912,f889])).
fof(f49912,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))))))),
  inference(forward_demodulation,[],[f49911,f21723])).
fof(f49911,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))))))),
  inference(forward_demodulation,[],[f49910,f721])).
fof(f49910,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))))))),
  inference(forward_demodulation,[],[f42808,f716])).
fof(f42808,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))))))),
  inference(forward_demodulation,[],[f42494,f721])).
fof(f42494,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24('const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))))))),
  inference(superposition,[],[f714,f21723])).
fof(f47509,plain,(
  ( ! [X17,X16] : (~p(s(bool,X17)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | ~p(s(bool,X16))) ) | $spl66),
  inference(subsumption_resolution,[],[f47486,f647])).
fof(f47486,plain,(
  ( ! [X17,X16] : (s(bool,X16) != s(bool,X17) | ~p(s(bool,X17)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | ~p(s(bool,X16))) ) | $spl66),
  inference(superposition,[],[f4797,f647])).
fof(f47510,plain,(
  ( ! [X19,X18] : (~p(s(bool,X19)) | ~p(s(bool,X18)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))))) ) | $spl66),
  inference(subsumption_resolution,[],[f47487,f647])).
fof(f47487,plain,(
  ( ! [X19,X18] : (s(bool,X18) != s(bool,X19) | ~p(s(bool,X19)) | ~p(s(bool,X18)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))))) ) | $spl66),
  inference(superposition,[],[f4797,f647])).
fof(f873,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) | $spl1),
  inference(cnf_transformation,[],[f873_D])).
fof(f873_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f875,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | $spl2),
  inference(cnf_transformation,[],[f875_D])).
fof(f875_D,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f7645,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))))),
  inference(resolution,[],[f743,f686])).
fof(f7646,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))))),
  inference(resolution,[],[f743,f686])).
fof(f7777,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))))),
  inference(resolution,[],[f749,f691])).
fof(f7778,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))))),
  inference(resolution,[],[f749,f696])).
fof(f872,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) | $spl0),
  inference(cnf_transformation,[],[f872_D])).
fof(f872_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f49905,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13))))))) | $spl192),
  inference(cnf_transformation,[],[f49905_D])).
fof(f49905_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13))))))) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f49903,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13))))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))))),
  inference(forward_demodulation,[],[f12042,f721])).
fof(f12042,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13)))))))),
  inference(forward_demodulation,[],[f12041,f721])).
fof(f12041,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13)))))))),
  inference(forward_demodulation,[],[f11747,f703])).
fof(f703,plain,(
  ( ! [X0] : (s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))) )),
  inference(cnf_transformation,[],[f336])).
fof(f336,plain,(
  ! [X0] : s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : s(bool,'const.integer.int_le_2'(s('type.integer.int',X6),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X6)))))),
  file('thm.integer.INT_LE_MUL','thm.integer.INT_NEG_GE|48|')).
fof(f11747,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13)))))))),
  inference(superposition,[],[f701,f2886])).
fof(f2886,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13))))),
  inference(resolution,[],[f749,f713])).
fof(f701,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f334])).
fof(f334,plain,(
  ! [X0] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X6)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X6)))))))),
  file('thm.integer.INT_LE_MUL','thm.integer.INT_LE_NEGTOTAL')).
fof(f49587,plain,(
  $false | $spl2),
  inference(trivial_inequality_removal,[],[f49323])).
fof(f49323,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int','const.integer.int_|48|_0') | $spl2),
  inference(superposition,[],[f31949,f49163])).
fof(f49163,plain,(
  ( ! [X50] : (s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X50),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl2),
  inference(forward_demodulation,[],[f49127,f721])).
fof(f49127,plain,(
  ( ! [X50] : (s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X50),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl2),
  inference(trivial_inequality_removal,[],[f49071])).
fof(f49071,plain,(
  ( ! [X50] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X50),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl2),
  inference(superposition,[],[f34115,f517])).
fof(f517,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(cnf_transformation,[],[f275])).
fof(f275,plain,(
  ! [X0] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X6] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X6),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.integer.INT_LE_MUL','thm.integer.INT_MUL_RZERO')).
fof(f34115,plain,(
  ( ! [X4] : (s('type.integer.int',X4) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',X4)) ) | $spl2),
  inference(superposition,[],[f30207,f712])).
fof(f30207,plain,(
  ( ! [X203] : (s('type.integer.int',X203) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',X203)))) ) | $spl2),
  inference(subsumption_resolution,[],[f30206,f875])).
fof(f30206,plain,(
  ( ! [X203] : (s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int',sK13) | s('type.integer.int',X203) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',X203)))) ) | $spl2),
  inference(forward_demodulation,[],[f29936,f721])).
fof(f29936,plain,(
  ( ! [X203] : (s('type.integer.int',X203) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',X203))) | s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl2),
  inference(superposition,[],[f710,f875])).
fof(f31949,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | $spl2),
  inference(forward_demodulation,[],[f3028,f875])).
fof(f3028,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))),
  inference(resolution,[],[f742,f687])).
fof(f49031,plain,(
  $false | ($spl2 | $spl78)),
  inference(subsumption_resolution,[],[f48942,f48831])).
fof(f48831,plain,(
  ( ! [X8] : (p(s(bool,X8))) ) | ($spl2 | $spl78)),
  inference(subsumption_resolution,[],[f47516,f46912])).
fof(f46912,plain,(
  ( ! [X161] : (p(s(bool,X161)) | s(bool,X161) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl2),
  inference(forward_demodulation,[],[f2978,f875])).
fof(f47516,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X8))) ) | ($spl2 | $spl78)),
  inference(forward_demodulation,[],[f8902,f875])).
fof(f48942,plain,(
  ( ! [X298,X299] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X298),s('type.integer.int',X299))))) ) | ($spl2 | $spl78)),
  inference(resolution,[],[f48831,f689])).
fof(f689,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f490])).
fof(f490,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) & (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f326])).
fof(f326,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(flattening,[],[f325])).
fof(f325,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6,X7] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) <=> ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int',X6)))))),
  file('thm.integer.INT_LE_MUL','thm.integer.int_le')).
fof(f49030,plain,(
  $false | ($spl2 | $spl78)),
  inference(subsumption_resolution,[],[f48941,f48831])).
fof(f48941,plain,(
  ( ! [X296,X297] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X296),s('type.integer.int',X297))))) ) | ($spl2 | $spl78)),
  inference(resolution,[],[f48831,f685])).
fof(f685,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f487])).
fof(f487,plain,(
  ! [X0,X1] : ((p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) & (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f323])).
fof(f323,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(flattening,[],[f322])).
fof(f322,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f98])).
fof(f98,axiom,(
  ! [X6,X7] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) <=> p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int',X6)))))),
  file('thm.integer.INT_LE_MUL','thm.integer.INT_NOT_LE')).
fof(f49029,plain,(
  $false | ($spl2 | $spl78)),
  inference(subsumption_resolution,[],[f48940,f48831])).
fof(f48940,plain,(
  ( ! [X294,X295] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X294),s('type.integer.int',X295))))) ) | ($spl2 | $spl78)),
  inference(resolution,[],[f48831,f683])).
fof(f683,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f486])).
fof(f486,plain,(
  ! [X0,X1] : ((p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) & (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f321])).
fof(f321,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(flattening,[],[f320])).
fof(f320,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f91])).
fof(f91,axiom,(
  ! [X6,X7] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) <=> p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X7),s('type.integer.int',X6)))))),
  file('thm.integer.INT_LE_MUL','thm.integer.INT_NOT_LT')).
fof(f49028,plain,(
  $false | ($spl2 | $spl78)),
  inference(subsumption_resolution,[],[f48939,f48831])).
fof(f48939,plain,(
  ( ! [X292,X293] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X292),s('type.integer.int',X293))))) ) | ($spl2 | $spl78)),
  inference(resolution,[],[f48831,f681])).
fof(f681,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f389])).
fof(f389,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(ennf_transformation,[],[f319])).
fof(f319,plain,(
  ! [X0,X1] : ~(p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X6,X7] : ~(p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X7),s('type.integer.int',X6)))))),
  file('thm.integer.INT_LE_MUL','thm.integer.INT_LET_ANTISYM')).
fof(f49027,plain,(
  $false | ($spl2 | $spl78)),
  inference(subsumption_resolution,[],[f48938,f48831])).
fof(f48938,plain,(
  ( ! [X290,X291] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X290),s('type.integer.int',X291))))) ) | ($spl2 | $spl78)),
  inference(resolution,[],[f48831,f680])).
fof(f49026,plain,(
  $false | ($spl2 | $spl78)),
  inference(subsumption_resolution,[],[f48935,f48831])).
fof(f48935,plain,(
  ( ! [X282,X283] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X282),s('type.integer.int',X283))))) ) | ($spl2 | $spl78)),
  inference(resolution,[],[f48831,f648])).
fof(f648,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f378])).
fof(f378,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(ennf_transformation,[],[f300])).
fof(f300,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) => ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(flattening,[],[f299])).
fof(f299,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) => ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f193])).
fof(f193,axiom,(
  ! [X6,X7] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) => ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int',X6)))))),
  file('thm.integer.INT_LE_MUL','thm.integer.INT_LT_GT')).
fof(f49025,plain,(
  $false | ($spl2 | $spl78)),
  inference(subsumption_resolution,[],[f48934,f48831])).
fof(f48934,plain,(
  ( ! [X280,X281] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X280),s('type.integer.int',X281))))) ) | ($spl2 | $spl78)),
  inference(resolution,[],[f48831,f648])).
fof(f49024,plain,(
  $false | ($spl2 | $spl78)),
  inference(subsumption_resolution,[],[f48933,f48831])).
fof(f48933,plain,(
  ( ! [X279,X278] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X278),s('type.integer.int',X279))))) ) | ($spl2 | $spl78)),
  inference(resolution,[],[f48831,f617])).
fof(f617,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f373])).
fof(f373,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(ennf_transformation,[],[f289])).
fof(f289,plain,(
  ! [X0,X1] : ~(p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f142])).
fof(f142,axiom,(
  ! [X6,X7] : ~(p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int',X6)))))),
  file('thm.integer.INT_LE_MUL','thm.integer.INT_LT_ANTISYM')).
fof(f49023,plain,(
  $false | ($spl2 | $spl78)),
  inference(subsumption_resolution,[],[f48932,f48831])).
fof(f48932,plain,(
  ( ! [X277,X276] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X276),s('type.integer.int',X277))))) ) | ($spl2 | $spl78)),
  inference(resolution,[],[f48831,f617])).
fof(f49022,plain,(
  $false | ($spl2 | $spl78)),
  inference(subsumption_resolution,[],[f48925,f48831])).
fof(f48925,plain,(
  ( ! [X265,X264] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X264),s('type.integer.int',X265))))) ) | ($spl2 | $spl78)),
  inference(resolution,[],[f48831,f689])).
fof(f49019,plain,(
  $false | ($spl2 | $spl78)),
  inference(subsumption_resolution,[],[f48922,f48831])).
fof(f48922,plain,(
  ( ! [X259,X258] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X258),s('type.integer.int',X259))))) ) | ($spl2 | $spl78)),
  inference(resolution,[],[f48831,f685])).
fof(f49018,plain,(
  $false | ($spl2 | $spl78)),
  inference(subsumption_resolution,[],[f48921,f48831])).
fof(f48921,plain,(
  ( ! [X257,X256] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X256),s('type.integer.int',X257))))) ) | ($spl2 | $spl78)),
  inference(resolution,[],[f48831,f683])).
fof(f49017,plain,(
  $false | ($spl2 | $spl78)),
  inference(subsumption_resolution,[],[f48920,f48831])).
fof(f48920,plain,(
  ( ! [X255,X254] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X254),s('type.integer.int',X255))))) ) | ($spl2 | $spl78)),
  inference(resolution,[],[f48831,f681])).
fof(f49016,plain,(
  $false | ($spl2 | $spl78)),
  inference(subsumption_resolution,[],[f48919,f48831])).
fof(f48919,plain,(
  ( ! [X253,X252] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X252),s('type.integer.int',X253))))) ) | ($spl2 | $spl78)),
  inference(resolution,[],[f48831,f680])).
fof(f48928,plain,(
  $false | ($spl2 | $spl78)),
  inference(resolution,[],[f48831,f30415])).
fof(f30415,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) | $spl2),
  inference(forward_demodulation,[],[f30414,f514])).
fof(f514,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f272])).
fof(f272,plain,(
  ! [X0,X1] : s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0)))),
  inference(rectify,[],[f73])).
fof(f73,axiom,(
  ! [X7,X6] : s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X6),s('type.integer.int',X7))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X7),s('type.integer.int',X6)))),
  file('thm.integer.INT_LE_MUL','thm.integer.INT_MUL_SYM')).
fof(f30414,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))))) | $spl2),
  inference(forward_demodulation,[],[f742,f875])).
fof(f48929,plain,(
  $false | ($spl2 | $spl21 | $spl78)),
  inference(resolution,[],[f48831,f1013])).
fof(f48948,plain,(
  $false | ($spl2 | $spl78)),
  inference(resolution,[],[f48831,f667])).
fof(f48951,plain,(
  $false | ($spl2 | $spl78)),
  inference(resolution,[],[f48831,f30559])).
fof(f30559,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) | $spl2),
  inference(resolution,[],[f30415,f692])).
fof(f692,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f492])).
fof(f48952,plain,(
  $false | ($spl2 | $spl78)),
  inference(resolution,[],[f48831,f889])).
fof(f48736,plain,(
  ( ! [X231,X229,X230] : (s(bool,X229) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X230)) | p(s(bool,X231)) | s(bool,X229) != s(bool,X231)) ) | $spl2),
  inference(resolution,[],[f46912,f735])).
fof(f48666,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl2),
  inference(resolution,[],[f46912,f526])).
fof(f48663,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X16) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl2),
  inference(resolution,[],[f46912,f525])).
fof(f48662,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl2),
  inference(resolution,[],[f46912,f524])).
fof(f48661,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl2),
  inference(resolution,[],[f46912,f522])).
fof(f48659,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl2),
  inference(resolution,[],[f46912,f520])).
fof(f48584,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | $spl188),
  inference(cnf_transformation,[],[f48584_D])).
fof(f48584_D,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f48597,plain,(
  ( ! [X16] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | ~p(s(bool,X16))) ) | ($spl2 | $spl24)),
  inference(resolution,[],[f46738,f31169])).
fof(f31169,plain,(
  ( ! [X134,X133] : (sP9(X133,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)),X134) | ~p(s(bool,X134))) ) | $spl24),
  inference(resolution,[],[f1017,f610])).
fof(f46738,plain,(
  ( ! [X212,X213,X214] : (~sP9(X213,X212,X214) | s(bool,X212) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl2),
  inference(resolution,[],[f44363,f614])).
fof(f44363,plain,(
  ( ! [X154] : (~p(s(bool,X154)) | s(bool,X154) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl2),
  inference(forward_demodulation,[],[f833,f875])).
fof(f48596,plain,(
  ( ! [X15] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | ~p(s(bool,X15))) ) | ($spl2 | $spl24)),
  inference(resolution,[],[f46738,f31172])).
fof(f31172,plain,(
  ( ! [X140,X139] : (sP9(X139,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)),X140) | ~p(s(bool,X139))) ) | $spl24),
  inference(resolution,[],[f1017,f611])).
fof(f48277,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | $spl184),
  inference(cnf_transformation,[],[f48277_D])).
fof(f48277_D,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f48593,plain,(
  ( ! [X12] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | ~p(s(bool,X12))) ) | $spl2),
  inference(resolution,[],[f46738,f957])).
fof(f48592,plain,(
  ( ! [X11] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | ~p(s(bool,X11))) ) | $spl2),
  inference(resolution,[],[f46738,f960])).
fof(f48582,plain,(
  ( ! [X19] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | ~p(s(bool,X19))) ) | ($spl2 | $spl24)),
  inference(resolution,[],[f46731,f31161])).
fof(f31161,plain,(
  ( ! [X118,X117] : (sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)),X117,X118) | ~p(s(bool,X118))) ) | $spl24),
  inference(resolution,[],[f1017,f600])).
fof(f46731,plain,(
  ( ! [X191,X192,X193] : (~sP8(X191,X192,X193) | s(bool,X191) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl2),
  inference(resolution,[],[f44363,f605])).
fof(f48581,plain,(
  ( ! [X18] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | ~p(s(bool,X18))) ) | ($spl2 | $spl24)),
  inference(resolution,[],[f46731,f31163])).
fof(f31163,plain,(
  ( ! [X121,X122] : (sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)),X121,X122) | ~p(s(bool,X121))) ) | $spl24),
  inference(resolution,[],[f1017,f601])).
fof(f48578,plain,(
  ( ! [X15] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | ~p(s(bool,X15))) ) | $spl2),
  inference(resolution,[],[f46731,f949])).
fof(f48577,plain,(
  ( ! [X14] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | ~p(s(bool,X14))) ) | $spl2),
  inference(resolution,[],[f46731,f951])).
fof(f48279,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | $spl186),
  inference(cnf_transformation,[],[f48279_D])).
fof(f48279_D,plain,(
  s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f48501,plain,(
  ( ! [X16] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | ~p(s(bool,X16))) ) | $spl24),
  inference(resolution,[],[f44551,f31169])).
fof(f48500,plain,(
  ( ! [X15] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | ~p(s(bool,X15))) ) | $spl24),
  inference(resolution,[],[f44551,f31172])).
fof(f48499,plain,(
  ( ! [X14] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | ~p(s(bool,X14))) ) | $spl2),
  inference(resolution,[],[f44551,f29692])).
fof(f29692,plain,(
  ( ! [X130,X131] : (sP9(X130,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X131) | ~p(s(bool,X131))) ) | $spl2),
  inference(backward_demodulation,[],[f875,f817])).
fof(f48498,plain,(
  ( ! [X13] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | ~p(s(bool,X13))) ) | $spl2),
  inference(resolution,[],[f44551,f29695])).
fof(f29695,plain,(
  ( ! [X136,X137] : (sP9(X136,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X137) | ~p(s(bool,X136))) ) | $spl2),
  inference(backward_demodulation,[],[f875,f820])).
fof(f48486,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f48397,f883])).
fof(f48397,plain,(
  ( ! [X310,X309] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X309),s('type.integer.int',X310))))) ) | $spl16),
  inference(resolution,[],[f883,f699])).
fof(f699,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f332])).
fof(f332,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f92])).
fof(f92,axiom,(
  ! [X6,X7] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X7),s('type.integer.int',X6)))))),
  file('thm.integer.INT_LE_MUL','thm.integer.INT_LTE_TOTAL')).
fof(f48485,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f48396,f883])).
fof(f48396,plain,(
  ( ! [X308,X307] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X307),s('type.integer.int',X308))))) ) | $spl16),
  inference(resolution,[],[f883,f698])).
fof(f698,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f331])).
fof(f331,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f95])).
fof(f95,axiom,(
  ! [X6,X7] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int',X6)))))),
  file('thm.integer.INT_LE_MUL','thm.integer.INT_LET_TOTAL')).
fof(f48484,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f48393,f883])).
fof(f48393,plain,(
  ( ! [X302,X301] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X301),s('type.integer.int',X302))))) ) | $spl16),
  inference(resolution,[],[f883,f690])).
fof(f690,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f490])).
fof(f48483,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f48392,f883])).
fof(f48392,plain,(
  ( ! [X300,X299] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X299),s('type.integer.int',X300))))) ) | $spl16),
  inference(resolution,[],[f883,f684])).
fof(f684,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f487])).
fof(f48482,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f48391,f883])).
fof(f48391,plain,(
  ( ! [X298,X297] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X297),s('type.integer.int',X298))))) ) | $spl16),
  inference(resolution,[],[f883,f682])).
fof(f682,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f486])).
fof(f48479,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f48383,f883])).
fof(f48383,plain,(
  ( ! [X283] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X283))))))) ) | $spl16),
  inference(resolution,[],[f883,f701])).
fof(f48478,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f48382,f883])).
fof(f48382,plain,(
  ( ! [X282] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X282))))) ) | $spl16),
  inference(resolution,[],[f883,f701])).
fof(f48477,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f48380,f883])).
fof(f48380,plain,(
  ( ! [X280,X281] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X280),s('type.integer.int',X281))))) ) | $spl16),
  inference(resolution,[],[f883,f700])).
fof(f48476,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f48379,f883])).
fof(f48379,plain,(
  ( ! [X279,X278] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X278),s('type.integer.int',X279))))) ) | $spl16),
  inference(resolution,[],[f883,f700])).
fof(f48475,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f48378,f883])).
fof(f48378,plain,(
  ( ! [X277,X276] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X276),s('type.integer.int',X277))))) ) | $spl16),
  inference(resolution,[],[f883,f699])).
fof(f48474,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f48377,f883])).
fof(f48377,plain,(
  ( ! [X275,X274] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X274),s('type.integer.int',X275))))) ) | $spl16),
  inference(resolution,[],[f883,f698])).
fof(f48473,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f48372,f883])).
fof(f48372,plain,(
  ( ! [X265,X264] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X264),s('type.integer.int',X265))))) ) | $spl16),
  inference(resolution,[],[f883,f690])).
fof(f48472,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f48369,f883])).
fof(f48369,plain,(
  ( ! [X259,X258] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X258),s('type.integer.int',X259))))) ) | $spl16),
  inference(resolution,[],[f883,f684])).
fof(f48471,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f48368,f883])).
fof(f48368,plain,(
  ( ! [X257,X256] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X256),s('type.integer.int',X257))))) ) | $spl16),
  inference(resolution,[],[f883,f682])).
fof(f48381,plain,(
  $false | $spl16),
  inference(resolution,[],[f883,f704])).
fof(f48384,plain,(
  $false | ($spl2 | $spl16)),
  inference(resolution,[],[f883,f30565])).
fof(f30565,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl2),
  inference(resolution,[],[f30415,f700])).
fof(f48385,plain,(
  $false | $spl16),
  inference(resolution,[],[f883,f744])).
fof(f48401,plain,(
  $false | ($spl16 | $spl42)),
  inference(resolution,[],[f883,f2740])).
fof(f48404,plain,(
  $false | ($spl16 | $spl24)),
  inference(resolution,[],[f883,f1017])).
fof(f48275,plain,(
  ( ! [X19] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | ~p(s(bool,X19))) ) | $spl24),
  inference(resolution,[],[f44544,f31161])).
fof(f48274,plain,(
  ( ! [X18] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | ~p(s(bool,X18))) ) | $spl24),
  inference(resolution,[],[f44544,f31163])).
fof(f48273,plain,(
  ( ! [X17] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | ~p(s(bool,X17))) ) | $spl2),
  inference(resolution,[],[f44544,f29688])).
fof(f29688,plain,(
  ( ! [X114,X115] : (sP8('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X114,X115) | ~p(s(bool,X115))) ) | $spl2),
  inference(backward_demodulation,[],[f875,f809])).
fof(f48272,plain,(
  ( ! [X16] : (s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | ~p(s(bool,X16))) ) | $spl2),
  inference(resolution,[],[f44544,f29690])).
fof(f29690,plain,(
  ( ! [X118,X119] : (sP8('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X118,X119) | ~p(s(bool,X118))) ) | $spl2),
  inference(backward_demodulation,[],[f875,f811])).
fof(f48064,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl182),
  inference(cnf_transformation,[],[f48064_D])).
fof(f48064_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f47990,plain,(
  ( ! [X243,X241,X242] : (s(bool,X241) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | p(s(bool,X242)) | ~p(s(bool,X243)) | s(bool,X241) != s(bool,X243)) ) | $spl24),
  inference(resolution,[],[f31185,f735])).
fof(f31185,plain,(
  ( ! [X157] : (~p(s(bool,X157)) | s(bool,X157) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl24),
  inference(resolution,[],[f1017,f647])).
fof(f47916,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl24),
  inference(resolution,[],[f31185,f526])).
fof(f47914,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) ) | $spl24),
  inference(resolution,[],[f31185,f525])).
fof(f47912,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl24),
  inference(resolution,[],[f31185,f524])).
fof(f47911,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl24),
  inference(resolution,[],[f31185,f522])).
fof(f47909,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl24),
  inference(resolution,[],[f31185,f520])).
fof(f47819,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl180),
  inference(cnf_transformation,[],[f47819_D])).
fof(f47819_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f47742,plain,(
  ( ! [X231,X229,X230] : (s(bool,X229) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X230)) | p(s(bool,X231)) | s(bool,X229) != s(bool,X231)) ) | $spl21),
  inference(resolution,[],[f31050,f735])).
fof(f31050,plain,(
  ( ! [X159] : (p(s(bool,X159)) | s(bool,X159) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl21),
  inference(resolution,[],[f1013,f644])).
fof(f47672,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) ) | $spl21),
  inference(resolution,[],[f31050,f526])).
fof(f47669,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X16) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) ) | $spl21),
  inference(resolution,[],[f31050,f525])).
fof(f47668,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl21),
  inference(resolution,[],[f31050,f524])).
fof(f47667,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) ) | $spl21),
  inference(resolution,[],[f31050,f522])).
fof(f47665,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl21),
  inference(resolution,[],[f31050,f520])).
fof(f47357,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) | $spl178),
  inference(cnf_transformation,[],[f47357_D])).
fof(f47357_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X0) != s(bool,X2) | p(s(bool,X2))) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f47280,plain,(
  ( ! [X231,X229,X230] : (s(bool,X229) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X230)) | p(s(bool,X231)) | s(bool,X229) != s(bool,X231)) )),
  inference(resolution,[],[f4212,f735])).
fof(f47210,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X29) | s(bool,X30) != s(bool,X31)) )),
  inference(resolution,[],[f4212,f526])).
fof(f47207,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,X16) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X16) != s(bool,X19)) )),
  inference(resolution,[],[f4212,f525])).
fof(f47206,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f4212,f524])).
fof(f47205,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X9) | s(bool,X10) != s(bool,X11)) )),
  inference(resolution,[],[f4212,f522])).
fof(f47203,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f4212,f520])).
fof(f47122,plain,(
  $false | ($spl2 | $spl58)),
  inference(subsumption_resolution,[],[f47033,f46917])).
fof(f46917,plain,(
  ( ! [X2] : (~p(s(bool,X2))) ) | ($spl2 | $spl58)),
  inference(subsumption_resolution,[],[f46916,f44363])).
fof(f46916,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X2))) ) | ($spl2 | $spl58)),
  inference(forward_demodulation,[],[f3581,f875])).
fof(f47033,plain,(
  ( ! [X308,X307] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X307),s('type.integer.int',X308))))) ) | ($spl2 | $spl58)),
  inference(resolution,[],[f46917,f699])).
fof(f47121,plain,(
  $false | ($spl2 | $spl58)),
  inference(subsumption_resolution,[],[f47032,f46917])).
fof(f47032,plain,(
  ( ! [X306,X305] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X305),s('type.integer.int',X306))))) ) | ($spl2 | $spl58)),
  inference(resolution,[],[f46917,f698])).
fof(f47120,plain,(
  $false | ($spl2 | $spl58)),
  inference(subsumption_resolution,[],[f47029,f46917])).
fof(f47029,plain,(
  ( ! [X300,X299] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X299),s('type.integer.int',X300))))) ) | ($spl2 | $spl58)),
  inference(resolution,[],[f46917,f690])).
fof(f47119,plain,(
  $false | ($spl2 | $spl58)),
  inference(subsumption_resolution,[],[f47028,f46917])).
fof(f47028,plain,(
  ( ! [X298,X297] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X297),s('type.integer.int',X298))))) ) | ($spl2 | $spl58)),
  inference(resolution,[],[f46917,f684])).
fof(f47118,plain,(
  $false | ($spl2 | $spl58)),
  inference(subsumption_resolution,[],[f47027,f46917])).
fof(f47027,plain,(
  ( ! [X296,X295] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X295),s('type.integer.int',X296))))) ) | ($spl2 | $spl58)),
  inference(resolution,[],[f46917,f682])).
fof(f47115,plain,(
  $false | ($spl2 | $spl58)),
  inference(subsumption_resolution,[],[f47019,f46917])).
fof(f47019,plain,(
  ( ! [X281] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X281))))))) ) | ($spl2 | $spl58)),
  inference(resolution,[],[f46917,f701])).
fof(f47114,plain,(
  $false | ($spl2 | $spl58)),
  inference(subsumption_resolution,[],[f47018,f46917])).
fof(f47018,plain,(
  ( ! [X280] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X280))))) ) | ($spl2 | $spl58)),
  inference(resolution,[],[f46917,f701])).
fof(f47113,plain,(
  $false | ($spl2 | $spl58)),
  inference(subsumption_resolution,[],[f47016,f46917])).
fof(f47016,plain,(
  ( ! [X279,X278] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X278),s('type.integer.int',X279))))) ) | ($spl2 | $spl58)),
  inference(resolution,[],[f46917,f700])).
fof(f47112,plain,(
  $false | ($spl2 | $spl58)),
  inference(subsumption_resolution,[],[f47015,f46917])).
fof(f47015,plain,(
  ( ! [X277,X276] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X276),s('type.integer.int',X277))))) ) | ($spl2 | $spl58)),
  inference(resolution,[],[f46917,f700])).
fof(f47111,plain,(
  $false | ($spl2 | $spl58)),
  inference(subsumption_resolution,[],[f47014,f46917])).
fof(f47014,plain,(
  ( ! [X275,X274] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X274),s('type.integer.int',X275))))) ) | ($spl2 | $spl58)),
  inference(resolution,[],[f46917,f699])).
fof(f47110,plain,(
  $false | ($spl2 | $spl58)),
  inference(subsumption_resolution,[],[f47013,f46917])).
fof(f47013,plain,(
  ( ! [X273,X272] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X272),s('type.integer.int',X273))))) ) | ($spl2 | $spl58)),
  inference(resolution,[],[f46917,f698])).
fof(f47109,plain,(
  $false | ($spl2 | $spl58)),
  inference(subsumption_resolution,[],[f47008,f46917])).
fof(f47008,plain,(
  ( ! [X263,X262] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X262),s('type.integer.int',X263))))) ) | ($spl2 | $spl58)),
  inference(resolution,[],[f46917,f690])).
fof(f47108,plain,(
  $false | ($spl2 | $spl58)),
  inference(subsumption_resolution,[],[f47005,f46917])).
fof(f47005,plain,(
  ( ! [X257,X256] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X256),s('type.integer.int',X257))))) ) | ($spl2 | $spl58)),
  inference(resolution,[],[f46917,f684])).
fof(f47107,plain,(
  $false | ($spl2 | $spl58)),
  inference(subsumption_resolution,[],[f47004,f46917])).
fof(f47004,plain,(
  ( ! [X255,X254] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X254),s('type.integer.int',X255))))) ) | ($spl2 | $spl58)),
  inference(resolution,[],[f46917,f682])).
fof(f47017,plain,(
  $false | ($spl2 | $spl58)),
  inference(resolution,[],[f46917,f704])).
fof(f47020,plain,(
  $false | ($spl2 | $spl58)),
  inference(resolution,[],[f46917,f30565])).
fof(f47021,plain,(
  $false | ($spl2 | $spl58)),
  inference(resolution,[],[f46917,f744])).
fof(f47037,plain,(
  $false | ($spl2 | $spl42 | $spl58)),
  inference(resolution,[],[f46917,f2740])).
fof(f47040,plain,(
  $false | ($spl2 | $spl24 | $spl58)),
  inference(resolution,[],[f46917,f1017])).
fof(f46751,plain,(
  ( ! [X243,X241,X242] : (s(bool,X241) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X242)) | ~p(s(bool,X243)) | s(bool,X241) != s(bool,X243)) ) | $spl2),
  inference(resolution,[],[f44363,f735])).
fof(f46677,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) ) | $spl2),
  inference(resolution,[],[f44363,f526])).
fof(f46675,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) ) | $spl2),
  inference(resolution,[],[f44363,f525])).
fof(f46673,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) ) | $spl2),
  inference(resolution,[],[f44363,f524])).
fof(f46672,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) ) | $spl2),
  inference(resolution,[],[f44363,f522])).
fof(f46670,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) ) | $spl2),
  inference(resolution,[],[f44363,f520])).
fof(f43836,plain,(
  ( ! [X35] : (s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl140),
  inference(cnf_transformation,[],[f43836_D])).
fof(f43836_D,plain,(
  ( ! [X35] : (s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f45198,plain,(
  ( ! [X35,X33,X34] : (sP4('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X33,X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl42),
  inference(forward_demodulation,[],[f45173,f721])).
fof(f45173,plain,(
  ( ! [X35,X33,X34] : (sP4('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')),X33,X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl42),
  inference(superposition,[],[f23422,f711])).
fof(f23422,plain,(
  ( ! [X109,X110] : (sP4('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')),X109,X110)) ) | $spl42),
  inference(resolution,[],[f2740,f579])).
fof(f43833,plain,(
  ( ! [X32] : (s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int',X32)))) ) | $spl136),
  inference(cnf_transformation,[],[f43833_D])).
fof(f43833_D,plain,(
  ( ! [X32] : (s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int',X32)))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f45196,plain,(
  ( ! [X30,X31,X32] : (sP4('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X30,X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int',X32)))) ) | $spl42),
  inference(forward_demodulation,[],[f45172,f721])).
fof(f45172,plain,(
  ( ! [X30,X31,X32] : (sP4('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')),X30,X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int',X32)))) ) | $spl42),
  inference(superposition,[],[f23422,f709])).
fof(f45138,plain,(
  ( ! [X35,X33,X34] : (sP4(X33,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl42),
  inference(forward_demodulation,[],[f45114,f721])).
fof(f45114,plain,(
  ( ! [X35,X33,X34] : (sP4(X33,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')),X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl42),
  inference(superposition,[],[f23421,f711])).
fof(f23421,plain,(
  ( ! [X107,X108] : (sP4(X107,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')),X108)) ) | $spl42),
  inference(resolution,[],[f2740,f578])).
fof(f45136,plain,(
  ( ! [X30,X31,X32] : (sP4(X30,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int',X32)))) ) | $spl42),
  inference(forward_demodulation,[],[f45113,f721])).
fof(f45113,plain,(
  ( ! [X30,X31,X32] : (sP4(X30,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')),X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int',X32)))) ) | $spl42),
  inference(superposition,[],[f23421,f709])).
fof(f45079,plain,(
  ( ! [X35,X33,X34] : (sP4(X33,X34,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl42),
  inference(forward_demodulation,[],[f45055,f721])).
fof(f45055,plain,(
  ( ! [X35,X33,X34] : (sP4(X33,X34,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0'))) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl42),
  inference(superposition,[],[f23420,f711])).
fof(f23420,plain,(
  ( ! [X105,X106] : (sP4(X105,X106,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl42),
  inference(resolution,[],[f2740,f577])).
fof(f45077,plain,(
  ( ! [X30,X31,X32] : (sP4(X30,X31,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int',X32)))) ) | $spl42),
  inference(forward_demodulation,[],[f45054,f721])).
fof(f45054,plain,(
  ( ! [X30,X31,X32] : (sP4(X30,X31,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0'))) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int',X32)))) ) | $spl42),
  inference(superposition,[],[f23420,f709])).
fof(f45020,plain,(
  ( ! [X35,X33,X34] : (sP5(X33,X34,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl42),
  inference(forward_demodulation,[],[f44997,f721])).
fof(f44997,plain,(
  ( ! [X35,X33,X34] : (sP5(X33,X34,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0'))) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl42),
  inference(superposition,[],[f23419,f711])).
fof(f23419,plain,(
  ( ! [X103,X104] : (sP5(X103,X104,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl42),
  inference(resolution,[],[f2740,f573])).
fof(f45018,plain,(
  ( ! [X30,X31,X32] : (sP5(X30,X31,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int',X32)))) ) | $spl42),
  inference(forward_demodulation,[],[f44996,f721])).
fof(f44996,plain,(
  ( ! [X30,X31,X32] : (sP5(X30,X31,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0'))) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int',X32)))) ) | $spl42),
  inference(superposition,[],[f23419,f709])).
fof(f44962,plain,(
  ( ! [X35,X33,X34] : (sP6(X33,X34,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl42),
  inference(forward_demodulation,[],[f44939,f721])).
fof(f44939,plain,(
  ( ! [X35,X33,X34] : (sP6(X33,X34,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0'))) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl42),
  inference(superposition,[],[f23416,f711])).
fof(f23416,plain,(
  ( ! [X97,X98] : (sP6(X97,X98,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl42),
  inference(resolution,[],[f2740,f569])).
fof(f44960,plain,(
  ( ! [X30,X31,X32] : (sP6(X30,X31,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int',X32)))) ) | $spl42),
  inference(forward_demodulation,[],[f44938,f721])).
fof(f44938,plain,(
  ( ! [X30,X31,X32] : (sP6(X30,X31,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0'))) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int',X32)))) ) | $spl42),
  inference(superposition,[],[f23416,f709])).
fof(f44904,plain,(
  ( ! [X35,X33,X34] : (sP2(X33,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl42),
  inference(forward_demodulation,[],[f44881,f721])).
fof(f44881,plain,(
  ( ! [X35,X33,X34] : (sP2(X33,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')),X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl42),
  inference(superposition,[],[f23403,f711])).
fof(f23403,plain,(
  ( ! [X72,X71] : (sP2(X71,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')),X72)) ) | $spl42),
  inference(resolution,[],[f2740,f553])).
fof(f44902,plain,(
  ( ! [X30,X31,X32] : (sP2(X30,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int',X32)))) ) | $spl42),
  inference(forward_demodulation,[],[f44880,f721])).
fof(f44880,plain,(
  ( ! [X30,X31,X32] : (sP2(X30,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')),X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int',X32)))) ) | $spl42),
  inference(superposition,[],[f23403,f709])).
fof(f44845,plain,(
  ( ! [X35,X33,X34] : (sP2(X33,X34,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl42),
  inference(forward_demodulation,[],[f44822,f721])).
fof(f44822,plain,(
  ( ! [X35,X33,X34] : (sP2(X33,X34,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0'))) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl42),
  inference(superposition,[],[f23402,f711])).
fof(f23402,plain,(
  ( ! [X70,X69] : (sP2(X69,X70,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl42),
  inference(resolution,[],[f2740,f552])).
fof(f44843,plain,(
  ( ! [X30,X31,X32] : (sP2(X30,X31,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int',X32)))) ) | $spl42),
  inference(forward_demodulation,[],[f44821,f721])).
fof(f44821,plain,(
  ( ! [X30,X31,X32] : (sP2(X30,X31,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0'))) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int',X32)))) ) | $spl42),
  inference(superposition,[],[f23402,f709])).
fof(f44717,plain,(
  ( ! [X35,X33,X34] : (sP0(X33,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl42),
  inference(forward_demodulation,[],[f44694,f721])).
fof(f44694,plain,(
  ( ! [X35,X33,X34] : (sP0(X33,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')),X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl42),
  inference(superposition,[],[f23389,f711])).
fof(f23389,plain,(
  ( ! [X43,X44] : (sP0(X43,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')),X44)) ) | $spl42),
  inference(resolution,[],[f2740,f536])).
fof(f44715,plain,(
  ( ! [X30,X31,X32] : (sP0(X30,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int',X32)))) ) | $spl42),
  inference(forward_demodulation,[],[f44693,f721])).
fof(f44693,plain,(
  ( ! [X30,X31,X32] : (sP0(X30,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')),X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int',X32)))) ) | $spl42),
  inference(superposition,[],[f23389,f709])).
fof(f44636,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) | $spl158),
  inference(cnf_transformation,[],[f44636_D])).
fof(f44636_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X0) != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f44564,plain,(
  ( ! [X243,X241,X242] : (s(bool,X241) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | p(s(bool,X242)) | ~p(s(bool,X243)) | s(bool,X241) != s(bool,X243)) )),
  inference(resolution,[],[f973,f735])).
fof(f44490,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,X31)) )),
  inference(resolution,[],[f973,f526])).
fof(f44488,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,X20) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X20) != s(bool,X22)) )),
  inference(resolution,[],[f973,f525])).
fof(f44486,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X12) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X12) != s(bool,X14)) )),
  inference(resolution,[],[f973,f524])).
fof(f44485,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X8) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f973,f522])).
fof(f44483,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X0) != s(bool,X2)) )),
  inference(resolution,[],[f973,f520])).
fof(f44480,plain,(
  $false | ($spl2 | $spl88 | $spl97)),
  inference(subsumption_resolution,[],[f44479,f704])).
fof(f44479,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl2 | $spl88 | $spl97)),
  inference(forward_demodulation,[],[f44478,f17521])).
fof(f44478,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl2 | $spl97)),
  inference(forward_demodulation,[],[f19260,f875])).
fof(f44345,plain,(
  ( ! [X35,X33,X34] : (~sP9(X33,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl2),
  inference(forward_demodulation,[],[f44331,f721])).
fof(f44331,plain,(
  ( ! [X35,X33,X34] : (~sP9(X33,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl2),
  inference(superposition,[],[f30635,f711])).
fof(f30635,plain,(
  ( ! [X146,X145] : (~sP9(X145,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))),X146)) ) | $spl2),
  inference(resolution,[],[f30415,f614])).
fof(f44343,plain,(
  ( ! [X30,X31,X32] : (~sP9(X30,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int',X32)))) ) | $spl2),
  inference(forward_demodulation,[],[f44330,f721])).
fof(f44330,plain,(
  ( ! [X30,X31,X32] : (~sP9(X30,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int',X32)))) ) | $spl2),
  inference(superposition,[],[f30635,f709])).
fof(f44282,plain,(
  ( ! [X35,X33,X34] : (~sP8('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X33,X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl2),
  inference(forward_demodulation,[],[f44268,f721])).
fof(f44268,plain,(
  ( ! [X35,X33,X34] : (~sP8('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X33,X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl2),
  inference(superposition,[],[f30628,f711])).
fof(f30628,plain,(
  ( ! [X132,X131] : (~sP8('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))),X131,X132)) ) | $spl2),
  inference(resolution,[],[f30415,f605])).
fof(f44280,plain,(
  ( ! [X30,X31,X32] : (~sP8('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X30,X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int',X32)))) ) | $spl2),
  inference(forward_demodulation,[],[f44267,f721])).
fof(f44267,plain,(
  ( ! [X30,X31,X32] : (~sP8('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X30,X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int',X32)))) ) | $spl2),
  inference(superposition,[],[f30628,f709])).
fof(f44219,plain,(
  ( ! [X35,X33,X34] : (sP5('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X33,X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl2),
  inference(forward_demodulation,[],[f44202,f721])).
fof(f44202,plain,(
  ( ! [X35,X33,X34] : (sP5('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X33,X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl2),
  inference(superposition,[],[f30614,f711])).
fof(f30614,plain,(
  ( ! [X103,X104] : (sP5('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))),X103,X104)) ) | $spl2),
  inference(resolution,[],[f30415,f575])).
fof(f44217,plain,(
  ( ! [X30,X31,X32] : (sP5('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X30,X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int',X32)))) ) | $spl2),
  inference(forward_demodulation,[],[f44201,f721])).
fof(f44201,plain,(
  ( ! [X30,X31,X32] : (sP5('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X30,X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int',X32)))) ) | $spl2),
  inference(superposition,[],[f30614,f709])).
fof(f44159,plain,(
  ( ! [X35,X33,X34] : (sP5(X33,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl2),
  inference(forward_demodulation,[],[f44142,f721])).
fof(f44142,plain,(
  ( ! [X35,X33,X34] : (sP5(X33,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl2),
  inference(superposition,[],[f30613,f711])).
fof(f30613,plain,(
  ( ! [X101,X102] : (sP5(X101,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))),X102)) ) | $spl2),
  inference(resolution,[],[f30415,f574])).
fof(f44157,plain,(
  ( ! [X30,X31,X32] : (sP5(X30,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int',X32)))) ) | $spl2),
  inference(forward_demodulation,[],[f44141,f721])).
fof(f44141,plain,(
  ( ! [X30,X31,X32] : (sP5(X30,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int',X32)))) ) | $spl2),
  inference(superposition,[],[f30613,f709])).
fof(f44099,plain,(
  ( ! [X35,X33,X34] : (sP6('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X33,X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl2),
  inference(forward_demodulation,[],[f44082,f721])).
fof(f44082,plain,(
  ( ! [X35,X33,X34] : (sP6('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X33,X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl2),
  inference(superposition,[],[f30611,f711])).
fof(f30611,plain,(
  ( ! [X97,X98] : (sP6('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))),X97,X98)) ) | $spl2),
  inference(resolution,[],[f30415,f571])).
fof(f44097,plain,(
  ( ! [X30,X31,X32] : (sP6('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X30,X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int',X32)))) ) | $spl2),
  inference(forward_demodulation,[],[f44081,f721])).
fof(f44081,plain,(
  ( ! [X30,X31,X32] : (sP6('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X30,X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int',X32)))) ) | $spl2),
  inference(superposition,[],[f30611,f709])).
fof(f44039,plain,(
  ( ! [X35,X33,X34] : (sP6(X33,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl2),
  inference(forward_demodulation,[],[f44023,f721])).
fof(f44023,plain,(
  ( ! [X35,X33,X34] : (sP6(X33,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl2),
  inference(superposition,[],[f30610,f711])).
fof(f30610,plain,(
  ( ! [X95,X96] : (sP6(X95,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))),X96)) ) | $spl2),
  inference(resolution,[],[f30415,f570])).
fof(f44037,plain,(
  ( ! [X30,X31,X32] : (sP6(X30,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int',X32)))) ) | $spl2),
  inference(forward_demodulation,[],[f44022,f721])).
fof(f44022,plain,(
  ( ! [X30,X31,X32] : (sP6(X30,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int',X32)))) ) | $spl2),
  inference(superposition,[],[f30610,f709])).
fof(f43980,plain,(
  ( ! [X35,X33,X34] : (sP2('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X33,X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl2),
  inference(forward_demodulation,[],[f43964,f721])).
fof(f43964,plain,(
  ( ! [X35,X33,X34] : (sP2('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X33,X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl2),
  inference(superposition,[],[f30598,f711])).
fof(f30598,plain,(
  ( ! [X72,X71] : (sP2('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))),X71,X72)) ) | $spl2),
  inference(resolution,[],[f30415,f554])).
fof(f43978,plain,(
  ( ! [X30,X31,X32] : (sP2('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X30,X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int',X32)))) ) | $spl2),
  inference(forward_demodulation,[],[f43963,f721])).
fof(f43963,plain,(
  ( ! [X30,X31,X32] : (sP2('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X30,X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int',X32)))) ) | $spl2),
  inference(superposition,[],[f30598,f709])).
fof(f43920,plain,(
  ( ! [X35,X33,X34] : (sP0('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X33,X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl2),
  inference(forward_demodulation,[],[f43905,f721])).
fof(f43905,plain,(
  ( ! [X35,X33,X34] : (sP0('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X33,X34) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl2),
  inference(superposition,[],[f30582,f711])).
fof(f30582,plain,(
  ( ! [X39,X40] : (sP0('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))),X39,X40)) ) | $spl2),
  inference(resolution,[],[f30415,f537])).
fof(f43918,plain,(
  ( ! [X30,X31,X32] : (sP0('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')),X30,X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int',X32)))) ) | $spl2),
  inference(forward_demodulation,[],[f43904,f721])).
fof(f43904,plain,(
  ( ! [X30,X31,X32] : (sP0('const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X30,X31) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int',X32)))) ) | $spl2),
  inference(superposition,[],[f30582,f709])).
fof(f43857,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1))) ) | $spl138),
  inference(subsumption_resolution,[],[f43840,f704])).
fof(f43840,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))))) ) | $spl138),
  inference(resolution,[],[f43834,f534])).
fof(f43835,plain,(
  ( ! [X35,X33,X34] : (sP0(X33,X34,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl2),
  inference(forward_demodulation,[],[f43819,f721])).
fof(f43819,plain,(
  ( ! [X35,X33,X34] : (sP0(X33,X34,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X35) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X35),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl2),
  inference(superposition,[],[f30581,f711])).
fof(f30581,plain,(
  ( ! [X37,X38] : (sP0(X37,X38,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl2),
  inference(resolution,[],[f30415,f535])).
fof(f43832,plain,(
  ( ! [X30,X31,X32] : (sP0(X30,X31,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int',X32)))) ) | $spl2),
  inference(forward_demodulation,[],[f43818,f721])).
fof(f43818,plain,(
  ( ! [X30,X31,X32] : (sP0(X30,X31,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X32) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int',X32)))) ) | $spl2),
  inference(superposition,[],[f30581,f709])).
fof(f43769,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl134),
  inference(cnf_transformation,[],[f43769_D])).
fof(f43769_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f43715,plain,(
  ( ! [X156,X157] : (p(s(bool,X156)) | p(s(bool,X157)) | s(bool,X157) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl2),
  inference(resolution,[],[f30565,f735])).
fof(f43645,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X26) != s(bool,X27)) ) | $spl2),
  inference(resolution,[],[f30565,f526])).
fof(f43642,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl2),
  inference(resolution,[],[f30565,f525])).
fof(f43641,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl2),
  inference(resolution,[],[f30565,f524])).
fof(f43640,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X11) != s(bool,X12)) ) | $spl2),
  inference(resolution,[],[f30565,f522])).
fof(f43638,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl2),
  inference(resolution,[],[f30565,f520])).
fof(f43615,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl132),
  inference(cnf_transformation,[],[f43615_D])).
fof(f43615_D,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f43560,plain,(
  ( ! [X163,X164] : (p(s(bool,X163)) | ~p(s(bool,X164)) | s(bool,X164) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl2),
  inference(resolution,[],[f30559,f735])).
fof(f43486,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))))) | s(bool,X24) != s(bool,X26)) ) | $spl2),
  inference(resolution,[],[f30559,f526])).
fof(f43484,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl2),
  inference(resolution,[],[f30559,f525])).
fof(f43482,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl2),
  inference(resolution,[],[f30559,f524])).
fof(f43481,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))))) | s(bool,X9) != s(bool,X11)) ) | $spl2),
  inference(resolution,[],[f30559,f522])).
fof(f43479,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl2),
  inference(resolution,[],[f30559,f520])).
fof(f40693,plain,(
  $false | ($spl2 | $spl76 | $spl88)),
  inference(subsumption_resolution,[],[f33920,f17521])).
fof(f33920,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | ($spl2 | $spl76)),
  inference(forward_demodulation,[],[f29320,f875])).
fof(f29320,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | $spl76),
  inference(resolution,[],[f8102,f695])).
fof(f40097,plain,(
  $false | ($spl2 | $spl34 | $spl75)),
  inference(subsumption_resolution,[],[f40096,f2426])).
fof(f40096,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl2 | $spl75)),
  inference(forward_demodulation,[],[f8099,f875])).
fof(f8099,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl75),
  inference(cnf_transformation,[],[f8099_D])).
fof(f8099_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) <=> ~$spl75),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl75])])).
fof(f23506,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl102),
  inference(cnf_transformation,[],[f23506_D])).
fof(f23506_D,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f39302,plain,(
  ( ! [X158,X157] : (p(s(bool,X157)) | p(s(bool,X158)) | s(bool,X158) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl42),
  inference(resolution,[],[f2740,f735])).
fof(f39232,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X27) != s(bool,X28)) ) | $spl42),
  inference(resolution,[],[f2740,f526])).
fof(f39229,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X19) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl42),
  inference(resolution,[],[f2740,f525])).
fof(f39228,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl42),
  inference(resolution,[],[f2740,f524])).
fof(f39227,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X12) != s(bool,X13)) ) | $spl42),
  inference(resolution,[],[f2740,f522])).
fof(f39225,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl42),
  inference(resolution,[],[f2740,f520])).
fof(f37824,plain,(
  $false | ($spl2 | $spl23 | $spl116)),
  inference(subsumption_resolution,[],[f37504,f1014])).
fof(f37504,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | ($spl2 | $spl116)),
  inference(superposition,[],[f30210,f24467])).
fof(f30210,plain,(
  ( ! [X206] : (s('type.integer.int',X206) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X206),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl2),
  inference(subsumption_resolution,[],[f30209,f875])).
fof(f30209,plain,(
  ( ! [X206] : (s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int',sK13) | s('type.integer.int',X206) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X206),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl2),
  inference(forward_demodulation,[],[f29939,f721])).
fof(f29939,plain,(
  ( ! [X206] : (s('type.integer.int',X206) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X206),s('type.integer.int','const.integer.int_|48|_0'))) | s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl2),
  inference(superposition,[],[f712,f875])).
fof(f37823,plain,(
  $false | ($spl114 | $spl116)),
  inference(subsumption_resolution,[],[f37822,f37323])).
fof(f37323,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))) | $spl114),
  inference(forward_demodulation,[],[f37291,f721])).
fof(f37291,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))) | $spl114),
  inference(superposition,[],[f24465,f717])).
fof(f37822,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))) | $spl116),
  inference(forward_demodulation,[],[f37503,f721])).
fof(f37503,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))) | $spl116),
  inference(superposition,[],[f718,f24467])).
fof(f37821,plain,(
  $false | ($spl23 | $spl116)),
  inference(subsumption_resolution,[],[f37820,f1014])).
fof(f37820,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | $spl116),
  inference(forward_demodulation,[],[f37502,f721])).
fof(f37502,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl116),
  inference(superposition,[],[f719,f24467])).
fof(f37819,plain,(
  $false | ($spl114 | $spl116)),
  inference(subsumption_resolution,[],[f37501,f24465])).
fof(f37501,plain,(
  ( ! [X39] : (s('type.integer.int',X39) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X39),s('type.integer.int',sK12)))) ) | $spl116),
  inference(superposition,[],[f734,f24467])).
fof(f37818,plain,(
  $false | ($spl114 | $spl116)),
  inference(subsumption_resolution,[],[f37500,f24465])).
fof(f37500,plain,(
  ( ! [X38] : (s('type.integer.int',X38) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X38),s('type.integer.int',sK12)))) ) | $spl116),
  inference(superposition,[],[f734,f24467])).
fof(f37815,plain,(
  $false | ($spl23 | $spl116)),
  inference(subsumption_resolution,[],[f37814,f1014])).
fof(f37814,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | $spl116),
  inference(forward_demodulation,[],[f37766,f721])).
fof(f37766,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl116),
  inference(trivial_inequality_removal,[],[f37482])).
fof(f37482,plain,(
  ( ! [X4] : (s('type.integer.int',X4) != s('type.integer.int',X4) | s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl116),
  inference(superposition,[],[f709,f24467])).
fof(f37777,plain,(
  $false | ($spl2 | $spl23 | $spl116)),
  inference(subsumption_resolution,[],[f37460,f1014])).
fof(f37460,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | ($spl2 | $spl116)),
  inference(superposition,[],[f24467,f30210])).
fof(f37776,plain,(
  $false | ($spl114 | $spl116)),
  inference(subsumption_resolution,[],[f37775,f37323])).
fof(f37775,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))) | $spl116),
  inference(forward_demodulation,[],[f37459,f721])).
fof(f37459,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12))) | $spl116),
  inference(superposition,[],[f24467,f718])).
fof(f37774,plain,(
  $false | ($spl23 | $spl116)),
  inference(subsumption_resolution,[],[f37773,f1014])).
fof(f37773,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | $spl116),
  inference(forward_demodulation,[],[f37458,f721])).
fof(f37458,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl116),
  inference(superposition,[],[f24467,f719])).
fof(f37772,plain,(
  $false | ($spl114 | $spl116)),
  inference(subsumption_resolution,[],[f37457,f24465])).
fof(f37457,plain,(
  ( ! [X14] : (s('type.integer.int',X14) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X14),s('type.integer.int',sK12)))) ) | $spl116),
  inference(superposition,[],[f24467,f734])).
fof(f37771,plain,(
  $false | ($spl114 | $spl116)),
  inference(subsumption_resolution,[],[f37456,f24465])).
fof(f37456,plain,(
  ( ! [X13] : (s('type.integer.int',X13) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X13),s('type.integer.int',sK12)))) ) | $spl116),
  inference(superposition,[],[f24467,f734])).
fof(f37765,plain,(
  $false | ($spl114 | $spl116)),
  inference(trivial_inequality_removal,[],[f37505])).
fof(f37505,plain,(
  s('type.integer.int',sK12) != s('type.integer.int',sK12) | ($spl114 | $spl116)),
  inference(superposition,[],[f24465,f24467])).
fof(f36970,plain,(
  $false | ($spl23 | $spl112)),
  inference(subsumption_resolution,[],[f36969,f1014])).
fof(f36969,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | $spl112),
  inference(forward_demodulation,[],[f36920,f721])).
fof(f36920,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl112),
  inference(trivial_inequality_removal,[],[f36643])).
fof(f36643,plain,(
  ( ! [X6] : (s('type.integer.int',X6) != s('type.integer.int',X6) | s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl112),
  inference(superposition,[],[f711,f24463])).
fof(f36964,plain,(
  $false | ($spl2 | $spl23 | $spl112)),
  inference(subsumption_resolution,[],[f36636,f1014])).
fof(f36636,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | ($spl2 | $spl112)),
  inference(superposition,[],[f30207,f24463])).
fof(f36962,plain,(
  $false | ($spl23 | $spl112)),
  inference(subsumption_resolution,[],[f36961,f1014])).
fof(f36961,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | $spl112),
  inference(forward_demodulation,[],[f36634,f721])).
fof(f36634,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl112),
  inference(superposition,[],[f720,f24463])).
fof(f36926,plain,(
  $false | ($spl2 | $spl23 | $spl112)),
  inference(subsumption_resolution,[],[f36606,f1014])).
fof(f36606,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | ($spl2 | $spl112)),
  inference(superposition,[],[f24463,f30207])).
fof(f36925,plain,(
  $false | ($spl23 | $spl112)),
  inference(subsumption_resolution,[],[f36924,f1014])).
fof(f36924,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | $spl112),
  inference(forward_demodulation,[],[f36605,f721])).
fof(f36605,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl112),
  inference(superposition,[],[f24463,f720])).
fof(f2740,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl42),
  inference(cnf_transformation,[],[f2740_D])).
fof(f2740_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f2742,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) | $spl44),
  inference(cnf_transformation,[],[f2742_D])).
fof(f2742_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f33206,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl2),
  inference(forward_demodulation,[],[f33205,f721])).
fof(f33205,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl2),
  inference(forward_demodulation,[],[f33204,f514])).
fof(f33204,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))))) | $spl2),
  inference(forward_demodulation,[],[f33203,f514])).
fof(f33203,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))))) | $spl2),
  inference(forward_demodulation,[],[f33202,f721])).
fof(f33202,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))))) | $spl2),
  inference(forward_demodulation,[],[f33201,f716])).
fof(f33201,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))))) | $spl2),
  inference(subsumption_resolution,[],[f33172,f721])).
fof(f33172,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))))) | $spl2),
  inference(superposition,[],[f31949,f714])).
fof(f32566,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl2),
  inference(forward_demodulation,[],[f32565,f721])).
fof(f32565,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) | $spl2),
  inference(forward_demodulation,[],[f32564,f721])).
fof(f32564,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) | $spl2),
  inference(forward_demodulation,[],[f32563,f716])).
fof(f32563,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) | $spl2),
  inference(subsumption_resolution,[],[f32554,f721])).
fof(f32554,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) | $spl2),
  inference(superposition,[],[f30556,f714])).
fof(f30556,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | $spl2),
  inference(resolution,[],[f30415,f687])).
fof(f9553,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | p(s(bool,X8))) ) | $spl82),
  inference(cnf_transformation,[],[f9553_D])).
fof(f9553_D,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | p(s(bool,X8))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f31187,plain,(
  ( ! [X159,X160] : (p(s(bool,X159)) | p(s(bool,X160)) | s(bool,X160) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl24),
  inference(resolution,[],[f1017,f735])).
fof(f31117,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X29) != s(bool,X30)) ) | $spl24),
  inference(resolution,[],[f1017,f526])).
fof(f31114,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X21) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl24),
  inference(resolution,[],[f1017,f525])).
fof(f31113,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl24),
  inference(resolution,[],[f1017,f524])).
fof(f31112,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X14) != s(bool,X15)) ) | $spl24),
  inference(resolution,[],[f1017,f522])).
fof(f31110,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl24),
  inference(resolution,[],[f1017,f520])).
fof(f4797,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X2))) ) | $spl66),
  inference(cnf_transformation,[],[f4797_D])).
fof(f4797_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X2))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f31052,plain,(
  ( ! [X161,X162] : (p(s(bool,X161)) | ~p(s(bool,X162)) | s(bool,X162) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl21),
  inference(resolution,[],[f1013,f735])).
fof(f30978,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X22) != s(bool,X24)) ) | $spl21),
  inference(resolution,[],[f1013,f526])).
fof(f30976,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl21),
  inference(resolution,[],[f1013,f525])).
fof(f30974,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl21),
  inference(resolution,[],[f1013,f524])).
fof(f30973,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X7) != s(bool,X9)) ) | $spl21),
  inference(resolution,[],[f1013,f522])).
fof(f30971,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl21),
  inference(resolution,[],[f1013,f520])).
fof(f30713,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl2),
  inference(forward_demodulation,[],[f30712,f721])).
fof(f30712,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) | $spl2),
  inference(forward_demodulation,[],[f30711,f721])).
fof(f30711,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) | $spl2),
  inference(forward_demodulation,[],[f30710,f716])).
fof(f30710,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) | $spl2),
  inference(subsumption_resolution,[],[f30709,f704])).
fof(f30709,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) | $spl2),
  inference(forward_demodulation,[],[f30688,f721])).
fof(f30688,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) | $spl2),
  inference(superposition,[],[f30415,f714])).
fof(f2725,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl38),
  inference(cnf_transformation,[],[f2725_D])).
fof(f2725_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f30648,plain,(
  ( ! [X161,X162] : (p(s(bool,X161)) | ~p(s(bool,X162)) | s(bool,X162) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl2),
  inference(resolution,[],[f30415,f735])).
fof(f30574,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))))) | s(bool,X22) != s(bool,X24)) ) | $spl2),
  inference(resolution,[],[f30415,f526])).
fof(f30572,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl2),
  inference(resolution,[],[f30415,f525])).
fof(f30570,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl2),
  inference(resolution,[],[f30415,f524])).
fof(f30569,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))))) | s(bool,X7) != s(bool,X9)) ) | $spl2),
  inference(resolution,[],[f30415,f522])).
fof(f30567,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl2),
  inference(resolution,[],[f30415,f520])).
fof(f4256,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X4))) ) | $spl62),
  inference(cnf_transformation,[],[f4256_D])).
fof(f4256_D,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X4))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f30511,plain,(
  ( ! [X163,X164] : (p(s(bool,X163)) | ~p(s(bool,X164)) | s(bool,X164) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) )),
  inference(resolution,[],[f889,f735])).
fof(f30437,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X24) != s(bool,X26)) )),
  inference(resolution,[],[f889,f526])).
fof(f30435,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) )),
  inference(resolution,[],[f889,f525])).
fof(f30433,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) )),
  inference(resolution,[],[f889,f524])).
fof(f30432,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f889,f522])).
fof(f30430,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) )),
  inference(resolution,[],[f889,f520])).
fof(f1018,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | p(s(bool,X5))) ) | $spl26),
  inference(cnf_transformation,[],[f1018_D])).
fof(f1018_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | p(s(bool,X5))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f30373,plain,(
  ( ! [X156,X157] : (p(s(bool,X156)) | p(s(bool,X157)) | s(bool,X157) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f744,f735])).
fof(f30303,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X26) != s(bool,X27)) )),
  inference(resolution,[],[f744,f526])).
fof(f30300,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f744,f525])).
fof(f30299,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f744,f524])).
fof(f30298,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X11) != s(bool,X12)) )),
  inference(resolution,[],[f744,f522])).
fof(f30296,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f744,f520])).
fof(f30007,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f29700,f704])).
fof(f29700,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl1 | $spl2)),
  inference(backward_demodulation,[],[f875,f873])).
fof(f877,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | $spl4),
  inference(cnf_transformation,[],[f877_D])).
fof(f877_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f29665,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))))),
  inference(forward_demodulation,[],[f29664,f721])).
fof(f29664,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))))),
  inference(forward_demodulation,[],[f12089,f2886])).
fof(f12089,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13))))),
  inference(forward_demodulation,[],[f12088,f721])).
fof(f12088,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13))))),
  inference(forward_demodulation,[],[f12087,f2886])).
fof(f12087,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13))))),
  inference(subsumption_resolution,[],[f12086,f749])).
fof(f12086,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13))))),
  inference(forward_demodulation,[],[f12085,f721])).
fof(f12085,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13))))),
  inference(forward_demodulation,[],[f11781,f716])).
fof(f11781,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13))))),
  inference(superposition,[],[f714,f2886])).
fof(f29663,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))))),
  inference(forward_demodulation,[],[f29662,f721])).
fof(f29662,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))))),
  inference(forward_demodulation,[],[f29661,f2886])).
fof(f29661,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13))))),
  inference(subsumption_resolution,[],[f29660,f749])).
fof(f29660,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13))))),
  inference(forward_demodulation,[],[f29659,f2886])).
fof(f29659,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13)))),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13))))),
  inference(forward_demodulation,[],[f29658,f721])).
fof(f29658,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13)))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13))))),
  inference(forward_demodulation,[],[f12090,f716])).
fof(f12090,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13))))),
  inference(forward_demodulation,[],[f11782,f721])).
fof(f11782,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13))))),
  inference(superposition,[],[f714,f2886])).
fof(f1016,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | $spl25),
  inference(cnf_transformation,[],[f1016_D])).
fof(f1016_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f876,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | $spl5),
  inference(cnf_transformation,[],[f876_D])).
fof(f876_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f24286,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))))),
  inference(resolution,[],[f3031,f509])).
fof(f509,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f365])).
fof(f365,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))))),
  inference(flattening,[],[f364])).
fof(f364,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',X1))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))))),
  inference(ennf_transformation,[],[f267])).
fof(f267,plain,(
  ! [X0,X1] : ((p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',X0)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',X1))))) => p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))))),
  inference(rectify,[],[f78])).
fof(f78,axiom,(
  ! [X6,X7] : ((p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',X6)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',X7))))) => p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X6),s('type.integer.int',X7)))))))),
  file('thm.integer.INT_LE_MUL','thm.integer.INT_LT_MUL')).
fof(f3031,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))))),
  inference(resolution,[],[f742,f692])).
fof(f7648,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))))),
  inference(resolution,[],[f743,f691])).
fof(f7649,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13)),
  inference(resolution,[],[f743,f696])).
fof(f7770,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13)),
  inference(resolution,[],[f749,f649])).
fof(f7771,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13)),
  inference(resolution,[],[f749,f649])).
fof(f29649,plain,(
  ( ! [X202] : (s('type.integer.int',X202) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK13),s('type.integer.int',X202)))) ) | $spl130),
  inference(cnf_transformation,[],[f29649_D])).
fof(f29649_D,plain,(
  ( ! [X202] : (s('type.integer.int',X202) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK13),s('type.integer.int',X202)))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f29657,plain,(
  ( ! [X10] : (s('type.integer.int',X10) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK13),s('type.integer.int',X10))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13)) )),
  inference(forward_demodulation,[],[f11865,f2886])).
fof(f11865,plain,(
  ( ! [X10] : (s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | s('type.integer.int',X10) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13)))),s('type.integer.int',X10)))) )),
  inference(forward_demodulation,[],[f11584,f721])).
fof(f11584,plain,(
  ( ! [X10] : (s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X10) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13)))),s('type.integer.int',X10)))) )),
  inference(superposition,[],[f2886,f709])).
fof(f29645,plain,(
  ( ! [X208] : (s('type.integer.int',X208) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X208),s('type.integer.int',sK13)))) ) | $spl126),
  inference(cnf_transformation,[],[f29645_D])).
fof(f29645_D,plain,(
  ( ! [X208] : (s('type.integer.int',X208) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X208),s('type.integer.int',sK13)))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f29656,plain,(
  ( ! [X11] : (s('type.integer.int',X11) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X11),s('type.integer.int',sK13))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13)) )),
  inference(forward_demodulation,[],[f11868,f2886])).
fof(f11868,plain,(
  ( ! [X11] : (s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | s('type.integer.int',X11) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X11),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13))))))) )),
  inference(forward_demodulation,[],[f11585,f721])).
fof(f11585,plain,(
  ( ! [X11] : (s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X11) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X11),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13))))))) )),
  inference(superposition,[],[f2886,f711])).
fof(f29655,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13)),
  inference(forward_demodulation,[],[f29654,f721])).
fof(f29654,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13)),
  inference(forward_demodulation,[],[f29653,f2886])).
fof(f29653,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13)))))))),
  inference(subsumption_resolution,[],[f29652,f749])).
fof(f29652,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13)))))))),
  inference(forward_demodulation,[],[f29651,f2886])).
fof(f29651,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13)))),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13)))))))),
  inference(forward_demodulation,[],[f29650,f721])).
fof(f29650,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13)))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13)))))))),
  inference(forward_demodulation,[],[f11871,f716])).
fof(f11871,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13)))))))),
  inference(forward_demodulation,[],[f11586,f721])).
fof(f11586,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13)))))))),
  inference(superposition,[],[f2886,f714])).
fof(f29648,plain,(
  ( ! [X202] : (s('type.integer.int',X202) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK13),s('type.integer.int',X202))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13)) )),
  inference(forward_demodulation,[],[f12072,f2886])).
fof(f12072,plain,(
  ( ! [X202] : (s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | s('type.integer.int',X202) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13)))),s('type.integer.int',X202)))) )),
  inference(forward_demodulation,[],[f11770,f721])).
fof(f11770,plain,(
  ( ! [X202] : (s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X202) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13)))),s('type.integer.int',X202)))) )),
  inference(superposition,[],[f709,f2886])).
fof(f874,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int',sK13) | $spl3),
  inference(cnf_transformation,[],[f874_D])).
fof(f874_D,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int',sK13) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f29647,plain,(
  ( ! [X205] : (s('type.integer.int',X205) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK13),s('type.integer.int',X205)))) ) | $spl128),
  inference(cnf_transformation,[],[f29647_D])).
fof(f29647_D,plain,(
  ( ! [X205] : (s('type.integer.int',X205) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK13),s('type.integer.int',X205)))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f29646,plain,(
  ( ! [X205] : (s('type.integer.int',X205) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK13),s('type.integer.int',X205))) | s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int',sK13)) )),
  inference(forward_demodulation,[],[f12076,f2886])).
fof(f12076,plain,(
  ( ! [X205] : (s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int',sK13) | s('type.integer.int',X205) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13)))),s('type.integer.int',X205)))) )),
  inference(forward_demodulation,[],[f11773,f721])).
fof(f11773,plain,(
  ( ! [X205] : (s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X205) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13)))),s('type.integer.int',X205)))) )),
  inference(superposition,[],[f710,f2886])).
fof(f29644,plain,(
  ( ! [X208] : (s('type.integer.int',X208) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X208),s('type.integer.int',sK13))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13)) )),
  inference(forward_demodulation,[],[f12078,f2886])).
fof(f12078,plain,(
  ( ! [X208] : (s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | s('type.integer.int',X208) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X208),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13))))))) )),
  inference(forward_demodulation,[],[f11776,f721])).
fof(f11776,plain,(
  ( ! [X208] : (s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X208) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X208),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13))))))) )),
  inference(superposition,[],[f711,f2886])).
fof(f29643,plain,(
  ( ! [X211] : (s('type.integer.int',X211) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X211),s('type.integer.int',sK13)))) ) | $spl124),
  inference(cnf_transformation,[],[f29643_D])).
fof(f29643_D,plain,(
  ( ! [X211] : (s('type.integer.int',X211) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X211),s('type.integer.int',sK13)))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f29642,plain,(
  ( ! [X211] : (s('type.integer.int',X211) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X211),s('type.integer.int',sK13))) | s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int',sK13)) )),
  inference(forward_demodulation,[],[f12082,f2886])).
fof(f12082,plain,(
  ( ! [X211] : (s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int',sK13) | s('type.integer.int',X211) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X211),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13))))))) )),
  inference(forward_demodulation,[],[f11779,f721])).
fof(f11779,plain,(
  ( ! [X211] : (s('type.integer.int',sK13) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X211) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X211),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13))))))) )),
  inference(superposition,[],[f712,f2886])).
fof(f29641,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13)),
  inference(forward_demodulation,[],[f29640,f721])).
fof(f29640,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK13)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13)),
  inference(forward_demodulation,[],[f29639,f2886])).
fof(f29639,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13)))))))),
  inference(subsumption_resolution,[],[f29638,f749])).
fof(f29638,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13)))))))),
  inference(forward_demodulation,[],[f29637,f2886])).
fof(f29637,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13)))),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13)))))))),
  inference(forward_demodulation,[],[f29636,f721])).
fof(f29636,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13)))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13)))))))),
  inference(forward_demodulation,[],[f12091,f716])).
fof(f12091,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13)))))))),
  inference(forward_demodulation,[],[f11783,f721])).
fof(f11783,plain,(
  s('type.integer.int',sK13) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK13)))))))),
  inference(superposition,[],[f714,f2886])).
fof(f29571,plain,(
  ( ! [X163,X164] : (p(s(bool,X163)) | ~p(s(bool,X164)) | s(bool,X164) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl77),
  inference(resolution,[],[f8101,f735])).
fof(f29497,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))))) | s(bool,X24) != s(bool,X26)) ) | $spl77),
  inference(resolution,[],[f8101,f526])).
fof(f29495,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl77),
  inference(resolution,[],[f8101,f525])).
fof(f29493,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl77),
  inference(resolution,[],[f8101,f524])).
fof(f29492,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))))) | s(bool,X9) != s(bool,X11)) ) | $spl77),
  inference(resolution,[],[f8101,f522])).
fof(f29490,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl77),
  inference(resolution,[],[f8101,f520])).
fof(f29461,plain,(
  $false | ($spl22 | $spl76)),
  inference(subsumption_resolution,[],[f29321,f24489])).
fof(f24489,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | $spl22),
  inference(backward_demodulation,[],[f1015,f742])).
fof(f29321,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | $spl76),
  inference(resolution,[],[f8102,f697])).
fof(f697,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f390])).
fof(f390,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))),
  inference(ennf_transformation,[],[f330])).
fof(f330,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) => p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))),
  inference(rectify,[],[f69])).
fof(f69,axiom,(
  ! [X6,X7] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) => p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X6),s('type.integer.int',X7)))))),
  file('thm.integer.INT_LE_MUL','thm.integer.INT_LT_IMP_LE')).
fof(f29460,plain,(
  $false | ($spl22 | $spl76)),
  inference(subsumption_resolution,[],[f29319,f24489])).
fof(f29319,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | $spl76),
  inference(resolution,[],[f8102,f694])).
fof(f694,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f494])).
fof(f29459,plain,(
  $false | ($spl22 | $spl76)),
  inference(subsumption_resolution,[],[f29318,f24489])).
fof(f29318,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | $spl76),
  inference(resolution,[],[f8102,f692])).
fof(f29458,plain,(
  $false | ($spl74 | $spl76)),
  inference(subsumption_resolution,[],[f29310,f8100])).
fof(f29310,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl76),
  inference(resolution,[],[f8102,f648])).
fof(f29457,plain,(
  $false | ($spl74 | $spl76)),
  inference(subsumption_resolution,[],[f29309,f8100])).
fof(f29309,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl76),
  inference(resolution,[],[f8102,f648])).
fof(f29456,plain,(
  $false | ($spl74 | $spl76)),
  inference(subsumption_resolution,[],[f29308,f8100])).
fof(f29308,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl76),
  inference(resolution,[],[f8102,f617])).
fof(f29455,plain,(
  $false | ($spl74 | $spl76)),
  inference(subsumption_resolution,[],[f29307,f8100])).
fof(f29307,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl76),
  inference(resolution,[],[f8102,f617])).
fof(f29234,plain,(
  ( ! [X158,X157] : (p(s(bool,X157)) | p(s(bool,X158)) | s(bool,X158) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl74),
  inference(resolution,[],[f8100,f735])).
fof(f29164,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X27) != s(bool,X28)) ) | $spl74),
  inference(resolution,[],[f8100,f526])).
fof(f29161,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X19) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl74),
  inference(resolution,[],[f8100,f525])).
fof(f29160,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl74),
  inference(resolution,[],[f8100,f524])).
fof(f29159,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X12) != s(bool,X13)) ) | $spl74),
  inference(resolution,[],[f8100,f522])).
fof(f29157,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl74),
  inference(resolution,[],[f8100,f520])).
fof(f25271,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl22),
  inference(forward_demodulation,[],[f25270,f721])).
fof(f25270,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | $spl22),
  inference(forward_demodulation,[],[f25269,f721])).
fof(f25269,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | $spl22),
  inference(forward_demodulation,[],[f25268,f716])).
fof(f25268,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | $spl22),
  inference(subsumption_resolution,[],[f25267,f704])).
fof(f25267,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | $spl22),
  inference(forward_demodulation,[],[f25246,f721])).
fof(f25246,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | $spl22),
  inference(superposition,[],[f24489,f714])).
fof(f25206,plain,(
  ( ! [X161,X162] : (p(s(bool,X161)) | ~p(s(bool,X162)) | s(bool,X162) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl22),
  inference(resolution,[],[f24489,f735])).
fof(f25132,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))))) | s(bool,X22) != s(bool,X24)) ) | $spl22),
  inference(resolution,[],[f24489,f526])).
fof(f25130,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl22),
  inference(resolution,[],[f24489,f525])).
fof(f25128,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl22),
  inference(resolution,[],[f24489,f524])).
fof(f25127,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))))) | s(bool,X7) != s(bool,X9)) ) | $spl22),
  inference(resolution,[],[f24489,f522])).
fof(f25125,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl22),
  inference(resolution,[],[f24489,f520])).
fof(f24488,plain,(
  ( ! [X164] : (~p(s(bool,X164)) | s(bool,X164) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))))) ) | $spl120),
  inference(cnf_transformation,[],[f24488_D])).
fof(f24488_D,plain,(
  ( ! [X164] : (~p(s(bool,X164)) | s(bool,X164) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f24299,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))))) )),
  inference(resolution,[],[f3031,f520])).
fof(f24301,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f3031,f522])).
fof(f24302,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))))) )),
  inference(resolution,[],[f3031,f524])).
fof(f24304,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))))) )),
  inference(resolution,[],[f3031,f525])).
fof(f24306,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) | s(bool,X24) != s(bool,X26)) )),
  inference(resolution,[],[f3031,f526])).
fof(f24380,plain,(
  ( ! [X163,X164] : (p(s(bool,X163)) | ~p(s(bool,X164)) | s(bool,X164) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))))) )),
  inference(resolution,[],[f3031,f735])).
fof(f24485,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))))),
  inference(forward_demodulation,[],[f24484,f721])).
fof(f24484,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))))),
  inference(forward_demodulation,[],[f21728,f4120])).
fof(f21728,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))),
  inference(forward_demodulation,[],[f21727,f721])).
fof(f21727,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))),
  inference(forward_demodulation,[],[f21726,f4120])).
fof(f21726,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))),
  inference(subsumption_resolution,[],[f21725,f889])).
fof(f21725,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))),
  inference(forward_demodulation,[],[f21724,f721])).
fof(f21724,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))),
  inference(forward_demodulation,[],[f21418,f716])).
fof(f21418,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))),
  inference(superposition,[],[f714,f4120])).
fof(f24483,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))))),
  inference(forward_demodulation,[],[f24482,f721])).
fof(f24482,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))))),
  inference(forward_demodulation,[],[f24481,f4120])).
fof(f24481,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))),
  inference(subsumption_resolution,[],[f24480,f889])).
fof(f24480,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))),
  inference(forward_demodulation,[],[f24479,f4120])).
fof(f24479,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))),
  inference(forward_demodulation,[],[f24478,f721])).
fof(f24478,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))),
  inference(forward_demodulation,[],[f21729,f716])).
fof(f21729,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))),
  inference(forward_demodulation,[],[f21419,f721])).
fof(f21419,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))),
  inference(superposition,[],[f714,f4120])).
fof(f19284,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))))),
  inference(resolution,[],[f744,f686])).
fof(f19285,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))))),
  inference(resolution,[],[f744,f686])).
fof(f19287,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))))),
  inference(resolution,[],[f744,f691])).
fof(f19288,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12)),
  inference(resolution,[],[f744,f696])).
fof(f19409,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12)),
  inference(resolution,[],[f889,f649])).
fof(f19410,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12)),
  inference(resolution,[],[f889,f649])).
fof(f19416,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))))),
  inference(resolution,[],[f889,f691])).
fof(f19417,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))))),
  inference(resolution,[],[f889,f696])).
fof(f24477,plain,(
  ( ! [X10] : (s('type.integer.int',X10) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK12),s('type.integer.int',X10))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12)) )),
  inference(forward_demodulation,[],[f21504,f4120])).
fof(f21504,plain,(
  ( ! [X10] : (s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | s('type.integer.int',X10) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))),s('type.integer.int',X10)))) )),
  inference(forward_demodulation,[],[f21221,f721])).
fof(f21221,plain,(
  ( ! [X10] : (s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X10) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))),s('type.integer.int',X10)))) )),
  inference(superposition,[],[f4120,f709])).
fof(f24476,plain,(
  ( ! [X11] : (s('type.integer.int',X11) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X11),s('type.integer.int',sK12))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12)) )),
  inference(forward_demodulation,[],[f21507,f4120])).
fof(f21507,plain,(
  ( ! [X11] : (s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | s('type.integer.int',X11) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X11),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))) )),
  inference(forward_demodulation,[],[f21222,f721])).
fof(f21222,plain,(
  ( ! [X11] : (s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X11) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X11),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))) )),
  inference(superposition,[],[f4120,f711])).
fof(f24475,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12)),
  inference(forward_demodulation,[],[f24474,f721])).
fof(f24474,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12)),
  inference(forward_demodulation,[],[f24473,f4120])).
fof(f24473,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))),
  inference(subsumption_resolution,[],[f24472,f889])).
fof(f24472,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))),
  inference(forward_demodulation,[],[f24471,f4120])).
fof(f24471,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))),
  inference(forward_demodulation,[],[f24470,f721])).
fof(f24470,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))),
  inference(forward_demodulation,[],[f21510,f716])).
fof(f21510,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))),
  inference(forward_demodulation,[],[f21223,f721])).
fof(f21223,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))),
  inference(superposition,[],[f4120,f714])).
fof(f24468,plain,(
  ( ! [X202] : (s('type.integer.int',X202) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK12),s('type.integer.int',X202))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12)) )),
  inference(forward_demodulation,[],[f21711,f4120])).
fof(f21711,plain,(
  ( ! [X202] : (s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | s('type.integer.int',X202) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))),s('type.integer.int',X202)))) )),
  inference(forward_demodulation,[],[f21407,f721])).
fof(f21407,plain,(
  ( ! [X202] : (s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X202) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))),s('type.integer.int',X202)))) )),
  inference(superposition,[],[f709,f4120])).
fof(f24466,plain,(
  ( ! [X205] : (s('type.integer.int',X205) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK12),s('type.integer.int',X205))) | s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int',sK12)) )),
  inference(forward_demodulation,[],[f21715,f4120])).
fof(f21715,plain,(
  ( ! [X205] : (s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int',sK12) | s('type.integer.int',X205) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))),s('type.integer.int',X205)))) )),
  inference(forward_demodulation,[],[f21410,f721])).
fof(f21410,plain,(
  ( ! [X205] : (s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X205) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))),s('type.integer.int',X205)))) )),
  inference(superposition,[],[f710,f4120])).
fof(f24464,plain,(
  ( ! [X208] : (s('type.integer.int',X208) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X208),s('type.integer.int',sK12))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12)) )),
  inference(forward_demodulation,[],[f21717,f4120])).
fof(f21717,plain,(
  ( ! [X208] : (s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | s('type.integer.int',X208) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X208),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))) )),
  inference(forward_demodulation,[],[f21413,f721])).
fof(f21413,plain,(
  ( ! [X208] : (s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X208) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X208),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))) )),
  inference(superposition,[],[f711,f4120])).
fof(f24462,plain,(
  ( ! [X211] : (s('type.integer.int',X211) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X211),s('type.integer.int',sK12))) | s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int',sK12)) )),
  inference(forward_demodulation,[],[f21721,f4120])).
fof(f21721,plain,(
  ( ! [X211] : (s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int',sK12) | s('type.integer.int',X211) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X211),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))) )),
  inference(forward_demodulation,[],[f21416,f721])).
fof(f21416,plain,(
  ( ! [X211] : (s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X211) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X211),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))) )),
  inference(superposition,[],[f712,f4120])).
fof(f24461,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12)),
  inference(forward_demodulation,[],[f24460,f721])).
fof(f24460,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12)),
  inference(forward_demodulation,[],[f24459,f4120])).
fof(f24459,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))),
  inference(subsumption_resolution,[],[f24458,f889])).
fof(f24458,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))),
  inference(forward_demodulation,[],[f24457,f4120])).
fof(f24457,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))),
  inference(forward_demodulation,[],[f24456,f721])).
fof(f24456,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))),
  inference(forward_demodulation,[],[f21730,f716])).
fof(f21730,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))),
  inference(forward_demodulation,[],[f21420,f721])).
fof(f21420,plain,(
  s('type.integer.int',sK12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK24(sK12)))))))),
  inference(superposition,[],[f714,f4120])).
fof(f24448,plain,(
  $false | ($spl4 | $spl24)),
  inference(subsumption_resolution,[],[f24447,f1017])).
fof(f24447,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | $spl4),
  inference(subsumption_resolution,[],[f24286,f877])).
fof(f24201,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl110),
  inference(cnf_transformation,[],[f24201_D])).
fof(f24201_D,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f24134,plain,(
  ( ! [X158,X157] : (p(s(bool,X157)) | p(s(bool,X158)) | s(bool,X158) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) )),
  inference(resolution,[],[f3026,f735])).
fof(f3026,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0'))))),
  inference(resolution,[],[f742,f682])).
fof(f24064,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X27) != s(bool,X28)) )),
  inference(resolution,[],[f3026,f526])).
fof(f24061,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X19) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) )),
  inference(resolution,[],[f3026,f525])).
fof(f24060,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) )),
  inference(resolution,[],[f3026,f524])).
fof(f24059,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X12) != s(bool,X13)) )),
  inference(resolution,[],[f3026,f522])).
fof(f24057,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) )),
  inference(resolution,[],[f3026,f520])).
fof(f23974,plain,(
  ( ! [X158,X157] : (p(s(bool,X157)) | p(s(bool,X158)) | s(bool,X158) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl42),
  inference(resolution,[],[f2740,f735])).
fof(f23904,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X27) != s(bool,X28)) ) | $spl42),
  inference(resolution,[],[f2740,f526])).
fof(f23901,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X19) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl42),
  inference(resolution,[],[f2740,f525])).
fof(f23900,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl42),
  inference(resolution,[],[f2740,f524])).
fof(f23899,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X12) != s(bool,X13)) ) | $spl42),
  inference(resolution,[],[f2740,f522])).
fof(f23897,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl42),
  inference(resolution,[],[f2740,f520])).
fof(f23750,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl108),
  inference(cnf_transformation,[],[f23750_D])).
fof(f23750_D,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f23691,plain,(
  ( ! [X163,X164] : (p(s(bool,X163)) | ~p(s(bool,X164)) | s(bool,X164) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl43),
  inference(resolution,[],[f2739,f735])).
fof(f2739,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl43),
  inference(cnf_transformation,[],[f2739_D])).
fof(f2739_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f23617,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X24) != s(bool,X26)) ) | $spl43),
  inference(resolution,[],[f2739,f526])).
fof(f23615,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl43),
  inference(resolution,[],[f2739,f525])).
fof(f23613,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl43),
  inference(resolution,[],[f2739,f524])).
fof(f23612,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X9) != s(bool,X11)) ) | $spl43),
  inference(resolution,[],[f2739,f522])).
fof(f23610,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl43),
  inference(resolution,[],[f2739,f520])).
fof(f23749,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl107),
  inference(cnf_transformation,[],[f23749_D])).
fof(f23749_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) <=> ~$spl107),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl107])])).
fof(f23747,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | $spl104),
  inference(cnf_transformation,[],[f23747_D])).
fof(f23747_D,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f23607,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl43),
  inference(resolution,[],[f2739,f696])).
fof(f23606,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl43),
  inference(resolution,[],[f2739,f691])).
fof(f23600,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | $spl43),
  inference(resolution,[],[f2739,f649])).
fof(f23599,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | $spl43),
  inference(resolution,[],[f2739,f649])).
fof(f23451,plain,(
  ( ! [X158,X157] : (p(s(bool,X157)) | p(s(bool,X158)) | s(bool,X158) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl42),
  inference(resolution,[],[f2740,f735])).
fof(f23381,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X27) != s(bool,X28)) ) | $spl42),
  inference(resolution,[],[f2740,f526])).
fof(f23378,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X19) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl42),
  inference(resolution,[],[f2740,f525])).
fof(f23377,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl42),
  inference(resolution,[],[f2740,f524])).
fof(f23376,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X12) != s(bool,X13)) ) | $spl42),
  inference(resolution,[],[f2740,f522])).
fof(f23374,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl42),
  inference(resolution,[],[f2740,f520])).
fof(f3202,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl54),
  inference(cnf_transformation,[],[f3202_D])).
fof(f3202_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f3204,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))))) | $spl56),
  inference(cnf_transformation,[],[f3204_D])).
fof(f3204_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f21210,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0'))))),
  inference(forward_demodulation,[],[f21209,f721])).
fof(f21209,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))))),
  inference(forward_demodulation,[],[f21208,f721])).
fof(f21208,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))))),
  inference(forward_demodulation,[],[f21207,f716])).
fof(f21207,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))))),
  inference(subsumption_resolution,[],[f21198,f721])).
fof(f21198,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))))),
  inference(superposition,[],[f3028,f714])).
fof(f20189,plain,(
  ( ! [X159,X160] : (p(s(bool,X159)) | p(s(bool,X160)) | s(bool,X160) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl24),
  inference(resolution,[],[f1017,f735])).
fof(f20119,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X29) != s(bool,X30)) ) | $spl24),
  inference(resolution,[],[f1017,f526])).
fof(f20116,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X21) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl24),
  inference(resolution,[],[f1017,f525])).
fof(f20115,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl24),
  inference(resolution,[],[f1017,f524])).
fof(f20114,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X14) != s(bool,X15)) ) | $spl24),
  inference(resolution,[],[f1017,f522])).
fof(f20112,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl24),
  inference(resolution,[],[f1017,f520])).
fof(f20054,plain,(
  ( ! [X161,X162] : (p(s(bool,X161)) | ~p(s(bool,X162)) | s(bool,X162) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl21),
  inference(resolution,[],[f1013,f735])).
fof(f19980,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X22) != s(bool,X24)) ) | $spl21),
  inference(resolution,[],[f1013,f526])).
fof(f19978,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl21),
  inference(resolution,[],[f1013,f525])).
fof(f19976,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl21),
  inference(resolution,[],[f1013,f524])).
fof(f19975,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X7) != s(bool,X9)) ) | $spl21),
  inference(resolution,[],[f1013,f522])).
fof(f19973,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl21),
  inference(resolution,[],[f1013,f520])).
fof(f19745,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0'))))),
  inference(forward_demodulation,[],[f19744,f721])).
fof(f19744,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))))),
  inference(forward_demodulation,[],[f19743,f721])).
fof(f19743,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))))),
  inference(forward_demodulation,[],[f19742,f716])).
fof(f19742,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))))),
  inference(subsumption_resolution,[],[f19741,f704])).
fof(f19741,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))))),
  inference(forward_demodulation,[],[f19720,f721])).
fof(f19720,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))))),
  inference(superposition,[],[f742,f714])).
fof(f3187,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))))) ) | $spl50),
  inference(cnf_transformation,[],[f3187_D])).
fof(f3187_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f19668,plain,(
  ( ! [X161,X162] : (p(s(bool,X161)) | ~p(s(bool,X162)) | s(bool,X162) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))))) )),
  inference(resolution,[],[f742,f735])).
fof(f19594,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) | s(bool,X22) != s(bool,X24)) )),
  inference(resolution,[],[f742,f526])).
fof(f19592,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))))) )),
  inference(resolution,[],[f742,f525])).
fof(f19590,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))))) )),
  inference(resolution,[],[f742,f524])).
fof(f19589,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) | s(bool,X7) != s(bool,X9)) )),
  inference(resolution,[],[f742,f522])).
fof(f19587,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))))) )),
  inference(resolution,[],[f742,f520])).
fof(f19501,plain,(
  ( ! [X163,X164] : (p(s(bool,X163)) | ~p(s(bool,X164)) | s(bool,X164) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) )),
  inference(resolution,[],[f889,f735])).
fof(f19427,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X24) != s(bool,X26)) )),
  inference(resolution,[],[f889,f526])).
fof(f19425,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) )),
  inference(resolution,[],[f889,f525])).
fof(f19423,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) )),
  inference(resolution,[],[f889,f524])).
fof(f19422,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f889,f522])).
fof(f19420,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) )),
  inference(resolution,[],[f889,f520])).
fof(f19366,plain,(
  ( ! [X156,X157] : (p(s(bool,X156)) | p(s(bool,X157)) | s(bool,X157) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f744,f735])).
fof(f19296,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X26) != s(bool,X27)) )),
  inference(resolution,[],[f744,f526])).
fof(f19293,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f744,f525])).
fof(f19292,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f744,f524])).
fof(f19291,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X11) != s(bool,X12)) )),
  inference(resolution,[],[f744,f522])).
fof(f19289,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f744,f520])).
fof(f19263,plain,(
  ( ! [X164] : (~p(s(bool,X164)) | s(bool,X164) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl100),
  inference(cnf_transformation,[],[f19263_D])).
fof(f19263_D,plain,(
  ( ! [X164] : (~p(s(bool,X164)) | s(bool,X164) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f19105,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl75),
  inference(resolution,[],[f8099,f520])).
fof(f19107,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X9) != s(bool,X11)) ) | $spl75),
  inference(resolution,[],[f8099,f522])).
fof(f19108,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl75),
  inference(resolution,[],[f8099,f524])).
fof(f19110,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl75),
  inference(resolution,[],[f8099,f525])).
fof(f19112,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X24) != s(bool,X26)) ) | $spl75),
  inference(resolution,[],[f8099,f526])).
fof(f19186,plain,(
  ( ! [X163,X164] : (p(s(bool,X163)) | ~p(s(bool,X164)) | s(bool,X164) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl75),
  inference(resolution,[],[f8099,f735])).
fof(f8123,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))))),
  inference(resolution,[],[f744,f686])).
fof(f8124,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))))),
  inference(resolution,[],[f744,f686])).
fof(f8913,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))))),
  inference(resolution,[],[f889,f691])).
fof(f8914,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))))),
  inference(resolution,[],[f889,f696])).
fof(f19262,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | $spl98),
  inference(cnf_transformation,[],[f19262_D])).
fof(f19262_D,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f19094,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | $spl75),
  inference(resolution,[],[f8099,f649])).
fof(f19095,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | $spl75),
  inference(resolution,[],[f8099,f649])).
fof(f19260,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl97),
  inference(cnf_transformation,[],[f19260_D])).
fof(f19260_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) <=> ~$spl97),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl97])])).
fof(f19101,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl75),
  inference(resolution,[],[f8099,f691])).
fof(f19102,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl75),
  inference(resolution,[],[f8099,f696])).
fof(f19249,plain,(
  $false | ($spl22 | $spl75)),
  inference(subsumption_resolution,[],[f19104,f9555])).
fof(f9555,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | $spl22),
  inference(backward_demodulation,[],[f1015,f742])).
fof(f19104,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | $spl75),
  inference(resolution,[],[f8099,f699])).
fof(f19248,plain,(
  $false | ($spl22 | $spl75)),
  inference(subsumption_resolution,[],[f19103,f9555])).
fof(f19103,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | $spl75),
  inference(resolution,[],[f8099,f698])).
fof(f19245,plain,(
  $false | ($spl22 | $spl75)),
  inference(subsumption_resolution,[],[f19100,f9555])).
fof(f19100,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | $spl75),
  inference(resolution,[],[f8099,f690])).
fof(f19244,plain,(
  $false | ($spl22 | $spl75)),
  inference(subsumption_resolution,[],[f19099,f9555])).
fof(f19099,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | $spl75),
  inference(resolution,[],[f8099,f684])).
fof(f19243,plain,(
  $false | ($spl22 | $spl75)),
  inference(subsumption_resolution,[],[f19098,f9555])).
fof(f19098,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | $spl75),
  inference(resolution,[],[f8099,f682])).
fof(f19020,plain,(
  ( ! [X158,X157] : (p(s(bool,X157)) | p(s(bool,X158)) | s(bool,X158) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl74),
  inference(resolution,[],[f8100,f735])).
fof(f18950,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X27) != s(bool,X28)) ) | $spl74),
  inference(resolution,[],[f8100,f526])).
fof(f18947,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X19) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl74),
  inference(resolution,[],[f8100,f525])).
fof(f18946,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl74),
  inference(resolution,[],[f8100,f524])).
fof(f18945,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X12) != s(bool,X13)) ) | $spl74),
  inference(resolution,[],[f8100,f522])).
fof(f18943,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl74),
  inference(resolution,[],[f8100,f520])).
fof(f17361,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl86),
  inference(cnf_transformation,[],[f17361_D])).
fof(f17361_D,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f17676,plain,(
  ( ! [X158,X157] : (p(s(bool,X157)) | p(s(bool,X158)) | s(bool,X158) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl34),
  inference(resolution,[],[f2426,f735])).
fof(f17606,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X27) != s(bool,X28)) ) | $spl34),
  inference(resolution,[],[f2426,f526])).
fof(f17603,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X19) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl34),
  inference(resolution,[],[f2426,f525])).
fof(f17602,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl34),
  inference(resolution,[],[f2426,f524])).
fof(f17601,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X12) != s(bool,X13)) ) | $spl34),
  inference(resolution,[],[f2426,f522])).
fof(f17599,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl34),
  inference(resolution,[],[f2426,f520])).
fof(f17524,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl92),
  inference(cnf_transformation,[],[f17524_D])).
fof(f17524_D,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f17475,plain,(
  ( ! [X163,X164] : (p(s(bool,X163)) | ~p(s(bool,X164)) | s(bool,X164) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl35),
  inference(resolution,[],[f2425,f735])).
fof(f2425,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl35),
  inference(cnf_transformation,[],[f2425_D])).
fof(f2425_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f17401,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X24) != s(bool,X26)) ) | $spl35),
  inference(resolution,[],[f2425,f526])).
fof(f17399,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl35),
  inference(resolution,[],[f2425,f525])).
fof(f17397,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl35),
  inference(resolution,[],[f2425,f524])).
fof(f17396,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X9) != s(bool,X11)) ) | $spl35),
  inference(resolution,[],[f2425,f522])).
fof(f17394,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl35),
  inference(resolution,[],[f2425,f520])).
fof(f17523,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl91),
  inference(cnf_transformation,[],[f17523_D])).
fof(f17523_D,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) <=> ~$spl91),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl91])])).
fof(f17521,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | $spl88),
  inference(cnf_transformation,[],[f17521_D])).
fof(f17521_D,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f17391,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl35),
  inference(resolution,[],[f2425,f696])).
fof(f17390,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl35),
  inference(resolution,[],[f2425,f691])).
fof(f2428,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))) | $spl36),
  inference(cnf_transformation,[],[f2428_D])).
fof(f2428_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f17384,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | $spl35),
  inference(resolution,[],[f2425,f649])).
fof(f17383,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | $spl35),
  inference(resolution,[],[f2425,f649])).
fof(f17316,plain,(
  ( ! [X158,X157] : (p(s(bool,X157)) | p(s(bool,X158)) | s(bool,X158) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl34),
  inference(resolution,[],[f2426,f735])).
fof(f17246,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X27) != s(bool,X28)) ) | $spl34),
  inference(resolution,[],[f2426,f526])).
fof(f17243,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X19) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl34),
  inference(resolution,[],[f2426,f525])).
fof(f17242,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl34),
  inference(resolution,[],[f2426,f524])).
fof(f17241,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X12) != s(bool,X13)) ) | $spl34),
  inference(resolution,[],[f2426,f522])).
fof(f17239,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X5) != s(bool,X7) | s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl34),
  inference(resolution,[],[f2426,f520])).
fof(f12752,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl22),
  inference(forward_demodulation,[],[f12751,f721])).
fof(f12751,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | $spl22),
  inference(forward_demodulation,[],[f12750,f721])).
fof(f12750,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | $spl22),
  inference(forward_demodulation,[],[f12749,f716])).
fof(f12749,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | $spl22),
  inference(subsumption_resolution,[],[f12740,f721])).
fof(f12740,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | $spl22),
  inference(superposition,[],[f10094,f714])).
fof(f10094,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | $spl22),
  inference(resolution,[],[f9555,f687])).
fof(f11146,plain,(
  $false | $spl10),
  inference(equality_resolution,[],[f880])).
fof(f10252,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl22),
  inference(forward_demodulation,[],[f10251,f721])).
fof(f10251,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | $spl22),
  inference(forward_demodulation,[],[f10250,f721])).
fof(f10250,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | $spl22),
  inference(forward_demodulation,[],[f10249,f716])).
fof(f10249,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | $spl22),
  inference(subsumption_resolution,[],[f10248,f704])).
fof(f10248,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | $spl22),
  inference(forward_demodulation,[],[f10226,f721])).
fof(f10226,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | $spl22),
  inference(superposition,[],[f9555,f714])).
fof(f10186,plain,(
  ( ! [X161,X162] : (p(s(bool,X161)) | ~p(s(bool,X162)) | s(bool,X162) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl22),
  inference(resolution,[],[f9555,f735])).
fof(f10112,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))))) | s(bool,X22) != s(bool,X24)) ) | $spl22),
  inference(resolution,[],[f9555,f526])).
fof(f10110,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl22),
  inference(resolution,[],[f9555,f525])).
fof(f8086,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl72),
  inference(cnf_transformation,[],[f8086_D])).
fof(f8086_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f881,plain,(
  ( ! [X8,X7] : (s(bool,X7) != s(bool,X8) | p(s(bool,X7)) | ~p(s(bool,X8))) ) | $spl12),
  inference(cnf_transformation,[],[f881_D])).
fof(f881_D,plain,(
  ( ! [X8,X7] : (s(bool,X7) != s(bool,X8) | p(s(bool,X7)) | ~p(s(bool,X8))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f10109,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))))) | s(bool,X13) != s(bool,X14)) ) | $spl22),
  inference(resolution,[],[f9555,f524])).
fof(f10108,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl22),
  inference(resolution,[],[f9555,f524])).
fof(f10107,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))))) | s(bool,X7) != s(bool,X9)) ) | $spl22),
  inference(resolution,[],[f9555,f522])).
fof(f10106,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))))) | s(bool,X4) != s(bool,X5)) ) | $spl22),
  inference(resolution,[],[f9555,f520])).
fof(f10105,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl22),
  inference(resolution,[],[f9555,f520])).
fof(f9826,plain,(
  $false | ($spl21 | $spl22)),
  inference(subsumption_resolution,[],[f9566,f704])).
fof(f9566,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl21 | $spl22)),
  inference(backward_demodulation,[],[f1015,f1013])).
fof(f8126,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))))),
  inference(resolution,[],[f744,f691])).
fof(f8127,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12)),
  inference(resolution,[],[f744,f696])).
fof(f8906,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12)),
  inference(resolution,[],[f889,f649])).
fof(f8907,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12)),
  inference(resolution,[],[f889,f649])).
fof(f9515,plain,(
  ( ! [X159,X160] : (p(s(bool,X159)) | p(s(bool,X160)) | s(bool,X160) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl24),
  inference(resolution,[],[f1017,f735])).
fof(f9445,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X29) != s(bool,X30)) ) | $spl24),
  inference(resolution,[],[f1017,f526])).
fof(f9554,plain,(
  ( ! [X12] : (s(bool,X12) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl84),
  inference(cnf_transformation,[],[f9554_D])).
fof(f9554_D,plain,(
  ( ! [X12] : (s(bool,X12) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f9444,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26) | s(bool,X27) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl24),
  inference(resolution,[],[f1017,f526])).
fof(f9442,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X21) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl24),
  inference(resolution,[],[f1017,f525])).
fof(f9441,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl24),
  inference(resolution,[],[f1017,f524])).
fof(f9440,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X14) != s(bool,X15)) ) | $spl24),
  inference(resolution,[],[f1017,f522])).
fof(f9439,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X11) | s(bool,X12) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl24),
  inference(resolution,[],[f1017,f522])).
fof(f9438,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl24),
  inference(resolution,[],[f1017,f520])).
fof(f9379,plain,(
  ( ! [X161,X162] : (p(s(bool,X161)) | ~p(s(bool,X162)) | s(bool,X162) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl21),
  inference(resolution,[],[f1013,f735])).
fof(f9305,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X22) != s(bool,X24)) ) | $spl21),
  inference(resolution,[],[f1013,f526])).
fof(f9303,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl21),
  inference(resolution,[],[f1013,f525])).
fof(f4798,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl68),
  inference(cnf_transformation,[],[f4798_D])).
fof(f4798_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f9302,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X13) != s(bool,X14)) ) | $spl21),
  inference(resolution,[],[f1013,f524])).
fof(f9301,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl21),
  inference(resolution,[],[f1013,f524])).
fof(f9300,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X7) != s(bool,X9)) ) | $spl21),
  inference(resolution,[],[f1013,f522])).
fof(f9299,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X4) != s(bool,X5)) ) | $spl21),
  inference(resolution,[],[f1013,f520])).
fof(f9298,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl21),
  inference(resolution,[],[f1013,f520])).
fof(f8998,plain,(
  ( ! [X163,X164] : (p(s(bool,X163)) | ~p(s(bool,X164)) | s(bool,X164) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) )),
  inference(resolution,[],[f889,f735])).
fof(f8924,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X24) != s(bool,X26)) )),
  inference(resolution,[],[f889,f526])).
fof(f8922,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) )),
  inference(resolution,[],[f889,f525])).
fof(f4257,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl64),
  inference(cnf_transformation,[],[f4257_D])).
fof(f4257_D,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f8921,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X15) != s(bool,X16)) )),
  inference(resolution,[],[f889,f524])).
fof(f8920,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) )),
  inference(resolution,[],[f889,f524])).
fof(f8919,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f889,f522])).
fof(f8918,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | p(s(bool,X7)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X6) != s(bool,X7)) )),
  inference(resolution,[],[f889,f520])).
fof(f8917,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) )),
  inference(resolution,[],[f889,f520])).
fof(f8864,plain,(
  ( ! [X159,X160] : (p(s(bool,X159)) | p(s(bool,X160)) | s(bool,X160) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl4),
  inference(resolution,[],[f877,f735])).
fof(f8794,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X28) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X29) != s(bool,X30)) ) | $spl4),
  inference(resolution,[],[f877,f526])).
fof(f8903,plain,(
  ( ! [X12] : (s(bool,X12) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl80),
  inference(cnf_transformation,[],[f8903_D])).
fof(f8903_D,plain,(
  ( ! [X12] : (s(bool,X12) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f8793,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X26) | s(bool,X27) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl4),
  inference(resolution,[],[f877,f526])).
fof(f8791,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | ~p(s(bool,X20)) | p(s(bool,X21)) | s(bool,X21) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl4),
  inference(resolution,[],[f877,f525])).
fof(f8790,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X16) != s(bool,X18) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl4),
  inference(resolution,[],[f877,f524])).
fof(f8789,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X14) != s(bool,X15)) ) | $spl4),
  inference(resolution,[],[f877,f522])).
fof(f8788,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X11) | s(bool,X12) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl4),
  inference(resolution,[],[f877,f522])).
fof(f8787,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X7) != s(bool,X9) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl4),
  inference(resolution,[],[f877,f520])).
fof(f8728,plain,(
  ( ! [X161,X162] : (p(s(bool,X161)) | ~p(s(bool,X162)) | s(bool,X162) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl1),
  inference(resolution,[],[f873,f735])).
fof(f8654,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X22) != s(bool,X24)) ) | $spl1),
  inference(resolution,[],[f873,f526])).
fof(f8652,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl1),
  inference(resolution,[],[f873,f525])).
fof(f3582,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl60),
  inference(cnf_transformation,[],[f3582_D])).
fof(f3582_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f8651,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X13) != s(bool,X14)) ) | $spl1),
  inference(resolution,[],[f873,f524])).
fof(f8650,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl1),
  inference(resolution,[],[f873,f524])).
fof(f8649,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X7) != s(bool,X9)) ) | $spl1),
  inference(resolution,[],[f873,f522])).
fof(f8648,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X4) != s(bool,X5)) ) | $spl1),
  inference(resolution,[],[f873,f520])).
fof(f8647,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl1),
  inference(resolution,[],[f873,f520])).
fof(f8419,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0'))))),
  inference(forward_demodulation,[],[f8418,f721])).
fof(f8418,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))))),
  inference(forward_demodulation,[],[f8417,f721])).
fof(f8417,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))))),
  inference(forward_demodulation,[],[f8416,f716])).
fof(f8416,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))))),
  inference(subsumption_resolution,[],[f8415,f704])).
fof(f8415,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))))),
  inference(forward_demodulation,[],[f8393,f721])).
fof(f8393,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))))),
  inference(superposition,[],[f742,f714])).
fof(f8341,plain,(
  ( ! [X161,X162] : (p(s(bool,X161)) | ~p(s(bool,X162)) | s(bool,X162) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))))) )),
  inference(resolution,[],[f742,f735])).
fof(f8267,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) | s(bool,X22) != s(bool,X24)) )),
  inference(resolution,[],[f742,f526])).
fof(f8265,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))))) )),
  inference(resolution,[],[f742,f525])).
fof(f3188,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))))) ) | $spl52),
  inference(cnf_transformation,[],[f3188_D])).
fof(f3188_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f8264,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) | s(bool,X13) != s(bool,X14)) )),
  inference(resolution,[],[f742,f524])).
fof(f8263,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))))) )),
  inference(resolution,[],[f742,f524])).
fof(f8262,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) | s(bool,X7) != s(bool,X9)) )),
  inference(resolution,[],[f742,f522])).
fof(f8261,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) | s(bool,X4) != s(bool,X5)) )),
  inference(resolution,[],[f742,f520])).
fof(f8260,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))))) )),
  inference(resolution,[],[f742,f520])).
fof(f8205,plain,(
  ( ! [X156,X157] : (p(s(bool,X156)) | p(s(bool,X157)) | s(bool,X157) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f744,f735])).
fof(f8135,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X26) != s(bool,X27)) )),
  inference(resolution,[],[f744,f526])).
fof(f1019,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) | $spl28),
  inference(cnf_transformation,[],[f1019_D])).
fof(f1019_D,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f8134,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X23) | s(bool,X24) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f744,f526])).
fof(f8132,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f744,f525])).
fof(f8131,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f744,f524])).
fof(f8130,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X11) != s(bool,X12)) )),
  inference(resolution,[],[f744,f522])).
fof(f8129,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X8) | s(bool,X9) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f744,f522])).
fof(f8128,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f744,f520])).
fof(f8098,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl22),
  inference(forward_demodulation,[],[f8097,f721])).
fof(f8097,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | $spl22),
  inference(forward_demodulation,[],[f8096,f721])).
fof(f8096,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | $spl22),
  inference(forward_demodulation,[],[f8095,f716])).
fof(f8095,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | $spl22),
  inference(subsumption_resolution,[],[f8094,f704])).
fof(f8094,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | $spl22),
  inference(forward_demodulation,[],[f8070,f721])).
fof(f8070,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | $spl22),
  inference(superposition,[],[f7624,f714])).
fof(f7624,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) | $spl22),
  inference(forward_demodulation,[],[f742,f1015])).
fof(f8030,plain,(
  ( ! [X161,X162] : (p(s(bool,X161)) | ~p(s(bool,X162)) | s(bool,X162) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl22),
  inference(resolution,[],[f7624,f735])).
fof(f7956,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))))) | s(bool,X22) != s(bool,X24)) ) | $spl22),
  inference(resolution,[],[f7624,f526])).
fof(f7954,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl22),
  inference(resolution,[],[f7624,f525])).
fof(f7953,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))))) | s(bool,X13) != s(bool,X14)) ) | $spl22),
  inference(resolution,[],[f7624,f524])).
fof(f7952,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl22),
  inference(resolution,[],[f7624,f524])).
fof(f7951,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))))) | s(bool,X7) != s(bool,X9)) ) | $spl22),
  inference(resolution,[],[f7624,f522])).
fof(f7950,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))))) | s(bool,X4) != s(bool,X5)) ) | $spl22),
  inference(resolution,[],[f7624,f520])).
fof(f7949,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))))) ) | $spl22),
  inference(resolution,[],[f7624,f520])).
fof(f7862,plain,(
  ( ! [X163,X164] : (p(s(bool,X163)) | ~p(s(bool,X164)) | s(bool,X164) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) )),
  inference(resolution,[],[f749,f735])).
fof(f7788,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X24) != s(bool,X26)) )),
  inference(resolution,[],[f749,f526])).
fof(f7786,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) )),
  inference(resolution,[],[f749,f525])).
fof(f3023,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl48),
  inference(cnf_transformation,[],[f3023_D])).
fof(f3023_D,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f7785,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X15) != s(bool,X16)) )),
  inference(resolution,[],[f749,f524])).
fof(f7784,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) )),
  inference(resolution,[],[f749,f524])).
fof(f7783,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f749,f522])).
fof(f7782,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | p(s(bool,X7)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X6) != s(bool,X7)) )),
  inference(resolution,[],[f749,f520])).
fof(f7781,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) )),
  inference(resolution,[],[f749,f520])).
fof(f7727,plain,(
  ( ! [X156,X157] : (p(s(bool,X156)) | p(s(bool,X157)) | s(bool,X157) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f743,f735])).
fof(f7657,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X26) != s(bool,X27)) )),
  inference(resolution,[],[f743,f526])).
fof(f880,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) | $spl10),
  inference(cnf_transformation,[],[f880_D])).
fof(f880_D,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f7656,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X23) | s(bool,X24) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f743,f526])).
fof(f7654,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f743,f525])).
fof(f7653,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f743,f524])).
fof(f7652,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X11) != s(bool,X12)) )),
  inference(resolution,[],[f743,f522])).
fof(f7651,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X8) | s(bool,X9) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f743,f522])).
fof(f7650,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f743,f520])).
fof(f7618,plain,(
  $false | ($spl2 | $spl10)),
  inference(equality_resolution,[],[f7435])).
fof(f7435,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) ) | ($spl2 | $spl10)),
  inference(forward_demodulation,[],[f880,f875])).
fof(f2426,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl34),
  inference(cnf_transformation,[],[f2426_D])).
fof(f2426_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f7507,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl2 | $spl22)),
  inference(forward_demodulation,[],[f7506,f721])).
fof(f7506,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))) | ($spl2 | $spl22)),
  inference(forward_demodulation,[],[f7505,f721])).
fof(f7505,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))) | ($spl2 | $spl22)),
  inference(forward_demodulation,[],[f7504,f716])).
fof(f7504,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))) | ($spl2 | $spl22)),
  inference(subsumption_resolution,[],[f7495,f721])).
fof(f7495,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))) | ($spl2 | $spl22)),
  inference(superposition,[],[f5341,f714])).
fof(f5341,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | ($spl2 | $spl22)),
  inference(resolution,[],[f4810,f687])).
fof(f4810,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))) | ($spl2 | $spl22)),
  inference(backward_demodulation,[],[f1015,f3854])).
fof(f3854,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) | $spl2),
  inference(forward_demodulation,[],[f3584,f514])).
fof(f3584,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))))) | $spl2),
  inference(backward_demodulation,[],[f875,f742])).
fof(f6158,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f6071,f5941])).
fof(f5941,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl14),
  inference(equality_resolution,[],[f882])).
fof(f6071,plain,(
  ( ! [X308,X307] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X307),s('type.integer.int',X308))))) ) | $spl14),
  inference(resolution,[],[f5941,f699])).
fof(f6157,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f6070,f5941])).
fof(f6070,plain,(
  ( ! [X306,X305] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X305),s('type.integer.int',X306))))) ) | $spl14),
  inference(resolution,[],[f5941,f698])).
fof(f6156,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f6067,f5941])).
fof(f6067,plain,(
  ( ! [X300,X299] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X299),s('type.integer.int',X300))))) ) | $spl14),
  inference(resolution,[],[f5941,f690])).
fof(f6155,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f6066,f5941])).
fof(f6066,plain,(
  ( ! [X298,X297] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X297),s('type.integer.int',X298))))) ) | $spl14),
  inference(resolution,[],[f5941,f684])).
fof(f6154,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f6065,f5941])).
fof(f6065,plain,(
  ( ! [X296,X295] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X295),s('type.integer.int',X296))))) ) | $spl14),
  inference(resolution,[],[f5941,f682])).
fof(f6151,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f6059,f5941])).
fof(f6059,plain,(
  ( ! [X281] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X281))))))) ) | $spl14),
  inference(resolution,[],[f5941,f701])).
fof(f6150,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f6058,f5941])).
fof(f6058,plain,(
  ( ! [X280] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X280))))) ) | $spl14),
  inference(resolution,[],[f5941,f701])).
fof(f6149,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f6056,f5941])).
fof(f6056,plain,(
  ( ! [X279,X278] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X278),s('type.integer.int',X279))))) ) | $spl14),
  inference(resolution,[],[f5941,f700])).
fof(f6148,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f6055,f5941])).
fof(f6055,plain,(
  ( ! [X277,X276] : (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X276),s('type.integer.int',X277))))) ) | $spl14),
  inference(resolution,[],[f5941,f700])).
fof(f6147,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f6054,f5941])).
fof(f6054,plain,(
  ( ! [X275,X274] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X274),s('type.integer.int',X275))))) ) | $spl14),
  inference(resolution,[],[f5941,f699])).
fof(f6146,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f6053,f5941])).
fof(f6053,plain,(
  ( ! [X273,X272] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X272),s('type.integer.int',X273))))) ) | $spl14),
  inference(resolution,[],[f5941,f698])).
fof(f6145,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f6048,f5941])).
fof(f6048,plain,(
  ( ! [X263,X262] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X262),s('type.integer.int',X263))))) ) | $spl14),
  inference(resolution,[],[f5941,f690])).
fof(f6144,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f6045,f5941])).
fof(f6045,plain,(
  ( ! [X257,X256] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X256),s('type.integer.int',X257))))) ) | $spl14),
  inference(resolution,[],[f5941,f684])).
fof(f6143,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f6044,f5941])).
fof(f6044,plain,(
  ( ! [X255,X254] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X254),s('type.integer.int',X255))))) ) | $spl14),
  inference(resolution,[],[f5941,f682])).
fof(f6057,plain,(
  $false | $spl14),
  inference(resolution,[],[f5941,f704])).
fof(f5902,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5817,f5692])).
fof(f5692,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(equality_resolution,[],[f879])).
fof(f5817,plain,(
  ( ! [X300,X299] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X299),s('type.integer.int',X300))))) ) | $spl8),
  inference(resolution,[],[f5692,f689])).
fof(f5901,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5816,f5692])).
fof(f5816,plain,(
  ( ! [X298,X297] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X297),s('type.integer.int',X298))))) ) | $spl8),
  inference(resolution,[],[f5692,f685])).
fof(f5900,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5815,f5692])).
fof(f5815,plain,(
  ( ! [X296,X295] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X295),s('type.integer.int',X296))))) ) | $spl8),
  inference(resolution,[],[f5692,f683])).
fof(f5899,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5814,f5692])).
fof(f5814,plain,(
  ( ! [X294,X293] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X293),s('type.integer.int',X294))))) ) | $spl8),
  inference(resolution,[],[f5692,f681])).
fof(f5898,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5813,f5692])).
fof(f5813,plain,(
  ( ! [X292,X291] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X291),s('type.integer.int',X292))))) ) | $spl8),
  inference(resolution,[],[f5692,f680])).
fof(f5897,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5810,f5692])).
fof(f5810,plain,(
  ( ! [X283,X284] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X283),s('type.integer.int',X284))))) ) | $spl8),
  inference(resolution,[],[f5692,f648])).
fof(f5896,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5809,f5692])).
fof(f5809,plain,(
  ( ! [X282,X281] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X281),s('type.integer.int',X282))))) ) | $spl8),
  inference(resolution,[],[f5692,f648])).
fof(f5895,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5808,f5692])).
fof(f5808,plain,(
  ( ! [X280,X279] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X279),s('type.integer.int',X280))))) ) | $spl8),
  inference(resolution,[],[f5692,f617])).
fof(f5894,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5807,f5692])).
fof(f5807,plain,(
  ( ! [X277,X278] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X277),s('type.integer.int',X278))))) ) | $spl8),
  inference(resolution,[],[f5692,f617])).
fof(f5893,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5799,f5692])).
fof(f5799,plain,(
  ( ! [X261,X262] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X261),s('type.integer.int',X262))))) ) | $spl8),
  inference(resolution,[],[f5692,f689])).
fof(f5890,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5796,f5692])).
fof(f5796,plain,(
  ( ! [X255,X256] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X255),s('type.integer.int',X256))))) ) | $spl8),
  inference(resolution,[],[f5692,f685])).
fof(f5889,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5795,f5692])).
fof(f5795,plain,(
  ( ! [X253,X254] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X253),s('type.integer.int',X254))))) ) | $spl8),
  inference(resolution,[],[f5692,f683])).
fof(f5888,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5794,f5692])).
fof(f5794,plain,(
  ( ! [X251,X252] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X251),s('type.integer.int',X252))))) ) | $spl8),
  inference(resolution,[],[f5692,f681])).
fof(f5887,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f5793,f5692])).
fof(f5793,plain,(
  ( ! [X249,X250] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X249),s('type.integer.int',X250))))) ) | $spl8),
  inference(resolution,[],[f5692,f680])).
fof(f5802,plain,(
  $false | ($spl2 | $spl8 | $spl22)),
  inference(resolution,[],[f5692,f4810])).
fof(f5823,plain,(
  $false | $spl8),
  inference(resolution,[],[f5692,f667])).
fof(f5489,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl2 | $spl22)),
  inference(forward_demodulation,[],[f5488,f721])).
fof(f5488,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))) | ($spl2 | $spl22)),
  inference(forward_demodulation,[],[f5487,f721])).
fof(f5487,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))) | ($spl2 | $spl22)),
  inference(forward_demodulation,[],[f5486,f716])).
fof(f5486,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))) | ($spl2 | $spl22)),
  inference(subsumption_resolution,[],[f5485,f704])).
fof(f5485,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))) | ($spl2 | $spl22)),
  inference(forward_demodulation,[],[f5463,f721])).
fof(f5463,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))) | ($spl2 | $spl22)),
  inference(superposition,[],[f4810,f714])).
fof(f2411,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))) ) | $spl30),
  inference(cnf_transformation,[],[f2411_D])).
fof(f2411_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f5433,plain,(
  ( ! [X161,X162] : (p(s(bool,X161)) | ~p(s(bool,X162)) | s(bool,X162) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))) ) | ($spl2 | $spl22)),
  inference(resolution,[],[f4810,f735])).
fof(f5359,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))))) | s(bool,X22) != s(bool,X24)) ) | ($spl2 | $spl22)),
  inference(resolution,[],[f4810,f526])).
fof(f5357,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))) ) | ($spl2 | $spl22)),
  inference(resolution,[],[f4810,f525])).
fof(f2412,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))) ) | $spl32),
  inference(cnf_transformation,[],[f2412_D])).
fof(f2412_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f5356,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))))) | s(bool,X13) != s(bool,X14)) ) | ($spl2 | $spl22)),
  inference(resolution,[],[f4810,f524])).
fof(f5355,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))) ) | ($spl2 | $spl22)),
  inference(resolution,[],[f4810,f524])).
fof(f5354,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))))) | s(bool,X7) != s(bool,X9)) ) | ($spl2 | $spl22)),
  inference(resolution,[],[f4810,f522])).
fof(f5353,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))))) | s(bool,X4) != s(bool,X5)) ) | ($spl2 | $spl22)),
  inference(resolution,[],[f4810,f520])).
fof(f5352,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))) ) | ($spl2 | $spl22)),
  inference(resolution,[],[f4810,f520])).
fof(f5333,plain,(
  $false | ($spl22 | $spl24)),
  inference(subsumption_resolution,[],[f5332,f667])).
fof(f5332,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl22 | $spl24)),
  inference(forward_demodulation,[],[f1017,f1015])).
fof(f2449,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))))),
  inference(resolution,[],[f744,f686])).
fof(f2450,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))))),
  inference(resolution,[],[f744,f686])).
fof(f2452,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))))),
  inference(resolution,[],[f744,f691])).
fof(f2453,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12)),
  inference(resolution,[],[f744,f696])).
fof(f4122,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12)),
  inference(resolution,[],[f889,f649])).
fof(f4123,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12)),
  inference(resolution,[],[f889,f649])).
fof(f4129,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))))),
  inference(resolution,[],[f889,f691])).
fof(f4130,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))))),
  inference(resolution,[],[f889,f696])).
fof(f4759,plain,(
  ( ! [X161,X162] : (p(s(bool,X161)) | ~p(s(bool,X162)) | s(bool,X162) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl21),
  inference(resolution,[],[f1013,f735])).
fof(f4685,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X22) != s(bool,X24)) ) | $spl21),
  inference(resolution,[],[f1013,f526])).
fof(f4683,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl21),
  inference(resolution,[],[f1013,f525])).
fof(f4682,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X13) != s(bool,X14)) ) | $spl21),
  inference(resolution,[],[f1013,f524])).
fof(f4681,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl21),
  inference(resolution,[],[f1013,f524])).
fof(f4680,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X7) != s(bool,X9)) ) | $spl21),
  inference(resolution,[],[f1013,f522])).
fof(f4679,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X4) != s(bool,X5)) ) | $spl21),
  inference(resolution,[],[f1013,f520])).
fof(f4678,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl21),
  inference(resolution,[],[f1013,f520])).
fof(f4420,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl2),
  inference(forward_demodulation,[],[f4419,f721])).
fof(f4419,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) | $spl2),
  inference(forward_demodulation,[],[f4418,f721])).
fof(f4418,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) | $spl2),
  inference(forward_demodulation,[],[f4417,f716])).
fof(f4417,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) | $spl2),
  inference(subsumption_resolution,[],[f4416,f704])).
fof(f4416,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) | $spl2),
  inference(forward_demodulation,[],[f4394,f721])).
fof(f4394,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) | $spl2),
  inference(superposition,[],[f3854,f714])).
fof(f4354,plain,(
  ( ! [X161,X162] : (p(s(bool,X161)) | ~p(s(bool,X162)) | s(bool,X162) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl2),
  inference(resolution,[],[f3854,f735])).
fof(f4280,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))))) | s(bool,X22) != s(bool,X24)) ) | $spl2),
  inference(resolution,[],[f3854,f526])).
fof(f4278,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl2),
  inference(resolution,[],[f3854,f525])).
fof(f2726,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl40),
  inference(cnf_transformation,[],[f2726_D])).
fof(f2726_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f4277,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))))) | s(bool,X13) != s(bool,X14)) ) | $spl2),
  inference(resolution,[],[f3854,f524])).
fof(f4276,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl2),
  inference(resolution,[],[f3854,f524])).
fof(f4275,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))))) | s(bool,X7) != s(bool,X9)) ) | $spl2),
  inference(resolution,[],[f3854,f522])).
fof(f4274,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))))) | s(bool,X4) != s(bool,X5)) ) | $spl2),
  inference(resolution,[],[f3854,f520])).
fof(f4273,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl2),
  inference(resolution,[],[f3854,f520])).
fof(f4214,plain,(
  ( ! [X163,X164] : (p(s(bool,X163)) | ~p(s(bool,X164)) | s(bool,X164) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) )),
  inference(resolution,[],[f889,f735])).
fof(f4140,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X24) != s(bool,X26)) )),
  inference(resolution,[],[f889,f526])).
fof(f4138,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) )),
  inference(resolution,[],[f889,f525])).
fof(f4137,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X15) != s(bool,X16)) )),
  inference(resolution,[],[f889,f524])).
fof(f4136,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) )),
  inference(resolution,[],[f889,f524])).
fof(f4135,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f889,f522])).
fof(f4134,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | p(s(bool,X7)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X6) != s(bool,X7)) )),
  inference(resolution,[],[f889,f520])).
fof(f4133,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))) )),
  inference(resolution,[],[f889,f520])).
fof(f4118,plain,(
  $false | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f4117,f667])).
fof(f4117,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl2 | $spl4)),
  inference(forward_demodulation,[],[f877,f875])).
fof(f2763,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))))),
  inference(resolution,[],[f743,f686])).
fof(f2764,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))))),
  inference(resolution,[],[f743,f686])).
fof(f2766,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))))),
  inference(resolution,[],[f743,f691])).
fof(f2767,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13)),
  inference(resolution,[],[f743,f696])).
fof(f2888,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13)),
  inference(resolution,[],[f749,f649])).
fof(f2889,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13)),
  inference(resolution,[],[f749,f649])).
fof(f2895,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))))),
  inference(resolution,[],[f749,f691])).
fof(f2896,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))))),
  inference(resolution,[],[f749,f696])).
fof(f3543,plain,(
  ( ! [X161,X162] : (p(s(bool,X161)) | ~p(s(bool,X162)) | s(bool,X162) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl1),
  inference(resolution,[],[f873,f735])).
fof(f3469,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X22) != s(bool,X24)) ) | $spl1),
  inference(resolution,[],[f873,f526])).
fof(f3467,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl1),
  inference(resolution,[],[f873,f525])).
fof(f3466,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X13) != s(bool,X14)) ) | $spl1),
  inference(resolution,[],[f873,f524])).
fof(f3465,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl1),
  inference(resolution,[],[f873,f524])).
fof(f3464,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X7) != s(bool,X9)) ) | $spl1),
  inference(resolution,[],[f873,f522])).
fof(f3463,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X4) != s(bool,X5)) ) | $spl1),
  inference(resolution,[],[f873,f520])).
fof(f3462,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl1),
  inference(resolution,[],[f873,f520])).
fof(f3200,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0'))))),
  inference(forward_demodulation,[],[f3199,f721])).
fof(f3199,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))))),
  inference(forward_demodulation,[],[f3198,f721])).
fof(f3198,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))))),
  inference(forward_demodulation,[],[f3197,f716])).
fof(f3197,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))))),
  inference(subsumption_resolution,[],[f3196,f704])).
fof(f3196,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))))),
  inference(forward_demodulation,[],[f3172,f721])).
fof(f3172,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))))),
  inference(superposition,[],[f742,f714])).
fof(f3120,plain,(
  ( ! [X161,X162] : (p(s(bool,X161)) | ~p(s(bool,X162)) | s(bool,X162) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))))) )),
  inference(resolution,[],[f742,f735])).
fof(f3046,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) | s(bool,X22) != s(bool,X24)) )),
  inference(resolution,[],[f742,f526])).
fof(f3044,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))))) )),
  inference(resolution,[],[f742,f525])).
fof(f3043,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) | s(bool,X13) != s(bool,X14)) )),
  inference(resolution,[],[f742,f524])).
fof(f3042,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))))) )),
  inference(resolution,[],[f742,f524])).
fof(f3041,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) | s(bool,X7) != s(bool,X9)) )),
  inference(resolution,[],[f742,f522])).
fof(f3040,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13))))) | s(bool,X4) != s(bool,X5)) )),
  inference(resolution,[],[f742,f520])).
fof(f3039,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int',sK13)))))) )),
  inference(resolution,[],[f742,f520])).
fof(f2980,plain,(
  ( ! [X163,X164] : (p(s(bool,X163)) | ~p(s(bool,X164)) | s(bool,X164) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) )),
  inference(resolution,[],[f749,f735])).
fof(f2906,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X24) != s(bool,X26)) )),
  inference(resolution,[],[f749,f526])).
fof(f2904,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X19) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) )),
  inference(resolution,[],[f749,f525])).
fof(f2903,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X17) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X15) != s(bool,X16)) )),
  inference(resolution,[],[f749,f524])).
fof(f2902,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) )),
  inference(resolution,[],[f749,f524])).
fof(f2901,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X9) != s(bool,X11)) )),
  inference(resolution,[],[f749,f522])).
fof(f2900,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | p(s(bool,X7)) | s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))) | s(bool,X6) != s(bool,X7)) )),
  inference(resolution,[],[f749,f520])).
fof(f2899,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,X3) != s(bool,X5) | s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0')))) )),
  inference(resolution,[],[f749,f520])).
fof(f2845,plain,(
  ( ! [X156,X157] : (p(s(bool,X156)) | p(s(bool,X157)) | s(bool,X157) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f743,f735])).
fof(f2775,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X26) != s(bool,X27)) )),
  inference(resolution,[],[f743,f526])).
fof(f2774,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X23) | s(bool,X24) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f743,f526])).
fof(f2772,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f743,f525])).
fof(f2771,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f743,f524])).
fof(f2770,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X11) != s(bool,X12)) )),
  inference(resolution,[],[f743,f522])).
fof(f2769,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X8) | s(bool,X9) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f743,f522])).
fof(f2768,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f743,f520])).
fof(f2738,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl2),
  inference(forward_demodulation,[],[f2737,f721])).
fof(f2737,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) | $spl2),
  inference(forward_demodulation,[],[f2736,f721])).
fof(f2736,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) | $spl2),
  inference(forward_demodulation,[],[f2735,f716])).
fof(f2735,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) | $spl2),
  inference(subsumption_resolution,[],[f2734,f704])).
fof(f2734,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) | $spl2),
  inference(forward_demodulation,[],[f2710,f721])).
fof(f2710,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) | $spl2),
  inference(superposition,[],[f2264,f714])).
fof(f2264,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) | $spl2),
  inference(forward_demodulation,[],[f2263,f514])).
fof(f2263,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0')))))) | $spl2),
  inference(forward_demodulation,[],[f742,f875])).
fof(f2670,plain,(
  ( ! [X161,X162] : (p(s(bool,X161)) | ~p(s(bool,X162)) | s(bool,X162) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl2),
  inference(resolution,[],[f2264,f735])).
fof(f2596,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))))) | s(bool,X22) != s(bool,X24)) ) | $spl2),
  inference(resolution,[],[f2264,f526])).
fof(f2594,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl2),
  inference(resolution,[],[f2264,f525])).
fof(f2593,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))))) | s(bool,X13) != s(bool,X14)) ) | $spl2),
  inference(resolution,[],[f2264,f524])).
fof(f2592,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl2),
  inference(resolution,[],[f2264,f524])).
fof(f2591,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))))) | s(bool,X7) != s(bool,X9)) ) | $spl2),
  inference(resolution,[],[f2264,f522])).
fof(f2590,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))))) | s(bool,X4) != s(bool,X5)) ) | $spl2),
  inference(resolution,[],[f2264,f520])).
fof(f2589,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))) ) | $spl2),
  inference(resolution,[],[f2264,f520])).
fof(f2531,plain,(
  ( ! [X156,X157] : (p(s(bool,X156)) | p(s(bool,X157)) | s(bool,X157) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f744,f735])).
fof(f2461,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X26) != s(bool,X27)) )),
  inference(resolution,[],[f744,f526])).
fof(f2460,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X23) | s(bool,X24) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f744,f526])).
fof(f2458,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f744,f525])).
fof(f2457,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f744,f524])).
fof(f2456,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X11) != s(bool,X12)) )),
  inference(resolution,[],[f744,f522])).
fof(f2455,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X8) | s(bool,X9) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f744,f522])).
fof(f2454,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f744,f520])).
fof(f2424,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl2 | $spl22)),
  inference(forward_demodulation,[],[f2423,f721])).
fof(f2423,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))) | ($spl2 | $spl22)),
  inference(forward_demodulation,[],[f2422,f721])).
fof(f2422,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))) | ($spl2 | $spl22)),
  inference(forward_demodulation,[],[f2421,f716])).
fof(f2421,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))) | ($spl2 | $spl22)),
  inference(subsumption_resolution,[],[f2420,f704])).
fof(f2420,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))) | ($spl2 | $spl22)),
  inference(forward_demodulation,[],[f2396,f721])).
fof(f2396,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))) | ($spl2 | $spl22)),
  inference(superposition,[],[f2265,f714])).
fof(f2265,plain,(
  ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))) | ($spl2 | $spl22)),
  inference(forward_demodulation,[],[f2264,f1015])).
fof(f2366,plain,(
  ( ! [X161,X162] : (p(s(bool,X161)) | ~p(s(bool,X162)) | s(bool,X162) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))) ) | ($spl2 | $spl22)),
  inference(resolution,[],[f2265,f735])).
fof(f2292,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X23) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))))) | s(bool,X22) != s(bool,X24)) ) | ($spl2 | $spl22)),
  inference(resolution,[],[f2265,f526])).
fof(f2290,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | ~p(s(bool,X18)) | s(bool,X17) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))) ) | ($spl2 | $spl22)),
  inference(resolution,[],[f2265,f525])).
fof(f2289,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X15) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))))) | s(bool,X13) != s(bool,X14)) ) | ($spl2 | $spl22)),
  inference(resolution,[],[f2265,f524])).
fof(f2288,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,X12) | s(bool,X11) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))) ) | ($spl2 | $spl22)),
  inference(resolution,[],[f2265,f524])).
fof(f2287,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))))) | s(bool,X7) != s(bool,X9)) ) | ($spl2 | $spl22)),
  inference(resolution,[],[f2265,f522])).
fof(f2286,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X6) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))))) | s(bool,X4) != s(bool,X5)) ) | ($spl2 | $spl22)),
  inference(resolution,[],[f2265,f520])).
fof(f2285,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))) ) | ($spl2 | $spl22)),
  inference(resolution,[],[f2265,f520])).
fof(f1212,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f1127,f884])).
fof(f1127,plain,(
  ( ! [X300,X299] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X299),s('type.integer.int',X300))))) ) | $spl18),
  inference(resolution,[],[f884,f689])).
fof(f1211,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f1126,f884])).
fof(f1126,plain,(
  ( ! [X298,X297] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X297),s('type.integer.int',X298))))) ) | $spl18),
  inference(resolution,[],[f884,f685])).
fof(f1210,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f1125,f884])).
fof(f1125,plain,(
  ( ! [X296,X295] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X295),s('type.integer.int',X296))))) ) | $spl18),
  inference(resolution,[],[f884,f683])).
fof(f1209,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f1124,f884])).
fof(f1124,plain,(
  ( ! [X294,X293] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X293),s('type.integer.int',X294))))) ) | $spl18),
  inference(resolution,[],[f884,f681])).
fof(f1208,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f1123,f884])).
fof(f1123,plain,(
  ( ! [X292,X291] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X291),s('type.integer.int',X292))))) ) | $spl18),
  inference(resolution,[],[f884,f680])).
fof(f1207,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f1120,f884])).
fof(f1120,plain,(
  ( ! [X283,X284] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X283),s('type.integer.int',X284))))) ) | $spl18),
  inference(resolution,[],[f884,f648])).
fof(f1206,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f1119,f884])).
fof(f1119,plain,(
  ( ! [X282,X281] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X281),s('type.integer.int',X282))))) ) | $spl18),
  inference(resolution,[],[f884,f648])).
fof(f1205,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f1118,f884])).
fof(f1118,plain,(
  ( ! [X280,X279] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X279),s('type.integer.int',X280))))) ) | $spl18),
  inference(resolution,[],[f884,f617])).
fof(f1204,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f1117,f884])).
fof(f1117,plain,(
  ( ! [X277,X278] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X277),s('type.integer.int',X278))))) ) | $spl18),
  inference(resolution,[],[f884,f617])).
fof(f1203,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f1110,f884])).
fof(f1110,plain,(
  ( ! [X261,X262] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X261),s('type.integer.int',X262))))) ) | $spl18),
  inference(resolution,[],[f884,f689])).
fof(f1200,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f1107,f884])).
fof(f1107,plain,(
  ( ! [X255,X256] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X255),s('type.integer.int',X256))))) ) | $spl18),
  inference(resolution,[],[f884,f685])).
fof(f1199,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f1106,f884])).
fof(f1106,plain,(
  ( ! [X253,X254] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X253),s('type.integer.int',X254))))) ) | $spl18),
  inference(resolution,[],[f884,f683])).
fof(f1198,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f1105,f884])).
fof(f1105,plain,(
  ( ! [X251,X252] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X251),s('type.integer.int',X252))))) ) | $spl18),
  inference(resolution,[],[f884,f681])).
fof(f1197,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f1104,f884])).
fof(f1104,plain,(
  ( ! [X249,X250] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X249),s('type.integer.int',X250))))) ) | $spl18),
  inference(resolution,[],[f884,f680])).
fof(f1133,plain,(
  $false | $spl18),
  inference(resolution,[],[f884,f667])).
fof(f975,plain,(
  ( ! [X156,X157] : (p(s(bool,X156)) | p(s(bool,X157)) | s(bool,X157) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f744,f735])).
fof(f905,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X26) != s(bool,X27)) )),
  inference(resolution,[],[f744,f526])).
fof(f904,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X23) | s(bool,X24) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f744,f526])).
fof(f902,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f744,f525])).
fof(f901,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f744,f524])).
fof(f900,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))) | s(bool,X11) != s(bool,X12)) )),
  inference(resolution,[],[f744,f522])).
fof(f899,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X8) | s(bool,X9) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f744,f522])).
fof(f898,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) )),
  inference(resolution,[],[f744,f520])).
fof(f897,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12)),
  inference(resolution,[],[f744,f696])).
fof(f896,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))))),
  inference(resolution,[],[f744,f691])).
fof(f894,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))))),
  inference(resolution,[],[f744,f686])).
fof(f893,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK12),s('type.integer.int','const.integer.int_|48|_0'))))),
  inference(resolution,[],[f744,f686])).
fof(f835,plain,(
  ( ! [X156,X157] : (p(s(bool,X156)) | p(s(bool,X157)) | s(bool,X157) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f743,f735])).
fof(f765,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X26) != s(bool,X27)) )),
  inference(resolution,[],[f743,f526])).
fof(f764,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X23) | s(bool,X24) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f743,f526])).
fof(f762,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f743,f525])).
fof(f761,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f743,f524])).
fof(f760,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))) | s(bool,X11) != s(bool,X12)) )),
  inference(resolution,[],[f743,f522])).
fof(f759,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X7) != s(bool,X8) | s(bool,X9) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f743,f522])).
fof(f758,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) )),
  inference(resolution,[],[f743,f520])).
fof(f757,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13)),
  inference(resolution,[],[f743,f696])).
fof(f756,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK13))))),
  inference(resolution,[],[f743,f691])).
fof(f754,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))))),
  inference(resolution,[],[f743,f686])).
fof(f753,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK13) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',sK13),s('type.integer.int','const.integer.int_|48|_0'))))),
  inference(resolution,[],[f743,f686])).
% SZS output end Proof for thm.integer
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1150
% Passive clauses: 30513
% Generated clauses: 72270
% Final active clauses: 429
% Final passive clauses: 13919
% Input formulas: 261
% Initial clauses: 232
% 
% Selected by SInE selection: 82
% SInE iterations: 1
% 
% Duplicate literals: 695
% Trivial inequalities: 56
% Fw subsumption resolutions: 1693
% Fw demodulations: 16819
% Bw demodulations: 291
% 
% Simple tautologies: 2235
% Equational tautologies: 32
% Forward subsumptions: 20622
% Fw demodulations to eq. taut.: 66
% 
% Binary resolution: 14251
% Forward superposition: 24036
% Backward superposition: 11976
% Self superposition: 15
% Equality resolution: 19
% 
% Split clauses: 711
% Split components: 1503
% SAT solver clauses: 889
% SAT solver unit clauses: 65
% SAT solver binary clauses: 340
% 
% Sat splits: 711
% Sat splitting refutations: 172
% 
% Memory used [KB]: 17142
% Time elapsed: 1.316 s
% ------------------------------
% ------------------------------
% Success in time 13.06 s
