Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.pred_set
% (62541)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 113
% Passive clauses: 2404
% Generated clauses: 18323
% Final active clauses: 86
% Final passive clauses: 2161
% Input formulas: 10
% Initial clauses: 14
% 
% Duplicate literals: 1206
% Trivial inequalities: 214
% Fw subsumption resolutions: 2263
% Fw demodulations: 1457
% Bw demodulations: 23
% 
% Simple tautologies: 314
% Equational tautologies: 97
% Forward subsumptions: 7964
% 
% Binary resolution: 82
% Forward superposition: 7155
% Backward superposition: 5075
% Self superposition: 367
% Equality factoring: 267
% Equality resolution: 12
% 
% Split clauses: 959
% Split components: 1926
% SAT solver clauses: 2422
% SAT solver unit clauses: 316
% SAT solver binary clauses: 684
% 
% Sat splits: 959
% Sat splitting refutations: 14
% 
% TWLsolver clauses: 2412
% TWLsolver calls for satisfiability: 99
% 
% Memory used [KB]: 7419
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 294 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (62562)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 108
% Passive clauses: 398
% Generated clauses: 11288
% Final active clauses: 80
% Final passive clauses: 241
% Input formulas: 10
% Initial clauses: 14
% 
% Duplicate literals: 1168
% Trivial inequalities: 126
% Fw subsumption resolutions: 919
% Fw demodulations: 895
% Bw demodulations: 72
% Condensations: 1501
% Global subsumptions: 1
% 
% Simple tautologies: 179
% Equational tautologies: 62
% Forward subsumptions: 4777
% 
% Binary resolution: 198
% Factoring: 22
% Forward superposition: 3328
% Backward superposition: 2355
% Self superposition: 174
% Equality factoring: 146
% Equality resolution: 6
% 
% Split clauses: 174
% Split components: 349
% SAT solver clauses: 1067
% SAT solver unit clauses: 21
% SAT solver binary clauses: 66
% SAT solver learnt clauses: 2
% 
% Sat splits: 174
% Sat splitting refutations: 11
% 
% TWLsolver clauses: 1063
% TWLsolver calls for satisfiability: 128
% 
% Memory used [KB]: 6652
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f58478,plain,(
  $false),
  inference(subsumption_resolution,[],[f58477,f43])).
fof(f43,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.PSUBSET_TRANS',aHL_TRUTH)).
fof(f58477,plain,(
  ~p(s(bool,'T_0'))),
  inference(forward_demodulation,[],[f58476,f474])).
fof(f474,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))),
  inference(subsumption_resolution,[],[f467,f43])).
fof(f467,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))),
  inference(superposition,[],[f80,f63])).
fof(f63,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))),
  inference(subsumption_resolution,[],[f58,f44])).
fof(f44,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f14])).
fof(f14,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pred_set.PSUBSET_TRANS',aHL_FALSITY)).
fof(f58,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))),
  inference(superposition,[],[f40,f45])).
fof(f45,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.PSUBSET_TRANS',aHL_BOOL_CASES)).
fof(f40,plain,(
  p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),
  inference(cnf_transformation,[],[f36])).
fof(f36,plain,(
  p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2)))) & p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))) & ~p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f27])).
fof(f27,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))))),
  inference(flattening,[],[f26])).
fof(f26,plain,(
  ? [X0,X1,X2,X3] : ((p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))) & ~p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))))),
  inference(ennf_transformation,[],[f13])).
fof(f13,plain,(
  ~! [X0,X1,X2,X3] : ((p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))) => p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))))),
  inference(flattening,[],[f12])).
fof(f12,plain,(
  ~! [X0] : ! [X1,X2,X3] : ((p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))) => p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : ((p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))) & p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X9))))) => p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X9)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : ((p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))) & p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X9))))) => p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X9)))))),
  file('thm.pred_set.PSUBSET_TRANS','thm.pred_set.PSUBSET_TRANS')).
fof(f80,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) )),
  inference(subsumption_resolution,[],[f79,f44])).
fof(f79,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | ~p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) )),
  inference(superposition,[],[f49,f45])).
fof(f49,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & s(fun(X0,bool),X1) != s(fun(X0,bool),X2))) & (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & s(fun(X0,bool),X1) != s(fun(X0,bool),X2))) & ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) | p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(nnf_transformation,[],[f21])).
fof(f21,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & s(fun(X0,bool),X1) != s(fun(X0,bool),X2)))),
  inference(flattening,[],[f20])).
fof(f20,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & ~s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : ! [X7,X8] : (p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))) <=> (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))) & ~s(fun(X6,bool),X7) = s(fun(X6,bool),X8)))),
  file('thm.pred_set.PSUBSET_TRANS','thm.pred_set.PSUBSET_DEF')).
fof(f58476,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))))),
  inference(forward_demodulation,[],[f44913,f56656])).
fof(f56656,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),sK3)),
  inference(global_subsumption,[],[f42,f41,f40,f43,f45,f53,f54,f56,f63,f62,f78,f91,f111,f188,f232,f256,f423,f274,f427,f428,f429,f430,f431,f474,f544,f545,f551,f571,f556,f557,f558,f559,f560,f594,f619,f644,f712,f2821,f2663,f2659,f2653,f1382,f3009,f1558,f3034,f1644,f1645,f2171,f1687,f1690,f3064,f1727,f1771,f1977,f3078,f1894,f3428,f3514,f3515,f3516,f3517,f3518,f3519,f3520,f3521,f3522,f3523,f3524,f3525,f3526,f90,f89,f44,f8160,f9288,f10880,f14852,f17037,f17040,f17506,f20317,f20318,f20323,f20322,f20320,f20319,f20316,f20313,f12374,f12375,f12377,f12380,f27658,f27657,f27513,f36372,f36376,f36380,f34535,f34539,f34541,f85,f36374,f34537,f86,f88,f87,f3067,f3065,f426,f424,f52,f48,f539,f14901,f14899,f14864,f14826,f14708,f14662,f3066,f3063,f3056,f476,f425,f80,f51,f50,f49,f14665,f12474,f3054,f1729,f1689,f1646,f1560,f585,f275,f233,f2655,f511,f53887,f53885,f54287,f54290,f54999,f55000,f55001,f55366,f55368,f55369,f55370,f56253,f56254,f56255,f56256,f56267,f56297,f56567,f56535,f56654,f56614,f56612,f56458,f56457,f56456,f55367,f2661,f2657,f2609,f657,f650,f614,f475,f422,f84,f65,f56570])).
fof(f56570,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),sK3) | p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))),
  inference(subsumption_resolution,[],[f56490,f43])).
fof(f56490,plain,(
  ~p(s(bool,'T_0')) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),sK3) | p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))),
  inference(superposition,[],[f51,f56267])).
fof(f65,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),
  inference(subsumption_resolution,[],[f60,f44])).
fof(f60,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),
  inference(superposition,[],[f41,f45])).
fof(f84,plain,(
  s(fun(sK0,bool),sK2) != s(fun(sK0,bool),sK3)),
  inference(subsumption_resolution,[],[f82,f43])).
fof(f82,plain,(
  ~p(s(bool,'T_0')) | s(fun(sK0,bool),sK2) != s(fun(sK0,bool),sK3)),
  inference(superposition,[],[f50,f65])).
fof(f422,plain,(
  ( ! [X13] : (s(bool,X13) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(subsumption_resolution,[],[f268,f43])).
fof(f268,plain,(
  ( ! [X13] : (~p(s(bool,'T_0')) | s(bool,X13) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f42,f54])).
fof(f475,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),
  inference(subsumption_resolution,[],[f468,f43])).
fof(f468,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))),
  inference(superposition,[],[f80,f65])).
fof(f614,plain,(
  ( ! [X0] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),X0)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X0))))) )),
  inference(subsumption_resolution,[],[f598,f43])).
fof(f598,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),X0)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X0))))) )),
  inference(superposition,[],[f52,f475])).
fof(f650,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f645,f43])).
fof(f645,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f619,f54])).
fof(f657,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f652,f43])).
fof(f652,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f644,f54])).
fof(f2609,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X2)) )),
  inference(condensation,[],[f2608])).
fof(f2608,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,X81) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))) | s(bool,X81) = s(bool,X82) | s(bool,'T_0') = s(bool,X80) | s(bool,X82) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(duplicate_literal_removal,[],[f1361])).
fof(f1361,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,X81) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))) | s(bool,X81) = s(bool,X82) | s(bool,'T_0') = s(bool,X80) | s(bool,X82) = s(bool,X83) | s(bool,'T_0') = s(bool,X83) | s(bool,'T_0') = s(bool,X80)) )),
  inference(superposition,[],[f188,f422])).
fof(f2657,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))) | s(bool,X0) = s(bool,X1)) )),
  inference(condensation,[],[f2656])).
fof(f2656,plain,(
  ( ! [X146,X147,X148] : (s(bool,'T_0') = s(bool,X146) | s(bool,X147) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))) | s(bool,X147) = s(bool,X148) | s(bool,'T_0') = s(bool,X148)) )),
  inference(duplicate_literal_removal,[],[f1344])).
fof(f1344,plain,(
  ( ! [X146,X147,X148] : (s(bool,'T_0') = s(bool,X146) | s(bool,X147) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))) | s(bool,'T_0') = s(bool,X146) | s(bool,X147) = s(bool,X148) | s(bool,'T_0') = s(bool,X148)) )),
  inference(superposition,[],[f188,f475])).
fof(f2661,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))) | s(bool,X0) = s(bool,X1)) )),
  inference(condensation,[],[f2660])).
fof(f2660,plain,(
  ( ! [X142,X140,X141] : (s(bool,'T_0') = s(bool,X140) | s(bool,X141) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))) | s(bool,X141) = s(bool,X142) | s(bool,'T_0') = s(bool,X142)) )),
  inference(duplicate_literal_removal,[],[f1342])).
fof(f1342,plain,(
  ( ! [X142,X140,X141] : (s(bool,'T_0') = s(bool,X140) | s(bool,X141) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))) | s(bool,'T_0') = s(bool,X140) | s(bool,X141) = s(bool,X142) | s(bool,'T_0') = s(bool,X142)) )),
  inference(superposition,[],[f188,f65])).
fof(f55367,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X2),s(fun(sK0,bool),sK3)))) | s(bool,X3) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X2),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f55336,f43])).
fof(f55336,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X2),s(fun(sK0,bool),sK3)))) | s(bool,X3) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X2),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f424,f475])).
fof(f56456,plain,(
  ( ! [X2,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X1)))) | s(bool,X2) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f56431,f43])).
fof(f56431,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X1)))) | s(bool,X2) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f614,f476])).
fof(f56457,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X4)))) | s(bool,X5) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),X4))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f56433,f43])).
fof(f56433,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X4)))) | s(bool,X5) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),X4))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f614,f54])).
fof(f56458,plain,(
  ( ! [X23,X21,X22,X20] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X20)))) | s(bool,X21) = s(bool,X22) | s(bool,X22) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),X20))) | s(bool,X21) = s(bool,X23) | s(bool,'T_0') = s(bool,X23)) )),
  inference(subsumption_resolution,[],[f56439,f43])).
fof(f56439,plain,(
  ( ! [X23,X21,X22,X20] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X20)))) | s(bool,X21) = s(bool,X22) | s(bool,X22) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),X20))) | s(bool,X21) = s(bool,X23) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f614,f188])).
fof(f56612,plain,(
  ( ! [X0] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),X0)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X0))))) )),
  inference(subsumption_resolution,[],[f56491,f43])).
fof(f56491,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),X0)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X0))))) )),
  inference(superposition,[],[f52,f56267])).
fof(f56614,plain,(
  ( ! [X2,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),sK3)))) | s(bool,X2) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f56492,f43])).
fof(f56492,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),sK3)))) | s(bool,X2) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f424,f56267])).
fof(f56654,plain,(
  ( ! [X3] : (s(fun(sK0,bool),sK1) = s(fun(sK0,bool),sK3) | s(bool,X3) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f56493,f43])).
fof(f56493,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),sK3) | s(bool,X3) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f426,f56267])).
fof(f56535,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))) | s(bool,X0) = s(bool,X1)) )),
  inference(condensation,[],[f56534])).
fof(f56534,plain,(
  ( ! [X19,X20,X18] : (s(bool,'T_0') = s(bool,X18) | s(bool,X19) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))) | s(bool,X19) = s(bool,X20) | s(bool,'T_0') = s(bool,X20)) )),
  inference(duplicate_literal_removal,[],[f56502])).
fof(f56502,plain,(
  ( ! [X19,X20,X18] : (s(bool,'T_0') = s(bool,X18) | s(bool,X19) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))) | s(bool,X19) = s(bool,X20) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f188,f56267])).
fof(f56567,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK1)))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),sK3)),
  inference(subsumption_resolution,[],[f56488,f43])).
fof(f56488,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK1)))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),sK3)),
  inference(superposition,[],[f48,f56267])).
fof(f56297,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))) | s(bool,X0) = s(bool,X1)) )),
  inference(condensation,[],[f56296])).
fof(f56296,plain,(
  ( ! [X14,X15,X16] : (s(bool,X15) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))) | s(bool,'T_0') = s(bool,X14) | s(bool,X15) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) )),
  inference(subsumption_resolution,[],[f56277,f56])).
fof(f56277,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | s(bool,X15) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))) | s(bool,'T_0') = s(bool,X14) | s(bool,X15) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f56253,f188])).
fof(f56267,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))),
  inference(unit_resulting_resolution,[],[f56253,f56])).
fof(f56256,plain,(
  ( ! [X23,X21,X22,X20] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X20)))) | s(bool,X21) = s(bool,X22) | s(bool,X22) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X20))) | s(bool,X21) = s(bool,X23) | s(bool,'T_0') = s(bool,X23)) )),
  inference(subsumption_resolution,[],[f56236,f43])).
fof(f56236,plain,(
  ( ! [X23,X21,X22,X20] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X20)))) | s(bool,X21) = s(bool,X22) | s(bool,X22) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X20))) | s(bool,X21) = s(bool,X23) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f539,f188])).
fof(f56255,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X4)))) | s(bool,X5) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X4))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f56230,f43])).
fof(f56230,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X4)))) | s(bool,X5) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X4))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f539,f54])).
fof(f56254,plain,(
  ( ! [X2,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X1)))) | s(bool,X2) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f56228,f43])).
fof(f56228,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X1)))) | s(bool,X2) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f539,f476])).
fof(f56253,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))),
  inference(subsumption_resolution,[],[f56225,f43])).
fof(f56225,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))),
  inference(superposition,[],[f539,f475])).
fof(f55370,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X65] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X65,bool),X68),s(fun(X65,bool),X67)))) | s(bool,X69) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X65,bool),X68),s(fun(X65,bool),X66))) | s(bool,'T_0') = s(bool,X69) | s(bool,X70) = s(bool,X71) | s(bool,X71) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X65,bool),X66),s(fun(X65,bool),X67))) | s(bool,X70) = s(bool,X72) | s(bool,'T_0') = s(bool,X72)) )),
  inference(subsumption_resolution,[],[f55349,f43])).
fof(f55349,plain,(
  ( ! [X70,X68,X66,X72,X71,X69,X67,X65] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X65,bool),X68),s(fun(X65,bool),X67)))) | s(bool,X69) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X65,bool),X68),s(fun(X65,bool),X66))) | s(bool,'T_0') = s(bool,X69) | s(bool,X70) = s(bool,X71) | s(bool,X71) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X65,bool),X66),s(fun(X65,bool),X67))) | s(bool,X70) = s(bool,X72) | s(bool,'T_0') = s(bool,X72)) )),
  inference(superposition,[],[f424,f188])).
fof(f55369,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X25,bool),X28),s(fun(X25,bool),X27)))) | s(bool,X29) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X25,bool),X28),s(fun(X25,bool),X26))) | s(bool,'T_0') = s(bool,X29) | s(bool,X30) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X25,bool),X26),s(fun(X25,bool),X27))) | s(bool,'T_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f55343,f43])).
fof(f55343,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X25,bool),X28),s(fun(X25,bool),X27)))) | s(bool,X29) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X25,bool),X28),s(fun(X25,bool),X26))) | s(bool,'T_0') = s(bool,X29) | s(bool,X30) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X25,bool),X26),s(fun(X25,bool),X27))) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f424,f54])).
fof(f55368,plain,(
  ( ! [X14,X19,X17,X15,X18,X16] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X14,bool),X17),s(fun(X14,bool),X16)))) | s(bool,X18) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X14,bool),X17),s(fun(X14,bool),X15))) | s(bool,'T_0') = s(bool,X18) | s(bool,X19) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X14,bool),X15),s(fun(X14,bool),X16))) | s(bool,'T_0') = s(bool,X19)) )),
  inference(subsumption_resolution,[],[f55341,f43])).
fof(f55341,plain,(
  ( ! [X14,X19,X17,X15,X18,X16] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X14,bool),X17),s(fun(X14,bool),X16)))) | s(bool,X18) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X14,bool),X17),s(fun(X14,bool),X15))) | s(bool,'T_0') = s(bool,X18) | s(bool,X19) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X14,bool),X15),s(fun(X14,bool),X16))) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f424,f476])).
fof(f55366,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),sK2)))) | s(bool,X1) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(subsumption_resolution,[],[f55335,f43])).
fof(f55335,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),sK2)))) | s(bool,X1) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f424,f474])).
fof(f55001,plain,(
  ( ! [X57,X54,X52,X56,X55,X53,X51] : (s(fun(X51,bool),X52) = s(fun(X51,bool),X53) | s(bool,X54) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X51,bool),X53),s(fun(X51,bool),X52))) | s(bool,'T_0') = s(bool,X54) | s(bool,X55) = s(bool,X56) | s(bool,X56) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X51,bool),X52),s(fun(X51,bool),X53))) | s(bool,X55) = s(bool,X57) | s(bool,'T_0') = s(bool,X57)) )),
  inference(subsumption_resolution,[],[f54981,f43])).
fof(f54981,plain,(
  ( ! [X57,X54,X52,X56,X55,X53,X51] : (~p(s(bool,'T_0')) | s(fun(X51,bool),X52) = s(fun(X51,bool),X53) | s(bool,X54) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X51,bool),X53),s(fun(X51,bool),X52))) | s(bool,'T_0') = s(bool,X54) | s(bool,X55) = s(bool,X56) | s(bool,X56) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X51,bool),X52),s(fun(X51,bool),X53))) | s(bool,X55) = s(bool,X57) | s(bool,'T_0') = s(bool,X57)) )),
  inference(superposition,[],[f426,f188])).
fof(f55000,plain,(
  ( ! [X21,X19,X17,X20,X18] : (s(fun(X17,bool),X18) = s(fun(X17,bool),X19) | s(bool,X20) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X17,bool),X19),s(fun(X17,bool),X18))) | s(bool,'T_0') = s(bool,X20) | s(bool,X21) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X17,bool),X18),s(fun(X17,bool),X19))) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f54975,f43])).
fof(f54975,plain,(
  ( ! [X21,X19,X17,X20,X18] : (~p(s(bool,'T_0')) | s(fun(X17,bool),X18) = s(fun(X17,bool),X19) | s(bool,X20) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X17,bool),X19),s(fun(X17,bool),X18))) | s(bool,'T_0') = s(bool,X20) | s(bool,X21) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X17,bool),X18),s(fun(X17,bool),X19))) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f426,f54])).
fof(f54999,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s(bool,X11) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X8,bool),X10),s(fun(X8,bool),X9))) | s(bool,'T_0') = s(bool,X11) | s(bool,X12) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X8,bool),X9),s(fun(X8,bool),X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f54998,f43])).
fof(f54998,plain,(
  ( ! [X12,X10,X8,X11,X9] : (~p(s(bool,'T_0')) | s(bool,X11) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X8,bool),X10),s(fun(X8,bool),X9))) | s(bool,'T_0') = s(bool,X11) | s(bool,X12) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X8,bool),X9),s(fun(X8,bool),X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f54973,f423])).
fof(f54973,plain,(
  ( ! [X12,X10,X8,X11,X9] : (~p(s(bool,'T_0')) | s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,X11) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X8,bool),X10),s(fun(X8,bool),X9))) | s(bool,'T_0') = s(bool,X11) | s(bool,X12) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X8,bool),X9),s(fun(X8,bool),X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f426,f476])).
fof(f54290,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(fun(X1,bool),X2) = s(fun(X1,bool),X3) | s(bool,X0) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X1,bool),X2),s(fun(X1,bool),X3))) | s(bool,X4) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X1,bool),X2),s(fun(X1,bool),X3))) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(condensation,[],[f54289])).
fof(f54289,plain,(
  ( ! [X61,X66,X64,X62,X67,X65,X63] : (s(fun(X61,bool),X62) = s(fun(X61,bool),X63) | s(bool,X65) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X61,bool),X62),s(fun(X61,bool),X63))) | s(bool,'T_0') = s(bool,X65) | s(bool,X66) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X61,bool),X62),s(fun(X61,bool),X63))) | s(bool,'T_0') = s(bool,X64) | s(bool,X66) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(subsumption_resolution,[],[f54266,f56])).
fof(f54266,plain,(
  ( ! [X61,X66,X64,X62,X67,X65,X63] : (p(s(bool,X64)) | s(fun(X61,bool),X62) = s(fun(X61,bool),X63) | s(bool,X65) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X61,bool),X62),s(fun(X61,bool),X63))) | s(bool,'T_0') = s(bool,X65) | s(bool,X66) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X61,bool),X62),s(fun(X61,bool),X63))) | s(bool,'T_0') = s(bool,X64) | s(bool,X66) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(superposition,[],[f425,f188])).
fof(f54287,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),sK3) | s(bool,X0) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))) )),
  inference(condensation,[],[f54286])).
fof(f54286,plain,(
  ( ! [X2,X3] : (s(fun(sK0,bool),sK1) = s(fun(sK0,bool),sK3) | s(bool,X3) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f54251,f56])).
fof(f54251,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),sK3) | s(bool,X3) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f425,f422])).
fof(f53885,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (s(fun(X42,bool),X43) = s(fun(X42,bool),X44) | s(bool,X45) = s(bool,g(s(X42,sK4(X43,X44,bool,X42)))) | s(bool,'T_0') = s(bool,X45) | s(bool,X46) = s(bool,X47) | s(bool,X47) = s(bool,f(s(X42,sK4(X43,X44,bool,X42)))) | s(bool,X46) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(trivial_inequality_removal,[],[f53866])).
fof(f53866,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(X42,bool),X43) = s(fun(X42,bool),X44) | s(bool,X45) = s(bool,g(s(X42,sK4(X43,X44,bool,X42)))) | s(bool,'T_0') = s(bool,X45) | s(bool,X46) = s(bool,X47) | s(bool,X47) = s(bool,f(s(X42,sK4(X43,X44,bool,X42)))) | s(bool,X46) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f275,f188])).
fof(f53887,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,X11) = s(bool,g(s(X8,sK4(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11) | s(bool,X12) = s(bool,f(s(X8,sK4(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(trivial_inequality_removal,[],[f53860])).
fof(f53860,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,X11) = s(bool,g(s(X8,sK4(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11) | s(bool,X12) = s(bool,f(s(X8,sK4(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f275,f54])).
fof(f511,plain,(
  ( ! [X21,X22,X20] : (s(X22,f(s(bool,'F_0'))) != s(X22,g(s(bool,'F_0'))) | s(fun(bool,X22),X20) = s(fun(bool,X22),X21) | s(X22,f(s(bool,'T_0'))) != s(X22,g(s(bool,'T_0')))) )),
  inference(duplicate_literal_removal,[],[f510])).
fof(f510,plain,(
  ( ! [X21,X22,X20] : (s(X22,f(s(bool,'T_0'))) != s(X22,g(s(bool,'T_0'))) | s(fun(bool,X22),X20) = s(fun(bool,X22),X21) | s(fun(bool,X22),X20) = s(fun(bool,X22),X21) | s(X22,f(s(bool,'F_0'))) != s(X22,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f53,f89])).
fof(f2655,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,sK4(X2,X3,X4,bool)) | s(bool,X0) = s(bool,X1) | s(fun(bool,X4),X2) = s(fun(bool,X4),X3) | s(X4,f(s(bool,'F_0'))) != s(X4,g(s(bool,'F_0')))) )),
  inference(condensation,[],[f2654])).
fof(f2654,plain,(
  ( ! [X158,X156,X159,X157,X155,X160] : (s(bool,'T_0') = s(bool,X158) | s(bool,X159) = s(bool,sK4(X155,X156,X157,bool)) | s(bool,X159) = s(bool,X160) | s(bool,'T_0') = s(bool,X160) | s(fun(bool,X157),X155) = s(fun(bool,X157),X156) | s(X157,f(s(bool,'F_0'))) != s(X157,g(s(bool,'F_0')))) )),
  inference(duplicate_literal_removal,[],[f1347])).
fof(f1347,plain,(
  ( ! [X158,X156,X159,X157,X155,X160] : (s(bool,'T_0') = s(bool,X158) | s(bool,X159) = s(bool,sK4(X155,X156,X157,bool)) | s(bool,'T_0') = s(bool,X158) | s(bool,X159) = s(bool,X160) | s(bool,'T_0') = s(bool,X160) | s(fun(bool,X157),X155) = s(fun(bool,X157),X156) | s(X157,f(s(bool,'F_0'))) != s(X157,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f188,f89])).
fof(f233,plain,(
  ( ! [X52,X50,X53,X51] : (s(bool,X53) != s(bool,f(s(X50,sK4(X51,X52,bool,X50)))) | s(fun(X50,bool),X51) = s(fun(X50,bool),X52) | s(bool,'T_0') = s(bool,X53) | s(bool,'T_0') = s(bool,g(s(X50,sK4(X51,X52,bool,X50))))) )),
  inference(superposition,[],[f53,f54])).
fof(f275,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,'T_0') != s(bool,f(s(X42,sK4(X43,X44,bool,X42)))) | s(fun(X42,bool),X43) = s(fun(X42,bool),X44) | s(bool,X45) = s(bool,g(s(X42,sK4(X43,X44,bool,X42)))) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f53,f54])).
fof(f585,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK4(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,f(s(X0,sK4(X1,X2,bool,X0))))) )),
  inference(condensation,[],[f584])).
fof(f584,plain,(
  ( ! [X10,X8,X11,X9] : (s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK4(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,f(s(X8,sK4(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(subsumption_resolution,[],[f566,f45])).
fof(f566,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') != s(bool,X11) | s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK4(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,f(s(X8,sK4(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f90,f54])).
fof(f1560,plain,(
  ( ! [X191,X189,X192,X190,X188,X193] : (s(bool,X191) != s(bool,f(s(X188,sK4(X189,X190,bool,X188)))) | s(fun(X188,bool),X189) = s(fun(X188,bool),X190) | s(bool,X191) = s(bool,X192) | s(bool,'T_0') = s(bool,X192) | s(bool,X193) = s(bool,g(s(X188,sK4(X189,X190,bool,X188)))) | s(bool,'T_0') = s(bool,X193)) )),
  inference(superposition,[],[f53,f188])).
fof(f1646,plain,(
  ( ! [X191,X189,X192,X190,X188,X193] : (s(bool,X191) != s(bool,f(s(X188,sK4(X189,X190,bool,X188)))) | s(fun(X188,bool),X189) = s(fun(X188,bool),X190) | s(bool,X192) = s(bool,g(s(X188,sK4(X189,X190,bool,X188)))) | s(bool,'T_0') = s(bool,X191) | s(bool,X192) = s(bool,X193) | s(bool,'T_0') = s(bool,X193)) )),
  inference(superposition,[],[f53,f188])).
fof(f1689,plain,(
  ( ! [X189,X187,X185,X190,X188,X186] : (s(bool,X188) != s(bool,f(s(X185,sK4(X186,X187,bool,X185)))) | s(fun(X185,bool),X186) = s(fun(X185,bool),X187) | s(bool,X188) = s(bool,X189) | s(bool,'T_0') = s(bool,g(s(X185,sK4(X186,X187,bool,X185)))) | s(bool,X189) = s(bool,X190) | s(bool,'T_0') = s(bool,X190)) )),
  inference(superposition,[],[f53,f188])).
fof(f1729,plain,(
  ( ! [X154,X152,X151,X149,X153,X150] : (s(bool,'T_0') != s(bool,f(s(X149,sK4(X150,X151,bool,X149)))) | s(fun(X149,bool),X150) = s(fun(X149,bool),X151) | s(bool,X152) = s(bool,X153) | s(bool,X153) = s(bool,g(s(X149,sK4(X150,X151,bool,X149)))) | s(bool,X152) = s(bool,X154) | s(bool,'T_0') = s(bool,X154)) )),
  inference(superposition,[],[f53,f188])).
fof(f3054,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK4(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,X3) = s(bool,f(s(X0,sK4(X1,X2,bool,X0)))) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(condensation,[],[f3053])).
fof(f3053,plain,(
  ( ! [X85,X83,X81,X86,X84,X82] : (s(fun(X81,bool),X82) = s(fun(X81,bool),X83) | s(bool,'T_0') = s(bool,g(s(X81,sK4(X82,X83,bool,X81)))) | s(bool,X85) = s(bool,f(s(X81,sK4(X82,X83,bool,X81)))) | s(bool,'T_0') = s(bool,X84) | s(bool,X85) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f1622,f45])).
fof(f1622,plain,(
  ( ! [X85,X83,X81,X86,X84,X82] : (s(bool,'F_0') != s(bool,X84) | s(fun(X81,bool),X82) = s(fun(X81,bool),X83) | s(bool,'T_0') = s(bool,g(s(X81,sK4(X82,X83,bool,X81)))) | s(bool,X85) = s(bool,f(s(X81,sK4(X82,X83,bool,X81)))) | s(bool,'T_0') = s(bool,X84) | s(bool,X85) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f90,f188])).
fof(f12474,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,f(s(X0,sK4(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,X3) = s(bool,g(s(X0,sK4(X1,X2,bool,X0)))) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(condensation,[],[f12473])).
fof(f12473,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (s(bool,'T_0') = s(bool,X40) | s(fun(X37,bool),X38) = s(fun(X37,bool),X39) | s(bool,'T_0') = s(bool,f(s(X37,sK4(X38,X39,bool,X37)))) | s(bool,X41) = s(bool,g(s(X37,sK4(X38,X39,bool,X37)))) | s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(duplicate_literal_removal,[],[f12399])).
fof(f12399,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (s(bool,'T_0') = s(bool,X40) | s(fun(X37,bool),X38) = s(fun(X37,bool),X39) | s(bool,'T_0') = s(bool,f(s(X37,sK4(X38,X39,bool,X37)))) | s(bool,X41) = s(bool,g(s(X37,sK4(X38,X39,bool,X37)))) | s(bool,'T_0') = s(bool,X40) | s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f585,f188])).
fof(f14665,plain,(
  ( ! [X420,X422,X424,X421,X423,X425] : (s(bool,f(s(X420,sK4(X421,X422,bool,X420)))) != s(bool,'const.pred_set.PSUBSET_2'(s(fun(X423,bool),X424),s(fun(X423,bool),X425))) | s(fun(X420,bool),X421) = s(fun(X420,bool),X422) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X423,bool),X424),s(fun(X423,bool),X425))) | s(bool,'T_0') = s(bool,g(s(X420,sK4(X421,X422,bool,X420))))) )),
  inference(superposition,[],[f53,f476])).
fof(f50,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | s(fun(X0,bool),X1) != s(fun(X0,bool),X2)) )),
  inference(cnf_transformation,[],[f38])).
fof(f51,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f38])).
fof(f425,plain,(
  ( ! [X26,X24,X23,X25] : (p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X23,bool),X24),s(fun(X23,bool),X25)))) | s(fun(X23,bool),X24) = s(fun(X23,bool),X25) | s(bool,X26) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X23,bool),X24),s(fun(X23,bool),X25))) | s(bool,'T_0') = s(bool,X26)) )),
  inference(subsumption_resolution,[],[f271,f43])).
fof(f271,plain,(
  ( ! [X26,X24,X23,X25] : (~p(s(bool,'T_0')) | s(fun(X23,bool),X24) = s(fun(X23,bool),X25) | p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X23,bool),X24),s(fun(X23,bool),X25)))) | s(bool,X26) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X23,bool),X24),s(fun(X23,bool),X25))) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f51,f54])).
fof(f476,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) | s(bool,X3) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f469,f43])).
fof(f469,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) | s(bool,X3) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f80,f54])).
fof(f3056,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X1,bool),X2),s(fun(X1,bool),X3)))) | s(bool,X4) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X1,bool),X2),s(fun(X1,bool),X3))) | s(bool,X0) = s(bool,X4)) )),
  inference(condensation,[],[f3055])).
fof(f3055,plain,(
  ( ! [X154,X152,X151,X149,X153,X150] : (~p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X149,bool),X150),s(fun(X149,bool),X151)))) | s(bool,X153) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X149,bool),X150),s(fun(X149,bool),X151))) | s(bool,'T_0') = s(bool,X152) | s(bool,X153) = s(bool,X154) | s(bool,'T_0') = s(bool,X154)) )),
  inference(subsumption_resolution,[],[f1638,f56])).
fof(f1638,plain,(
  ( ! [X154,X152,X151,X149,X153,X150] : (p(s(bool,X152)) | ~p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X149,bool),X150),s(fun(X149,bool),X151)))) | s(bool,X153) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X149,bool),X150),s(fun(X149,bool),X151))) | s(bool,'T_0') = s(bool,X152) | s(bool,X153) = s(bool,X154) | s(bool,'T_0') = s(bool,X154)) )),
  inference(superposition,[],[f49,f188])).
fof(f3063,plain,(
  ( ! [X80,X83,X81,X79,X84,X82] : (s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X79,bool),X80),s(fun(X79,bool),X81))) | s(bool,X82) = s(bool,X83) | s(bool,X83) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X79,bool),X80),s(fun(X79,bool),X81))) | s(bool,X82) = s(bool,X84) | s(bool,'T_0') = s(bool,X84)) )),
  inference(subsumption_resolution,[],[f1716,f43])).
fof(f1716,plain,(
  ( ! [X80,X83,X81,X79,X84,X82] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X79,bool),X80),s(fun(X79,bool),X81))) | s(bool,X82) = s(bool,X83) | s(bool,X83) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X79,bool),X80),s(fun(X79,bool),X81))) | s(bool,X82) = s(bool,X84) | s(bool,'T_0') = s(bool,X84)) )),
  inference(superposition,[],[f80,f188])).
fof(f3066,plain,(
  ( ! [X111,X109,X112,X110,X108,X113] : (s(fun(X108,bool),X109) = s(fun(X108,bool),X110) | p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X108,bool),X109),s(fun(X108,bool),X110)))) | s(bool,X111) = s(bool,X112) | s(bool,X112) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X108,bool),X109),s(fun(X108,bool),X110))) | s(bool,X111) = s(bool,X113) | s(bool,'T_0') = s(bool,X113)) )),
  inference(subsumption_resolution,[],[f1722,f43])).
fof(f1722,plain,(
  ( ! [X111,X109,X112,X110,X108,X113] : (~p(s(bool,'T_0')) | s(fun(X108,bool),X109) = s(fun(X108,bool),X110) | p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X108,bool),X109),s(fun(X108,bool),X110)))) | s(bool,X111) = s(bool,X112) | s(bool,X112) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X108,bool),X109),s(fun(X108,bool),X110))) | s(bool,X111) = s(bool,X113) | s(bool,'T_0') = s(bool,X113)) )),
  inference(superposition,[],[f51,f188])).
fof(f14662,plain,(
  ( ! [X408,X405,X407,X409,X404,X406] : (s(X406,f(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X407,bool),X408),s(fun(X407,bool),X409))))) != s(X406,g(s(bool,'const.pred_set.PSUBSET_2'(s(fun(X407,bool),X408),s(fun(X407,bool),X409))))) | s(fun(bool,X406),X404) = s(fun(bool,X406),X405) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X407,bool),X408),s(fun(X407,bool),X409))) | s(bool,'T_0') = s(bool,sK4(X404,X405,X406,bool))) )),
  inference(superposition,[],[f53,f476])).
fof(f14708,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) | s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) )),
  inference(trivial_inequality_removal,[],[f14704])).
fof(f14704,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) | s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) )),
  inference(equality_factoring,[],[f476])).
fof(f14826,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) | s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X3,bool),X4),s(fun(X3,bool),X5))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X3,bool),X4),s(fun(X3,bool),X5)))) )),
  inference(condensation,[],[f14825])).
fof(f14825,plain,(
  ( ! [X101,X99,X97,X102,X100,X98,X96] : (s(bool,'const.pred_set.PSUBSET_2'(s(fun(X97,bool),X98),s(fun(X97,bool),X99))) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X100,bool),X101),s(fun(X100,bool),X102))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X100,bool),X101),s(fun(X100,bool),X102))) | s(bool,'T_0') = s(bool,X96) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X97,bool),X98),s(fun(X97,bool),X99)))) )),
  inference(duplicate_literal_removal,[],[f14521])).
fof(f14521,plain,(
  ( ! [X101,X99,X97,X102,X100,X98,X96] : (s(bool,'const.pred_set.PSUBSET_2'(s(fun(X97,bool),X98),s(fun(X97,bool),X99))) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X100,bool),X101),s(fun(X100,bool),X102))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X100,bool),X101),s(fun(X100,bool),X102))) | s(bool,'T_0') = s(bool,X96) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X97,bool),X98),s(fun(X97,bool),X99))) | s(bool,'T_0') = s(bool,X96)) )),
  inference(superposition,[],[f476,f476])).
fof(f14864,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X1,bool),X2),s(fun(X1,bool),X3))) | s(bool,X4) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X1,bool),X2),s(fun(X1,bool),X3))) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(condensation,[],[f14863])).
fof(f14863,plain,(
  ( ! [X47,X52,X50,X48,X46,X51,X49] : (s(bool,'T_0') = s(bool,X49) | s(bool,X50) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X46,bool),X47),s(fun(X46,bool),X48))) | s(bool,'T_0') = s(bool,X50) | s(bool,X51) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X46,bool),X47),s(fun(X46,bool),X48))) | s(bool,X51) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(duplicate_literal_removal,[],[f14490])).
fof(f14490,plain,(
  ( ! [X47,X52,X50,X48,X46,X51,X49] : (s(bool,'T_0') = s(bool,X49) | s(bool,X50) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X46,bool),X47),s(fun(X46,bool),X48))) | s(bool,'T_0') = s(bool,X50) | s(bool,X51) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X46,bool),X47),s(fun(X46,bool),X48))) | s(bool,'T_0') = s(bool,X49) | s(bool,X51) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f476,f188])).
fof(f14899,plain,(
  ( ! [X12,X10,X8,X11,X9] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X8,bool),X10),s(fun(X8,bool),X11)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X8,bool),X9),s(fun(X8,bool),X11)))) | s(bool,X12) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X8,bool),X9),s(fun(X8,bool),X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f14581,f43])).
fof(f14581,plain,(
  ( ! [X12,X10,X8,X11,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X8,bool),X10),s(fun(X8,bool),X11)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X8,bool),X9),s(fun(X8,bool),X11)))) | s(bool,X12) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X8,bool),X9),s(fun(X8,bool),X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f52,f476])).
fof(f14901,plain,(
  ( ! [X19,X17,X20,X18] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X17,bool),X19),s(fun(X17,bool),X18)))) | s(bool,X20) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X17,bool),X18),s(fun(X17,bool),X19))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(subsumption_resolution,[],[f14900,f43])).
fof(f14900,plain,(
  ( ! [X19,X17,X20,X18] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X17,bool),X19),s(fun(X17,bool),X18)))) | s(bool,X20) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X17,bool),X18),s(fun(X17,bool),X19))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(subsumption_resolution,[],[f14583,f423])).
fof(f14583,plain,(
  ( ! [X19,X17,X20,X18] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X17,bool),X19),s(fun(X17,bool),X18)))) | s(fun(X17,bool),X18) = s(fun(X17,bool),X19) | s(bool,X20) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X17,bool),X18),s(fun(X17,bool),X19))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f48,f476])).
fof(f539,plain,(
  ( ! [X0] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X0)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X0))))) )),
  inference(subsumption_resolution,[],[f523,f43])).
fof(f523,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X0)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X0))))) )),
  inference(superposition,[],[f52,f474])).
fof(f48,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) )),
  inference(cnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2))),
  inference(flattening,[],[f31])).
fof(f31,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2))),
  inference(ennf_transformation,[],[f19])).
fof(f19,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))))) => s(fun(X0,bool),X1) = s(fun(X0,bool),X2))),
  inference(flattening,[],[f18])).
fof(f18,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))))) => s(fun(X0,bool),X1) = s(fun(X0,bool),X2))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : ! [X7,X8] : ((p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X7))))) => s(fun(X6,bool),X7) = s(fun(X6,bool),X8))),
  file('thm.pred_set.PSUBSET_TRANS','thm.pred_set.SUBSET_ANTISYM')).
fof(f52,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))))),
  inference(ennf_transformation,[],[f23])).
fof(f23,plain,(
  ! [X0,X1,X2,X3] : ((p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))) => p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))))),
  inference(flattening,[],[f22])).
fof(f22,plain,(
  ! [X0] : ! [X1,X2,X3] : ((p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))) => p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X3)))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : ! [X7,X8,X9] : ((p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X9))))) => p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X9)))))),
  file('thm.pred_set.PSUBSET_TRANS','thm.pred_set.SUBSET_TRANS')).
fof(f424,plain,(
  ( ! [X21,X19,X22,X20,X18] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X18,bool),X20),s(fun(X18,bool),X21)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X18,bool),X19),s(fun(X18,bool),X21)))) | s(bool,X22) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X18,bool),X19),s(fun(X18,bool),X20))) | s(bool,'T_0') = s(bool,X22)) )),
  inference(subsumption_resolution,[],[f270,f43])).
fof(f270,plain,(
  ( ! [X21,X19,X22,X20,X18] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X18,bool),X20),s(fun(X18,bool),X21)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X18,bool),X19),s(fun(X18,bool),X21)))) | s(bool,X22) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X18,bool),X19),s(fun(X18,bool),X20))) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f52,f54])).
fof(f426,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X27,bool),X29),s(fun(X27,bool),X28)))) | s(fun(X27,bool),X28) = s(fun(X27,bool),X29) | s(bool,X30) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X27,bool),X28),s(fun(X27,bool),X29))) | s(bool,'T_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f272,f43])).
fof(f272,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X27,bool),X29),s(fun(X27,bool),X28)))) | s(fun(X27,bool),X28) = s(fun(X27,bool),X29) | s(bool,X30) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X27,bool),X28),s(fun(X27,bool),X29))) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f48,f54])).
fof(f3065,plain,(
  ( ! [X103,X101,X107,X105,X102,X106,X104] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X101,bool),X103),s(fun(X101,bool),X104)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X101,bool),X102),s(fun(X101,bool),X104)))) | s(bool,X105) = s(bool,X106) | s(bool,X106) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X101,bool),X102),s(fun(X101,bool),X103))) | s(bool,X105) = s(bool,X107) | s(bool,'T_0') = s(bool,X107)) )),
  inference(subsumption_resolution,[],[f1721,f43])).
fof(f1721,plain,(
  ( ! [X103,X101,X107,X105,X102,X106,X104] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X101,bool),X103),s(fun(X101,bool),X104)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X101,bool),X102),s(fun(X101,bool),X104)))) | s(bool,X105) = s(bool,X106) | s(bool,X106) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X101,bool),X102),s(fun(X101,bool),X103))) | s(bool,X105) = s(bool,X107) | s(bool,'T_0') = s(bool,X107)) )),
  inference(superposition,[],[f52,f188])).
fof(f3067,plain,(
  ( ! [X118,X116,X114,X119,X117,X115] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X114,bool),X116),s(fun(X114,bool),X115)))) | s(fun(X114,bool),X115) = s(fun(X114,bool),X116) | s(bool,X117) = s(bool,X118) | s(bool,X118) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X114,bool),X115),s(fun(X114,bool),X116))) | s(bool,X117) = s(bool,X119) | s(bool,'T_0') = s(bool,X119)) )),
  inference(subsumption_resolution,[],[f1723,f43])).
fof(f1723,plain,(
  ( ! [X118,X116,X114,X119,X117,X115] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X114,bool),X116),s(fun(X114,bool),X115)))) | s(fun(X114,bool),X115) = s(fun(X114,bool),X116) | s(bool,X117) = s(bool,X118) | s(bool,X118) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X114,bool),X115),s(fun(X114,bool),X116))) | s(bool,X117) = s(bool,X119) | s(bool,'T_0') = s(bool,X119)) )),
  inference(superposition,[],[f48,f188])).
fof(f87,plain,(
  s(bool,f(s(sK0,sK4(sK1,sK2,bool,sK0)))) != s(bool,g(s(sK0,sK4(sK1,sK2,bool,sK0))))),
  inference(unit_resulting_resolution,[],[f78,f53])).
fof(f88,plain,(
  s(bool,f(s(sK0,sK4(sK2,sK3,bool,sK0)))) != s(bool,g(s(sK0,sK4(sK2,sK3,bool,sK0))))),
  inference(unit_resulting_resolution,[],[f84,f53])).
fof(f86,plain,(
  s(bool,f(s(sK0,sK4(sK3,sK2,bool,sK0)))) != s(bool,g(s(sK0,sK4(sK3,sK2,bool,sK0))))),
  inference(unit_resulting_resolution,[],[f84,f53])).
fof(f34537,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) | s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK2))) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) )),
  inference(condensation,[],[f34536])).
fof(f34536,plain,(
  ( ! [X114,X112,X115,X113] : (s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK2))) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X113,bool),X114),s(fun(X113,bool),X115))) | s(bool,'T_0') = s(bool,X112) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X113,bool),X114),s(fun(X113,bool),X115)))) )),
  inference(duplicate_literal_removal,[],[f34429])).
fof(f34429,plain,(
  ( ! [X114,X112,X115,X113] : (s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK2))) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X113,bool),X114),s(fun(X113,bool),X115))) | s(bool,'T_0') = s(bool,X112) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X113,bool),X114),s(fun(X113,bool),X115))) | s(bool,'T_0') = s(bool,X112)) )),
  inference(superposition,[],[f650,f476])).
fof(f36374,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK4(X1,X2,bool,X0)))) | s(bool,f(s(X0,sK4(X1,X2,bool,X0)))) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) )),
  inference(condensation,[],[f36373])).
fof(f36373,plain,(
  ( ! [X118,X120,X119,X117] : (s(bool,f(s(X118,sK4(X119,X120,bool,X118)))) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK2))) | s(fun(X118,bool),X119) = s(fun(X118,bool),X120) | s(bool,'T_0') = s(bool,X117) | s(bool,'T_0') = s(bool,g(s(X118,sK4(X119,X120,bool,X118))))) )),
  inference(duplicate_literal_removal,[],[f36268])).
fof(f36268,plain,(
  ( ! [X118,X120,X119,X117] : (s(bool,f(s(X118,sK4(X119,X120,bool,X118)))) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK2))) | s(fun(X118,bool),X119) = s(fun(X118,bool),X120) | s(bool,'T_0') = s(bool,X117) | s(bool,'T_0') = s(bool,g(s(X118,sK4(X119,X120,bool,X118)))) | s(bool,'T_0') = s(bool,X117)) )),
  inference(superposition,[],[f233,f650])).
fof(f85,plain,(
  s(bool,f(s(sK0,sK4(sK2,sK1,bool,sK0)))) != s(bool,g(s(sK0,sK4(sK2,sK1,bool,sK0))))),
  inference(unit_resulting_resolution,[],[f78,f53])).
fof(f34541,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) | s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) )),
  inference(condensation,[],[f34540])).
fof(f34540,plain,(
  ( ! [X107,X105,X106,X104] : (s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X105,bool),X106),s(fun(X105,bool),X107))) | s(bool,'T_0') = s(bool,X104) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X105,bool),X106),s(fun(X105,bool),X107)))) )),
  inference(duplicate_literal_removal,[],[f34427])).
fof(f34427,plain,(
  ( ! [X107,X105,X106,X104] : (s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X105,bool),X106),s(fun(X105,bool),X107))) | s(bool,'T_0') = s(bool,X104) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X105,bool),X106),s(fun(X105,bool),X107))) | s(bool,'T_0') = s(bool,X104)) )),
  inference(superposition,[],[f551,f476])).
fof(f34539,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) | s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) )),
  inference(condensation,[],[f34538])).
fof(f34538,plain,(
  ( ! [X111,X109,X110,X108] : (s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X109,bool),X110),s(fun(X109,bool),X111))) | s(bool,'T_0') = s(bool,X108) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X109,bool),X110),s(fun(X109,bool),X111)))) )),
  inference(duplicate_literal_removal,[],[f34428])).
fof(f34428,plain,(
  ( ! [X111,X109,X110,X108] : (s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X109,bool),X110),s(fun(X109,bool),X111))) | s(bool,'T_0') = s(bool,X108) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X109,bool),X110),s(fun(X109,bool),X111))) | s(bool,'T_0') = s(bool,X108)) )),
  inference(superposition,[],[f594,f476])).
fof(f34535,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) | s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK2))) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) )),
  inference(condensation,[],[f34534])).
fof(f34534,plain,(
  ( ! [X118,X116,X119,X117] : (s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK2))) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X117,bool),X118),s(fun(X117,bool),X119))) | s(bool,'T_0') = s(bool,X116) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X117,bool),X118),s(fun(X117,bool),X119)))) )),
  inference(duplicate_literal_removal,[],[f34430])).
fof(f34430,plain,(
  ( ! [X118,X116,X119,X117] : (s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK2))) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X117,bool),X118),s(fun(X117,bool),X119))) | s(bool,'T_0') = s(bool,X116) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X117,bool),X118),s(fun(X117,bool),X119))) | s(bool,'T_0') = s(bool,X116)) )),
  inference(superposition,[],[f657,f476])).
fof(f36380,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK4(X1,X2,bool,X0)))) | s(bool,f(s(X0,sK4(X1,X2,bool,X0)))) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) )),
  inference(condensation,[],[f36379])).
fof(f36379,plain,(
  ( ! [X103,X105,X106,X104] : (s(bool,f(s(X104,sK4(X105,X106,bool,X104)))) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))) | s(fun(X104,bool),X105) = s(fun(X104,bool),X106) | s(bool,'T_0') = s(bool,X103) | s(bool,'T_0') = s(bool,g(s(X104,sK4(X105,X106,bool,X104))))) )),
  inference(duplicate_literal_removal,[],[f36265])).
fof(f36265,plain,(
  ( ! [X103,X105,X106,X104] : (s(bool,f(s(X104,sK4(X105,X106,bool,X104)))) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))) | s(fun(X104,bool),X105) = s(fun(X104,bool),X106) | s(bool,'T_0') = s(bool,X103) | s(bool,'T_0') = s(bool,g(s(X104,sK4(X105,X106,bool,X104)))) | s(bool,'T_0') = s(bool,X103)) )),
  inference(superposition,[],[f233,f551])).
fof(f36376,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK4(X1,X2,bool,X0)))) | s(bool,f(s(X0,sK4(X1,X2,bool,X0)))) != s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) )),
  inference(condensation,[],[f36375])).
fof(f36375,plain,(
  ( ! [X116,X114,X115,X113] : (s(bool,f(s(X114,sK4(X115,X116,bool,X114)))) != s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))) | s(fun(X114,bool),X115) = s(fun(X114,bool),X116) | s(bool,'T_0') = s(bool,X113) | s(bool,'T_0') = s(bool,g(s(X114,sK4(X115,X116,bool,X114))))) )),
  inference(duplicate_literal_removal,[],[f36267])).
fof(f36267,plain,(
  ( ! [X116,X114,X115,X113] : (s(bool,f(s(X114,sK4(X115,X116,bool,X114)))) != s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))) | s(fun(X114,bool),X115) = s(fun(X114,bool),X116) | s(bool,'T_0') = s(bool,X113) | s(bool,'T_0') = s(bool,g(s(X114,sK4(X115,X116,bool,X114)))) | s(bool,'T_0') = s(bool,X113)) )),
  inference(superposition,[],[f233,f594])).
fof(f36372,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK4(X1,X2,bool,X0)))) | s(bool,f(s(X0,sK4(X1,X2,bool,X0)))) != s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) )),
  inference(condensation,[],[f36371])).
fof(f36371,plain,(
  ( ! [X123,X121,X124,X122] : (s(bool,f(s(X122,sK4(X123,X124,bool,X122)))) != s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK2))) | s(fun(X122,bool),X123) = s(fun(X122,bool),X124) | s(bool,'T_0') = s(bool,X121) | s(bool,'T_0') = s(bool,g(s(X122,sK4(X123,X124,bool,X122))))) )),
  inference(duplicate_literal_removal,[],[f36269])).
fof(f36269,plain,(
  ( ! [X123,X121,X124,X122] : (s(bool,f(s(X122,sK4(X123,X124,bool,X122)))) != s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK2))) | s(fun(X122,bool),X123) = s(fun(X122,bool),X124) | s(bool,'T_0') = s(bool,X121) | s(bool,'T_0') = s(bool,g(s(X122,sK4(X123,X124,bool,X122)))) | s(bool,'T_0') = s(bool,X121)) )),
  inference(superposition,[],[f233,f657])).
fof(f27513,plain,(
  ( ! [X70,X68,X66,X69,X67,X65] : (s(X66,f(s(bool,X65))) != s(X66,g(s(bool,X65))) | s(fun(bool,X66),X67) = s(fun(bool,X66),X68) | s(X66,f(s(bool,'T_0'))) != s(X66,g(s(bool,'T_0'))) | s(bool,X65) = s(bool,X69) | s(bool,X69) = s(bool,X70) | s(bool,'T_0') = s(bool,X70) | s(bool,'T_0') = s(bool,'F_0')) )),
  inference(superposition,[],[f511,f188])).
fof(f27657,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,'T_0') = s(bool,g(s(bool,X16))) | s(fun(bool,bool),X14) = s(fun(bool,bool),X15) | s(bool,'T_0') = s(bool,f(s(bool,sK4(X14,X15,bool,bool)))) | s(bool,X16) = s(bool,X17) | s(bool,'T_0') = s(bool,X17) | s(bool,X18) = s(bool,sK4(X14,X15,bool,bool)) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f585,f188])).
fof(f27658,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s(bool,'T_0') = s(bool,g(s(bool,X21))) | s(fun(bool,bool),X19) = s(fun(bool,bool),X20) | s(bool,'T_0') = s(bool,f(s(bool,sK4(X19,X20,bool,bool)))) | s(bool,X22) = s(bool,sK4(X19,X20,bool,bool)) | s(bool,'T_0') = s(bool,X22) | s(bool,X21) = s(bool,X23) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f585,f188])).
fof(f12380,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s(bool,'T_0') = s(bool,g(s(bool,X48))) | s(fun(bool,bool),X46) = s(fun(bool,bool),X47) | s(bool,'T_0') = s(bool,f(s(bool,sK4(X46,X47,bool,bool)))) | s(bool,X49) = s(bool,sK4(X46,X47,bool,bool)) | s(bool,X49) = s(bool,X50) | s(bool,'T_0') = s(bool,X50) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f585,f188])).
fof(f12377,plain,(
  ( ! [X35,X33,X31,X34,X32] : (s(bool,'T_0') = s(bool,g(s(bool,X33))) | s(fun(bool,bool),X31) = s(fun(bool,bool),X32) | s(bool,'T_0') = s(bool,f(s(bool,sK4(X31,X32,bool,bool)))) | s(bool,X34) = s(bool,sK4(X31,X32,bool,bool)) | s(bool,'T_0') = s(bool,X33) | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f585,f188])).
fof(f12375,plain,(
  ( ! [X24,X23,X21,X25,X22] : (s(bool,'T_0') = s(bool,g(s(bool,X23))) | s(fun(bool,bool),X21) = s(fun(bool,bool),X22) | s(bool,'T_0') = s(bool,f(s(bool,sK4(X21,X22,bool,bool)))) | s(bool,X24) = s(bool,sK4(X21,X22,bool,bool)) | s(bool,'T_0') = s(bool,X24) | s(bool,X23) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f585,f188])).
fof(f12374,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s(bool,'T_0') = s(bool,g(s(bool,X18))) | s(fun(bool,bool),X16) = s(fun(bool,bool),X17) | s(bool,'T_0') = s(bool,f(s(bool,sK4(X16,X17,bool,bool)))) | s(bool,X18) = s(bool,X19) | s(bool,'T_0') = s(bool,X19) | s(bool,X20) = s(bool,sK4(X16,X17,bool,bool)) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f585,f188])).
fof(f20313,plain,(
  ( ! [X6,X8,X7,X9] : (s(bool,X8) != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,g(s(bool,sK4(X6,X7,bool,bool)))) | s(bool,X9) = s(bool,sK4(X6,X7,bool,bool)) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f233,f54])).
fof(f20316,plain,(
  ( ! [X21,X19,X20,X18] : (s(bool,X20) != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X18) = s(fun(bool,bool),X19) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,g(s(bool,sK4(X18,X19,bool,bool)))) | s(bool,X21) = s(bool,sK4(X18,X19,bool,bool)) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f233,f54])).
fof(f20319,plain,(
  ( ! [X39,X37,X35,X38,X36,X34] : (s(bool,X36) != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X34) = s(fun(bool,bool),X35) | s(bool,'T_0') = s(bool,X36) | s(bool,'T_0') = s(bool,g(s(bool,sK4(X34,X35,bool,bool)))) | s(bool,X37) = s(bool,X38) | s(bool,X38) = s(bool,sK4(X34,X35,bool,bool)) | s(bool,X37) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f233,f188])).
fof(f20320,plain,(
  ( ! [X45,X43,X41,X44,X42,X40] : (s(bool,X43) != s(bool,f(s(bool,X42))) | s(fun(bool,bool),X40) = s(fun(bool,bool),X41) | s(bool,'T_0') = s(bool,X43) | s(bool,'T_0') = s(bool,g(s(bool,sK4(X40,X41,bool,bool)))) | s(bool,X44) = s(bool,sK4(X40,X41,bool,bool)) | s(bool,'T_0') = s(bool,X42) | s(bool,X44) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f233,f188])).
fof(f20322,plain,(
  ( ! [X57,X54,X52,X56,X55,X53] : (s(bool,X54) != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X52) = s(fun(bool,bool),X53) | s(bool,'T_0') = s(bool,X54) | s(bool,'T_0') = s(bool,g(s(bool,sK4(X52,X53,bool,bool)))) | s(bool,X55) = s(bool,X56) | s(bool,X56) = s(bool,X57) | s(bool,'T_0') = s(bool,X57) | s(bool,X55) = s(bool,sK4(X52,X53,bool,bool))) )),
  inference(superposition,[],[f233,f188])).
fof(f20323,plain,(
  ( ! [X61,X59,X62,X60,X58,X63] : (s(bool,X61) != s(bool,f(s(bool,X60))) | s(fun(bool,bool),X58) = s(fun(bool,bool),X59) | s(bool,'T_0') = s(bool,X61) | s(bool,'T_0') = s(bool,g(s(bool,sK4(X58,X59,bool,bool)))) | s(bool,X62) = s(bool,sK4(X58,X59,bool,bool)) | s(bool,X62) = s(bool,X63) | s(bool,'T_0') = s(bool,X63) | s(bool,'T_0') = s(bool,X60)) )),
  inference(superposition,[],[f233,f188])).
fof(f20318,plain,(
  ( ! [X30,X28,X33,X31,X29,X32] : (s(bool,X31) != s(bool,f(s(bool,X30))) | s(fun(bool,bool),X28) = s(fun(bool,bool),X29) | s(bool,'T_0') = s(bool,X31) | s(bool,'T_0') = s(bool,g(s(bool,sK4(X28,X29,bool,bool)))) | s(bool,X32) = s(bool,sK4(X28,X29,bool,bool)) | s(bool,'T_0') = s(bool,X32) | s(bool,X30) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f233,f188])).
fof(f20317,plain,(
  ( ! [X26,X24,X23,X27,X25,X22] : (s(bool,X25) != s(bool,f(s(bool,X24))) | s(fun(bool,bool),X22) = s(fun(bool,bool),X23) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,g(s(bool,sK4(X22,X23,bool,bool)))) | s(bool,X24) = s(bool,X26) | s(bool,'T_0') = s(bool,X26) | s(bool,X27) = s(bool,sK4(X22,X23,bool,bool)) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f233,f188])).
fof(f17506,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X1,f(s(bool,X0))) != s(X1,g(s(bool,X0))) | s(fun(bool,X1),X2) = s(fun(bool,X1),X3) | s(bool,X4) = s(bool,sK4(X2,X3,X1,bool)) | s(bool,'T_0') = s(bool,X4) | s(bool,X5) = s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7) | s(bool,X0) = s(bool,X5)) )),
  inference(superposition,[],[f274,f188])).
fof(f17040,plain,(
  ( ! [X70,X68,X66,X71,X69,X67] : (s(bool,X67) != s(bool,f(s(bool,X66))) | s(fun(bool,bool),X68) = s(fun(bool,bool),X69) | s(bool,'T_0') = s(bool,X66) | s(bool,'T_0') = s(bool,sK4(X68,X69,bool,bool)) | s(bool,X70) = s(bool,g(s(bool,X66))) | s(bool,X70) = s(bool,X71) | s(bool,'T_0') = s(bool,X71) | s(bool,'T_0') = s(bool,X67)) )),
  inference(superposition,[],[f232,f188])).
fof(f17037,plain,(
  ( ! [X52,X50,X48,X53,X51,X49] : (s(bool,X49) != s(bool,f(s(bool,X48))) | s(fun(bool,bool),X50) = s(fun(bool,bool),X51) | s(bool,'T_0') = s(bool,X48) | s(bool,'T_0') = s(bool,sK4(X50,X51,bool,bool)) | s(bool,X52) = s(bool,g(s(bool,X48))) | s(bool,'T_0') = s(bool,X49) | s(bool,X52) = s(bool,X53) | s(bool,'T_0') = s(bool,X53)) )),
  inference(superposition,[],[f232,f188])).
fof(f14852,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) | s(bool,'F_0') = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) )),
  inference(condensation,[],[f14851])).
fof(f14851,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X1,bool),X2),s(fun(X1,bool),X3))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X1,bool),X2),s(fun(X1,bool),X3))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f14504])).
fof(f14504,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X1,bool),X2),s(fun(X1,bool),X3))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(X1,bool),X2),s(fun(X1,bool),X3))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f476,f45])).
fof(f10880,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(bool,X0))) != s(X1,g(s(bool,X0))) | s(fun(bool,X1),X2) = s(fun(bool,X1),X3) | s(X1,f(s(bool,'T_0'))) != s(X1,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f511,f45])).
fof(f9288,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK2))) | s(bool,X0) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK2)))) )),
  inference(equality_factoring,[],[f650])).
fof(f8160,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))) | s(bool,X0) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))) )),
  inference(equality_factoring,[],[f551])).
fof(f89,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,sK4(X0,X1,X2,bool)) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f53,f45])).
fof(f90,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK4(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK4(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f53,f45])).
fof(f3526,plain,(
  ( ! [X26,X27,X25] : (s(bool,X25) != s(bool,f(s(sK0,sK4(sK2,sK3,bool,sK0)))) | s(bool,X25) = s(bool,X26) | s(bool,X26) = s(bool,X27) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,g(s(sK0,sK4(sK2,sK3,bool,sK0))))) )),
  inference(superposition,[],[f88,f188])).
fof(f3525,plain,(
  ( ! [X24,X23,X22] : (s(bool,X22) != s(bool,f(s(sK0,sK4(sK2,sK3,bool,sK0)))) | s(bool,X23) = s(bool,g(s(sK0,sK4(sK2,sK3,bool,sK0)))) | s(bool,X23) = s(bool,X24) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f88,f188])).
fof(f3524,plain,(
  ( ! [X21,X19,X20] : (s(bool,'T_0') != s(bool,f(s(sK0,sK4(sK2,sK3,bool,sK0)))) | s(bool,X19) = s(bool,X20) | s(bool,X20) = s(bool,X21) | s(bool,'T_0') = s(bool,X21) | s(bool,X19) = s(bool,g(s(sK0,sK4(sK2,sK3,bool,sK0))))) )),
  inference(superposition,[],[f88,f188])).
fof(f3523,plain,(
  ( ! [X17,X18,X16] : (s(bool,X16) != s(bool,f(s(sK0,sK4(sK2,sK3,bool,sK0)))) | s(bool,X16) = s(bool,X17) | s(bool,'T_0') = s(bool,g(s(sK0,sK4(sK2,sK3,bool,sK0)))) | s(bool,X17) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f88,f188])).
fof(f3522,plain,(
  ( ! [X14,X15,X13] : (s(bool,X13) != s(bool,f(s(sK0,sK4(sK2,sK3,bool,sK0)))) | s(bool,X14) = s(bool,g(s(sK0,sK4(sK2,sK3,bool,sK0)))) | s(bool,'T_0') = s(bool,X13) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f88,f188])).
fof(f3521,plain,(
  ( ! [X12,X10,X11] : (s(bool,'T_0') != s(bool,f(s(sK0,sK4(sK2,sK3,bool,sK0)))) | s(bool,X10) = s(bool,X11) | s(bool,X11) = s(bool,g(s(sK0,sK4(sK2,sK3,bool,sK0)))) | s(bool,X10) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f88,f188])).
fof(f3520,plain,(
  ( ! [X8,X7,X9] : (s(bool,X7) != s(bool,f(s(sK0,sK4(sK2,sK3,bool,sK0)))) | s(bool,X8) = s(bool,g(s(sK0,sK4(sK2,sK3,bool,sK0)))) | s(bool,'T_0') = s(bool,X8) | s(bool,X7) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f88,f188])).
fof(f3519,plain,(
  ( ! [X6,X4,X5] : (s(bool,X4) != s(bool,f(s(sK0,sK4(sK2,sK3,bool,sK0)))) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X5) | s(bool,X6) = s(bool,g(s(sK0,sK4(sK2,sK3,bool,sK0)))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f88,f188])).
fof(f3518,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,f(s(sK0,sK4(sK2,sK3,bool,sK0)))) | s(bool,X3) = s(bool,g(s(sK0,sK4(sK2,sK3,bool,sK0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f88,f54])).
fof(f3517,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(sK0,sK4(sK2,sK3,bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK4(sK2,sK3,bool,sK0)))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f88,f54])).
fof(f3516,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,f(s(sK0,sK4(sK2,sK3,bool,sK0)))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,g(s(sK0,sK4(sK2,sK3,bool,sK0))))) )),
  inference(superposition,[],[f88,f54])).
fof(f3515,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK0,sK4(sK2,sK3,bool,sK0)))) | s(bool,X0) = s(bool,g(s(sK0,sK4(sK2,sK3,bool,sK0)))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f88,f54])).
fof(f3514,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK4(sK2,sK3,bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK4(sK2,sK3,bool,sK0))))),
  inference(superposition,[],[f88,f45])).
fof(f3428,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X1,bool),X2),s(fun(X1,bool),X2))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(equality_resolution,[],[f423])).
fof(f1894,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X2) = s(bool,X3) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2)) )),
  inference(equality_factoring,[],[f188])).
fof(f3078,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f1887,f54])).
fof(f1887,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,X2)) )),
  inference(equality_factoring,[],[f188])).
fof(f1977,plain,(
  ( ! [X2,X3,X1] : (s(bool,X1) = s(bool,X2) | s(bool,X2) != s(bool,X3) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X3)) )),
  inference(condensation,[],[f1860])).
fof(f1860,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) != s(bool,X2) | s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(equality_factoring,[],[f188])).
fof(f1771,plain,(
  ( ! [X187,X185,X186,X184,X183] : (s(bool,'F_0') != s(bool,f(s(bool,X185))) | s(fun(bool,bool),X183) = s(fun(bool,bool),X184) | s(bool,'T_0') = s(bool,g(s(bool,sK4(X183,X184,bool,bool)))) | s(bool,X186) = s(bool,sK4(X183,X184,bool,bool)) | s(bool,X186) = s(bool,X187) | s(bool,'T_0') = s(bool,X187) | s(bool,'T_0') = s(bool,X185)) )),
  inference(superposition,[],[f90,f188])).
fof(f1727,plain,(
  ( ! [X142,X140,X138,X143,X141,X139] : (s(X140,f(s(bool,'T_0'))) != s(X140,g(s(bool,'T_0'))) | s(fun(bool,X140),X138) = s(fun(bool,X140),X139) | s(bool,X141) = s(bool,X142) | s(bool,X142) = s(bool,sK4(X138,X139,X140,bool)) | s(bool,X141) = s(bool,X143) | s(bool,'T_0') = s(bool,X143)) )),
  inference(superposition,[],[f53,f188])).
fof(f3064,plain,(
  ( ! [X90,X88,X87,X85,X89,X86] : (s(fun(X85,bool),X86) != s(fun(X85,bool),X87) | s(bool,X88) = s(bool,X89) | s(bool,X89) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X85,bool),X86),s(fun(X85,bool),X87))) | s(bool,X88) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f1717,f43])).
fof(f1717,plain,(
  ( ! [X90,X88,X87,X85,X89,X86] : (~p(s(bool,'T_0')) | s(fun(X85,bool),X86) != s(fun(X85,bool),X87) | s(bool,X88) = s(bool,X89) | s(bool,X89) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X85,bool),X86),s(fun(X85,bool),X87))) | s(bool,X88) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f50,f188])).
fof(f1690,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X0)) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X2) | s(bool,X1) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f43,f188])).
fof(f1687,plain,(
  ( ! [X177,X175,X178,X176,X174,X179] : (s(X176,f(s(bool,X177))) != s(X176,g(s(bool,X177))) | s(fun(bool,X176),X174) = s(fun(bool,X176),X175) | s(bool,X177) = s(bool,X178) | s(bool,'T_0') = s(bool,sK4(X174,X175,X176,bool)) | s(bool,X178) = s(bool,X179) | s(bool,'T_0') = s(bool,X179)) )),
  inference(superposition,[],[f53,f188])).
fof(f2171,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X2)) )),
  inference(condensation,[],[f2170])).
fof(f2170,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,X25)) | s(bool,'T_0') = s(bool,X24) | s(bool,X25) = s(bool,X26) | s(bool,X26) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(duplicate_literal_removal,[],[f1652])).
fof(f1652,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,X25)) | s(bool,'T_0') = s(bool,X24) | s(bool,X25) = s(bool,X26) | s(bool,'T_0') = s(bool,X24) | s(bool,X26) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f56,f188])).
fof(f1645,plain,(
  ( ! [X187,X185,X186,X184,X183] : (s(bool,'F_0') != s(bool,f(s(bool,X185))) | s(fun(bool,bool),X183) = s(fun(bool,bool),X184) | s(bool,'T_0') = s(bool,g(s(bool,sK4(X183,X184,bool,bool)))) | s(bool,X186) = s(bool,sK4(X183,X184,bool,bool)) | s(bool,'T_0') = s(bool,X185) | s(bool,X186) = s(bool,X187) | s(bool,'T_0') = s(bool,X187)) )),
  inference(superposition,[],[f90,f188])).
fof(f1644,plain,(
  ( ! [X177,X182,X180,X178,X181,X179] : (s(X179,f(s(bool,X180))) != s(X179,g(s(bool,X180))) | s(fun(bool,X179),X177) = s(fun(bool,X179),X178) | s(bool,X181) = s(bool,sK4(X177,X178,X179,bool)) | s(bool,'T_0') = s(bool,X180) | s(bool,X181) = s(bool,X182) | s(bool,'T_0') = s(bool,X182)) )),
  inference(superposition,[],[f53,f188])).
fof(f3034,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK4(X1,X2,bool,X0)))) | s(bool,'F_0') != s(bool,X3) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(condensation,[],[f3033])).
fof(f3033,plain,(
  ( ! [X85,X83,X81,X86,X84,X82] : (s(bool,'F_0') != s(bool,X84) | s(fun(X81,bool),X82) = s(fun(X81,bool),X83) | s(bool,'T_0') = s(bool,g(s(X81,sK4(X82,X83,bool,X81)))) | s(bool,'T_0') = s(bool,X85) | s(bool,X84) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f1579,f233])).
fof(f1579,plain,(
  ( ! [X85,X83,X81,X86,X84,X82] : (s(bool,'F_0') != s(bool,X84) | s(fun(X81,bool),X82) = s(fun(X81,bool),X83) | s(bool,'T_0') = s(bool,g(s(X81,sK4(X82,X83,bool,X81)))) | s(bool,X85) = s(bool,f(s(X81,sK4(X82,X83,bool,X81)))) | s(bool,'T_0') = s(bool,X85) | s(bool,X84) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f90,f188])).
fof(f1558,plain,(
  ( ! [X177,X182,X180,X178,X181,X179] : (s(X179,f(s(bool,X180))) != s(X179,g(s(bool,X180))) | s(fun(bool,X179),X177) = s(fun(bool,X179),X178) | s(bool,X180) = s(bool,X181) | s(bool,'T_0') = s(bool,X181) | s(bool,X182) = s(bool,sK4(X177,X178,X179,bool)) | s(bool,'T_0') = s(bool,X182)) )),
  inference(superposition,[],[f53,f188])).
fof(f3009,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK4(X1,X2,bool,X0)))) | s(bool,'F_0') != s(bool,X3) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(condensation,[],[f3008])).
fof(f3008,plain,(
  ( ! [X85,X83,X81,X86,X84,X82] : (s(bool,'F_0') != s(bool,X84) | s(fun(X81,bool),X82) = s(fun(X81,bool),X83) | s(bool,'T_0') = s(bool,g(s(X81,sK4(X82,X83,bool,X81)))) | s(bool,X84) = s(bool,X85) | s(bool,'T_0') = s(bool,X85) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f1536,f233])).
fof(f1536,plain,(
  ( ! [X85,X83,X81,X86,X84,X82] : (s(bool,'F_0') != s(bool,X84) | s(fun(X81,bool),X82) = s(fun(X81,bool),X83) | s(bool,'T_0') = s(bool,g(s(X81,sK4(X82,X83,bool,X81)))) | s(bool,X84) = s(bool,X85) | s(bool,'T_0') = s(bool,X85) | s(bool,X86) = s(bool,f(s(X81,sK4(X82,X83,bool,X81)))) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f90,f188])).
fof(f1382,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X2) = s(bool,X3) | s(bool,X1) = s(bool,X3) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4) | s(bool,X5) = s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7) | s(bool,X0) = s(bool,X5)) )),
  inference(superposition,[],[f188,f188])).
fof(f2653,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X2)) )),
  inference(condensation,[],[f2652])).
fof(f2652,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X0) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(duplicate_literal_removal,[],[f1348])).
fof(f1348,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X0) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f188,f45])).
fof(f2659,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | s(bool,X0) = s(bool,X1)) )),
  inference(condensation,[],[f2658])).
fof(f2658,plain,(
  ( ! [X144,X145,X143] : (s(bool,'T_0') = s(bool,X143) | s(bool,X144) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | s(bool,X144) = s(bool,X145) | s(bool,'T_0') = s(bool,X145)) )),
  inference(duplicate_literal_removal,[],[f1343])).
fof(f1343,plain,(
  ( ! [X144,X145,X143] : (s(bool,'T_0') = s(bool,X143) | s(bool,X144) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X143) | s(bool,X144) = s(bool,X145) | s(bool,'T_0') = s(bool,X145)) )),
  inference(superposition,[],[f188,f474])).
fof(f2663,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | s(bool,X0) = s(bool,X1)) )),
  inference(condensation,[],[f2662])).
fof(f2662,plain,(
  ( ! [X132,X130,X131] : (s(bool,'T_0') = s(bool,X130) | s(bool,X131) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | s(bool,X131) = s(bool,X132) | s(bool,'T_0') = s(bool,X132)) )),
  inference(duplicate_literal_removal,[],[f1339])).
fof(f1339,plain,(
  ( ! [X132,X130,X131] : (s(bool,'T_0') = s(bool,X130) | s(bool,X131) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X130) | s(bool,X131) = s(bool,X132) | s(bool,'T_0') = s(bool,X132)) )),
  inference(superposition,[],[f188,f63])).
fof(f2821,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X2) | s(bool,X3) = s(bool,X4) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X5) | s(bool,X1) = s(bool,X3)) )),
  inference(condensation,[],[f2820])).
fof(f2820,plain,(
  ( ! [X90,X88,X87,X85,X89,X86,X84] : (s(bool,X84) = s(bool,X85) | s(bool,X85) = s(bool,X86) | s(bool,'T_0') = s(bool,X86) | s(bool,'T_0') = s(bool,X87) | s(bool,X88) = s(bool,X89) | s(bool,X89) = s(bool,X90) | s(bool,'T_0') = s(bool,X90) | s(bool,X84) = s(bool,X88)) )),
  inference(duplicate_literal_removal,[],[f1260])).
fof(f1260,plain,(
  ( ! [X90,X88,X87,X85,X89,X86,X84] : (s(bool,X84) = s(bool,X85) | s(bool,X85) = s(bool,X86) | s(bool,'T_0') = s(bool,X86) | s(bool,'T_0') = s(bool,X87) | s(bool,'T_0') = s(bool,X87) | s(bool,X88) = s(bool,X89) | s(bool,X89) = s(bool,X90) | s(bool,'T_0') = s(bool,X90) | s(bool,X84) = s(bool,X88)) )),
  inference(superposition,[],[f188,f188])).
fof(f712,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))) | s(bool,X0) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3)))) )),
  inference(equality_factoring,[],[f422])).
fof(f644,plain,(
  ~p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK2))))),
  inference(unit_resulting_resolution,[],[f619,f49])).
fof(f619,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK2))))),
  inference(subsumption_resolution,[],[f618,f84])).
fof(f618,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK2)))) | s(fun(sK0,bool),sK2) = s(fun(sK0,bool),sK3)),
  inference(subsumption_resolution,[],[f600,f43])).
fof(f600,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK2)))) | s(fun(sK0,bool),sK2) = s(fun(sK0,bool),sK3)),
  inference(superposition,[],[f48,f475])).
fof(f594,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f589,f43])).
fof(f589,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f545,f54])).
fof(f560,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X14) = s(fun(bool,bool),X15) | s(bool,'T_0') = s(bool,g(s(bool,sK4(X14,X15,bool,bool)))) | s(bool,'T_0') = s(bool,sK4(X14,X15,bool,bool))) )),
  inference(superposition,[],[f90,f45])).
fof(f559,plain,(
  ( ! [X12,X13,X11] : (s(bool,'F_0') != s(bool,f(s(bool,X13))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,'T_0') = s(bool,g(s(bool,sK4(X11,X12,bool,bool)))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,sK4(X11,X12,bool,bool))) )),
  inference(superposition,[],[f90,f54])).
fof(f558,plain,(
  ( ! [X10,X8,X9] : (s(bool,'F_0') != s(bool,f(s(bool,X10))) | s(fun(bool,bool),X8) = s(fun(bool,bool),X9) | s(bool,'T_0') = s(bool,g(s(bool,sK4(X8,X9,bool,bool)))) | s(bool,'T_0') = s(bool,sK4(X8,X9,bool,bool)) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f90,f54])).
fof(f557,plain,(
  ( ! [X6,X7,X5] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X5) = s(fun(bool,bool),X6) | s(bool,'T_0') = s(bool,g(s(bool,sK4(X5,X6,bool,bool)))) | s(bool,X7) = s(bool,sK4(X5,X6,bool,bool)) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f90,f54])).
fof(f556,plain,(
  ( ! [X4,X2,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,g(s(bool,sK4(X2,X3,bool,bool)))) | s(bool,X4) = s(bool,sK4(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f90,f54])).
fof(f571,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK4(X0,X1,bool,bool)))) | s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,'F_0')))) )),
  inference(duplicate_literal_removal,[],[f555])).
fof(f555,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK4(X0,X1,bool,bool)))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f90,f89])).
fof(f551,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f546,f43])).
fof(f546,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f544,f54])).
fof(f545,plain,(
  ~p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),
  inference(unit_resulting_resolution,[],[f544,f49])).
fof(f544,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))),
  inference(subsumption_resolution,[],[f543,f78])).
fof(f543,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),sK2)),
  inference(subsumption_resolution,[],[f525,f43])).
fof(f525,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),sK2)),
  inference(superposition,[],[f48,f474])).
fof(f431,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK4(sK1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK4(sK1,sK2,bool,sK0))))),
  inference(superposition,[],[f87,f45])).
fof(f430,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,f(s(sK0,sK4(sK1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,g(s(sK0,sK4(sK1,sK2,bool,sK0))))) )),
  inference(superposition,[],[f87,f54])).
fof(f429,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(sK0,sK4(sK1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK4(sK1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f87,f54])).
fof(f428,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,f(s(sK0,sK4(sK1,sK2,bool,sK0)))) | s(bool,X1) = s(bool,g(s(sK0,sK4(sK1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f87,f54])).
fof(f427,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK0,sK4(sK1,sK2,bool,sK0)))) | s(bool,X0) = s(bool,g(s(sK0,sK4(sK1,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f87,f54])).
fof(f274,plain,(
  ( ! [X39,X41,X38,X40] : (s(X40,f(s(bool,'T_0'))) != s(X40,g(s(bool,'T_0'))) | s(fun(bool,X40),X38) = s(fun(bool,X40),X39) | s(bool,X41) = s(bool,sK4(X38,X39,X40,bool)) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f53,f54])).
fof(f423,plain,(
  ( ! [X14,X17,X15,X16] : (s(fun(X14,bool),X15) != s(fun(X14,bool),X16) | s(bool,X17) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X14,bool),X15),s(fun(X14,bool),X16))) | s(bool,'T_0') = s(bool,X17)) )),
  inference(subsumption_resolution,[],[f269,f43])).
fof(f269,plain,(
  ( ! [X14,X17,X15,X16] : (~p(s(bool,'T_0')) | s(fun(X14,bool),X15) != s(fun(X14,bool),X16) | s(bool,X17) = s(bool,'const.pred_set.PSUBSET_2'(s(fun(X14,bool),X15),s(fun(X14,bool),X16))) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f50,f54])).
fof(f256,plain,(
  ( ! [X4,X3] : (p(s(bool,X3)) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f43,f54])).
fof(f232,plain,(
  ( ! [X47,X48,X46,X49] : (s(X48,f(s(bool,X49))) != s(X48,g(s(bool,X49))) | s(fun(bool,X48),X46) = s(fun(bool,X48),X47) | s(bool,'T_0') = s(bool,X49) | s(bool,'T_0') = s(bool,sK4(X46,X47,X48,bool))) )),
  inference(superposition,[],[f53,f54])).
fof(f188,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f54,f54])).
fof(f111,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK4(sK3,sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK4(sK3,sK2,bool,sK0))))),
  inference(superposition,[],[f86,f45])).
fof(f91,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK4(sK2,sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK4(sK2,sK1,bool,sK0))))),
  inference(superposition,[],[f85,f45])).
fof(f78,plain,(
  s(fun(sK0,bool),sK1) != s(fun(sK0,bool),sK2)),
  inference(subsumption_resolution,[],[f76,f43])).
fof(f76,plain,(
  ~p(s(bool,'T_0')) | s(fun(sK0,bool),sK1) != s(fun(sK0,bool),sK2)),
  inference(superposition,[],[f50,f63])).
fof(f62,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f45])).
fof(f56,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f44,f45])).
fof(f54,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f45,f45])).
fof(f53,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK4(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK4(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK4(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK4(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f35])).
fof(f35,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f25])).
fof(f25,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f24])).
fof(f24,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.pred_set.PSUBSET_TRANS',aHL_EXT)).
fof(f41,plain,(
  p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))),
  inference(cnf_transformation,[],[f36])).
fof(f42,plain,(
  ~p(s(bool,'const.pred_set.PSUBSET_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),sK3))))),
  inference(cnf_transformation,[],[f36])).
fof(f44913,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK2))))),
  inference(subsumption_resolution,[],[f44912,f84])).
fof(f44912,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK2)))) | s(fun(sK0,bool),sK2) = s(fun(sK0,bool),sK3)),
  inference(subsumption_resolution,[],[f44884,f43])).
fof(f44884,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK3),s(fun(sK0,bool),sK2)))) | s(fun(sK0,bool),sK2) = s(fun(sK0,bool),sK3)),
  inference(superposition,[],[f48,f475])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 258
% Passive clauses: 1037
% Generated clauses: 67404
% Final active clauses: 48
% Final passive clauses: 132
% Input formulas: 10
% Initial clauses: 14
% 
% Duplicate literals: 2831
% Trivial inequalities: 213
% Fw subsumption resolutions: 2259
% Fw demodulations: 17268
% Bw demodulations: 541
% Condensations: 3717
% Global subsumptions: 713
% 
% Simple tautologies: 371
% Equational tautologies: 4025
% Forward subsumptions: 32703
% Bw demodulations to eq. taut.: 23
% 
% Unit resulting resolution: 18201
% Forward superposition: 7566
% Backward superposition: 4489
% Self superposition: 213
% Equality factoring: 170
% Equality resolution: 10
% 
% Split clauses: 169
% Split components: 370
% SAT solver clauses: 3421
% SAT solver unit clauses: 51
% SAT solver binary clauses: 455
% 
% Sat splits: 169
% Sat splitting refutations: 1331
% 
% Memory used [KB]: 19317
% Time elapsed: 2.081 s
% ------------------------------
% ------------------------------
% Success in time 3.374 s
