Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.real
% (84684)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1327
% Passive clauses: 13273
% Generated clauses: 30734
% Final active clauses: 1281
% Final passive clauses: 11535
% Input formulas: 69
% Initial clauses: 253
% 
% Duplicate literals: 384
% Trivial inequalities: 109
% Fw subsumption resolutions: 3003
% Fw demodulations: 1759
% Bw demodulations: 43
% 
% Simple tautologies: 270
% Equational tautologies: 35
% Forward subsumptions: 11122
% Fw demodulations to eq. taut.: 1
% 
% Binary resolution: 14121
% Factoring: 671
% Forward superposition: 2992
% Backward superposition: 7035
% Self superposition: 35
% Equality factoring: 26
% Equality resolution: 7
% 
% Split clauses: 697
% Split components: 1407
% SAT solver clauses: 1075
% SAT solver unit clauses: 15
% SAT solver binary clauses: 214
% 
% Sat splits: 697
% Sat splitting refutations: 11
% 
% TWLsolver clauses: 1075
% TWLsolver calls for satisfiability: 493
% 
% Memory used [KB]: 16119
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.real
% (84717)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 531
% Passive clauses: 3393
% Generated clauses: 14452
% Final active clauses: 461
% Final passive clauses: 2715
% Input formulas: 69
% Initial clauses: 253
% 
% Duplicate literals: 253
% Trivial inequalities: 42
% Fw subsumption resolutions: 1017
% Fw demodulations: 1117
% Bw demodulations: 169
% Condensations: 247
% Global subsumptions: 18
% 
% Simple tautologies: 313
% Equational tautologies: 17
% Forward subsumptions: 4174
% Fw demodulations to eq. taut.: 1
% Bw demodulations to eq. taut.: 34
% 
% Binary resolution: 3238
% Factoring: 25
% Forward superposition: 2473
% Backward superposition: 2018
% Self superposition: 35
% Equality factoring: 26
% Equality resolution: 2
% 
% Split clauses: 142
% Split components: 308
% SAT solver clauses: 4068
% SAT solver unit clauses: 253
% SAT solver binary clauses: 88
% SAT solver learnt clauses: 36
% SAT solver learnt literals: 10
% 
% Sat splits: 142
% Sat splitting refutations: 10
% 
% TWLsolver clauses: 3996
% TWLsolver calls for satisfiability: 2568
% 
% Memory used [KB]: 17270
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.real
% (84739)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 3026
% Passive clauses: 21480
% Generated clauses: 135734
% Final active clauses: 2269
% Final passive clauses: 5914
% Input formulas: 69
% Initial clauses: 253
% Discarded non-redundant clauses: 16708
% 
% Duplicate literals: 158
% Trivial inequalities: 36
% Fw subsumption resolutions: 2499
% Fw demodulations: 32481
% Bw demodulations: 1264
% Condensations: 299
% Global subsumptions: 43
% 
% Simple tautologies: 90
% Equational tautologies: 116
% Forward subsumptions: 68544
% Fw demodulations to eq. taut.: 3
% Bw demodulations to eq. taut.: 9
% 
% Unit resulting resolution: 90431
% Factoring: 1
% Forward superposition: 5954
% Backward superposition: 1724
% Self superposition: 52
% Equality factoring: 26
% Equality resolution: 1
% 
% Split clauses: 142
% Split components: 289
% SAT solver clauses: 23871
% SAT solver unit clauses: 3065
% SAT solver binary clauses: 10929
% 
% Sat splits: 142
% Sat splitting refutations: 2061
% 
% Memory used [KB]: 64604
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.real
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.real
% SZS output start Proof for thm.real
fof(f83405,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f341,f585,f340,f586,f583,f339,f587,f338,f588,f590,f592,f595,f596,f597,f598,f599,f600,f601,f602,f603,f604,f605,f606,f607,f608,f609,f610,f611,f612,f613,f614,f615,f616,f617,f618,f619,f620,f621,f622,f623,f761,f624,f625,f626,f627,f628,f629,f768,f630,f774,f775,f776,f777,f778,f779,f1071,f1205,f1072,f1206,f1073,f1074,f1075,f1207,f1076,f1077,f1208,f1078,f1079,f1080,f1209,f1081,f1082,f1210,f1083,f1084,f1212,f1085,f1086,f1213,f1087,f1088,f1215,f1089,f1090,f1216,f1091,f1092,f1218,f1093,f1094,f1095,f1224,f1096,f1225,f1097,f1227,f1098,f1229,f1364,f1365,f1366,f1367,f1368,f1369,f1370,f1371,f1372,f1373,f1374,f1375,f1376,f1377,f1378,f1379,f1380,f1381,f1382,f1383,f1384,f1385,f1386,f1387,f1388,f1389,f1390,f1391,f1392,f1393,f1394,f1395,f1396,f1397,f1398,f1399,f1400,f1401,f1402,f1403,f1404,f1405,f1406,f1407,f1408,f1409,f1897,f1898,f1899,f1900,f1901,f1902,f1903,f1904,f1905,f1906,f1907,f1908,f1909,f1910,f1911,f1912,f2039,f1913,f1914,f2041,f1915,f1916,f2042,f1917,f1918,f1919,f1220,f1920,f2045,f1223,f1921,f2046,f1922,f2047,f1923,f2049,f1924,f2051,f2430,f3832,f3834,f3837,f3838,f3841,f3843,f3845,f3846,f3849,f3851,f3853,f3854,f3855,f3856,f3857,f3858,f3859,f3860,f3863,f3864,f3867,f3869,f3870,f3999,f4001,f4000,f4723,f758,f4724,f4970,f3978,f4972,f3979,f4975,f4976,f4979,f4981,f4983,f4984,f4987,f4989,f4991,f4992,f4995,f4996,f5001,f5002,f5003,f3984,f5004,f5005,f5007,f3986,f5008,f5150,f5152,f5155,f5156,f5159,f5161,f5163,f5164,f5167,f5169,f5171,f5172,f5175,f5176,f5181,f5182,f5183,f5184,f5187,f5188,f5309,f5438,f5558,f5560,f887,f5562,f5563,f5808,f5809,f1513,f5810,f5811,f1515,f5812,f5813,f5814,f5815,f5818,f5820,f5822,f5823,f5826,f5828,f5830,f5831,f5834,f5835,f5840,f5841,f5842,f1520,f5843,f5844,f5846,f1521,f5847,f5850,f1523,f5851,f1525,f5852,f1526,f5853,f6166,f7340,f7486,f7342,f7487,f7345,f7346,f7349,f7351,f7353,f7354,f7357,f7359,f7361,f7362,f7365,f7366,f7371,f7372,f7373,f7492,f7374,f7377,f7495,f7378,f7548,f7694,f7550,f7695,f7553,f7554,f7557,f7559,f7561,f7562,f7565,f7567,f7569,f7570,f7573,f7574,f7579,f7580,f7581,f7700,f7582,f7585,f7703,f7586,f7838,f7840,f7851,f7852,f7855,f7857,f7859,f7860,f7863,f7864,f7869,f7870,f7871,f7872,f7875,f7876,f8035,f8181,f8037,f8182,f8040,f8183,f8041,f8044,f8046,f8056,f8057,f8060,f8061,f8184,f8066,f8067,f8068,f8185,f8069,f8072,f8187,f8073,f8242,f8244,f8247,f8248,f8251,f8253,f8255,f8256,f8259,f8261,f8267,f8268,f8273,f8274,f8275,f8392,f8276,f8279,f8280,f8641,f8643,f8646,f8647,f8650,f8652,f8654,f8655,f8658,f8660,f8662,f8663,f8672,f8673,f8674,f8675,f8678,f8679,f8825,f8827,f8830,f8831,f8834,f8836,f8838,f8839,f8842,f8844,f8846,f8847,f8850,f8851,f8974,f8858,f8975,f8859,f8862,f8977,f8863,f9030,f9032,f9035,f9036,f9039,f9041,f9043,f9044,f9047,f9049,f9051,f9052,f9055,f9056,f9061,f9062,f9063,f9064,f7583,f7375,f765,f9425,f5296,f9427,f5297,f9430,f9431,f9434,f9436,f9438,f9439,f9442,f9444,f9446,f9447,f9450,f9451,f9456,f9457,f9458,f5302,f9459,f9462,f5304,f9463,f12502,f737,f12504,f738,f12507,f12508,f12509,f12510,f12511,f12512,f12513,f12514,f12515,f12516,f12517,f12518,f12519,f12520,f12521,f19561,f888,f19562,f889,f890,f20605,f20606,f20753,f20607,f20608,f20754,f20609,f20610,f20611,f20612,f20615,f20617,f20619,f20620,f20623,f20625,f20627,f20628,f20631,f20632,f20637,f20638,f20639,f20640,f20643,f20644,f20650,f1527,f20892,f7984,f20894,f7985,f20905,f20906,f20909,f20911,f20913,f20914,f20917,f20918,f7987,f754,f20923,f20924,f20925,f7988,f20926,f20929,f7990,f20930,f21042,f20912,f20910,f20908,f20907,f7986,f20626,f20624,f20622,f20621,f21157,f21158,f21159,f21160,f9445,f9443,f9441,f9440,f7568,f7566,f7564,f7563,f7360,f7358,f7356,f7355,f5829,f5827,f5825,f5824,f4990,f4988,f4986,f4985,f21045,f21044,f21043,f21288,f8787,f21290,f8788,f21293,f21294,f21297,f21299,f21301,f21302,f21303,f21304,f21305,f21306,f21307,f21308,f21309,f21310,f21319,f21320,f21321,f21322,f21325,f8793,f21326,f21567,f1511,f21568,f21714,f21569,f1514,f21570,f21715,f21571,f1516,f21572,f1517,f21573,f21574,f21577,f21579,f21581,f21582,f21583,f21584,f21585,f21586,f21587,f21588,f21589,f21590,f21599,f21600,f21601,f21602,f762,f21605,f21606,f21874,f21580,f21578,f21576,f21575,f21716,f21300,f21298,f21296,f21295,f8789,f20618,f20616,f20614,f20613,f20755,f21989,f21990,f21991,f21992,f740,f9437,f9435,f9433,f9432,f5298,f7560,f7558,f7556,f7555,f7696,f7352,f7350,f7348,f7347,f7488,f5821,f5819,f5817,f5816,f1518,f4982,f4980,f4978,f4977,f3980,f21877,f21876,f21875,f22258,f22259,f22260,f22570,f21592,f21591,f21312,f21311,f20630,f20629,f22683,f9449,f9448,f7572,f7571,f7364,f7363,f5833,f5832,f4994,f4993,f22571,f22739,f21596,f21717,f21595,f21316,f8790,f21315,f20634,f20756,f20633,f22852,f9453,f5300,f9452,f7576,f7698,f7575,f7368,f7490,f7367,f5837,f1519,f5836,f4998,f3982,f4997,f22740,f31495,f31630,f31496,f31631,f31500,f31501,f31632,f31505,f31506,f31633,f31514,f31519,f20757,f31520,f21718,f31699,f31700,f31704,f31705,f31709,f31710,f31723,f31724,f31852,f31983,f31985,f31988,f31989,f31990,f31991,f31992,f31993,f31994,f31995,f31996,f31997,f31998,f31999,f32000,f32001,f32002,f32003,f32006,f32007,f32010,f32011,f8391,f32021,f763,f32275,f9176,f32277,f9177,f32280,f32281,f32282,f32283,f32284,f32285,f32286,f32287,f32288,f32289,f32290,f9178,f32291,f32292,f32293,f32294,f32295,f32298,f32299,f32302,f32303,f32306,f32308,f9180,f32611,f32022,f32315,f32314,f32725,f32023,f21608,f21607,f21328,f20646,f20645,f32726,f9465,f7588,f7380,f5849,f5010,f32613,f769,f32297,f9179,f32296,f32752,f31834,f32753,f31835,f32757,f32758,f31836,f32762,f32763,f31837,f32764,f8394,f32776,f20758,f32777,f21719,f45193,f45309,f45194,f45310,f45200,f45314,f45312,f45542,f45543,f49143,f49184,f49187,f49188,f49207,f49208,f49209,f49186,f49153,f49152,f49348,f49151,f49150,f49606,f2031,f49607,f2032,f49608,f49609,f49610,f2033,f49611,f49612,f2034,f49613,f49614,f49615,f2035,f49616,f49617,f2036,f49618,f2037,f49619,f2038,f49627,f2044,f49960,f49961,f49950,f49952,f49974,f49975,f49964,f49965,f49351,f1512,f65440,f65557,f66026,f66032,f43565,f65555,f66037,f76069,f76189,f76070,f76190,f76080,f76192,f76553,f76554,f78669,f78779,f78670,f78781,f79691,f31508,f8791,f79874,f79545,f79888,f79547,f79889,f79550,f79551,f79552,f79553,f79554,f79555,f79556,f79557,f79558,f79559,f79560,f79890,f79561,f79562,f79563,f79564,f79565,f79566,f79567,f749,f79568,f79569,f79572,f750,f79573,f79892,f79576,f756,f79578,f79894,f759,f79584,f79896,f766,f79696,f79697,f79698,f79699,f79700,f79701,f79704,f79714,f79900,f79715,f79719,f79720,f79721,f79722,f79723,f79724,f79725,f79726,f79727,f79728,f79729,f79730,f79736,f79737,f79739,f79741,f79743,f79745,f79746,f79747,f79748,f79749,f79750,f79751,f79758,f79759,f79763,f79764,f79765,f79770,f79771,f79772,f79773,f79774,f79776,f79777,f79778,f79779,f79780,f79781,f79782,f79783,f79794,f79795,f79800,f79801,f79802,f79803,f79804,f79805,f79806,f79807,f79808,f79809,f79810,f79811,f79812,f79813,f79818,f79819,f79820,f79823,f79824,f79825,f79826,f79827,f79833,f79834,f79835,f79836,f79837,f79838,f79839,f79840,f79841,f79842,f79843,f79844,f79845,f79846,f79847,f79848,f79850,f79851,f79854,f79855,f79856,f79857,f79858,f79859,f79860,f79861,f79862,f79863,f79864,f79867,f79869,f79870,f79904,f79873,f79898,f79922,f79923,f79911,f79937,f79938,f79926,f79952,f79955,f79941,f79954,f80171,f80170,f80169,f80168,f80167,f80166,f80165,f80164,f80157,f80156,f80155,f80154,f80153,f80152,f80151,f80150,f80134,f80133,f80132,f80131,f80130,f80129,f80128,f80127,f80120,f80102,f80101,f80100,f80052,f80245,f80246,f80247,f80248,f80249,f80250,f80251,f80252,f80253,f80254,f755,f80255,f752,f753,f80267,f76194,f80379,f80570,f81200,f81392,f81393,f81381,f81407,f81408,f81396,f79906,f81674,f8388,f81676,f8389,f81679,f81680,f81681,f81682,f81683,f81684,f81685,f81686,f81687,f81688,f81689,f8390,f81690,f81691,f81692,f81693,f81694,f81865,f8971,f81867,f8972,f81870,f81871,f81872,f81873,f81874,f81875,f81876,f81877,f81878,f81879,f81880,f8973,f81881,f81882,f81883,f81884,f81885,f82140,f82248,f82141,f82249,f82250,f82477,f82478,f82466,f82492,f82493,f82481,f79902,f82538,f82545,f82625,f82630,f582,f82645,f82791,f82647,f82792,f82650,f82651,f82652,f82653,f82654,f82655,f743,f82656,f739,f82657,f744,f742,f82658,f82659,f82660,f82793,f736,f82661,f741,f82662,f82663,f747,f82664,f745,f82665,f748,f746,f83054,f83056,f83040,f83073,f83074,f83059,f83055,f83091,f83093,f83077,f83121,f83122,f83107,f83092,f83139,f83141,f83125,f83158,f83159,f83144,f83140,f83266,f83267,f83255,f83281,f83282,f83270,f1211,f1221,f83283,f83287,f83292,f83288,f49349,f83289,f79908,f20752,f83293])).
fof(f83293,plain,(
  $false | ($spl42 | $spl88)),
  inference(resolution,[],[f1211,f753])).
fof(f20752,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) | $spl296),
  inference(cnf_transformation,[],[f20752_D])).
fof(f20752_D,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f79908,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(sK17)),s('type.realax.real',sK17)))) | $spl399),
  inference(cnf_transformation,[],[f79908_D])).
fof(f79908_D,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(sK17)),s('type.realax.real',sK17)))) <=> ~$spl399),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl399])])).
fof(f83289,plain,(
  ( ! [X11] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(sK17)),s('type.realax.real',sK17)))) | p(s(bool,X11))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f55387,f640])).
fof(f640,plain,(
  ( ! [X103,X102] : (~sP0(happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK14)),X102,X103) | ~p(s(bool,X103)) | p(s(bool,X102))) )),
  inference(resolution,[],[f336,f382])).
fof(f382,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f242])).
fof(f242,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(rectify,[],[f241])).
fof(f241,plain,(
  ! [X2,X0,X1] : ((~sP0(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP0(X2,X0,X1)))),
  inference(flattening,[],[f240])).
fof(f240,plain,(
  ! [X2,X0,X1] : ((~sP0(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP0(X2,X0,X1)))),
  inference(nnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X2,X0,X1] : (sP0(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f336,plain,(
  p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK14))))),
  inference(cnf_transformation,[],[f218])).
fof(f218,plain,(
  p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK14)))) & ! [X7] : (~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X7)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',sK15))))) & (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ! [X4] : ((p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK16(X4))))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(X4)),s('type.realax.real',X4))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X4)))))) & (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | (! [X3] : (~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X3)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',sK17))))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK17))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14,sK15,sK16,sK17])],[f217])).
fof(f217,plain,(
  ? [X0,X1] : (? [X8] : p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X8)))) & ? [X6] : ! [X7] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X7)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X6))))) & (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0)))))) | ! [X4] : (? [X5] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X5)))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X4))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X4)))))) & (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0)))))) | ? [X2] : (! [X3] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X2))))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))))))),
  inference(rectify,[],[f216])).
fof(f216,plain,(
  ? [X0,X1] : (? [X4] : p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X4)))) & ? [X2] : ! [X3] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X2))))) & (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0)))))) | ! [X5] : (? [X6] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X6)))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X5))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X5)))))) & (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0)))))) | ? [X5] : (! [X6] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X6)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X5))))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X5)))))))),
  inference(flattening,[],[f215])).
fof(f215,plain,(
  ? [X0,X1] : (? [X4] : p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X4)))) & ? [X2] : ! [X3] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X2))))) & ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0)))))) | ! [X5] : (? [X6] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X6)))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X5))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X5)))))) & (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0)))))) | ? [X5] : (! [X6] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X6)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X5))))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X5))))))))),
  inference(nnf_transformation,[],[f156])).
fof(f156,plain,(
  ? [X0,X1] : (? [X4] : p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X4)))) & ? [X2] : ! [X3] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X2))))) & (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0)))))) <~> ! [X5] : (? [X6] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X6)))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X5))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X5)))))))),
  inference(flattening,[],[f155])).
fof(f155,plain,(
  ? [X0,X1] : ((? [X4] : p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X4)))) & ? [X2] : ! [X3] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X2)))))) & (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0)))))) <~> ! [X5] : (? [X6] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X6)))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X5))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X5)))))))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1] : ((? [X4] : p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X4)))) & ? [X2] : ! [X3] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X2)))))) => (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0)))))) <=> ! [X5] : (! [X6] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X6)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X5))))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X5)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ((? [X8] : p(s(bool,happ(s(fun('type.realax.real',bool),X6),s('type.realax.real',X8)))) & ? [X8] : ! [X9] : (p(s(bool,happ(s(fun('type.realax.real',bool),X6),s('type.realax.real',X9)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',X8)))))) => (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X6)))))) <=> ! [X8] : (! [X9] : (p(s(bool,happ(s(fun('type.realax.real',bool),X6),s('type.realax.real',X9)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',X8))))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X8)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ((? [X8] : p(s(bool,happ(s(fun('type.realax.real',bool),X6),s('type.realax.real',X8)))) & ? [X8] : ! [X9] : (p(s(bool,happ(s(fun('type.realax.real',bool),X6),s('type.realax.real',X9)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',X8)))))) => (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X6)))))) <=> ! [X8] : (! [X9] : (p(s(bool,happ(s(fun('type.realax.real',bool),X6),s('type.realax.real',X9)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',X8))))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X8)))))))),
  file('thm.real.REAL_LE_SUP','thm.real.REAL_LE_SUP')).
fof(f55387,plain,(
  ( ! [X35,X36] : (sP0(X35,X36,'const.real.real_lte_2'(s('type.realax.real',sK16(sK17)),s('type.realax.real',sK17)))) ) | ($spl3 | $spl6)),
  inference(subsumption_resolution,[],[f55172,f563])).
fof(f563,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f146])).
fof(f146,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.real.REAL_LE_SUP',aHL_FALSITY)).
fof(f55172,plain,(
  ( ! [X35,X36] : (p(s(bool,'F_0')) | sP0(X35,X36,'const.real.real_lte_2'(s('type.realax.real',sK16(sK17)),s('type.realax.real',sK17)))) ) | ($spl3 | $spl6)),
  inference(superposition,[],[f383,f49371])).
fof(f49371,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(sK17)),s('type.realax.real',sK17))) | ($spl3 | $spl6)),
  inference(resolution,[],[f585,f1004])).
fof(f1004,plain,(
  ( ! [X249] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X249)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(X249)),s('type.realax.real',X249)))) ) | $spl6),
  inference(resolution,[],[f587,f547])).
fof(f547,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f329])).
fof(f329,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f131])).
fof(f131,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.real.REAL_LE_SUP','thm.bool.EQ_CLAUSES|split|2')).
fof(f383,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f242])).
fof(f49349,plain,(
  ( ! [X11] : (~p(s(bool,X11))) ) | $spl352),
  inference(cnf_transformation,[],[f49349_D])).
fof(f49349_D,plain,(
  ( ! [X11] : (~p(s(bool,X11))) ) <=> ~$spl352),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl352])])).
fof(f83288,plain,(
  ( ! [X10] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(sK17)),s('type.realax.real',sK17)))) | ~p(s(bool,X10))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f55387,f49380])).
fof(f49380,plain,(
  ( ! [X12,X11] : (~sP0(X11,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK17)),X12) | ~p(s(bool,X12)) | ~p(s(bool,X11))) ) | $spl3),
  inference(resolution,[],[f585,f382])).
fof(f83292,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,'const.real.real_lte_2'(s('type.realax.real',sK16(sK17)),s('type.realax.real',sK17))) | p(s(bool,X9))) ) | $spl440),
  inference(cnf_transformation,[],[f83292_D])).
fof(f83292_D,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,'const.real.real_lte_2'(s('type.realax.real',sK16(sK17)),s('type.realax.real',sK17))) | p(s(bool,X9))) ) <=> ~$spl440),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl440])])).
fof(f83287,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,'const.real.real_lte_2'(s('type.realax.real',sK16(sK17)),s('type.realax.real',sK17))) | p(s(bool,X9)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(sK17)),s('type.realax.real',sK17))))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f55387,f381])).
fof(f381,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~sP0(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP0(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP0(X0,X1,X2) | sP1(X0,X1,X2)))),
  inference(rectify,[],[f238])).
fof(f238,plain,(
  ! [X2,X0,X1] : ((~sP1(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP0(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP0(X2,X0,X1) | sP1(X2,X0,X1)))),
  inference(flattening,[],[f237])).
fof(f237,plain,(
  ! [X2,X0,X1] : ((~sP1(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP0(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP0(X2,X0,X1)) | sP1(X2,X0,X1)))),
  inference(nnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X2,X0,X1] : (sP1(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP0(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f83283,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(sK17)),s('type.realax.real',sK17))))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f55387,f382])).
fof(f1221,plain,(
  ( ! [X62] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X62),s('type.realax.real',sK15))))) ) | $spl102),
  inference(cnf_transformation,[],[f1221_D])).
fof(f1221_D,plain,(
  ( ! [X62] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X62),s('type.realax.real',sK15))))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f1211,plain,(
  ( ! [X38] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X38,sK12)),s('type.realax.real',sK15))))) ) | $spl88),
  inference(cnf_transformation,[],[f1211_D])).
fof(f1211_D,plain,(
  ( ! [X38] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X38,sK12)),s('type.realax.real',sK15))))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f83270,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X4,sK12)),s('type.realax.real',sK15)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',sK15))))) ) | $spl36),
  inference(resolution,[],[f31927,f1230])).
fof(f1230,plain,(
  ( ! [X263,X264] : (~sP11(X264,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X263))) | p(s(bool,X264)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X263),s('type.realax.real',sK15))))) )),
  inference(subsumption_resolution,[],[f1166,f541])).
fof(f541,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f326])).
fof(f326,plain,(
  ! [X0] : ((~p(s(bool,X0)) | ~p(s(bool,'F_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(flattening,[],[f325])).
fof(f325,plain,(
  ! [X0] : (((~p(s(bool,X0)) | ~p(s(bool,'F_0'))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0] : ((p(s(bool,X0)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X12] : ((p(s(bool,X12)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  file('thm.real.REAL_LE_SUP','thm.bool.AND_CLAUSES|split|3')).
fof(f1166,plain,(
  ( ! [X263,X264] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X263),s('type.realax.real',sK15)))) | p(s(bool,'F_0')) | p(s(bool,X264)) | ~sP11(X264,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X263)))) )),
  inference(resolution,[],[f337,f485])).
fof(f485,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP11(X0,X1)) )),
  inference(cnf_transformation,[],[f290])).
fof(f290,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(flattening,[],[f289])).
fof(f289,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(nnf_transformation,[],[f213])).
fof(f213,plain,(
  ! [X0,X1] : (sP11(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f337,plain,(
  ( ! [X7] : (~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X7)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',sK15))))) )),
  inference(cnf_transformation,[],[f218])).
fof(f31927,plain,(
  ( ! [X212,X211] : (sP11('const.real.real_lte_2'(s('type.realax.real',sK24(X211,sK12)),s('type.realax.real',sK15)),X212)) ) | $spl36),
  inference(resolution,[],[f8240,f487])).
fof(f487,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f290])).
fof(f8240,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X0,sK12)),s('type.realax.real',sK15))))) ) | $spl36),
  inference(resolution,[],[f749,f337])).
fof(f83282,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X2,sK12)),s('type.realax.real',sK15)))) | p(s(bool,X3))) ) | $spl36),
  inference(subsumption_resolution,[],[f83269,f541])).
fof(f83269,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X2,sK12)),s('type.realax.real',sK15)))) | p(s(bool,X3))) ) | $spl36),
  inference(resolution,[],[f31927,f485])).
fof(f83281,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X0,sK12)),s('type.realax.real',sK15)))) | p(s(bool,X1))) ) | $spl36),
  inference(subsumption_resolution,[],[f83268,f541])).
fof(f83268,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X0,sK12)),s('type.realax.real',sK15)))) | p(s(bool,X1))) ) | $spl36),
  inference(resolution,[],[f31927,f489])).
fof(f489,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP11(X0,X1)) )),
  inference(cnf_transformation,[],[f292])).
fof(f292,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | (p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(flattening,[],[f291])).
fof(f291,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0')))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(nnf_transformation,[],[f214])).
fof(f214,plain,(
  ! [X0,X1] : (sP11(X0,X1) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(definition_folding,[],[f179,f213])).
fof(f179,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f178])).
fof(f178,plain,(
  ! [X0,X1] : (((p(s(bool,X1)) | p(s(bool,X0))) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(ennf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(flattening,[],[f105])).
fof(f105,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X18,X17] : ((~(p(s(bool,X17)) | p(s(bool,X18))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X17)) => p(s(bool,'F_0'))) => (~p(s(bool,X18)) => p(s(bool,'F_0')))))),
  file('thm.real.REAL_LE_SUP','thm.sat.sat_9_prove_ltb')).
fof(f83255,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X4,sK12)),s('type.realax.real',sK15)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',sK15))))) ) | $spl36),
  inference(resolution,[],[f31926,f1231])).
fof(f1231,plain,(
  ( ! [X265,X266] : (~sP11(happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X265)),X266) | p(s(bool,X266)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X265),s('type.realax.real',sK15))))) )),
  inference(subsumption_resolution,[],[f1167,f541])).
fof(f1167,plain,(
  ( ! [X265,X266] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X265),s('type.realax.real',sK15)))) | p(s(bool,'F_0')) | p(s(bool,X266)) | ~sP11(happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X265)),X266)) )),
  inference(resolution,[],[f337,f485])).
fof(f31926,plain,(
  ( ! [X210,X209] : (sP11(X209,'const.real.real_lte_2'(s('type.realax.real',sK24(X210,sK12)),s('type.realax.real',sK15)))) ) | $spl36),
  inference(resolution,[],[f8240,f486])).
fof(f486,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f290])).
fof(f83267,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X3,sK12)),s('type.realax.real',sK15))))) ) | $spl36),
  inference(subsumption_resolution,[],[f83254,f541])).
fof(f83254,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X2)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X3,sK12)),s('type.realax.real',sK15))))) ) | $spl36),
  inference(resolution,[],[f31926,f485])).
fof(f83266,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X1,sK12)),s('type.realax.real',sK15))))) ) | $spl36),
  inference(subsumption_resolution,[],[f83253,f541])).
fof(f83253,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X1,sK12)),s('type.realax.real',sK15))))) ) | $spl36),
  inference(resolution,[],[f31926,f489])).
fof(f83140,plain,(
  ( ! [X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',sK25(X1,sK12)))))) ) | $spl438),
  inference(cnf_transformation,[],[f83140_D])).
fof(f83140_D,plain,(
  ( ! [X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',sK25(X1,sK12)))))) ) <=> ~$spl438),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl438])])).
fof(f83144,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',sK25(X4,sK12))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',sK15))))) ) | $spl46),
  inference(resolution,[],[f11786,f1230])).
fof(f11786,plain,(
  ( ! [X210,X209] : (sP11('const.real.real_lte_2'(s('type.realax.real',X209),s('type.realax.real',sK25(X209,sK12))),X210)) ) | $spl46),
  inference(resolution,[],[f4594,f487])).
fof(f4594,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',sK25(X0,sK12)))))) ) | $spl46),
  inference(resolution,[],[f755,f366])).
fof(f366,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) )),
  inference(cnf_transformation,[],[f232])).
fof(f232,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) & (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))))),
  inference(nnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) <=> ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0)))))),
  inference(flattening,[],[f81])).
fof(f81,plain,(
  ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) <=> ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0)))))),
  inference(rectify,[],[f63])).
fof(f63,axiom,(
  ! [X7,X8] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X8)))) <=> ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X8),s('type.realax.real',X7)))))),
  file('thm.real.REAL_LE_SUP','thm.real.real_lte')).
fof(f83159,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',sK25(X2,sK12))))) | p(s(bool,X3))) ) | $spl46),
  inference(subsumption_resolution,[],[f83143,f541])).
fof(f83143,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',sK25(X2,sK12))))) | p(s(bool,X3))) ) | $spl46),
  inference(resolution,[],[f11786,f485])).
fof(f83158,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',sK25(X0,sK12))))) | p(s(bool,X1))) ) | $spl46),
  inference(subsumption_resolution,[],[f83142,f541])).
fof(f83142,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',sK25(X0,sK12))))) | p(s(bool,X1))) ) | $spl46),
  inference(resolution,[],[f11786,f489])).
fof(f83125,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',sK25(X4,sK12))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',sK15))))) ) | $spl46),
  inference(resolution,[],[f11785,f1231])).
fof(f11785,plain,(
  ( ! [X208,X207] : (sP11(X207,'const.real.real_lte_2'(s('type.realax.real',X208),s('type.realax.real',sK25(X208,sK12))))) ) | $spl46),
  inference(resolution,[],[f4594,f486])).
fof(f83141,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',sK25(X3,sK12)))))) ) | $spl46),
  inference(subsumption_resolution,[],[f83124,f541])).
fof(f83124,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X2)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',sK25(X3,sK12)))))) ) | $spl46),
  inference(resolution,[],[f11785,f485])).
fof(f83139,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',sK25(X1,sK12)))))) ) | $spl46),
  inference(subsumption_resolution,[],[f83123,f541])).
fof(f83123,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',sK25(X1,sK12)))))) ) | $spl46),
  inference(resolution,[],[f11785,f489])).
fof(f83092,plain,(
  ( ! [X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',sK24(X1,sK12)))))) ) | $spl436),
  inference(cnf_transformation,[],[f83092_D])).
fof(f83092_D,plain,(
  ( ! [X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',sK24(X1,sK12)))))) ) <=> ~$spl436),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl436])])).
fof(f83107,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real',sK24(X4,sK12))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',sK15))))) ) | $spl42),
  inference(resolution,[],[f11543,f1230])).
fof(f11543,plain,(
  ( ! [X213,X214] : (sP11('const.realax.real_lt_2'(s('type.realax.real',X213),s('type.realax.real',sK24(X213,sK12))),X214)) ) | $spl42),
  inference(resolution,[],[f4468,f487])).
fof(f4468,plain,(
  ( ! [X20] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X20),s('type.realax.real',sK24(X20,sK12)))))) ) | $spl42),
  inference(resolution,[],[f753,f366])).
fof(f83122,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X2),s('type.realax.real',sK24(X2,sK12))))) | p(s(bool,X3))) ) | $spl42),
  inference(subsumption_resolution,[],[f83106,f541])).
fof(f83106,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X2),s('type.realax.real',sK24(X2,sK12))))) | p(s(bool,X3))) ) | $spl42),
  inference(resolution,[],[f11543,f485])).
fof(f83121,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',sK24(X0,sK12))))) | p(s(bool,X1))) ) | $spl42),
  inference(subsumption_resolution,[],[f83105,f541])).
fof(f83105,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',sK24(X0,sK12))))) | p(s(bool,X1))) ) | $spl42),
  inference(resolution,[],[f11543,f489])).
fof(f83077,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real',sK24(X4,sK12))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',sK15))))) ) | $spl42),
  inference(resolution,[],[f11542,f1231])).
fof(f11542,plain,(
  ( ! [X212,X211] : (sP11(X211,'const.realax.real_lt_2'(s('type.realax.real',X212),s('type.realax.real',sK24(X212,sK12))))) ) | $spl42),
  inference(resolution,[],[f4468,f486])).
fof(f83093,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X3),s('type.realax.real',sK24(X3,sK12)))))) ) | $spl42),
  inference(subsumption_resolution,[],[f83076,f541])).
fof(f83076,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X2)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X3),s('type.realax.real',sK24(X3,sK12)))))) ) | $spl42),
  inference(resolution,[],[f11542,f485])).
fof(f83091,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',sK24(X1,sK12)))))) ) | $spl42),
  inference(subsumption_resolution,[],[f83075,f541])).
fof(f83075,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',sK24(X1,sK12)))))) ) | $spl42),
  inference(resolution,[],[f11542,f489])).
fof(f83055,plain,(
  ( ! [X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',sK29(X1,sK12)))))) ) | $spl434),
  inference(cnf_transformation,[],[f83055_D])).
fof(f83055_D,plain,(
  ( ! [X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',sK29(X1,sK12)))))) ) <=> ~$spl434),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl434])])).
fof(f83059,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',sK29(X4,sK12))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',sK15))))) ) | $spl62),
  inference(resolution,[],[f8585,f1230])).
fof(f8585,plain,(
  ( ! [X210,X209] : (sP11('const.real.real_lte_2'(s('type.realax.real',X209),s('type.realax.real',sK29(X209,sK12))),X210)) ) | $spl62),
  inference(resolution,[],[f3156,f487])).
fof(f3156,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',sK29(X0,sK12)))))) ) | $spl62),
  inference(resolution,[],[f766,f366])).
fof(f83074,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',sK29(X2,sK12))))) | p(s(bool,X3))) ) | $spl62),
  inference(subsumption_resolution,[],[f83058,f541])).
fof(f83058,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',sK29(X2,sK12))))) | p(s(bool,X3))) ) | $spl62),
  inference(resolution,[],[f8585,f485])).
fof(f83073,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',sK29(X0,sK12))))) | p(s(bool,X1))) ) | $spl62),
  inference(subsumption_resolution,[],[f83057,f541])).
fof(f83057,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',sK29(X0,sK12))))) | p(s(bool,X1))) ) | $spl62),
  inference(resolution,[],[f8585,f489])).
fof(f83040,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',sK29(X4,sK12))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',sK15))))) ) | $spl62),
  inference(resolution,[],[f8584,f1231])).
fof(f8584,plain,(
  ( ! [X208,X207] : (sP11(X207,'const.real.real_lte_2'(s('type.realax.real',X208),s('type.realax.real',sK29(X208,sK12))))) ) | $spl62),
  inference(resolution,[],[f3156,f486])).
fof(f83056,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',sK29(X3,sK12)))))) ) | $spl62),
  inference(subsumption_resolution,[],[f83039,f541])).
fof(f83039,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X2)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',sK29(X3,sK12)))))) ) | $spl62),
  inference(resolution,[],[f8584,f485])).
fof(f83054,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',sK29(X1,sK12)))))) ) | $spl62),
  inference(subsumption_resolution,[],[f83038,f541])).
fof(f83038,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',sK29(X1,sK12)))))) ) | $spl62),
  inference(resolution,[],[f8584,f489])).
fof(f746,plain,(
  ( ! [X41] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X41,sK12)),s('type.realax.real',X41))))) ) | $spl30),
  inference(cnf_transformation,[],[f746_D])).
fof(f746_D,plain,(
  ( ! [X41] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X41,sK12)),s('type.realax.real',X41))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f748,plain,(
  ( ! [X50] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X50),s('type.realax.real',sK23(X50,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X50),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) | $spl34),
  inference(cnf_transformation,[],[f748_D])).
fof(f748_D,plain,(
  ( ! [X50] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X50),s('type.realax.real',sK23(X50,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X50),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f82665,plain,(
  ( ! [X61,X62] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X61),s('type.realax.real',sK23(X61,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X61),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X62,sK12)),s('type.realax.real',X62))))) ) | $spl70),
  inference(resolution,[],[f888,f357])).
fof(f357,plain,(
  ( ! [X6,X4,X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',sK23(X1,X0))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X4,X0)),s('type.realax.real',X4)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X6))))) )),
  inference(cnf_transformation,[],[f226])).
fof(f226,plain,(
  ! [X0] : (! [X6] : ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X6)))) | ! [X4] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',sK22(X4,X0))))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X4,X0)),s('type.realax.real',X4))))) | ! [X1] : ((! [X3] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X3))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0))))))) & (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0)))))) | (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',sK23(X1,X0))))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',sK23(X1,X0)))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK22,sK23])],[f225])).
fof(f225,plain,(
  ! [X0] : (! [X6] : ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X6)))) | ! [X4] : ? [X5] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X5)))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X4))))) | ! [X1] : ((! [X3] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X3))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0))))))) & (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0)))))) | ? [X2] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X2)))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X2))))))))),
  inference(rectify,[],[f224])).
fof(f224,plain,(
  ! [X0] : (! [X3] : ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) | ! [X1] : ? [X2] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X2)))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X1))))) | ! [X4] : ((! [X5] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X5)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real',X5))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0))))))) & (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0)))))) | ? [X5] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X5)))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real',X5))))))))),
  inference(nnf_transformation,[],[f164])).
fof(f164,plain,(
  ! [X0] : (! [X3] : ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) | ! [X1] : ? [X2] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X2)))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X1))))) | ! [X4] : (? [X5] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X5)))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real',X5))))) <=> p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0))))))))),
  inference(flattening,[],[f163])).
fof(f163,plain,(
  ! [X0] : ((! [X3] : ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) | ! [X1] : ? [X2] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X2)))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X1)))))) | ! [X4] : (? [X5] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X5)))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real',X5))))) <=> p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0))))))))),
  inference(ennf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0] : ((? [X3] : p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) & ? [X1] : ! [X2] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X2)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X1)))))) => ! [X4] : (? [X5] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X5)))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real',X5))))) <=> p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0))))))))),
  inference(rectify,[],[f61])).
fof(f61,axiom,(
  ! [X10] : ((? [X7] : p(s(bool,happ(s(fun('type.realax.real',bool),X10),s('type.realax.real',X7)))) & ? [X9] : ! [X7] : (p(s(bool,happ(s(fun('type.realax.real',bool),X10),s('type.realax.real',X7)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X9)))))) => ! [X8] : (? [X7] : (p(s(bool,happ(s(fun('type.realax.real',bool),X10),s('type.realax.real',X7)))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X8),s('type.realax.real',X7))))) <=> p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X8),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X10))))))))),
  file('thm.real.REAL_LE_SUP','thm.real.REAL_SUP_LE')).
fof(f745,plain,(
  ( ! [X35] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X35,sK12)))))) ) | $spl28),
  inference(cnf_transformation,[],[f745_D])).
fof(f745_D,plain,(
  ( ! [X35] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X35,sK12)))))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f82664,plain,(
  ( ! [X59,X60] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X59),s('type.realax.real',sK23(X59,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X59),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X60,sK12)))))) ) | $spl70),
  inference(resolution,[],[f888,f356])).
fof(f356,plain,(
  ( ! [X6,X4,X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',sK23(X1,X0))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',sK22(X4,X0))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X6))))) )),
  inference(cnf_transformation,[],[f226])).
fof(f747,plain,(
  ( ! [X46] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X46,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X46),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) | $spl32),
  inference(cnf_transformation,[],[f747_D])).
fof(f747_D,plain,(
  ( ! [X46] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X46,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X46),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f82663,plain,(
  ( ! [X57,X58] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X57,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X57),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X58,sK12)),s('type.realax.real',X58))))) ) | $spl70),
  inference(resolution,[],[f888,f355])).
fof(f355,plain,(
  ( ! [X6,X4,X0,X1] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',sK23(X1,X0))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X4,X0)),s('type.realax.real',X4)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X6))))) )),
  inference(cnf_transformation,[],[f226])).
fof(f82662,plain,(
  ( ! [X56,X55] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X55,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X55),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X56,sK12)))))) ) | $spl70),
  inference(resolution,[],[f888,f354])).
fof(f354,plain,(
  ( ! [X6,X4,X0,X1] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',sK23(X1,X0))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',sK22(X4,X0))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X6))))) )),
  inference(cnf_transformation,[],[f226])).
fof(f741,plain,(
  ( ! [X17,X18] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X18)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real',X18))))) ) | $spl20),
  inference(cnf_transformation,[],[f741_D])).
fof(f741_D,plain,(
  ( ! [X17,X18] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X18)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real',X18))))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f82661,plain,(
  ( ! [X54,X52,X53] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X52),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X52),s('type.realax.real',X53)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X53)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X54,sK12)),s('type.realax.real',X54))))) ) | $spl70),
  inference(resolution,[],[f888,f353])).
fof(f353,plain,(
  ( ! [X6,X4,X0,X3,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X3)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X4,X0)),s('type.realax.real',X4)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X6))))) )),
  inference(cnf_transformation,[],[f226])).
fof(f736,plain,(
  ( ! [X5] : (~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X5))))) ) | $spl10),
  inference(cnf_transformation,[],[f736_D])).
fof(f736_D,plain,(
  ( ! [X5] : (~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X5))))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f82793,plain,(
  ( ! [X19,X20] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X19),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X19),s('type.realax.real',sK18(X20,sK13,sK12)))))) ) | $spl432),
  inference(cnf_transformation,[],[f82793_D])).
fof(f82793_D,plain,(
  ( ! [X19,X20] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X19),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X19),s('type.realax.real',sK18(X20,sK13,sK12)))))) ) <=> ~$spl432),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl432])])).
fof(f82660,plain,(
  ( ! [X50,X48,X51,X49] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X48),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X48),s('type.realax.real',sK18(X49,sK13,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X50,sK12)),s('type.realax.real',X50)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X51))))) ) | $spl70),
  inference(resolution,[],[f888,f353])).
fof(f82659,plain,(
  ( ! [X47,X45,X46] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X45),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X45),s('type.realax.real',X46)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X46)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X47,sK12)))))) ) | $spl70),
  inference(resolution,[],[f888,f352])).
fof(f352,plain,(
  ( ! [X6,X4,X0,X3,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X3)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) | p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',sK22(X4,X0))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X6))))) )),
  inference(cnf_transformation,[],[f226])).
fof(f82658,plain,(
  ( ! [X43,X41,X44,X42] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X41),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X41),s('type.realax.real',sK18(X42,sK13,sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X43,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X44))))) ) | $spl70),
  inference(resolution,[],[f888,f352])).
fof(f742,plain,(
  ( ! [X21] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X21,sK12)),s('type.realax.real',X21))))) ) | $spl22),
  inference(cnf_transformation,[],[f742_D])).
fof(f742_D,plain,(
  ( ! [X21] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X21,sK12)),s('type.realax.real',X21))))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f744,plain,(
  ( ! [X30] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real',sK21(X30,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) | $spl26),
  inference(cnf_transformation,[],[f744_D])).
fof(f744_D,plain,(
  ( ! [X30] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real',sK21(X30,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f82657,plain,(
  ( ! [X39,X40] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X39),s('type.realax.real',sK21(X39,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X39),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X40,sK12)),s('type.realax.real',X40))))) ) | $spl70),
  inference(resolution,[],[f888,f351])).
fof(f351,plain,(
  ( ! [X6,X4,X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',sK21(X1,X0))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X4,X0)),s('type.realax.real',X4)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X6))))) )),
  inference(cnf_transformation,[],[f223])).
fof(f223,plain,(
  ! [X0] : (! [X6] : ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X6)))) | ! [X4] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',sK20(X4,X0))))) & ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X4,X0)),s('type.realax.real',X4))))) | ! [X1] : ((! [X3] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X3))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0))))))) & (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0)))))) | (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',sK21(X1,X0))))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',sK21(X1,X0)))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20,sK21])],[f222])).
fof(f222,plain,(
  ! [X0] : (! [X6] : ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X6)))) | ! [X4] : ? [X5] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X5)))) & ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X5),s('type.realax.real',X4))))) | ! [X1] : ((! [X3] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X3))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0))))))) & (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0)))))) | ? [X2] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X2)))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X2))))))))),
  inference(rectify,[],[f221])).
fof(f221,plain,(
  ! [X0] : (! [X3] : ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) | ! [X1] : ? [X2] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X2)))) & ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X2),s('type.realax.real',X1))))) | ! [X4] : ((! [X5] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X5)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real',X5))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0))))))) & (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0)))))) | ? [X5] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X5)))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real',X5))))))))),
  inference(nnf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0] : (! [X3] : ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) | ! [X1] : ? [X2] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X2)))) & ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X2),s('type.realax.real',X1))))) | ! [X4] : (? [X5] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X5)))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real',X5))))) <=> p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0))))))))),
  inference(flattening,[],[f161])).
fof(f161,plain,(
  ! [X0] : ((! [X3] : ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) | ! [X1] : ? [X2] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X2)))) & ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X2),s('type.realax.real',X1)))))) | ! [X4] : (? [X5] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X5)))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real',X5))))) <=> p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0))))))))),
  inference(ennf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0] : ((? [X3] : p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) & ? [X1] : ! [X2] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X2)))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X2),s('type.realax.real',X1)))))) => ! [X4] : (? [X5] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X5)))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real',X5))))) <=> p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0))))))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X10] : ((? [X7] : p(s(bool,happ(s(fun('type.realax.real',bool),X10),s('type.realax.real',X7)))) & ? [X9] : ! [X7] : (p(s(bool,happ(s(fun('type.realax.real',bool),X10),s('type.realax.real',X7)))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X7),s('type.realax.real',X9)))))) => ! [X8] : (? [X7] : (p(s(bool,happ(s(fun('type.realax.real',bool),X10),s('type.realax.real',X7)))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X8),s('type.realax.real',X7))))) <=> p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X8),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X10))))))))),
  file('thm.real.REAL_LE_SUP','thm.real.REAL_SUP')).
fof(f739,plain,(
  ( ! [X15] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X15,sK12)))))) ) | $spl16),
  inference(cnf_transformation,[],[f739_D])).
fof(f739_D,plain,(
  ( ! [X15] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X15,sK12)))))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f82656,plain,(
  ( ! [X37,X38] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X37),s('type.realax.real',sK21(X37,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X37),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X38,sK12)))))) ) | $spl70),
  inference(resolution,[],[f888,f350])).
fof(f350,plain,(
  ( ! [X6,X4,X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',sK21(X1,X0))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',sK20(X4,X0))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X6))))) )),
  inference(cnf_transformation,[],[f223])).
fof(f743,plain,(
  ( ! [X26] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X26,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X26),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) | $spl24),
  inference(cnf_transformation,[],[f743_D])).
fof(f743_D,plain,(
  ( ! [X26] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X26,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X26),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f82655,plain,(
  ( ! [X35,X36] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X35,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X35),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X36,sK12)),s('type.realax.real',X36))))) ) | $spl70),
  inference(resolution,[],[f888,f349])).
fof(f349,plain,(
  ( ! [X6,X4,X0,X1] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',sK21(X1,X0))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X4,X0)),s('type.realax.real',X4)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X6))))) )),
  inference(cnf_transformation,[],[f223])).
fof(f82654,plain,(
  ( ! [X33,X34] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X33,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X33),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X34,sK12)))))) ) | $spl70),
  inference(resolution,[],[f888,f348])).
fof(f348,plain,(
  ( ! [X6,X4,X0,X1] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',sK21(X1,X0))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',sK20(X4,X0))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X6))))) )),
  inference(cnf_transformation,[],[f223])).
fof(f82653,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real',X31)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X31)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X32,sK12)),s('type.realax.real',X32))))) ) | $spl70),
  inference(resolution,[],[f888,f347])).
fof(f347,plain,(
  ( ! [X6,X4,X0,X3,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X3)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X4,X0)),s('type.realax.real',X4)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X6))))) )),
  inference(cnf_transformation,[],[f223])).
fof(f82652,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X26),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X26),s('type.realax.real',sK18(X27,sK13,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X28,sK12)),s('type.realax.real',X28)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X29))))) ) | $spl70),
  inference(resolution,[],[f888,f347])).
fof(f82651,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X23),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X23),s('type.realax.real',X24)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X24)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X25,sK12)))))) ) | $spl70),
  inference(resolution,[],[f888,f346])).
fof(f346,plain,(
  ( ! [X6,X4,X0,X3,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X3)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) | p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',sK20(X4,X0))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X6))))) )),
  inference(cnf_transformation,[],[f223])).
fof(f82650,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X19),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X19),s('type.realax.real',sK18(X20,sK13,sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X21,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X22))))) ) | $spl70),
  inference(resolution,[],[f888,f346])).
fof(f82792,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X15,X13,sK12)),s('type.realax.real',X15)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X13),s('type.realax.real',sK18(X14,sK13,sK12)))))) ) | $spl430),
  inference(cnf_transformation,[],[f82792_D])).
fof(f82792_D,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X15,X13,sK12)),s('type.realax.real',X15)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X13),s('type.realax.real',sK18(X14,sK13,sK12)))))) ) <=> ~$spl430),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl430])])).
fof(f82647,plain,(
  ( ! [X14,X15,X13,X16] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X13),s('type.realax.real',sK18(X14,sK13,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X15,X13,sK12)),s('type.realax.real',X15)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X16))))) ) | $spl70),
  inference(resolution,[],[f888,f343])).
fof(f343,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X3,X1,X0)),s('type.realax.real',X3)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X5))))) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1] : (! [X5] : ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X5)))) | ! [X3] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',sK18(X3,X1,X0))))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X3,X1,X0)),s('type.realax.real',X3))))) | ! [X2] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK18])],[f158])).
fof(f158,plain,(
  ! [X0,X1] : (! [X5] : ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X5)))) | ! [X3] : ? [X4] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X4)))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',X3))))) | ! [X2] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0)))))))),
  inference(flattening,[],[f157])).
fof(f157,plain,(
  ! [X0,X1] : ((! [X5] : ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X5)))) | ! [X3] : ? [X4] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X4)))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',X3))))) | ! [X2] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0)))))))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1] : ((? [X5] : p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X5)))) & ? [X3] : ! [X4] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X4)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',X3))))) & ? [X2] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X2)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0)))))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6,X7] : ((? [X11] : p(s(bool,happ(s(fun('type.realax.real',bool),X6),s('type.realax.real',X11)))) & ? [X9] : ! [X11] : (p(s(bool,happ(s(fun('type.realax.real',bool),X6),s('type.realax.real',X11)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',X9))))) & ? [X11] : (p(s(bool,happ(s(fun('type.realax.real',bool),X6),s('type.realax.real',X11)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X11)))))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X6)))))))),
  file('thm.real.REAL_LE_SUP','thm.real.REAL_IMP_LE_SUP')).
fof(f82791,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X8,X6,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',sK18(X7,sK13,sK12)))))) ) | $spl428),
  inference(cnf_transformation,[],[f82791_D])).
fof(f82791_D,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X8,X6,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',sK18(X7,sK13,sK12)))))) ) <=> ~$spl428),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl428])])).
fof(f82645,plain,(
  ( ! [X6,X8,X7,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',sK18(X7,sK13,sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X8,X6,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X9))))) ) | $spl70),
  inference(resolution,[],[f888,f342])).
fof(f342,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X2)))) | p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',sK18(X3,X1,X0))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X5))))) )),
  inference(cnf_transformation,[],[f219])).
fof(f582,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | $spl0),
  inference(cnf_transformation,[],[f582_D])).
fof(f582_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f82630,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))),X5) | ~sP6(X5,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))),X4) | ~sP8(X4,X5,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) ) | $spl426),
  inference(cnf_transformation,[],[f82630_D])).
fof(f82630_D,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))),X5) | ~sP6(X5,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))),X4) | ~sP8(X4,X5,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) ) <=> ~$spl426),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl426])])).
fof(f82625,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))),X5) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~sP8(X4,X5,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))) | ~sP6(X5,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))),X4)) ) | $spl1),
  inference(resolution,[],[f836,f431])).
fof(f431,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X1)) | ~sP8(X0,X2,X1) | ~sP7(X1,X2,X0) | ~sP6(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f262])).
fof(f262,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | (sP6(X2,X1,X0) & sP7(X1,X2,X0) & sP8(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X2,X1,X0) | ~sP7(X1,X2,X0) | ~sP8(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f261])).
fof(f261,plain,(
  ! [X2,X1,X0] : ((~sP9(X2,X1,X0) | (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | ~sP7(X1,X0,X2) | ~sP8(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP9(X2,X1,X0)))),
  inference(flattening,[],[f260])).
fof(f260,plain,(
  ! [X2,X1,X0] : ((~sP9(X2,X1,X0) | (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | ~sP7(X1,X0,X2) | ~sP8(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP9(X2,X1,X0)))),
  inference(nnf_transformation,[],[f209])).
fof(f209,plain,(
  ! [X2,X1,X0] : (sP9(X2,X1,X0) <=> (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f836,plain,(
  ( ! [X121,X122] : (sP7('const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))),X121,X122)) ) | $spl1),
  inference(resolution,[],[f583,f440])).
fof(f440,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f268])).
fof(f268,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f267])).
fof(f267,plain,(
  ! [X1,X0,X2] : ((~sP7(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP7(X1,X0,X2)))),
  inference(flattening,[],[f266])).
fof(f266,plain,(
  ! [X1,X0,X2] : ((~sP7(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP7(X1,X0,X2)))),
  inference(nnf_transformation,[],[f207])).
fof(f207,plain,(
  ! [X1,X0,X2] : (sP7(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f82545,plain,(
  ( ! [X6,X7] : (sP9('const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))),X6,X7) | ~sP6(X7,X6,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))) | ~sP7(X6,X7,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) ) | $spl422),
  inference(cnf_transformation,[],[f82545_D])).
fof(f82545_D,plain,(
  ( ! [X6,X7] : (sP9('const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))),X6,X7) | ~sP6(X7,X6,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))) | ~sP7(X6,X7,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) ) <=> ~$spl422),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl422])])).
fof(f82538,plain,(
  ( ! [X6,X7] : (sP9('const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))),X6,X7) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~sP7(X6,X7,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))) | ~sP6(X7,X6,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) ) | $spl1),
  inference(resolution,[],[f833,f432])).
fof(f432,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X0)) | ~sP8(X0,X2,X1) | ~sP7(X1,X2,X0) | ~sP6(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f262])).
fof(f833,plain,(
  ( ! [X116,X115] : (sP8('const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))),X115,X116)) ) | $spl1),
  inference(resolution,[],[f583,f436])).
fof(f436,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f265])).
fof(f265,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f264])).
fof(f264,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f263])).
fof(f263,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f208])).
fof(f208,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f79902,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK16(sK17))))) | $spl392),
  inference(cnf_transformation,[],[f79902_D])).
fof(f79902_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK16(sK17))))) <=> ~$spl392),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl392])])).
fof(f82481,plain,(
  ( ! [X2] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK16(sK17))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',sK15))))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f55574,f1230])).
fof(f55574,plain,(
  ( ! [X139] : (sP11('const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK16(sK17))),X139)) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f55377,f487])).
fof(f55377,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK16(sK17))))) | ($spl3 | $spl6)),
  inference(subsumption_resolution,[],[f55155,f562])).
fof(f562,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f196])).
fof(f196,plain,(
  ! [X0] : (~p(s(bool,'F_0')) | p(s(bool,X0)))),
  inference(ennf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0] : (p(s(bool,'F_0')) => p(s(bool,X0)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X12] : (p(s(bool,'F_0')) => p(s(bool,X12)))),
  file('thm.real.REAL_LE_SUP','thm.bool.FALSITY')).
fof(f55155,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK16(sK17))))) | ($spl3 | $spl6)),
  inference(superposition,[],[f363,f49371])).
fof(f363,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1] : ((p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) & (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1))))))),
  inference(nnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) <=> p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0)))))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) <=> p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0)))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X7,X8] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X7),s('type.realax.real',X8)))) <=> p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real',X7)))))),
  file('thm.real.REAL_LE_SUP','thm.real.REAL_NOT_LT')).
fof(f82493,plain,(
  ( ! [X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK16(sK17))))) | p(s(bool,X1))) ) | ($spl3 | $spl6)),
  inference(subsumption_resolution,[],[f82480,f541])).
fof(f82480,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK16(sK17))))) | p(s(bool,X1))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f55574,f485])).
fof(f82492,plain,(
  ( ! [X0] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK16(sK17))))) | p(s(bool,X0))) ) | ($spl3 | $spl6)),
  inference(subsumption_resolution,[],[f82479,f541])).
fof(f82479,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK16(sK17))))) | p(s(bool,X0))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f55574,f489])).
fof(f82466,plain,(
  ( ! [X2] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK16(sK17))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',sK15))))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f55573,f1231])).
fof(f55573,plain,(
  ( ! [X138] : (sP11(X138,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK16(sK17))))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f55377,f486])).
fof(f82478,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK16(sK17)))))) ) | ($spl3 | $spl6)),
  inference(subsumption_resolution,[],[f82465,f541])).
fof(f82465,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK16(sK17)))))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f55573,f485])).
fof(f82477,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK16(sK17)))))) ) | ($spl3 | $spl6)),
  inference(subsumption_resolution,[],[f82464,f541])).
fof(f82464,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK16(sK17)))))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f55573,f489])).
fof(f82250,plain,(
  ( ! [X12] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X12,sK14,sK12)),s('type.realax.real',X12))))) ) | $spl408),
  inference(cnf_transformation,[],[f82250_D])).
fof(f82250_D,plain,(
  ( ! [X12] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X12,sK14,sK12)),s('type.realax.real',X12))))) ) <=> ~$spl408),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl408])])).
fof(f82249,plain,(
  ( ! [X8] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',X8)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X8))))) ) | $spl406),
  inference(cnf_transformation,[],[f82249_D])).
fof(f82249_D,plain,(
  ( ! [X8] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',X8)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X8))))) ) <=> ~$spl406),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl406])])).
fof(f82141,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',X11)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X11)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X12,sK14,sK12)),s('type.realax.real',X12)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X13))))) ) | $spl41),
  inference(resolution,[],[f751,f343])).
fof(f751,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | $spl41),
  inference(cnf_transformation,[],[f751_D])).
fof(f751_D,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f82248,plain,(
  ( ! [X9] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X9,sK14,sK12)))))) ) | $spl404),
  inference(cnf_transformation,[],[f82248_D])).
fof(f82248_D,plain,(
  ( ! [X9] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X9,sK14,sK12)))))) ) <=> ~$spl404),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl404])])).
fof(f82140,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',X8)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X8)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X9,sK14,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X10))))) ) | $spl41),
  inference(resolution,[],[f751,f342])).
fof(f81885,plain,(
  ( ! [X61,X62] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X61),s('type.realax.real',sK23(X61,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X61),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X62,sK12)),s('type.realax.real',X62))))) ) | $spl48),
  inference(resolution,[],[f756,f357])).
fof(f81884,plain,(
  ( ! [X59,X60] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X59),s('type.realax.real',sK23(X59,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X59),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X60,sK12)))))) ) | $spl48),
  inference(resolution,[],[f756,f356])).
fof(f81883,plain,(
  ( ! [X57,X58] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X57,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X57),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X58,sK12)),s('type.realax.real',X58))))) ) | $spl48),
  inference(resolution,[],[f756,f355])).
fof(f81882,plain,(
  ( ! [X56,X55] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X55,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X55),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X56,sK12)))))) ) | $spl48),
  inference(resolution,[],[f756,f354])).
fof(f81881,plain,(
  ( ! [X54,X52,X53] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X52),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X52),s('type.realax.real',X53)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X53)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X54,sK12)),s('type.realax.real',X54))))) ) | $spl48),
  inference(resolution,[],[f756,f353])).
fof(f8973,plain,(
  ( ! [X17,X18] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real',sK26(X18,sK12)))))) ) | $spl278),
  inference(cnf_transformation,[],[f8973_D])).
fof(f8973_D,plain,(
  ( ! [X17,X18] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real',sK26(X18,sK12)))))) ) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f81880,plain,(
  ( ! [X50,X48,X51,X49] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X48),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X48),s('type.realax.real',sK26(X49,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X50,sK12)),s('type.realax.real',X50)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X51))))) ) | $spl48),
  inference(resolution,[],[f756,f353])).
fof(f81879,plain,(
  ( ! [X47,X45,X46] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X45),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X45),s('type.realax.real',X46)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X46)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X47,sK12)))))) ) | $spl48),
  inference(resolution,[],[f756,f352])).
fof(f81878,plain,(
  ( ! [X43,X41,X44,X42] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X41),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X41),s('type.realax.real',sK26(X42,sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X43,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X44))))) ) | $spl48),
  inference(resolution,[],[f756,f352])).
fof(f81877,plain,(
  ( ! [X39,X40] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X39),s('type.realax.real',sK21(X39,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X39),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X40,sK12)),s('type.realax.real',X40))))) ) | $spl48),
  inference(resolution,[],[f756,f351])).
fof(f81876,plain,(
  ( ! [X37,X38] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X37),s('type.realax.real',sK21(X37,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X37),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X38,sK12)))))) ) | $spl48),
  inference(resolution,[],[f756,f350])).
fof(f81875,plain,(
  ( ! [X35,X36] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X35,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X35),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X36,sK12)),s('type.realax.real',X36))))) ) | $spl48),
  inference(resolution,[],[f756,f349])).
fof(f81874,plain,(
  ( ! [X33,X34] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X33,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X33),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X34,sK12)))))) ) | $spl48),
  inference(resolution,[],[f756,f348])).
fof(f81873,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real',X31)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X31)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X32,sK12)),s('type.realax.real',X32))))) ) | $spl48),
  inference(resolution,[],[f756,f347])).
fof(f81872,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X26),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X26),s('type.realax.real',sK26(X27,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X28,sK12)),s('type.realax.real',X28)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X29))))) ) | $spl48),
  inference(resolution,[],[f756,f347])).
fof(f81871,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X23),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X23),s('type.realax.real',X24)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X24)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X25,sK12)))))) ) | $spl48),
  inference(resolution,[],[f756,f346])).
fof(f81870,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X19),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X19),s('type.realax.real',sK26(X20,sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X21,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X22))))) ) | $spl48),
  inference(resolution,[],[f756,f346])).
fof(f8972,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X13,X11,sK12)),s('type.realax.real',X13)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK26(X12,sK12)))))) ) | $spl276),
  inference(cnf_transformation,[],[f8972_D])).
fof(f8972_D,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X13,X11,sK12)),s('type.realax.real',X13)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK26(X12,sK12)))))) ) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f81867,plain,(
  ( ! [X14,X15,X13,X16] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X13),s('type.realax.real',sK26(X14,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X15,X13,sK12)),s('type.realax.real',X15)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X16))))) ) | $spl48),
  inference(resolution,[],[f756,f343])).
fof(f8971,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X6,X4,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',sK26(X5,sK12)))))) ) | $spl274),
  inference(cnf_transformation,[],[f8971_D])).
fof(f8971_D,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X6,X4,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',sK26(X5,sK12)))))) ) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f81865,plain,(
  ( ! [X6,X8,X7,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',sK26(X7,sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X8,X6,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X9))))) ) | $spl48),
  inference(resolution,[],[f756,f342])).
fof(f81694,plain,(
  ( ! [X61,X62] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X61),s('type.realax.real',sK23(X61,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X61),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X62,sK12)),s('type.realax.real',X62))))) ) | $spl36),
  inference(resolution,[],[f749,f357])).
fof(f81693,plain,(
  ( ! [X59,X60] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X59),s('type.realax.real',sK23(X59,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X59),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X60,sK12)))))) ) | $spl36),
  inference(resolution,[],[f749,f356])).
fof(f81692,plain,(
  ( ! [X57,X58] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X57,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X57),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X58,sK12)),s('type.realax.real',X58))))) ) | $spl36),
  inference(resolution,[],[f749,f355])).
fof(f81691,plain,(
  ( ! [X56,X55] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X55,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X55),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X56,sK12)))))) ) | $spl36),
  inference(resolution,[],[f749,f354])).
fof(f81690,plain,(
  ( ! [X54,X52,X53] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X52),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X52),s('type.realax.real',X53)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X53)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X54,sK12)),s('type.realax.real',X54))))) ) | $spl36),
  inference(resolution,[],[f749,f353])).
fof(f8390,plain,(
  ( ! [X17,X18] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real',sK24(X18,sK12)))))) ) | $spl254),
  inference(cnf_transformation,[],[f8390_D])).
fof(f8390_D,plain,(
  ( ! [X17,X18] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real',sK24(X18,sK12)))))) ) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f81689,plain,(
  ( ! [X50,X48,X51,X49] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X48),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X48),s('type.realax.real',sK24(X49,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X50,sK12)),s('type.realax.real',X50)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X51))))) ) | $spl36),
  inference(resolution,[],[f749,f353])).
fof(f81688,plain,(
  ( ! [X47,X45,X46] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X45),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X45),s('type.realax.real',X46)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X46)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X47,sK12)))))) ) | $spl36),
  inference(resolution,[],[f749,f352])).
fof(f81687,plain,(
  ( ! [X43,X41,X44,X42] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X41),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X41),s('type.realax.real',sK24(X42,sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X43,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X44))))) ) | $spl36),
  inference(resolution,[],[f749,f352])).
fof(f81686,plain,(
  ( ! [X39,X40] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X39),s('type.realax.real',sK21(X39,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X39),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X40,sK12)),s('type.realax.real',X40))))) ) | $spl36),
  inference(resolution,[],[f749,f351])).
fof(f81685,plain,(
  ( ! [X37,X38] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X37),s('type.realax.real',sK21(X37,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X37),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X38,sK12)))))) ) | $spl36),
  inference(resolution,[],[f749,f350])).
fof(f81684,plain,(
  ( ! [X35,X36] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X35,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X35),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X36,sK12)),s('type.realax.real',X36))))) ) | $spl36),
  inference(resolution,[],[f749,f349])).
fof(f81683,plain,(
  ( ! [X33,X34] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X33,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X33),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X34,sK12)))))) ) | $spl36),
  inference(resolution,[],[f749,f348])).
fof(f81682,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real',X31)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X31)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X32,sK12)),s('type.realax.real',X32))))) ) | $spl36),
  inference(resolution,[],[f749,f347])).
fof(f81681,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X26),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X26),s('type.realax.real',sK24(X27,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X28,sK12)),s('type.realax.real',X28)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X29))))) ) | $spl36),
  inference(resolution,[],[f749,f347])).
fof(f81680,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X23),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X23),s('type.realax.real',X24)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X24)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X25,sK12)))))) ) | $spl36),
  inference(resolution,[],[f749,f346])).
fof(f81679,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X19),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X19),s('type.realax.real',sK24(X20,sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X21,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X22))))) ) | $spl36),
  inference(resolution,[],[f749,f346])).
fof(f8389,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X13,X11,sK12)),s('type.realax.real',X13)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK24(X12,sK12)))))) ) | $spl252),
  inference(cnf_transformation,[],[f8389_D])).
fof(f8389_D,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X13,X11,sK12)),s('type.realax.real',X13)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK24(X12,sK12)))))) ) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f81676,plain,(
  ( ! [X14,X15,X13,X16] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X13),s('type.realax.real',sK24(X14,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X15,X13,sK12)),s('type.realax.real',X15)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X16))))) ) | $spl36),
  inference(resolution,[],[f749,f343])).
fof(f8388,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X6,X4,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',sK24(X5,sK12)))))) ) | $spl250),
  inference(cnf_transformation,[],[f8388_D])).
fof(f8388_D,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X6,X4,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',sK24(X5,sK12)))))) ) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f81674,plain,(
  ( ! [X6,X8,X7,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',sK24(X7,sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X8,X6,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X9))))) ) | $spl36),
  inference(resolution,[],[f749,f342])).
fof(f79906,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK31(sK12))))) | $spl396),
  inference(cnf_transformation,[],[f79906_D])).
fof(f79906_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK31(sK12))))) <=> ~$spl396),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl396])])).
fof(f81396,plain,(
  ( ! [X2] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK31(sK12))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',sK15))))) ) | ($spl6 | $spl134)),
  inference(resolution,[],[f30853,f1230])).
fof(f30853,plain,(
  ( ! [X135] : (sP11('const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK31(sK12))),X135)) ) | ($spl6 | $spl134)),
  inference(resolution,[],[f30771,f487])).
fof(f30771,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK31(sK12))))) | ($spl6 | $spl134)),
  inference(duplicate_literal_removal,[],[f30663])).
fof(f30663,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK31(sK12))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK31(sK12))))) | ($spl6 | $spl134)),
  inference(resolution,[],[f1521,f587])).
fof(f81408,plain,(
  ( ! [X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK31(sK12))))) | p(s(bool,X1))) ) | ($spl6 | $spl134)),
  inference(subsumption_resolution,[],[f81395,f541])).
fof(f81395,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK31(sK12))))) | p(s(bool,X1))) ) | ($spl6 | $spl134)),
  inference(resolution,[],[f30853,f485])).
fof(f81407,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK31(sK12))))) | p(s(bool,X0))) ) | ($spl6 | $spl134)),
  inference(subsumption_resolution,[],[f81394,f541])).
fof(f81394,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK31(sK12))))) | p(s(bool,X0))) ) | ($spl6 | $spl134)),
  inference(resolution,[],[f30853,f489])).
fof(f81381,plain,(
  ( ! [X2] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK31(sK12))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',sK15))))) ) | ($spl6 | $spl134)),
  inference(resolution,[],[f30852,f1231])).
fof(f30852,plain,(
  ( ! [X134] : (sP11(X134,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK31(sK12))))) ) | ($spl6 | $spl134)),
  inference(resolution,[],[f30771,f486])).
fof(f81393,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK31(sK12)))))) ) | ($spl6 | $spl134)),
  inference(subsumption_resolution,[],[f81380,f541])).
fof(f81380,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK31(sK12)))))) ) | ($spl6 | $spl134)),
  inference(resolution,[],[f30852,f485])).
fof(f81392,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK31(sK12)))))) ) | ($spl6 | $spl134)),
  inference(subsumption_resolution,[],[f81379,f541])).
fof(f81379,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK31(sK12)))))) ) | ($spl6 | $spl134)),
  inference(resolution,[],[f30852,f489])).
fof(f81200,plain,(
  $false | ($spl52 | $spl160)),
  inference(subsumption_resolution,[],[f81179,f569])).
fof(f569,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.real.REAL_LE_SUP',aHL_TRUTH)).
fof(f81179,plain,(
  ~p(s(bool,'T_0')) | ($spl52 | $spl160)),
  inference(superposition,[],[f759,f4811])).
fof(f4811,plain,(
  ( ! [X224] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X224,sK12)),s('type.realax.real',sK17)))) ) | $spl160),
  inference(resolution,[],[f2039,f551])).
fof(f551,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f331])).
fof(f331,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.real.REAL_LE_SUP','thm.bool.EQ_CLAUSES|split|1')).
fof(f80570,plain,(
  $false | ($spl42 | $spl102)),
  inference(resolution,[],[f753,f1221])).
fof(f80379,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK16(sK15)))) | $spl402),
  inference(cnf_transformation,[],[f80379_D])).
fof(f80379_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK16(sK15)))) <=> ~$spl402),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl402])])).
fof(f76194,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) | $spl368),
  inference(cnf_transformation,[],[f76194_D])).
fof(f76194_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) <=> ~$spl368),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl368])])).
fof(f80267,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK16(sK15)))) | $spl6),
  inference(resolution,[],[f587,f1180])).
fof(f1180,plain,(
  ( ! [X289] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X289),s('type.realax.real',sK15)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X289)))) )),
  inference(resolution,[],[f337,f547])).
fof(f753,plain,(
  ( ! [X59] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X59,sK12)),s('type.realax.real',X59))))) ) | $spl42),
  inference(cnf_transformation,[],[f753_D])).
fof(f753_D,plain,(
  ( ! [X59] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X59,sK12)),s('type.realax.real',X59))))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f752,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | $spl40),
  inference(cnf_transformation,[],[f752_D])).
fof(f752_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f80255,plain,(
  ( ! [X61,X60] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X60,sK12)),s('type.realax.real',X60)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X61))))) )),
  inference(subsumption_resolution,[],[f12265,f579])).
fof(f579,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f512])).
fof(f512,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f307])).
fof(f307,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f116])).
fof(f116,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X17] : ! [X7] : (s(X17,X7) = s(X17,X7) <=> p(s(bool,'T_0')))),
  file('thm.real.REAL_LE_SUP','thm.bool.REFL_CLAUSE')).
fof(f12265,plain,(
  ( ! [X61,X60] : (~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X60,sK12)),s('type.realax.real',X60)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X61))))) )),
  inference(superposition,[],[f359,f715])).
fof(f715,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK14)))),
  inference(resolution,[],[f336,f551])).
fof(f359,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X2,X0)),s('type.realax.real',X2)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X4))))) )),
  inference(cnf_transformation,[],[f228])).
fof(f228,plain,(
  ! [X0] : (! [X4] : ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X4)))) | ! [X2] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',sK24(X2,X0))))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X2,X0)),s('type.realax.real',X2))))) | ! [X1] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X1)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK24])],[f227])).
fof(f227,plain,(
  ! [X0] : (! [X4] : ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X4)))) | ! [X2] : ? [X3] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X2))))) | ! [X1] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X1)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0))))))))),
  inference(rectify,[],[f166])).
fof(f166,plain,(
  ! [X0] : (! [X3] : ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) | ! [X1] : ? [X2] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X2)))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X1))))) | ! [X4] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X4)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0))))))))),
  inference(flattening,[],[f165])).
fof(f165,plain,(
  ! [X0] : ((! [X3] : ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) | ! [X1] : ? [X2] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X2)))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X1)))))) | ! [X4] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X4)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0))))))))),
  inference(ennf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : ((? [X3] : p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) & ? [X1] : ! [X2] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X2)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X1)))))) => ! [X4] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X4)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0))))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X10] : ((? [X7] : p(s(bool,happ(s(fun('type.realax.real',bool),X10),s('type.realax.real',X7)))) & ? [X9] : ! [X7] : (p(s(bool,happ(s(fun('type.realax.real',bool),X10),s('type.realax.real',X7)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X9)))))) => ! [X8] : (p(s(bool,happ(s(fun('type.realax.real',bool),X10),s('type.realax.real',X8)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X10))))))))),
  file('thm.real.REAL_LE_SUP','thm.real.REAL_SUP_UBOUND_LE')).
fof(f755,plain,(
  ( ! [X67] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X67,sK12)),s('type.realax.real',X67))))) ) | $spl46),
  inference(cnf_transformation,[],[f755_D])).
fof(f755_D,plain,(
  ( ! [X67] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X67,sK12)),s('type.realax.real',X67))))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f80254,plain,(
  ( ! [X68,X69] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X68,sK12)),s('type.realax.real',X68)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X69))))) )),
  inference(subsumption_resolution,[],[f12269,f579])).
fof(f12269,plain,(
  ( ! [X68,X69] : (~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X68,sK12)),s('type.realax.real',X68)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X69))))) )),
  inference(superposition,[],[f361,f715])).
fof(f361,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X2,X0)),s('type.realax.real',X2)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X4))))) )),
  inference(cnf_transformation,[],[f230])).
fof(f230,plain,(
  ! [X0] : (! [X4] : ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X4)))) | ! [X2] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',sK25(X2,X0))))) & ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X2,X0)),s('type.realax.real',X2))))) | ! [X1] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X1)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK25])],[f229])).
fof(f229,plain,(
  ! [X0] : (! [X4] : ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X4)))) | ! [X2] : ? [X3] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) & ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X3),s('type.realax.real',X2))))) | ! [X1] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X1)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0))))))))),
  inference(rectify,[],[f168])).
fof(f168,plain,(
  ! [X0] : (! [X3] : ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) | ! [X1] : ? [X2] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X2)))) & ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X2),s('type.realax.real',X1))))) | ! [X4] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X4)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0))))))))),
  inference(flattening,[],[f167])).
fof(f167,plain,(
  ! [X0] : ((! [X3] : ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) | ! [X1] : ? [X2] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X2)))) & ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X2),s('type.realax.real',X1)))))) | ! [X4] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X4)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0))))))))),
  inference(ennf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : ((? [X3] : p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) & ? [X1] : ! [X2] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X2)))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X2),s('type.realax.real',X1)))))) => ! [X4] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X4)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0))))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X10] : ((? [X7] : p(s(bool,happ(s(fun('type.realax.real',bool),X10),s('type.realax.real',X7)))) & ? [X9] : ! [X7] : (p(s(bool,happ(s(fun('type.realax.real',bool),X10),s('type.realax.real',X7)))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X7),s('type.realax.real',X9)))))) => ! [X8] : (p(s(bool,happ(s(fun('type.realax.real',bool),X10),s('type.realax.real',X8)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X10))))))))),
  file('thm.real.REAL_LE_SUP','thm.real.REAL_SUP_UBOUND')).
fof(f80253,plain,(
  $false | $spl296),
  inference(subsumption_resolution,[],[f80136,f20752])).
fof(f80136,plain,(
  ( ! [X447,X448] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X447),s('type.realax.real',X448))))) ) | $spl296),
  inference(resolution,[],[f20752,f364])).
fof(f364,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) )),
  inference(cnf_transformation,[],[f231])).
fof(f80252,plain,(
  $false | $spl296),
  inference(subsumption_resolution,[],[f80135,f20752])).
fof(f80135,plain,(
  ( ! [X445,X446] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X445),s('type.realax.real',X446))))) ) | $spl296),
  inference(resolution,[],[f20752,f365])).
fof(f365,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) )),
  inference(cnf_transformation,[],[f232])).
fof(f80251,plain,(
  $false | ($spl3 | $spl296)),
  inference(subsumption_resolution,[],[f80117,f20752])).
fof(f80117,plain,(
  ( ! [X419] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X419),s('type.realax.real',sK17))))) ) | ($spl3 | $spl296)),
  inference(resolution,[],[f20752,f49374])).
fof(f49374,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X0)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',sK17))))) ) | $spl3),
  inference(resolution,[],[f585,f362])).
fof(f362,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) )),
  inference(cnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))),
  inference(flattening,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))),
  inference(ennf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X7,X8,X9] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X8)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real',X9))))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X9)))))),
  file('thm.real.REAL_LE_SUP','thm.real.REAL_LE_TRANS')).
fof(f80250,plain,(
  $false | $spl296),
  inference(subsumption_resolution,[],[f80112,f20752])).
fof(f80112,plain,(
  ( ! [X401,X402] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X401),s('type.realax.real',X402))))) ) | $spl296),
  inference(resolution,[],[f20752,f364])).
fof(f80249,plain,(
  $false | $spl296),
  inference(subsumption_resolution,[],[f80111,f20752])).
fof(f80111,plain,(
  ( ! [X399,X400] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X399),s('type.realax.real',X400))))) ) | $spl296),
  inference(resolution,[],[f20752,f365])).
fof(f80248,plain,(
  $false | $spl296),
  inference(subsumption_resolution,[],[f80054,f20752])).
fof(f80054,plain,(
  ~p(s(bool,'T_0')) | $spl296),
  inference(resolution,[],[f20752,f567])).
fof(f567,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f334])).
fof(f334,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f148])).
fof(f148,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f50])).
fof(f50,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.real.REAL_LE_SUP','thm.bool.NOT_CLAUSES|split|2')).
fof(f80247,plain,(
  $false | $spl296),
  inference(subsumption_resolution,[],[f80053,f20752])).
fof(f80053,plain,(
  ~p(s(bool,'T_0')) | $spl296),
  inference(resolution,[],[f20752,f565])).
fof(f565,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f333])).
fof(f333,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f147])).
fof(f147,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f51])).
fof(f51,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.real.REAL_LE_SUP','thm.bool.NOT_CLAUSES|split|1')).
fof(f80246,plain,(
  $false | $spl296),
  inference(subsumption_resolution,[],[f80050,f20752])).
fof(f80050,plain,(
  ( ! [X224] : (~p(s(bool,X224))) ) | $spl296),
  inference(resolution,[],[f20752,f560])).
fof(f560,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0] : ((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(flattening,[],[f141])).
fof(f141,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X12] : ((p(s(bool,X12)) => p(s(bool,'F_0'))) => ~p(s(bool,X12)))),
  file('thm.real.REAL_LE_SUP','thm.bool.IMP_F')).
fof(f80245,plain,(
  $false | $spl296),
  inference(subsumption_resolution,[],[f80046,f20752])).
fof(f80046,plain,(
  ( ! [X220] : (~p(s(bool,X220))) ) | $spl296),
  inference(resolution,[],[f20752,f517])).
fof(f517,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f310])).
fof(f310,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f309])).
fof(f309,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f120])).
fof(f120,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X12] : ((p(s(bool,X12)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X12)))),
  file('thm.real.REAL_LE_SUP','thm.bool.IMP_CLAUSES|split|4')).
fof(f80052,plain,(
  $false | $spl296),
  inference(resolution,[],[f20752,f563])).
fof(f80100,plain,(
  $false | ($spl296 | $spl361)),
  inference(resolution,[],[f20752,f66037])).
fof(f80101,plain,(
  $false | ($spl101 | $spl296)),
  inference(resolution,[],[f20752,f1220])).
fof(f80102,plain,(
  $false | ($spl105 | $spl296)),
  inference(resolution,[],[f20752,f1223])).
fof(f80120,plain,(
  $false | ($spl296 | $spl375)),
  inference(resolution,[],[f20752,f78781])).
fof(f80127,plain,(
  $false | ($spl234 | $spl296)),
  inference(resolution,[],[f20752,f21878])).
fof(f21878,plain,(
  ( ! [X7] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK28(sK12)),s('type.realax.real',sK20(X7,sK12)))))) ) | $spl234),
  inference(resolution,[],[f7988,f365])).
fof(f80128,plain,(
  $false | ($spl246 | $spl296)),
  inference(resolution,[],[f20752,f22133])).
fof(f22133,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK28(sK12)),s('type.realax.real',sK22(X0,sK12)))))) ) | $spl246),
  inference(resolution,[],[f8185,f365])).
fof(f80129,plain,(
  $false | ($spl258 | $spl296)),
  inference(resolution,[],[f20752,f22416])).
fof(f22416,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK28(sK12)),s('type.realax.real',sK24(X0,sK12)))))) ) | $spl258),
  inference(resolution,[],[f8392,f365])).
fof(f80130,plain,(
  $false | ($spl186 | $spl296)),
  inference(resolution,[],[f20752,f11029])).
fof(f11029,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK28(sK12)),s('type.realax.real',sK27(sK12))))) | $spl186),
  inference(resolution,[],[f3984,f365])).
fof(f80131,plain,(
  $false | ($spl198 | $spl296)),
  inference(resolution,[],[f20752,f12111])).
fof(f12111,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK28(sK12)),s('type.realax.real',sK30(sK12))))) | $spl198),
  inference(resolution,[],[f5302,f365])).
fof(f80132,plain,(
  $false | ($spl210 | $spl296)),
  inference(resolution,[],[f20752,f16796])).
fof(f16796,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK28(sK12)),s('type.realax.real',sK38(sK12,'type.realax.real'))))) | $spl210),
  inference(resolution,[],[f7492,f365])).
fof(f80133,plain,(
  $false | ($spl222 | $spl296)),
  inference(resolution,[],[f20752,f17190])).
fof(f17190,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK28(sK12)),s('type.realax.real',sK39(sK12,'type.realax.real'))))) | $spl222),
  inference(resolution,[],[f7700,f365])).
fof(f80134,plain,(
  $false | $spl296),
  inference(resolution,[],[f20752,f575])).
fof(f575,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X0))))) )),
  inference(cnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0] : ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X0))))),
  inference(flattening,[],[f153])).
fof(f153,plain,(
  ! [X0] : ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X0))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X7] : ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X7),s('type.realax.real',X7))))),
  file('thm.real.REAL_LE_SUP','thm.realax.REAL_LT_REFL')).
fof(f80150,plain,(
  $false | ($spl40 | $spl296)),
  inference(resolution,[],[f20752,f12840])).
fof(f12840,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))),s('type.realax.real',sK14)))) | $spl40),
  inference(resolution,[],[f752,f365])).
fof(f80151,plain,(
  $false | ($spl44 | $spl296)),
  inference(resolution,[],[f20752,f21161])).
fof(f21161,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK25(X0,sK12)))))) ) | $spl44),
  inference(resolution,[],[f8639,f365])).
fof(f8639,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK25(X0,sK12)),s('type.realax.real',sK15))))) ) | $spl44),
  inference(resolution,[],[f754,f337])).
fof(f80152,plain,(
  $false | ($spl96 | $spl296)),
  inference(resolution,[],[f20752,f6023])).
fof(f6023,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK29(X0,sK12)))))) ) | $spl96),
  inference(resolution,[],[f1216,f365])).
fof(f80153,plain,(
  $false | ($spl110 | $spl296)),
  inference(resolution,[],[f20752,f3308])).
fof(f3308,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK38(sK12,'type.realax.real'))))) | $spl110),
  inference(resolution,[],[f1227,f365])).
fof(f80154,plain,(
  $false | ($spl112 | $spl296)),
  inference(resolution,[],[f20752,f3447])).
fof(f3447,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK39(sK12,'type.realax.real'))))) | $spl112),
  inference(resolution,[],[f1229,f365])).
fof(f80155,plain,(
  $false | ($spl160 | $spl296)),
  inference(resolution,[],[f20752,f4725])).
fof(f4725,plain,(
  ( ! [X3] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK26(X3,sK12)))))) ) | $spl160),
  inference(resolution,[],[f2039,f365])).
fof(f80156,plain,(
  $false | ($spl174 | $spl296)),
  inference(resolution,[],[f20752,f3587])).
fof(f3587,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK38(sK12,'type.realax.real'))))) | $spl174),
  inference(resolution,[],[f2049,f365])).
fof(f80157,plain,(
  $false | ($spl176 | $spl296)),
  inference(resolution,[],[f20752,f3710])).
fof(f3710,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK39(sK12,'type.realax.real'))))) | $spl176),
  inference(resolution,[],[f2051,f365])).
fof(f80164,plain,(
  $false | ($spl236 | $spl296)),
  inference(resolution,[],[f20752,f22008])).
fof(f22008,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK31(sK12)),s('type.realax.real',sK20(X0,sK12)))))) ) | $spl236),
  inference(resolution,[],[f7990,f365])).
fof(f80165,plain,(
  $false | ($spl248 | $spl296)),
  inference(resolution,[],[f20752,f22261])).
fof(f22261,plain,(
  ( ! [X7] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK31(sK12)),s('type.realax.real',sK22(X7,sK12)))))) ) | $spl248),
  inference(resolution,[],[f8187,f365])).
fof(f80166,plain,(
  $false | ($spl272 | $spl296)),
  inference(resolution,[],[f20752,f22887])).
fof(f22887,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK31(sK12)),s('type.realax.real',sK25(X0,sK12)))))) ) | $spl272),
  inference(resolution,[],[f8793,f365])).
fof(f80167,plain,(
  $false | ($spl188 | $spl296)),
  inference(resolution,[],[f20752,f11163])).
fof(f11163,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK31(sK12)),s('type.realax.real',sK27(sK12))))) | $spl188),
  inference(resolution,[],[f3986,f365])).
fof(f80168,plain,(
  $false | ($spl296 | $spl338)),
  inference(resolution,[],[f20752,f32911])).
fof(f32911,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK31(sK12)),s('type.realax.real',sK29(X0,sK12)))))) ) | $spl338),
  inference(resolution,[],[f32725,f365])).
fof(f80169,plain,(
  $false | ($spl200 | $spl296)),
  inference(resolution,[],[f20752,f12609])).
fof(f12609,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK31(sK12)),s('type.realax.real',sK30(sK12))))) | $spl200),
  inference(resolution,[],[f5304,f365])).
fof(f80170,plain,(
  $false | ($spl212 | $spl296)),
  inference(resolution,[],[f20752,f16937])).
fof(f16937,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK31(sK12)),s('type.realax.real',sK38(sK12,'type.realax.real'))))) | $spl212),
  inference(resolution,[],[f7495,f365])).
fof(f80171,plain,(
  $false | ($spl224 | $spl296)),
  inference(resolution,[],[f20752,f17331])).
fof(f17331,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK31(sK12)),s('type.realax.real',sK39(sK12,'type.realax.real'))))) | $spl224),
  inference(resolution,[],[f7703,f365])).
fof(f79954,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK14),s('type.realax.real',sK13)))) | $spl400),
  inference(cnf_transformation,[],[f79954_D])).
fof(f79954_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK14),s('type.realax.real',sK13)))) <=> ~$spl400),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl400])])).
fof(f79941,plain,(
  ( ! [X2] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK14),s('type.realax.real',sK13)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',sK15))))) ) | ($spl1 | $spl40)),
  inference(resolution,[],[f45819,f1231])).
fof(f45819,plain,(
  ( ! [X138] : (sP11(X138,'const.realax.real_lt_2'(s('type.realax.real',sK14),s('type.realax.real',sK13)))) ) | ($spl1 | $spl40)),
  inference(resolution,[],[f45322,f486])).
fof(f45322,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK14),s('type.realax.real',sK13)))) | ($spl1 | $spl40)),
  inference(resolution,[],[f45201,f366])).
fof(f45201,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK14)))) | ($spl1 | $spl40)),
  inference(resolution,[],[f5566,f752])).
fof(f5566,plain,(
  ( ! [X14] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X14))))) ) | $spl1),
  inference(resolution,[],[f583,f362])).
fof(f79955,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK14),s('type.realax.real',sK13))))) ) | ($spl1 | $spl40)),
  inference(subsumption_resolution,[],[f79940,f541])).
fof(f79940,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK14),s('type.realax.real',sK13))))) ) | ($spl1 | $spl40)),
  inference(resolution,[],[f45819,f485])).
fof(f79952,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK14),s('type.realax.real',sK13))))) ) | ($spl1 | $spl40)),
  inference(subsumption_resolution,[],[f79939,f541])).
fof(f79939,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK14),s('type.realax.real',sK13))))) ) | ($spl1 | $spl40)),
  inference(resolution,[],[f45819,f489])).
fof(f79926,plain,(
  ( ! [X2] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',sK15))))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28575,f1230])).
fof(f28575,plain,(
  ( ! [X135] : (sP11('const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)),X135)) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28494,f487])).
fof(f28494,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f28386])).
fof(f28386,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) | ($spl6 | $spl8)),
  inference(resolution,[],[f1363,f587])).
fof(f1363,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(X0)),s('type.realax.real',sK15)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X0))))) ) | $spl8),
  inference(resolution,[],[f588,f337])).
fof(f79938,plain,(
  ( ! [X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) | p(s(bool,X1))) ) | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f79925,f541])).
fof(f79925,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) | p(s(bool,X1))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28575,f485])).
fof(f79937,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) | p(s(bool,X0))) ) | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f79924,f541])).
fof(f79924,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) | p(s(bool,X0))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28575,f489])).
fof(f79911,plain,(
  ( ! [X2] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',sK15))))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28574,f1231])).
fof(f28574,plain,(
  ( ! [X134] : (sP11(X134,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28494,f486])).
fof(f79923,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f79910,f541])).
fof(f79910,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28574,f485])).
fof(f79922,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl6 | $spl8)),
  inference(subsumption_resolution,[],[f79909,f541])).
fof(f79909,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28574,f489])).
fof(f79898,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13)))) | $spl389),
  inference(cnf_transformation,[],[f79898_D])).
fof(f79898_D,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13)))) <=> ~$spl389),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl389])])).
fof(f79873,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f41873])).
fof(f41873,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | ~p(s(bool,X10)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28630,f388])).
fof(f388,plain,(
  ( ! [X2,X0,X1] : (sP1(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP1(X2,X0,X1)) & (~sP1(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f243])).
fof(f243,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP1(X2,X0,X1)) & (~sP1(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP1(X2,X0,X1))),
  inference(definition_folding,[],[f172,f198,f197])).
fof(f172,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f84])).
fof(f84,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X11,X16,X6] : ((p(s(bool,X6)) <=> (p(s(bool,X16)) => p(s(bool,X11)))) <=> ((p(s(bool,X6)) | p(s(bool,X16))) & (p(s(bool,X6)) | ~p(s(bool,X11))) & (~p(s(bool,X16)) | p(s(bool,X11)) | ~p(s(bool,X6)))))),
  file('thm.real.REAL_LE_SUP','thm.sat.sat_27_prove_ltb')).
fof(f28630,plain,(
  ( ! [X4,X5] : (~sP1('const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13)),X4,X5) | ~p(s(bool,X4))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28501,f376])).
fof(f376,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | ~sP1(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f239])).
fof(f28501,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13)))) | ($spl6 | $spl8)),
  inference(resolution,[],[f28494,f365])).
fof(f79904,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK31(sK12)),s('type.realax.real',sK13)))) | $spl395),
  inference(cnf_transformation,[],[f79904_D])).
fof(f79904_D,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK31(sK12)),s('type.realax.real',sK13)))) <=> ~$spl395),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl395])])).
fof(f79870,plain,(
  ( ! [X2] : (p(s(bool,X2)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK31(sK12)),s('type.realax.real',sK13))))) ) | ($spl6 | $spl134)),
  inference(duplicate_literal_removal,[],[f42560])).
fof(f42560,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,X2)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK31(sK12)),s('type.realax.real',sK13))))) ) | ($spl6 | $spl134)),
  inference(resolution,[],[f30968,f471])).
fof(f471,plain,(
  ( ! [X0,X1] : (sP10(X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f285])).
fof(f285,plain,(
  ! [X0,X1] : ((~sP10(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1)))),
  inference(rectify,[],[f284])).
fof(f284,plain,(
  ! [X1,X0] : ((~sP10(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP10(X1,X0)))),
  inference(flattening,[],[f283])).
fof(f283,plain,(
  ! [X1,X0] : ((~sP10(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP10(X1,X0)))),
  inference(nnf_transformation,[],[f211])).
fof(f211,plain,(
  ! [X1,X0] : (sP10(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f30968,plain,(
  ( ! [X119] : (~sP10('const.realax.real_lt_2'(s('type.realax.real',sK31(sK12)),s('type.realax.real',sK13)),X119) | p(s(bool,X119))) ) | ($spl6 | $spl134)),
  inference(resolution,[],[f30779,f469])).
fof(f469,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | ~sP10(X0,X1)) )),
  inference(cnf_transformation,[],[f285])).
fof(f30779,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK31(sK12)),s('type.realax.real',sK13)))) | ($spl6 | $spl134)),
  inference(resolution,[],[f30771,f365])).
fof(f79869,plain,(
  ( ! [X1] : (p(s(bool,X1)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK31(sK12)),s('type.realax.real',sK13))))) ) | ($spl6 | $spl134)),
  inference(duplicate_literal_removal,[],[f42559])).
fof(f42559,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,X1)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK31(sK12)),s('type.realax.real',sK13))))) ) | ($spl6 | $spl134)),
  inference(resolution,[],[f30968,f476])).
fof(f476,plain,(
  ( ! [X0,X1] : (sP10(X1,X0) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f286])).
fof(f286,plain,(
  ! [X0,X1] : ((((p(s(bool,X1)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)))) | sP10(X1,X0)) & (~sP10(X1,X0) | ((~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X0)) | p(s(bool,X1))))))),
  inference(nnf_transformation,[],[f212])).
fof(f212,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> sP10(X1,X0))),
  inference(definition_folding,[],[f102,f211])).
fof(f102,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(flattening,[],[f101])).
fof(f101,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X16,X6] : ((p(s(bool,X6)) <=> ~p(s(bool,X16))) <=> ((p(s(bool,X6)) | p(s(bool,X16))) & (~p(s(bool,X16)) | ~p(s(bool,X6)))))),
  file('thm.real.REAL_LE_SUP','thm.sat.sat_29_prove_ltb')).
fof(f79867,plain,(
  ( ! [X6] : (p(s(bool,X6)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f40808])).
fof(f40808,plain,(
  ( ! [X6] : (p(s(bool,X6)) | p(s(bool,X6)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28536,f420])).
fof(f420,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f259])).
fof(f259,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP5(X2,X0,X1)) & (~sP5(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f258])).
fof(f258,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP5(X2,X0,X1)) & (~sP5(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f205])).
fof(f205,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP5(X2,X0,X1))),
  inference(definition_folding,[],[f89,f204,f203])).
fof(f203,plain,(
  ! [X0,X1,X2] : (sP4(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f204,plain,(
  ! [X2,X0,X1] : (sP5(X2,X0,X1) <=> (sP4(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f89,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f88])).
fof(f88,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X11,X16,X6] : ((p(s(bool,X6)) <=> (p(s(bool,X16)) & p(s(bool,X11)))) <=> ((p(s(bool,X6)) | ~p(s(bool,X16)) | ~p(s(bool,X11))) & (p(s(bool,X16)) | ~p(s(bool,X6))) & (p(s(bool,X11)) | ~p(s(bool,X6)))))),
  file('thm.real.REAL_LE_SUP','thm.sat.sat_23_prove_ltb')).
fof(f28536,plain,(
  ( ! [X66,X67] : (~sP5('const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)),X67,X66) | p(s(bool,X66))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28494,f410])).
fof(f410,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X2)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f255])).
fof(f255,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP4(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP4(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f254])).
fof(f254,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | (sP4(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP4(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP5(X2,X0,X1)))),
  inference(flattening,[],[f253])).
fof(f253,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | (sP4(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP4(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP5(X2,X0,X1)))),
  inference(nnf_transformation,[],[f204])).
fof(f79864,plain,(
  ( ! [X3] : (p(s(bool,X3)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(sK17)),s('type.realax.real',sK17))))) ) | ($spl3 | $spl6)),
  inference(duplicate_literal_removal,[],[f60729])).
fof(f60729,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,X3)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(sK17)),s('type.realax.real',sK17))))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f55459,f471])).
fof(f55459,plain,(
  ( ! [X262] : (~sP10('const.real.real_lte_2'(s('type.realax.real',sK16(sK17)),s('type.realax.real',sK17)),X262) | p(s(bool,X262))) ) | ($spl3 | $spl6)),
  inference(subsumption_resolution,[],[f55286,f541])).
fof(f55286,plain,(
  ( ! [X262] : (p(s(bool,'F_0')) | p(s(bool,X262)) | ~sP10('const.real.real_lte_2'(s('type.realax.real',sK16(sK17)),s('type.realax.real',sK17)),X262)) ) | ($spl3 | $spl6)),
  inference(superposition,[],[f469,f49371])).
fof(f79863,plain,(
  ( ! [X2] : (p(s(bool,X2)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(sK17)),s('type.realax.real',sK17))))) ) | ($spl3 | $spl6)),
  inference(duplicate_literal_removal,[],[f60728])).
fof(f60728,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,X2)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(sK17)),s('type.realax.real',sK17))))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f55459,f476])).
fof(f79862,plain,(
  ( ! [X17] : (~p(s(bool,X17)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl3 | $spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f58817])).
fof(f58817,plain,(
  ( ! [X17] : (~p(s(bool,X17)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) | ~p(s(bool,X17))) ) | ($spl3 | $spl6 | $spl8)),
  inference(resolution,[],[f49403,f28529])).
fof(f28529,plain,(
  ( ! [X52,X53] : (sP3(X52,X53,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))) | ~p(s(bool,X52))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28494,f404])).
fof(f404,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f252])).
fof(f252,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP3(X2,X0,X1)) & (~sP3(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f251])).
fof(f251,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP3(X2,X0,X1)) & (~sP3(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f202])).
fof(f202,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP3(X2,X0,X1))),
  inference(definition_folding,[],[f87,f201,f200])).
fof(f200,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f201,plain,(
  ! [X2,X0,X1] : (sP3(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP2(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f87,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f86])).
fof(f86,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X11,X16,X6] : ((p(s(bool,X6)) <=> (p(s(bool,X16)) | p(s(bool,X11)))) <=> ((p(s(bool,X6)) | ~p(s(bool,X16))) & (p(s(bool,X6)) | ~p(s(bool,X11))) & (p(s(bool,X16)) | p(s(bool,X11)) | ~p(s(bool,X6)))))),
  file('thm.real.REAL_LE_SUP','thm.sat.sat_25_prove_ltb')).
fof(f49403,plain,(
  ( ! [X57,X58] : (~sP3(X58,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK17)),X57) | ~p(s(bool,X58)) | p(s(bool,X57))) ) | $spl3),
  inference(resolution,[],[f585,f406])).
fof(f406,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP3(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f252])).
fof(f79861,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK16(sK17)))))) ) | ($spl3 | $spl6)),
  inference(duplicate_literal_removal,[],[f61854])).
fof(f61854,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK16(sK17))))) | ~p(s(bool,X3))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f55565,f474])).
fof(f474,plain,(
  ( ! [X0,X1] : (sP10(X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f285])).
fof(f55565,plain,(
  ( ! [X130] : (~sP10('const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK16(sK17))),X130) | ~p(s(bool,X130))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f55377,f470])).
fof(f470,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP10(X0,X1)) )),
  inference(cnf_transformation,[],[f285])).
fof(f79860,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK16(sK17)))))) ) | ($spl3 | $spl6)),
  inference(duplicate_literal_removal,[],[f61851])).
fof(f61851,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK16(sK17)))))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f55565,f475])).
fof(f475,plain,(
  ( ! [X0,X1] : (sP10(X1,X0) | ~p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f286])).
fof(f79859,plain,(
  ( ! [X8] : (p(s(bool,X8)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f41818])).
fof(f41818,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,X8)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28629,f387])).
fof(f387,plain,(
  ( ! [X2,X0,X1] : (sP1(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f244])).
fof(f28629,plain,(
  ( ! [X2,X3] : (~sP1('const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13)),X3,X2) | p(s(bool,X2))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28501,f375])).
fof(f375,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | p(s(bool,X0)) | ~sP1(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f239])).
fof(f79858,plain,(
  ( ! [X7] : (~p(s(bool,X7)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f40773])).
fof(f40773,plain,(
  ( ! [X7] : (~p(s(bool,X7)) | ~p(s(bool,X7)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28532,f393])).
fof(f393,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f247])).
fof(f247,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP2(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP2(X0,X1,X2) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f246])).
fof(f246,plain,(
  ! [X2,X0,X1] : ((~sP3(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP2(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP2(X2,X0,X1) | sP3(X2,X0,X1)))),
  inference(flattening,[],[f245])).
fof(f245,plain,(
  ! [X2,X0,X1] : ((~sP3(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP2(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP2(X2,X0,X1)) | sP3(X2,X0,X1)))),
  inference(nnf_transformation,[],[f201])).
fof(f28532,plain,(
  ( ! [X59,X58] : (sP3('const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)),X58,X59) | ~p(s(bool,X58))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28494,f405])).
fof(f405,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f252])).
fof(f79857,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f40771])).
fof(f40771,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) | ~p(s(bool,X4))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28532,f408])).
fof(f408,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~sP3(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f252])).
fof(f79856,plain,(
  ( ! [X17] : (~p(s(bool,X17)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl3 | $spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f58745])).
fof(f58745,plain,(
  ( ! [X17] : (~p(s(bool,X17)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) | ~p(s(bool,X17))) ) | ($spl3 | $spl6 | $spl8)),
  inference(resolution,[],[f49402,f28531])).
fof(f28531,plain,(
  ( ! [X57,X56] : (sP3(X56,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)),X57) | ~p(s(bool,X56))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28494,f405])).
fof(f49402,plain,(
  ( ! [X56,X55] : (~sP3(X56,X55,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK17))) | ~p(s(bool,X56)) | p(s(bool,X55))) ) | $spl3),
  inference(resolution,[],[f585,f406])).
fof(f79855,plain,(
  ( ! [X3] : (p(s(bool,X3)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK31(sK12)),s('type.realax.real',sK13))))) ) | ($spl6 | $spl134)),
  inference(duplicate_literal_removal,[],[f42517])).
fof(f42517,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,X3)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK31(sK12)),s('type.realax.real',sK13))))) ) | ($spl6 | $spl134)),
  inference(resolution,[],[f30967,f474])).
fof(f30967,plain,(
  ( ! [X118] : (~sP10(X118,'const.realax.real_lt_2'(s('type.realax.real',sK31(sK12)),s('type.realax.real',sK13))) | p(s(bool,X118))) ) | ($spl6 | $spl134)),
  inference(resolution,[],[f30779,f469])).
fof(f79854,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK31(sK12)),s('type.realax.real',sK13))))) ) | ($spl6 | $spl134)),
  inference(duplicate_literal_removal,[],[f42514])).
fof(f42514,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK31(sK12)),s('type.realax.real',sK13)))) | p(s(bool,X0))) ) | ($spl6 | $spl134)),
  inference(resolution,[],[f30967,f475])).
fof(f79851,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f30032])).
fof(f30032,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13)))) | ~p(s(bool,X4))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28690,f470])).
fof(f28690,plain,(
  ( ! [X121] : (sP10('const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13)),X121) | ~p(s(bool,X121))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28501,f474])).
fof(f79850,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f30029])).
fof(f30029,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | ~p(s(bool,X1)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28690,f477])).
fof(f477,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP10(X1,X0)) )),
  inference(cnf_transformation,[],[f286])).
fof(f79848,plain,(
  ( ! [X5] : (~p(s(bool,X5)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(sK17)),s('type.realax.real',sK17))))) ) | ($spl3 | $spl6)),
  inference(duplicate_literal_removal,[],[f61002])).
fof(f61002,plain,(
  ( ! [X5] : (~p(s(bool,X5)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(sK17)),s('type.realax.real',sK17)))) | ~p(s(bool,X5))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f55462,f470])).
fof(f55462,plain,(
  ( ! [X267] : (sP10('const.real.real_lte_2'(s('type.realax.real',sK16(sK17)),s('type.realax.real',sK17)),X267) | ~p(s(bool,X267))) ) | ($spl3 | $spl6)),
  inference(subsumption_resolution,[],[f55291,f517])).
fof(f55291,plain,(
  ( ! [X267] : (p(s(bool,'F_0')) | sP10('const.real.real_lte_2'(s('type.realax.real',sK16(sK17)),s('type.realax.real',sK17)),X267) | ~p(s(bool,X267))) ) | ($spl3 | $spl6)),
  inference(superposition,[],[f474,f49371])).
fof(f79847,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f30008])).
fof(f30008,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | ~p(s(bool,X4)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28689,f470])).
fof(f28689,plain,(
  ( ! [X120] : (sP10(X120,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))) | ~p(s(bool,X120))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28501,f471])).
fof(f79846,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f30005])).
fof(f30005,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13)))) | ~p(s(bool,X1))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28689,f477])).
fof(f79845,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) | p(s(bool,X2))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f40734])).
fof(f40734,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) | p(s(bool,X2)) | ~p(s(bool,X1))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28531,f406])).
fof(f79844,plain,(
  ( ! [X5] : (~p(s(bool,X5)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(sK17)),s('type.realax.real',sK17))))) ) | ($spl3 | $spl6)),
  inference(duplicate_literal_removal,[],[f60890])).
fof(f60890,plain,(
  ( ! [X5] : (~p(s(bool,X5)) | ~p(s(bool,X5)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(sK17)),s('type.realax.real',sK17))))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f55460,f470])).
fof(f55460,plain,(
  ( ! [X265] : (sP10(X265,'const.real.real_lte_2'(s('type.realax.real',sK16(sK17)),s('type.realax.real',sK17))) | ~p(s(bool,X265))) ) | ($spl3 | $spl6)),
  inference(subsumption_resolution,[],[f55289,f517])).
fof(f55289,plain,(
  ( ! [X265] : (p(s(bool,'F_0')) | sP10(X265,'const.real.real_lte_2'(s('type.realax.real',sK16(sK17)),s('type.realax.real',sK17))) | ~p(s(bool,X265))) ) | ($spl3 | $spl6)),
  inference(superposition,[],[f471,f49371])).
fof(f79843,plain,(
  ( ! [X2] : (p(s(bool,X2)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f29971])).
fof(f29971,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,X2)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28688,f471])).
fof(f28688,plain,(
  ( ! [X119] : (~sP10('const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13)),X119) | p(s(bool,X119))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28501,f469])).
fof(f79842,plain,(
  ( ! [X1] : (p(s(bool,X1)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f29970])).
fof(f29970,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,X1)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28688,f476])).
fof(f79841,plain,(
  ( ! [X6,X7] : (p(s(bool,X6)) | p(s(bool,X7)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f41768])).
fof(f41768,plain,(
  ( ! [X6,X7] : (p(s(bool,X6)) | p(s(bool,X7)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13)))) | p(s(bool,X6))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28628,f386])).
fof(f386,plain,(
  ( ! [X2,X0,X1] : (sP1(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f244])).
fof(f28628,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))) | p(s(bool,X0))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28501,f375])).
fof(f79840,plain,(
  ( ! [X9] : (p(s(bool,X9)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK16(sK17)))))) ) | ($spl3 | $spl6)),
  inference(duplicate_literal_removal,[],[f55596])).
fof(f55596,plain,(
  ( ! [X9] : (p(s(bool,X9)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK16(sK17))))) | p(s(bool,X9))) ) | ($spl3 | $spl6)),
  inference(superposition,[],[f55377,f498])).
fof(f498,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f180])).
fof(f180,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X14,X15] : ((p(s(bool,X14)) => p(s(bool,X15))) => ((p(s(bool,X15)) => p(s(bool,X14))) => s(bool,X14) = s(bool,X15)))),
  file('thm.real.REAL_LE_SUP','thm.bool.IMP_ANTISYM_AX')).
fof(f79839,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK16(sK17)))))) ) | ($spl3 | $spl6)),
  inference(duplicate_literal_removal,[],[f55595])).
fof(f55595,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,X8)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK16(sK17)))))) ) | ($spl3 | $spl6)),
  inference(superposition,[],[f55377,f498])).
fof(f79838,plain,(
  ( ! [X5] : (~p(s(bool,X5)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f40700])).
fof(f40700,plain,(
  ( ! [X5] : (~p(s(bool,X5)) | ~p(s(bool,X5)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28530,f392])).
fof(f392,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X0)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f247])).
fof(f28530,plain,(
  ( ! [X54,X55] : (sP3('const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)),X54,X55) | ~p(s(bool,X55))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28494,f404])).
fof(f79837,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f40698])).
fof(f40698,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) | ~p(s(bool,X2))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28530,f407])).
fof(f407,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X1)) | ~sP3(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f252])).
fof(f79836,plain,(
  ( ! [X3] : (p(s(bool,X3)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f29944])).
fof(f29944,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,X3)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28687,f474])).
fof(f28687,plain,(
  ( ! [X118] : (~sP10(X118,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))) | p(s(bool,X118))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28501,f469])).
fof(f79835,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f29941])).
fof(f29941,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13)))) | p(s(bool,X0))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28687,f475])).
fof(f79834,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f40662])).
fof(f40662,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) | ~p(s(bool,X1))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28529,f406])).
fof(f79833,plain,(
  ( ! [X6,X7] : (p(s(bool,X6)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) | p(s(bool,X7))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f40624])).
fof(f40624,plain,(
  ( ! [X6,X7] : (p(s(bool,X6)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) | p(s(bool,X7)) | p(s(bool,X6))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28523,f403])).
fof(f403,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f252])).
fof(f28523,plain,(
  ( ! [X41,X40] : (~sP3(X40,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)),X41) | p(s(bool,X40))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28494,f393])).
fof(f79827,plain,(
  ( ! [X6,X7] : (p(s(bool,X6)) | p(s(bool,X7)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f40584])).
fof(f40584,plain,(
  ( ! [X6,X7] : (p(s(bool,X6)) | p(s(bool,X7)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) | p(s(bool,X6))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28522,f403])).
fof(f28522,plain,(
  ( ! [X39,X38] : (~sP3(X38,X39,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))) | p(s(bool,X38))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28494,f392])).
fof(f79826,plain,(
  ( ! [X13] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) | p(s(bool,X13))) ) | ($spl3 | $spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f58673])).
fof(f58673,plain,(
  ( ! [X13] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) | p(s(bool,X13)) | p(s(bool,X13))) ) | ($spl3 | $spl6 | $spl8)),
  inference(resolution,[],[f49401,f28522])).
fof(f49401,plain,(
  ( ! [X54,X53] : (sP3(X53,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK17)),X54) | p(s(bool,X54)) | p(s(bool,X53))) ) | $spl3),
  inference(resolution,[],[f585,f403])).
fof(f79825,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(sK17)),s('type.realax.real',sK17))))) ) | ($spl3 | $spl6)),
  inference(duplicate_literal_removal,[],[f60887])).
fof(f60887,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(sK17)),s('type.realax.real',sK17)))) | ~p(s(bool,X2))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f55460,f477])).
fof(f79824,plain,(
  ( ! [X7] : (~p(s(bool,X7)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f40502])).
fof(f40502,plain,(
  ( ! [X7] : (~p(s(bool,X7)) | ~p(s(bool,X7)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28518,f376])).
fof(f28518,plain,(
  ( ! [X30,X31] : (sP1('const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)),X30,X31) | ~p(s(bool,X30))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28494,f388])).
fof(f79823,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f40500])).
fof(f40500,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) | ~p(s(bool,X4))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28518,f391])).
fof(f391,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~sP1(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f244])).
fof(f79820,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(sK17)),s('type.realax.real',sK17))))) ) | ($spl3 | $spl6)),
  inference(duplicate_literal_removal,[],[f60999])).
fof(f60999,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | ~p(s(bool,X2)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(sK17)),s('type.realax.real',sK17))))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f55462,f477])).
fof(f79819,plain,(
  ( ! [X11] : (~p(s(bool,X11)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f40464])).
fof(f40464,plain,(
  ( ! [X11] : (~p(s(bool,X11)) | ~p(s(bool,X11)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28517,f647])).
fof(f647,plain,(
  ( ! [X116,X117] : (~sP1(X117,X116,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK14))) | ~p(s(bool,X117)) | p(s(bool,X116))) )),
  inference(resolution,[],[f336,f389])).
fof(f389,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP1(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f244])).
fof(f28517,plain,(
  ( ! [X28,X29] : (sP1(X28,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)),X29) | ~p(s(bool,X28))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28494,f388])).
fof(f79818,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) | ~p(s(bool,X2))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f40458])).
fof(f40458,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) | ~p(s(bool,X2)) | ~p(s(bool,X1))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28517,f389])).
fof(f79813,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f29548])).
fof(f29548,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) | p(s(bool,X3))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28569,f469])).
fof(f28569,plain,(
  ( ! [X129] : (sP10(X129,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))) | p(s(bool,X129))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28494,f474])).
fof(f79812,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f29547])).
fof(f29547,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,X2)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28569,f478])).
fof(f478,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | ~sP10(X1,X0)) )),
  inference(cnf_transformation,[],[f286])).
fof(f79811,plain,(
  ( ! [X5] : (p(s(bool,X5)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f40423])).
fof(f40423,plain,(
  ( ! [X5] : (p(s(bool,X5)) | p(s(bool,X5)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28516,f375])).
fof(f28516,plain,(
  ( ! [X26,X27] : (sP1('const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)),X26,X27) | p(s(bool,X27))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28494,f387])).
fof(f79810,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f40421])).
fof(f40421,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) | p(s(bool,X2))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28516,f390])).
fof(f390,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | p(s(bool,X1)) | ~sP1(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f244])).
fof(f79809,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f29519])).
fof(f29519,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,X3)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28568,f469])).
fof(f28568,plain,(
  ( ! [X128] : (sP10('const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)),X128) | p(s(bool,X128))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28494,f471])).
fof(f79808,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f29518])).
fof(f29518,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) | p(s(bool,X2))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28568,f478])).
fof(f79807,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f29497])).
fof(f29497,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) | ~p(s(bool,X2))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28567,f471])).
fof(f28567,plain,(
  ( ! [X127] : (~sP10(X127,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))) | ~p(s(bool,X127))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28494,f470])).
fof(f79806,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f29496])).
fof(f29496,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) | ~p(s(bool,X1))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28567,f476])).
fof(f79805,plain,(
  ( ! [X13] : (p(s(bool,X13)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f40384])).
fof(f40384,plain,(
  ( ! [X13] : (p(s(bool,X13)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) | p(s(bool,X13))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28507,f643])).
fof(f643,plain,(
  ( ! [X109,X108] : (sP1(X108,X109,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK14))) | p(s(bool,X109)) | p(s(bool,X108))) )),
  inference(resolution,[],[f336,f386])).
fof(f28507,plain,(
  ( ! [X8,X9] : (~sP1(X8,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)),X9) | p(s(bool,X8))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28494,f376])).
fof(f79804,plain,(
  ( ! [X6,X7] : (p(s(bool,X6)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) | ~p(s(bool,X7))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f40380])).
fof(f40380,plain,(
  ( ! [X6,X7] : (p(s(bool,X6)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) | ~p(s(bool,X7)) | p(s(bool,X6))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28507,f386])).
fof(f79803,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f29477])).
fof(f29477,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) | ~p(s(bool,X3))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28566,f474])).
fof(f28566,plain,(
  ( ! [X126] : (~sP10('const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)),X126) | ~p(s(bool,X126))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28494,f470])).
fof(f79802,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f29474])).
fof(f29474,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28566,f475])).
fof(f79801,plain,(
  ( ! [X11] : (p(s(bool,X11)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f42307])).
fof(f42307,plain,(
  ( ! [X11] : (p(s(bool,X11)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13)))) | p(s(bool,X11))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28669,f679])).
fof(f679,plain,(
  ( ! [X180,X181] : (~sP5(X180,X181,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK14))) | ~p(s(bool,X181)) | p(s(bool,X180))) )),
  inference(resolution,[],[f336,f425])).
fof(f425,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP5(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f259])).
fof(f28669,plain,(
  ( ! [X83,X82] : (sP5(X82,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13)),X83) | p(s(bool,X82))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28501,f421])).
fof(f421,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f259])).
fof(f79800,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13)))) | ~p(s(bool,X5))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f42303])).
fof(f42303,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | p(s(bool,X4)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13)))) | ~p(s(bool,X5))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28669,f425])).
fof(f79795,plain,(
  ( ! [X11] : (~p(s(bool,X11)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK31(sK12)),s('type.realax.real',sK13))))) ) | ($spl6 | $spl134)),
  inference(duplicate_literal_removal,[],[f31002])).
fof(f31002,plain,(
  ( ! [X11] : (~p(s(bool,X11)) | ~p(s(bool,X11)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK31(sK12)),s('type.realax.real',sK13))))) ) | ($spl6 | $spl134)),
  inference(superposition,[],[f30779,f501])).
fof(f501,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f181])).
fof(f79794,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK31(sK12)),s('type.realax.real',sK13))))) ) | ($spl6 | $spl134)),
  inference(duplicate_literal_removal,[],[f31001])).
fof(f31001,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK31(sK12)),s('type.realax.real',sK13)))) | ~p(s(bool,X10))) ) | ($spl6 | $spl134)),
  inference(superposition,[],[f30779,f501])).
fof(f79783,plain,(
  ( ! [X9] : (p(s(bool,X9)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f42256])).
fof(f42256,plain,(
  ( ! [X9] : (p(s(bool,X9)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13)))) | p(s(bool,X9))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28668,f411])).
fof(f411,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f255])).
fof(f28668,plain,(
  ( ! [X80,X81] : (sP5('const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13)),X80,X81) | p(s(bool,X80))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28501,f421])).
fof(f79782,plain,(
  ( ! [X2] : (p(s(bool,X2)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f42252])).
fof(f42252,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,X2)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28668,f424])).
fof(f424,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X2)) | ~sP5(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f259])).
fof(f79781,plain,(
  ( ! [X9] : (p(s(bool,X9)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK31(sK12)))))) ) | ($spl6 | $spl134)),
  inference(duplicate_literal_removal,[],[f30875])).
fof(f30875,plain,(
  ( ! [X9] : (p(s(bool,X9)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK31(sK12))))) | p(s(bool,X9))) ) | ($spl6 | $spl134)),
  inference(superposition,[],[f30771,f498])).
fof(f79780,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK31(sK12)))))) ) | ($spl6 | $spl134)),
  inference(duplicate_literal_removal,[],[f30874])).
fof(f30874,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,X8)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK31(sK12)))))) ) | ($spl6 | $spl134)),
  inference(superposition,[],[f30771,f498])).
fof(f79779,plain,(
  ( ! [X11] : (p(s(bool,X11)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f42218])).
fof(f42218,plain,(
  ( ! [X11] : (p(s(bool,X11)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13)))) | p(s(bool,X11))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28667,f678])).
fof(f678,plain,(
  ( ! [X178,X179] : (~sP5(X178,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK14)),X179) | ~p(s(bool,X179)) | p(s(bool,X178))) )),
  inference(resolution,[],[f336,f425])).
fof(f28667,plain,(
  ( ! [X78,X79] : (sP5(X78,X79,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))) | p(s(bool,X78))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28501,f420])).
fof(f79778,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | ~p(s(bool,X5)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f42214])).
fof(f42214,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | p(s(bool,X4)) | ~p(s(bool,X5)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28667,f425])).
fof(f79777,plain,(
  ( ! [X11] : (~p(s(bool,X11)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f28722])).
fof(f28722,plain,(
  ( ! [X11] : (~p(s(bool,X11)) | ~p(s(bool,X11)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl6 | $spl8)),
  inference(superposition,[],[f28501,f501])).
fof(f79776,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f28721])).
fof(f28721,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13)))) | ~p(s(bool,X10))) ) | ($spl6 | $spl8)),
  inference(superposition,[],[f28501,f501])).
fof(f79774,plain,(
  ( ! [X7] : (p(s(bool,X7)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f42166])).
fof(f42166,plain,(
  ( ! [X7] : (p(s(bool,X7)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13)))) | p(s(bool,X7))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28666,f410])).
fof(f28666,plain,(
  ( ! [X76,X77] : (sP5('const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13)),X76,X77) | p(s(bool,X77))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28501,f420])).
fof(f79773,plain,(
  ( ! [X9] : (p(s(bool,X9)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f28597])).
fof(f28597,plain,(
  ( ! [X9] : (p(s(bool,X9)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) | p(s(bool,X9))) ) | ($spl6 | $spl8)),
  inference(superposition,[],[f28494,f498])).
fof(f79772,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f28596])).
fof(f28596,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,X8)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl6 | $spl8)),
  inference(superposition,[],[f28494,f498])).
fof(f79771,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f42162])).
fof(f42162,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,X0)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28666,f423])).
fof(f423,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP5(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f259])).
fof(f79770,plain,(
  ( ! [X23] : (~p(s(bool,X23)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl3 | $spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f57659])).
fof(f57659,plain,(
  ( ! [X23] : (~p(s(bool,X23)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13)))) | ~p(s(bool,X23))) ) | ($spl3 | $spl6 | $spl8)),
  inference(resolution,[],[f49386,f28638])).
fof(f28638,plain,(
  ( ! [X21,X20] : (sP1(X20,X21,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))) | ~p(s(bool,X20))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28501,f387])).
fof(f49386,plain,(
  ( ! [X24,X23] : (~sP1(X24,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK17)),X23) | ~p(s(bool,X24)) | ~p(s(bool,X23))) ) | $spl3),
  inference(resolution,[],[f585,f389])).
fof(f79765,plain,(
  ( ! [X13] : (~p(s(bool,X13)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f42119])).
fof(f42119,plain,(
  ( ! [X13] : (~p(s(bool,X13)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13)))) | ~p(s(bool,X13))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28660,f674])).
fof(f674,plain,(
  ( ! [X171,X170] : (sP5(X170,X171,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK14))) | ~p(s(bool,X171)) | ~p(s(bool,X170))) )),
  inference(resolution,[],[f336,f422])).
fof(f422,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f259])).
fof(f28660,plain,(
  ( ! [X64,X65] : (~sP5(X64,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13)),X65) | ~p(s(bool,X64))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28501,f411])).
fof(f79764,plain,(
  ( ! [X10,X9] : (~p(s(bool,X9)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13)))) | ~p(s(bool,X10))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f42117])).
fof(f42117,plain,(
  ( ! [X10,X9] : (~p(s(bool,X9)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13)))) | ~p(s(bool,X10)) | ~p(s(bool,X9))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28660,f422])).
fof(f79763,plain,(
  ( ! [X13] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) | p(s(bool,X13))) ) | ($spl3 | $spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f58499])).
fof(f58499,plain,(
  ( ! [X13] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) | p(s(bool,X13)) | p(s(bool,X13))) ) | ($spl3 | $spl6 | $spl8)),
  inference(resolution,[],[f49400,f28523])).
fof(f49400,plain,(
  ( ! [X52,X51] : (sP3(X51,X52,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK17))) | p(s(bool,X52)) | p(s(bool,X51))) ) | $spl3),
  inference(resolution,[],[f585,f403])).
fof(f79759,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK16(sK17)))))) ) | ($spl3 | $spl6)),
  inference(duplicate_literal_removal,[],[f61912])).
fof(f61912,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK16(sK17))))) | ~p(s(bool,X2))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f55566,f471])).
fof(f55566,plain,(
  ( ! [X131] : (~sP10(X131,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK16(sK17)))) | ~p(s(bool,X131))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f55377,f470])).
fof(f79758,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK16(sK17)))))) ) | ($spl3 | $spl6)),
  inference(duplicate_literal_removal,[],[f61911])).
fof(f61911,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK16(sK17))))) | ~p(s(bool,X1))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f55566,f476])).
fof(f79751,plain,(
  ( ! [X13] : (~p(s(bool,X13)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f42041])).
fof(f42041,plain,(
  ( ! [X13] : (~p(s(bool,X13)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13)))) | ~p(s(bool,X13))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28659,f673])).
fof(f673,plain,(
  ( ! [X169,X168] : (sP5(X168,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK14)),X169) | ~p(s(bool,X169)) | ~p(s(bool,X168))) )),
  inference(resolution,[],[f336,f422])).
fof(f28659,plain,(
  ( ! [X62,X63] : (~sP5(X62,X63,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))) | ~p(s(bool,X62))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28501,f410])).
fof(f79750,plain,(
  ( ! [X10,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f42039])).
fof(f42039,plain,(
  ( ! [X10,X9] : (~p(s(bool,X9)) | ~p(s(bool,X10)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13)))) | ~p(s(bool,X9))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28659,f422])).
fof(f79749,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK31(sK12)))))) ) | ($spl6 | $spl134)),
  inference(duplicate_literal_removal,[],[f41063])).
fof(f41063,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK31(sK12))))) | p(s(bool,X3))) ) | ($spl6 | $spl134)),
  inference(resolution,[],[f30847,f469])).
fof(f30847,plain,(
  ( ! [X129] : (sP10(X129,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK31(sK12)))) | p(s(bool,X129))) ) | ($spl6 | $spl134)),
  inference(resolution,[],[f30771,f474])).
fof(f79748,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK31(sK12)))))) ) | ($spl6 | $spl134)),
  inference(duplicate_literal_removal,[],[f41062])).
fof(f41062,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,X2)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK31(sK12)))))) ) | ($spl6 | $spl134)),
  inference(resolution,[],[f30847,f478])).
fof(f79747,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f41999])).
fof(f41999,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | ~p(s(bool,X10)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28644,f405])).
fof(f28644,plain,(
  ( ! [X33,X32] : (~sP3('const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13)),X32,X33) | ~p(s(bool,X32))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28501,f393])).
fof(f79746,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK31(sK12)))))) ) | ($spl6 | $spl134)),
  inference(duplicate_literal_removal,[],[f41018])).
fof(f41018,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,X3)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK31(sK12)))))) ) | ($spl6 | $spl134)),
  inference(resolution,[],[f30846,f469])).
fof(f30846,plain,(
  ( ! [X128] : (sP10('const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK31(sK12))),X128) | p(s(bool,X128))) ) | ($spl6 | $spl134)),
  inference(resolution,[],[f30771,f471])).
fof(f79745,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK31(sK12)))))) ) | ($spl6 | $spl134)),
  inference(duplicate_literal_removal,[],[f41017])).
fof(f41017,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK31(sK12))))) | p(s(bool,X2))) ) | ($spl6 | $spl134)),
  inference(resolution,[],[f30846,f478])).
fof(f79743,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK31(sK12)),s('type.realax.real',sK13))))) ) | ($spl6 | $spl134)),
  inference(duplicate_literal_removal,[],[f42652])).
fof(f42652,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK31(sK12)),s('type.realax.real',sK13)))) | ~p(s(bool,X4))) ) | ($spl6 | $spl134)),
  inference(resolution,[],[f30970,f470])).
fof(f30970,plain,(
  ( ! [X121] : (sP10('const.realax.real_lt_2'(s('type.realax.real',sK31(sK12)),s('type.realax.real',sK13)),X121) | ~p(s(bool,X121))) ) | ($spl6 | $spl134)),
  inference(resolution,[],[f30779,f474])).
fof(f79741,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK31(sK12)),s('type.realax.real',sK13))))) ) | ($spl6 | $spl134)),
  inference(duplicate_literal_removal,[],[f42649])).
fof(f42649,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | ~p(s(bool,X1)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK31(sK12)),s('type.realax.real',sK13))))) ) | ($spl6 | $spl134)),
  inference(resolution,[],[f30970,f477])).
fof(f79739,plain,(
  ( ! [X8] : (~p(s(bool,X8)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f41958])).
fof(f41958,plain,(
  ( ! [X8] : (~p(s(bool,X8)) | ~p(s(bool,X8)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28643,f404])).
fof(f28643,plain,(
  ( ! [X30,X31] : (~sP3('const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13)),X31,X30) | ~p(s(bool,X30))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28501,f392])).
fof(f79737,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK31(sK12)))))) ) | ($spl6 | $spl134)),
  inference(duplicate_literal_removal,[],[f40928])).
fof(f40928,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,X2))) ) | ($spl6 | $spl134)),
  inference(resolution,[],[f30845,f471])).
fof(f30845,plain,(
  ( ! [X127] : (~sP10(X127,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK31(sK12)))) | ~p(s(bool,X127))) ) | ($spl6 | $spl134)),
  inference(resolution,[],[f30771,f470])).
fof(f79736,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK31(sK12)))))) ) | ($spl6 | $spl134)),
  inference(duplicate_literal_removal,[],[f40927])).
fof(f40927,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,X1))) ) | ($spl6 | $spl134)),
  inference(resolution,[],[f30845,f476])).
fof(f79730,plain,(
  ( ! [X4] : (p(s(bool,X4)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(sK17)),s('type.realax.real',sK17))))) ) | ($spl3 | $spl6)),
  inference(duplicate_literal_removal,[],[f60639])).
fof(f60639,plain,(
  ( ! [X4] : (p(s(bool,X4)) | p(s(bool,X4)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(sK17)),s('type.realax.real',sK17))))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f55458,f474])).
fof(f55458,plain,(
  ( ! [X261] : (~sP10(X261,'const.real.real_lte_2'(s('type.realax.real',sK16(sK17)),s('type.realax.real',sK17))) | p(s(bool,X261))) ) | ($spl3 | $spl6)),
  inference(subsumption_resolution,[],[f55285,f541])).
fof(f55285,plain,(
  ( ! [X261] : (p(s(bool,'F_0')) | p(s(bool,X261)) | ~sP10(X261,'const.real.real_lte_2'(s('type.realax.real',sK16(sK17)),s('type.realax.real',sK17)))) ) | ($spl3 | $spl6)),
  inference(superposition,[],[f469,f49371])).
fof(f79729,plain,(
  ( ! [X19] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13)))) | p(s(bool,X19))) ) | ($spl3 | $spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f57343])).
fof(f57343,plain,(
  ( ! [X19] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13)))) | p(s(bool,X19)) | p(s(bool,X19))) ) | ($spl3 | $spl6 | $spl8)),
  inference(resolution,[],[f49384,f28628])).
fof(f49384,plain,(
  ( ! [X19,X20] : (sP1(X19,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK17)),X20) | ~p(s(bool,X20)) | p(s(bool,X19))) ) | $spl3),
  inference(resolution,[],[f585,f386])).
fof(f79728,plain,(
  ( ! [X1] : (p(s(bool,X1)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(sK17)),s('type.realax.real',sK17))))) ) | ($spl3 | $spl6)),
  inference(duplicate_literal_removal,[],[f60636])).
fof(f60636,plain,(
  ( ! [X1] : (p(s(bool,X1)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(sK17)),s('type.realax.real',sK17)))) | p(s(bool,X1))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f55458,f475])).
fof(f79727,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK16(sK17)))))) ) | ($spl3 | $spl6)),
  inference(duplicate_literal_removal,[],[f61972])).
fof(f61972,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,X3)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK16(sK17)))))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f55567,f469])).
fof(f55567,plain,(
  ( ! [X132] : (sP10('const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK16(sK17))),X132) | p(s(bool,X132))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f55377,f471])).
fof(f79726,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK16(sK17)))))) ) | ($spl3 | $spl6)),
  inference(duplicate_literal_removal,[],[f61971])).
fof(f61971,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK16(sK17))))) | p(s(bool,X2))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f55567,f478])).
fof(f79725,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK31(sK12)))))) ) | ($spl6 | $spl134)),
  inference(duplicate_literal_removal,[],[f40886])).
fof(f40886,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,X3))) ) | ($spl6 | $spl134)),
  inference(resolution,[],[f30844,f474])).
fof(f30844,plain,(
  ( ! [X126] : (~sP10('const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK31(sK12))),X126) | ~p(s(bool,X126))) ) | ($spl6 | $spl134)),
  inference(resolution,[],[f30771,f470])).
fof(f79724,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK31(sK12)))))) ) | ($spl6 | $spl134)),
  inference(duplicate_literal_removal,[],[f40883])).
fof(f40883,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK31(sK12)))))) ) | ($spl6 | $spl134)),
  inference(resolution,[],[f30844,f475])).
fof(f79723,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f41909])).
fof(f41909,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13)))) | ~p(s(bool,X1))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28638,f389])).
fof(f79722,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK31(sK12)),s('type.realax.real',sK13))))) ) | ($spl6 | $spl134)),
  inference(duplicate_literal_removal,[],[f42606])).
fof(f42606,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | ~p(s(bool,X4)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK31(sK12)),s('type.realax.real',sK13))))) ) | ($spl6 | $spl134)),
  inference(resolution,[],[f30969,f470])).
fof(f30969,plain,(
  ( ! [X120] : (sP10(X120,'const.realax.real_lt_2'(s('type.realax.real',sK31(sK12)),s('type.realax.real',sK13))) | ~p(s(bool,X120))) ) | ($spl6 | $spl134)),
  inference(resolution,[],[f30779,f471])).
fof(f79721,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK31(sK12)),s('type.realax.real',sK13))))) ) | ($spl6 | $spl134)),
  inference(duplicate_literal_removal,[],[f42603])).
fof(f42603,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK31(sK12)),s('type.realax.real',sK13)))) | ~p(s(bool,X1))) ) | ($spl6 | $spl134)),
  inference(resolution,[],[f30969,f477])).
fof(f79720,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK16(sK17)))))) ) | ($spl3 | $spl6)),
  inference(duplicate_literal_removal,[],[f62095])).
fof(f62095,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK16(sK17))))) | p(s(bool,X3))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f55568,f469])).
fof(f55568,plain,(
  ( ! [X133] : (sP10(X133,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK16(sK17)))) | p(s(bool,X133))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f55377,f474])).
fof(f79719,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK16(sK17)))))) ) | ($spl3 | $spl6)),
  inference(duplicate_literal_removal,[],[f62094])).
fof(f62094,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,X2)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK16(sK17)))))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f55568,f478])).
fof(f79715,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f40848])).
fof(f40848,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,X8)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28537,f421])).
fof(f28537,plain,(
  ( ! [X68,X69] : (~sP5('const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)),X68,X69) | p(s(bool,X68))) ) | ($spl6 | $spl8)),
  inference(resolution,[],[f28494,f411])).
fof(f79900,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK14),s('type.realax.real',sK16(sK14))))) | $spl390),
  inference(cnf_transformation,[],[f79900_D])).
fof(f79900_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK14),s('type.realax.real',sK16(sK14))))) <=> ~$spl390),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl390])])).
fof(f79714,plain,(
  ( ! [X9] : (p(s(bool,X9)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK14),s('type.realax.real',sK16(sK14)))))) ) | ($spl1 | $spl6 | $spl40)),
  inference(duplicate_literal_removal,[],[f48944])).
fof(f48944,plain,(
  ( ! [X9] : (p(s(bool,X9)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK14),s('type.realax.real',sK16(sK14))))) | p(s(bool,X9))) ) | ($spl1 | $spl6 | $spl40)),
  inference(superposition,[],[f48725,f498])).
fof(f48725,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK14),s('type.realax.real',sK16(sK14))))) | ($spl1 | $spl6 | $spl40)),
  inference(subsumption_resolution,[],[f48503,f562])).
fof(f48503,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK14),s('type.realax.real',sK16(sK14))))) | ($spl1 | $spl6 | $spl40)),
  inference(superposition,[],[f363,f45321])).
fof(f45321,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(sK14)),s('type.realax.real',sK14))) | ($spl1 | $spl6 | $spl40)),
  inference(resolution,[],[f45201,f1004])).
fof(f79704,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK14),s('type.realax.real',sK16(sK14)))))) ) | ($spl1 | $spl6 | $spl40)),
  inference(duplicate_literal_removal,[],[f48943])).
fof(f48943,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,X8)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK14),s('type.realax.real',sK16(sK14)))))) ) | ($spl1 | $spl6 | $spl40)),
  inference(superposition,[],[f48725,f498])).
fof(f79701,plain,(
  ( ! [X12] : (~p(s(bool,X12)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl1 | $spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f41916])).
fof(f41916,plain,(
  ( ! [X12] : (~p(s(bool,X12)) | ~p(s(bool,X12)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl1 | $spl6 | $spl8)),
  inference(resolution,[],[f28638,f5578])).
fof(f5578,plain,(
  ( ! [X37,X38] : (~sP1(X38,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))),X37) | ~p(s(bool,X38)) | ~p(s(bool,X37))) ) | $spl1),
  inference(resolution,[],[f583,f389])).
fof(f79700,plain,(
  ( ! [X14] : (p(s(bool,X14)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13))))) ) | ($spl1 | $spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f41773])).
fof(f41773,plain,(
  ( ! [X14] : (p(s(bool,X14)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK15),s('type.realax.real',sK13)))) | p(s(bool,X14))) ) | ($spl1 | $spl6 | $spl8)),
  inference(resolution,[],[f28628,f5576])).
fof(f5576,plain,(
  ( ! [X33,X34] : (sP1(X33,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))),X34) | ~p(s(bool,X34)) | p(s(bool,X33))) ) | $spl1),
  inference(resolution,[],[f583,f386])).
fof(f79699,plain,(
  ( ! [X12] : (~p(s(bool,X12)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl1 | $spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f40741])).
fof(f40741,plain,(
  ( ! [X12] : (~p(s(bool,X12)) | ~p(s(bool,X12)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl1 | $spl6 | $spl8)),
  inference(resolution,[],[f28531,f5594])).
fof(f5594,plain,(
  ( ! [X70,X69] : (~sP3(X70,X69,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))) | ~p(s(bool,X70)) | p(s(bool,X69))) ) | $spl1),
  inference(resolution,[],[f583,f406])).
fof(f79698,plain,(
  ( ! [X12] : (~p(s(bool,X12)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl1 | $spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f40669])).
fof(f40669,plain,(
  ( ! [X12] : (~p(s(bool,X12)) | ~p(s(bool,X12)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl1 | $spl6 | $spl8)),
  inference(resolution,[],[f28529,f5595])).
fof(f5595,plain,(
  ( ! [X72,X71] : (~sP3(X72,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))),X71) | ~p(s(bool,X72)) | p(s(bool,X71))) ) | $spl1),
  inference(resolution,[],[f583,f406])).
fof(f79697,plain,(
  ( ! [X14] : (p(s(bool,X14)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl1 | $spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f40629])).
fof(f40629,plain,(
  ( ! [X14] : (p(s(bool,X14)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) | p(s(bool,X14))) ) | ($spl1 | $spl6 | $spl8)),
  inference(resolution,[],[f28523,f5592])).
fof(f5592,plain,(
  ( ! [X66,X65] : (sP3(X65,X66,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))) | p(s(bool,X66)) | p(s(bool,X65))) ) | $spl1),
  inference(resolution,[],[f583,f403])).
fof(f79696,plain,(
  ( ! [X14] : (p(s(bool,X14)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15))))) ) | ($spl1 | $spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f40589])).
fof(f40589,plain,(
  ( ! [X14] : (p(s(bool,X14)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) | p(s(bool,X14))) ) | ($spl1 | $spl6 | $spl8)),
  inference(resolution,[],[f28522,f5593])).
fof(f5593,plain,(
  ( ! [X68,X67] : (sP3(X67,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))),X68) | p(s(bool,X68)) | p(s(bool,X67))) ) | $spl1),
  inference(resolution,[],[f583,f403])).
fof(f766,plain,(
  ( ! [X81] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK29(X81,sK12)),s('type.realax.real',X81))))) ) | $spl62),
  inference(cnf_transformation,[],[f766_D])).
fof(f766_D,plain,(
  ( ! [X81] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK29(X81,sK12)),s('type.realax.real',X81))))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f79896,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK19(sK17,sK12)),s('type.realax.real',sK31(sK12))))) | $spl386),
  inference(cnf_transformation,[],[f79896_D])).
fof(f79896_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK19(sK17,sK12)),s('type.realax.real',sK31(sK12))))) <=> ~$spl386),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl386])])).
fof(f79584,plain,(
  ( ! [X87,X86] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK19(sK17,sK12)),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK29(X86,sK12)),s('type.realax.real',X86)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X87))))) ) | $spl372),
  inference(resolution,[],[f78779,f374])).
fof(f374,plain,(
  ( ! [X6,X4,X2,X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',sK31(X0))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X2)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK29(X4,X0)),s('type.realax.real',X4)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X6))))) )),
  inference(cnf_transformation,[],[f236])).
fof(f236,plain,(
  ! [X0] : ((! [X12] : ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X12)))) | ! [X10] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',sK26(X10,X0))))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X10,X0)),s('type.realax.real',X10))))) | (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',sK27(X0))))) & ! [X8] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X8)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X8),s('type.realax.real',sK28(X0)))))))) & (! [X6] : ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X6)))) | ! [X4] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',sK29(X4,X0))))) & ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK29(X4,X0)),s('type.realax.real',X4))))) | (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',sK30(X0))))) & ! [X2] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X2)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',sK31(X0)))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK26,sK27,sK28,sK29,sK30,sK31])],[f235])).
fof(f235,plain,(
  ! [X0] : ((! [X12] : ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X12)))) | ! [X10] : ? [X11] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X11)))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',X10))))) | (? [X9] : p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X9)))) & ? [X7] : ! [X8] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X8)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X8),s('type.realax.real',X7))))))) & (! [X6] : ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X6)))) | ! [X4] : ? [X5] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X5)))) & ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X5),s('type.realax.real',X4))))) | (? [X3] : p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) & ? [X1] : ! [X2] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X2)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X1))))))))),
  inference(rectify,[],[f234])).
fof(f234,plain,(
  ! [X0] : ((! [X3] : ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) | ! [X1] : ? [X2] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X2)))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X1))))) | (? [X6] : p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X6)))) & ? [X4] : ! [X5] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X5)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X5),s('type.realax.real',X4))))))) & (! [X6] : ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X6)))) | ! [X4] : ? [X5] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X5)))) & ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X5),s('type.realax.real',X4))))) | (? [X3] : p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) & ? [X1] : ! [X2] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X2)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X1))))))))),
  inference(flattening,[],[f233])).
fof(f233,plain,(
  ! [X0] : (((! [X3] : ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) | ! [X1] : ? [X2] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X2)))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X1)))))) | (? [X6] : p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X6)))) & ? [X4] : ! [X5] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X5)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X5),s('type.realax.real',X4))))))) & ((! [X6] : ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X6)))) | ! [X4] : ? [X5] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X5)))) & ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X5),s('type.realax.real',X4)))))) | (? [X3] : p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) & ? [X1] : ! [X2] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X2)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X1))))))))),
  inference(nnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0] : ((? [X3] : p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) & ? [X1] : ! [X2] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X2)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X1)))))) <=> (? [X6] : p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X6)))) & ? [X4] : ! [X5] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X5)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X5),s('type.realax.real',X4)))))))),
  inference(ennf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0] : ((? [X3] : p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) & ? [X1] : ! [X2] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X2)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X1)))))) <=> (? [X6] : p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X6)))) & ? [X4] : ! [X5] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X5)))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X5),s('type.realax.real',X4)))))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X10] : ((? [X7] : p(s(bool,happ(s(fun('type.realax.real',bool),X10),s('type.realax.real',X7)))) & ? [X9] : ! [X7] : (p(s(bool,happ(s(fun('type.realax.real',bool),X10),s('type.realax.real',X7)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X9)))))) <=> (? [X7] : p(s(bool,happ(s(fun('type.realax.real',bool),X10),s('type.realax.real',X7)))) & ? [X9] : ! [X7] : (p(s(bool,happ(s(fun('type.realax.real',bool),X10),s('type.realax.real',X7)))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X7),s('type.realax.real',X9)))))))),
  file('thm.real.REAL_LE_SUP','thm.real.SETOK_LE_LT')).
fof(f759,plain,(
  ( ! [X71] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X71,sK12)),s('type.realax.real',X71))))) ) | $spl52),
  inference(cnf_transformation,[],[f759_D])).
fof(f759_D,plain,(
  ( ! [X71] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X71,sK12)),s('type.realax.real',X71))))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f79894,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK19(sK17,sK12)),s('type.realax.real',sK28(sK12))))) | $spl384),
  inference(cnf_transformation,[],[f79894_D])).
fof(f79894_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK19(sK17,sK12)),s('type.realax.real',sK28(sK12))))) <=> ~$spl384),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl384])])).
fof(f79578,plain,(
  ( ! [X76,X77] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK19(sK17,sK12)),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X76,sK12)),s('type.realax.real',X76)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X77))))) ) | $spl372),
  inference(resolution,[],[f78779,f370])).
fof(f370,plain,(
  ( ! [X0,X12,X10,X8] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X8),s('type.realax.real',sK28(X0))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X8)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X10,X0)),s('type.realax.real',X10)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X12))))) )),
  inference(cnf_transformation,[],[f236])).
fof(f756,plain,(
  ( ! [X70] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X70,sK12)))))) ) | $spl48),
  inference(cnf_transformation,[],[f756_D])).
fof(f756_D,plain,(
  ( ! [X70] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X70,sK12)))))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f79576,plain,(
  ( ! [X72,X73] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK19(sK17,sK12)),s('type.realax.real',sK28(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X72,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X73))))) ) | $spl372),
  inference(resolution,[],[f78779,f369])).
fof(f369,plain,(
  ( ! [X0,X12,X10,X8] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X8),s('type.realax.real',sK28(X0))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X8)))) | p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',sK26(X10,X0))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X12))))) )),
  inference(cnf_transformation,[],[f236])).
fof(f79892,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK19(sK17,sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | $spl382),
  inference(cnf_transformation,[],[f79892_D])).
fof(f79892_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK19(sK17,sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) <=> ~$spl382),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl382])])).
fof(f79573,plain,(
  ( ! [X68,X69] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK19(sK17,sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X68,sK12)),s('type.realax.real',X68)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X69))))) ) | $spl372),
  inference(resolution,[],[f78779,f361])).
fof(f750,plain,(
  ( ! [X54] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X54),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X54))))) ) | $spl38),
  inference(cnf_transformation,[],[f750_D])).
fof(f750_D,plain,(
  ( ! [X54] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X54),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X54))))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f79572,plain,(
  ( ! [X66,X67] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X66),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X66)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X67,sK12)),s('type.realax.real',X67))))) ) | $spl372),
  inference(resolution,[],[f78779,f361])).
fof(f79569,plain,(
  ( ! [X61,X60] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK19(sK17,sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X60,sK12)),s('type.realax.real',X60)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X61))))) ) | $spl372),
  inference(resolution,[],[f78779,f359])).
fof(f79568,plain,(
  ( ! [X59,X58] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X58),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X58)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X59,sK12)),s('type.realax.real',X59))))) ) | $spl372),
  inference(resolution,[],[f78779,f359])).
fof(f749,plain,(
  ( ! [X55] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X55,sK12)))))) ) | $spl36),
  inference(cnf_transformation,[],[f749_D])).
fof(f749_D,plain,(
  ( ! [X55] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X55,sK12)))))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f79567,plain,(
  ( ! [X57,X56] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK19(sK17,sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X56,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X57))))) ) | $spl372),
  inference(resolution,[],[f78779,f358])).
fof(f358,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X1)))) | p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',sK24(X2,X0))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X4))))) )),
  inference(cnf_transformation,[],[f228])).
fof(f79566,plain,(
  ( ! [X54,X55] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X54),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X54)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X55,sK12)))))) ) | $spl372),
  inference(resolution,[],[f78779,f358])).
fof(f79565,plain,(
  ( ! [X52,X53] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X52),s('type.realax.real',sK23(X52,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X52),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X53,sK12)),s('type.realax.real',X53))))) ) | $spl372),
  inference(resolution,[],[f78779,f357])).
fof(f79564,plain,(
  ( ! [X50,X51] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X50),s('type.realax.real',sK23(X50,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X50),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X51,sK12)))))) ) | $spl372),
  inference(resolution,[],[f78779,f356])).
fof(f79563,plain,(
  ( ! [X48,X49] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X48,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X48),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X49,sK12)),s('type.realax.real',X49))))) ) | $spl372),
  inference(resolution,[],[f78779,f355])).
fof(f79562,plain,(
  ( ! [X47,X46] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X46,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X46),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X47,sK12)))))) ) | $spl372),
  inference(resolution,[],[f78779,f354])).
fof(f79561,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X43),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X43),s('type.realax.real',X44)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X44)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X45,sK12)),s('type.realax.real',X45))))) ) | $spl372),
  inference(resolution,[],[f78779,f353])).
fof(f79890,plain,(
  ( ! [X14] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real',sK19(sK17,sK12)))))) ) | $spl380),
  inference(cnf_transformation,[],[f79890_D])).
fof(f79890_D,plain,(
  ( ! [X14] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real',sK19(sK17,sK12)))))) ) <=> ~$spl380),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl380])])).
fof(f79560,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X40),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X40),s('type.realax.real',sK19(sK17,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X41,sK12)),s('type.realax.real',X41)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X42))))) ) | $spl372),
  inference(resolution,[],[f78779,f353])).
fof(f79559,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X37),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X37),s('type.realax.real',X38)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X38)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X39,sK12)))))) ) | $spl372),
  inference(resolution,[],[f78779,f352])).
fof(f79558,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X34),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X34),s('type.realax.real',sK19(sK17,sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X35,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X36))))) ) | $spl372),
  inference(resolution,[],[f78779,f352])).
fof(f79557,plain,(
  ( ! [X33,X32] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X32),s('type.realax.real',sK21(X32,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X32),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X33,sK12)),s('type.realax.real',X33))))) ) | $spl372),
  inference(resolution,[],[f78779,f351])).
fof(f79556,plain,(
  ( ! [X30,X31] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real',sK21(X30,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X31,sK12)))))) ) | $spl372),
  inference(resolution,[],[f78779,f350])).
fof(f79555,plain,(
  ( ! [X28,X29] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X28,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X28),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X29,sK12)),s('type.realax.real',X29))))) ) | $spl372),
  inference(resolution,[],[f78779,f349])).
fof(f79554,plain,(
  ( ! [X26,X27] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X26,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X26),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X27,sK12)))))) ) | $spl372),
  inference(resolution,[],[f78779,f348])).
fof(f79553,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X23),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X23),s('type.realax.real',X24)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X24)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X25,sK12)),s('type.realax.real',X25))))) ) | $spl372),
  inference(resolution,[],[f78779,f347])).
fof(f79552,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X20),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X20),s('type.realax.real',sK19(sK17,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X21,sK12)),s('type.realax.real',X21)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X22))))) ) | $spl372),
  inference(resolution,[],[f78779,f347])).
fof(f79551,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real',X18)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X18)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X19,sK12)))))) ) | $spl372),
  inference(resolution,[],[f78779,f346])).
fof(f79550,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real',sK19(sK17,sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X15,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X16))))) ) | $spl372),
  inference(resolution,[],[f78779,f346])).
fof(f79889,plain,(
  ( ! [X10,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X10,X9,sK12)),s('type.realax.real',X10)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',sK19(sK17,sK12)))))) ) | $spl378),
  inference(cnf_transformation,[],[f79889_D])).
fof(f79889_D,plain,(
  ( ! [X10,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X10,X9,sK12)),s('type.realax.real',X10)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',sK19(sK17,sK12)))))) ) <=> ~$spl378),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl378])])).
fof(f79547,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',sK19(sK17,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X10,X9,sK12)),s('type.realax.real',X10)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X11))))) ) | $spl372),
  inference(resolution,[],[f78779,f343])).
fof(f79888,plain,(
  ( ! [X4,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X4,X3,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',sK19(sK17,sK12)))))) ) | $spl376),
  inference(cnf_transformation,[],[f79888_D])).
fof(f79888_D,plain,(
  ( ! [X4,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X4,X3,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',sK19(sK17,sK12)))))) ) <=> ~$spl376),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl376])])).
fof(f79545,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',sK19(sK17,sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X4,X3,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X5))))) ) | $spl372),
  inference(resolution,[],[f78779,f342])).
fof(f79874,plain,(
  $false | ($spl58 | $spl62 | $spl72)),
  inference(subsumption_resolution,[],[f28775,f763])).
fof(f28775,plain,(
  ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(sK13,sK12))))) | ($spl62 | $spl72)),
  inference(resolution,[],[f889,f3156])).
fof(f8791,plain,(
  ( ! [X83] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X83,sK12)),s('type.realax.real',sK28(sK12)))))) ) | $spl270),
  inference(cnf_transformation,[],[f8791_D])).
fof(f8791_D,plain,(
  ( ! [X83] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X83,sK12)),s('type.realax.real',sK28(sK12)))))) ) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f31508,plain,(
  ( ! [X47,X45,X46] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X45,sK12)),s('type.realax.real',sK28(sK12))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X46),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X46)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X47))))) ) | $spl56),
  inference(resolution,[],[f762,f360])).
fof(f360,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X1)))) | p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',sK25(X2,X0))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X4))))) )),
  inference(cnf_transformation,[],[f230])).
fof(f79691,plain,(
  $false | ($spl4 | $spl372 | $spl375)),
  inference(subsumption_resolution,[],[f79540,f78781])).
fof(f79540,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK19(sK17,sK12)),s('type.realax.real',sK17)))) | ($spl4 | $spl372)),
  inference(resolution,[],[f78779,f586])).
fof(f78781,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK19(sK17,sK12)),s('type.realax.real',sK17)))) | $spl375),
  inference(cnf_transformation,[],[f78781_D])).
fof(f78781_D,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK19(sK17,sK12)),s('type.realax.real',sK17)))) <=> ~$spl375),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl375])])).
fof(f78670,plain,(
  ( ! [X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK19(sK17,sK12)),s('type.realax.real',sK17)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X1))))) ) | ($spl0 | $spl3)),
  inference(resolution,[],[f65563,f345])).
fof(f345,plain,(
  ( ! [X0,X3,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0))),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK19(X1,X0)),s('type.realax.real',X1)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3))))) )),
  inference(cnf_transformation,[],[f220])).
fof(f220,plain,(
  ! [X0,X1] : (! [X3] : ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) | (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',sK19(X1,X0))))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK19(X1,X0)),s('type.realax.real',X1))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0))),s('type.realax.real',X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK19])],[f160])).
fof(f160,plain,(
  ! [X0,X1] : (! [X3] : ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) | ? [X2] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X2)))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X1))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0))),s('type.realax.real',X1)))))),
  inference(flattening,[],[f159])).
fof(f159,plain,(
  ! [X0,X1] : ((! [X3] : ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) | ? [X2] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X2)))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X1)))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0))),s('type.realax.real',X1)))))),
  inference(ennf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1] : ((? [X3] : p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3)))) & ! [X2] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X2)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X1)))))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0))),s('type.realax.real',X1)))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6,X7] : ((? [X11] : p(s(bool,happ(s(fun('type.realax.real',bool),X6),s('type.realax.real',X11)))) & ! [X11] : (p(s(bool,happ(s(fun('type.realax.real',bool),X6),s('type.realax.real',X11)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',X7)))))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X6))),s('type.realax.real',X7)))))),
  file('thm.real.REAL_LE_SUP','thm.real.REAL_IMP_SUP_LE')).
fof(f65563,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))),s('type.realax.real',sK17)))) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f65533,f579])).
fof(f65533,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))),s('type.realax.real',sK17)))) | ($spl0 | $spl3)),
  inference(superposition,[],[f49374,f49572])).
fof(f49572,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))) | $spl0),
  inference(resolution,[],[f582,f551])).
fof(f78779,plain,(
  p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK19(sK17,sK12))))) | $spl372),
  inference(cnf_transformation,[],[f78779_D])).
fof(f78779_D,plain,(
  p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK19(sK17,sK12))))) <=> ~$spl372),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl372])])).
fof(f78669,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK19(sK17,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X0))))) ) | ($spl0 | $spl3)),
  inference(resolution,[],[f65563,f344])).
fof(f344,plain,(
  ( ! [X0,X3,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),X0))),s('type.realax.real',X1)))) | p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',sK19(X1,X0))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X3))))) )),
  inference(cnf_transformation,[],[f220])).
fof(f76554,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') != s(bool,X2)) ) | $spl370),
  inference(cnf_transformation,[],[f76554_D])).
fof(f76554_D,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') != s(bool,X2)) ) <=> ~$spl370),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl370])])).
fof(f76553,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) )),
  inference(duplicate_literal_removal,[],[f76540])).
fof(f76540,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) )),
  inference(resolution,[],[f49355,f448])).
fof(f448,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | s(bool,X0) != s(bool,X1) | ~sP9(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f271])).
fof(f271,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP9(X2,X1,X0)) & (~sP9(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f210])).
fof(f210,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP9(X2,X1,X0))),
  inference(definition_folding,[],[f91,f209,f208,f207,f206])).
fof(f206,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f91,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f90])).
fof(f90,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X11,X16,X6] : ((p(s(bool,X6)) <=> s(bool,X11) = s(bool,X16)) <=> ((p(s(bool,X6)) | p(s(bool,X16)) | p(s(bool,X11))) & (p(s(bool,X6)) | ~p(s(bool,X11)) | ~p(s(bool,X16))) & (p(s(bool,X16)) | ~p(s(bool,X11)) | ~p(s(bool,X6))) & (p(s(bool,X11)) | ~p(s(bool,X16)) | ~p(s(bool,X6)))))),
  file('thm.real.REAL_LE_SUP','thm.sat.sat_2|49|_prove_ltb')).
fof(f49355,plain,(
  ( ! [X4,X3] : (sP9(X4,X3,'F_0') | s(bool,'T_0') != s(bool,X3) | p(s(bool,X4))) )),
  inference(superposition,[],[f12226,f445])).
fof(f445,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f271])).
fof(f12226,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(backward_demodulation,[],[f715,f713])).
fof(f713,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK14)))),
  inference(resolution,[],[f336,f546])).
fof(f546,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f329])).
fof(f76192,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))) | $spl366),
  inference(cnf_transformation,[],[f76192_D])).
fof(f76192_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))) <=> ~$spl366),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl366])])).
fof(f76080,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK15)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))) | $spl0),
  inference(resolution,[],[f49488,f1180])).
fof(f49488,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))),s('type.realax.real',X0)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X0))))) ) | $spl0),
  inference(resolution,[],[f582,f362])).
fof(f76190,plain,(
  ( ! [X2] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK19(X2,sK12)),s('type.realax.real',X2))))) ) | $spl364),
  inference(cnf_transformation,[],[f76190_D])).
fof(f76190_D,plain,(
  ( ! [X2] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK19(X2,sK12)),s('type.realax.real',X2))))) ) <=> ~$spl364),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl364])])).
fof(f76070,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK19(X2,sK12)),s('type.realax.real',X2)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X3))))) ) | $spl0),
  inference(resolution,[],[f49488,f345])).
fof(f76189,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X0)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK19(X0,sK12)))))) ) | $spl362),
  inference(cnf_transformation,[],[f76189_D])).
fof(f76189_D,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X0)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK19(X0,sK12)))))) ) <=> ~$spl362),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl362])])).
fof(f76069,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X0)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK19(X0,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X1))))) ) | $spl0),
  inference(resolution,[],[f49488,f344])).
fof(f66037,plain,(
  ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK15)))) | $spl361),
  inference(cnf_transformation,[],[f66037_D])).
fof(f66037_D,plain,(
  ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK15)))) <=> ~$spl361),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl361])])).
fof(f65555,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK13))) | $spl356),
  inference(cnf_transformation,[],[f65555_D])).
fof(f65555_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK13))) <=> ~$spl356),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl356])])).
fof(f43565,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK13))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK15)))) | $spl72),
  inference(resolution,[],[f1180,f889])).
fof(f66032,plain,(
  $false | ($spl3 | $spl4 | $spl6 | $spl8)),
  inference(subsumption_resolution,[],[f66031,f563])).
fof(f66031,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f66030,f49456])).
fof(f49456,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK17))) | $spl3),
  inference(resolution,[],[f585,f547])).
fof(f66030,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK17)))) | ($spl3 | $spl4 | $spl6 | $spl8)),
  inference(subsumption_resolution,[],[f66007,f562])).
fof(f66007,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK17)))) | ($spl3 | $spl4 | $spl6 | $spl8)),
  inference(superposition,[],[f49597,f49371])).
fof(f49597,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(X0)),s('type.realax.real',sK17)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X0))))) ) | ($spl4 | $spl8)),
  inference(resolution,[],[f586,f588])).
fof(f66026,plain,(
  $false | ($spl3 | $spl4 | $spl6 | $spl8)),
  inference(subsumption_resolution,[],[f66025,f563])).
fof(f66025,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f66024,f49456])).
fof(f66024,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK17)))) | ($spl4 | $spl6 | $spl8)),
  inference(duplicate_literal_removal,[],[f65915])).
fof(f65915,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK17)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK17)))) | ($spl4 | $spl6 | $spl8)),
  inference(resolution,[],[f49597,f587])).
fof(f65557,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK15),s('type.realax.real',sK17)))) | $spl359),
  inference(cnf_transformation,[],[f65557_D])).
fof(f65557_D,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK15),s('type.realax.real',sK17)))) <=> ~$spl359),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl359])])).
fof(f65440,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK15),s('type.realax.real',sK17)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK13))) | $spl3),
  inference(resolution,[],[f49374,f1180])).
fof(f1512,plain,(
  ( ! [X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X1))))) ) | $spl116),
  inference(cnf_transformation,[],[f1512_D])).
fof(f1512_D,plain,(
  ( ! [X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X1))))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f49351,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK17)))) | $spl354),
  inference(cnf_transformation,[],[f49351_D])).
fof(f49351_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK17)))) <=> ~$spl354),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl354])])).
fof(f49965,plain,(
  ( ! [X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK17)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X3))))) ) | ($spl4 | $spl6)),
  inference(resolution,[],[f2179,f1029])).
fof(f1029,plain,(
  ( ! [X223,X224] : (~sP11(X224,'const.real.real_lte_2'(s('type.realax.real',sK16(X223)),s('type.realax.real',X223))) | p(s(bool,X224)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X223))))) ) | $spl6),
  inference(subsumption_resolution,[],[f990,f541])).
fof(f990,plain,(
  ( ! [X223,X224] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X223)))) | p(s(bool,'F_0')) | p(s(bool,X224)) | ~sP11(X224,'const.real.real_lte_2'(s('type.realax.real',sK16(X223)),s('type.realax.real',X223)))) ) | $spl6),
  inference(resolution,[],[f587,f485])).
fof(f2179,plain,(
  ( ! [X135] : (sP11('const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK17)),X135)) ) | $spl4),
  inference(resolution,[],[f1896,f487])).
fof(f1896,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK17)))) | $spl4),
  inference(resolution,[],[f586,f336])).
fof(f49964,plain,(
  ( ! [X2] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK17)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',sK15))))) ) | $spl4),
  inference(resolution,[],[f2179,f1230])).
fof(f49975,plain,(
  ( ! [X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK17)))) | p(s(bool,X1))) ) | $spl4),
  inference(subsumption_resolution,[],[f49963,f541])).
fof(f49963,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK17)))) | p(s(bool,X1))) ) | $spl4),
  inference(resolution,[],[f2179,f485])).
fof(f49974,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK17)))) | p(s(bool,X0))) ) | $spl4),
  inference(subsumption_resolution,[],[f49962,f541])).
fof(f49962,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK17)))) | p(s(bool,X0))) ) | $spl4),
  inference(resolution,[],[f2179,f489])).
fof(f49952,plain,(
  ( ! [X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK17)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X3))))) ) | ($spl4 | $spl6)),
  inference(resolution,[],[f2178,f1030])).
fof(f1030,plain,(
  ( ! [X225,X226] : (~sP11('const.real.real_lte_2'(s('type.realax.real',sK16(X225)),s('type.realax.real',X225)),X226) | p(s(bool,X226)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X225))))) ) | $spl6),
  inference(subsumption_resolution,[],[f991,f541])).
fof(f991,plain,(
  ( ! [X225,X226] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X225)))) | p(s(bool,'F_0')) | p(s(bool,X226)) | ~sP11('const.real.real_lte_2'(s('type.realax.real',sK16(X225)),s('type.realax.real',X225)),X226)) ) | $spl6),
  inference(resolution,[],[f587,f485])).
fof(f2178,plain,(
  ( ! [X134] : (sP11(X134,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK17)))) ) | $spl4),
  inference(resolution,[],[f1896,f486])).
fof(f49950,plain,(
  ( ! [X2] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK17)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',sK15))))) ) | $spl4),
  inference(resolution,[],[f2178,f1231])).
fof(f49961,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK17))))) ) | $spl4),
  inference(subsumption_resolution,[],[f49949,f541])).
fof(f49949,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK17))))) ) | $spl4),
  inference(resolution,[],[f2178,f485])).
fof(f49960,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK17))))) ) | $spl4),
  inference(subsumption_resolution,[],[f49948,f541])).
fof(f49948,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK17))))) ) | $spl4),
  inference(resolution,[],[f2178,f489])).
fof(f2044,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK30(sK12)),s('type.realax.real',sK17)))) | $spl166),
  inference(cnf_transformation,[],[f2044_D])).
fof(f2044_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK30(sK12)),s('type.realax.real',sK17)))) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f49627,plain,(
  ( ! [X66,X65] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK30(sK12)),s('type.realax.real',sK17)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK29(X65,sK12)),s('type.realax.real',X65)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X66))))) ) | $spl4),
  inference(resolution,[],[f586,f372])).
fof(f372,plain,(
  ( ! [X6,X4,X0] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',sK30(X0))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK29(X4,X0)),s('type.realax.real',X4)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X6))))) )),
  inference(cnf_transformation,[],[f236])).
fof(f2038,plain,(
  ( ! [X41] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK25(X41,sK12)),s('type.realax.real',sK17))))) ) | $spl158),
  inference(cnf_transformation,[],[f2038_D])).
fof(f2038_D,plain,(
  ( ! [X41] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK25(X41,sK12)),s('type.realax.real',sK17))))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f49619,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK25(X46,sK12)),s('type.realax.real',sK17)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X47),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X47)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X48))))) ) | $spl4),
  inference(resolution,[],[f586,f360])).
fof(f2037,plain,(
  ( ! [X38] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X38,sK12)),s('type.realax.real',sK17))))) ) | $spl156),
  inference(cnf_transformation,[],[f2037_D])).
fof(f2037_D,plain,(
  ( ! [X38] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X38,sK12)),s('type.realax.real',sK17))))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f49618,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X43,sK12)),s('type.realax.real',sK17)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X44),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X44)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X45))))) ) | $spl4),
  inference(resolution,[],[f586,f358])).
fof(f2036,plain,(
  ( ! [X32] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK23(X32,sK12)),s('type.realax.real',sK17)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X32),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) | $spl154),
  inference(cnf_transformation,[],[f2036_D])).
fof(f2036_D,plain,(
  ( ! [X32] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK23(X32,sK12)),s('type.realax.real',sK17)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X32),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f49617,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK23(X40,sK12)),s('type.realax.real',sK17)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X40),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X41,sK12)),s('type.realax.real',X41)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X42))))) ) | $spl4),
  inference(resolution,[],[f586,f355])).
fof(f49616,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK23(X37,sK12)),s('type.realax.real',sK17)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X37),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X38,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X39))))) ) | $spl4),
  inference(resolution,[],[f586,f354])).
fof(f2035,plain,(
  ( ! [X22] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X22,sK12)),s('type.realax.real',sK17))))) ) | $spl152),
  inference(cnf_transformation,[],[f2035_D])).
fof(f2035_D,plain,(
  ( ! [X22] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X22,sK12)),s('type.realax.real',sK17))))) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f49615,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X34,sK12)),s('type.realax.real',sK17)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X35),s('type.realax.real',sK23(X35,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X35),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X36))))) ) | $spl4),
  inference(resolution,[],[f586,f356])).
fof(f49614,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X31,sK12)),s('type.realax.real',sK17)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X32,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X32),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X33))))) ) | $spl4),
  inference(resolution,[],[f586,f354])).
fof(f49613,plain,(
  ( ! [X30,X28,X29,X27] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X27,sK12)),s('type.realax.real',sK17)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X28),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X28),s('type.realax.real',X29)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X29)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X30))))) ) | $spl4),
  inference(resolution,[],[f586,f352])).
fof(f2034,plain,(
  ( ! [X16] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK21(X16,sK12)),s('type.realax.real',sK17)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X16),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) | $spl150),
  inference(cnf_transformation,[],[f2034_D])).
fof(f2034_D,plain,(
  ( ! [X16] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK21(X16,sK12)),s('type.realax.real',sK17)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X16),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f49612,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK21(X24,sK12)),s('type.realax.real',sK17)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X24),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X25,sK12)),s('type.realax.real',X25)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X26))))) ) | $spl4),
  inference(resolution,[],[f586,f349])).
fof(f49611,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK21(X21,sK12)),s('type.realax.real',sK17)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X21),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X22,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X23))))) ) | $spl4),
  inference(resolution,[],[f586,f348])).
fof(f2033,plain,(
  ( ! [X6] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK20(X6,sK12)),s('type.realax.real',sK17))))) ) | $spl148),
  inference(cnf_transformation,[],[f2033_D])).
fof(f2033_D,plain,(
  ( ! [X6] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK20(X6,sK12)),s('type.realax.real',sK17))))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f49610,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK20(X18,sK12)),s('type.realax.real',sK17)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X19),s('type.realax.real',sK21(X19,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X19),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X20))))) ) | $spl4),
  inference(resolution,[],[f586,f350])).
fof(f49609,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK20(X15,sK12)),s('type.realax.real',sK17)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X16,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X16),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X17))))) ) | $spl4),
  inference(resolution,[],[f586,f348])).
fof(f49608,plain,(
  ( ! [X14,X12,X13,X11] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK20(X11,sK12)),s('type.realax.real',sK17)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X12),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X12),s('type.realax.real',X13)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X13)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X14))))) ) | $spl4),
  inference(resolution,[],[f586,f346])).
fof(f2032,plain,(
  ( ! [X4] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK19(X4,sK12)),s('type.realax.real',sK17)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))),s('type.realax.real',X4))))) ) | $spl146),
  inference(cnf_transformation,[],[f2032_D])).
fof(f2032_D,plain,(
  ( ! [X4] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK19(X4,sK12)),s('type.realax.real',sK17)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))),s('type.realax.real',X4))))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f49607,plain,(
  ( ! [X10,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK19(X9,sK12)),s('type.realax.real',sK17)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))),s('type.realax.real',X9)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X10))))) ) | $spl4),
  inference(resolution,[],[f586,f344])).
fof(f2031,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X0,X1,sK12)),s('type.realax.real',sK17)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) | $spl144),
  inference(cnf_transformation,[],[f2031_D])).
fof(f2031_D,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X0,X1,sK12)),s('type.realax.real',sK17)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f49606,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X5,X6,sK12)),s('type.realax.real',sK17)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X7)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X7)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X8))))) ) | $spl4),
  inference(resolution,[],[f586,f342])).
fof(f49150,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK17))))) ) | $spl4),
  inference(duplicate_literal_removal,[],[f2200])).
fof(f2200,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,X8)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK17))))) ) | $spl4),
  inference(superposition,[],[f1896,f498])).
fof(f49151,plain,(
  ( ! [X9] : (p(s(bool,X9)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK17))))) ) | $spl4),
  inference(duplicate_literal_removal,[],[f2201])).
fof(f2201,plain,(
  ( ! [X9] : (p(s(bool,X9)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK17)))) | p(s(bool,X9))) ) | $spl4),
  inference(superposition,[],[f1896,f498])).
fof(f49348,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK14)))) | $spl351),
  inference(cnf_transformation,[],[f49348_D])).
fof(f49348_D,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK14)))) <=> ~$spl351),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl351])])).
fof(f49152,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK14))))) ) | $spl4),
  inference(duplicate_literal_removal,[],[f2817])).
fof(f2817,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK14)))) | ~p(s(bool,X10))) ) | $spl4),
  inference(superposition,[],[f2105,f501])).
fof(f2105,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK14)))) | $spl4),
  inference(resolution,[],[f1896,f365])).
fof(f49153,plain,(
  ( ! [X11] : (~p(s(bool,X11)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK14))))) ) | $spl4),
  inference(duplicate_literal_removal,[],[f2818])).
fof(f2818,plain,(
  ( ! [X11] : (~p(s(bool,X11)) | ~p(s(bool,X11)) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK17),s('type.realax.real',sK14))))) ) | $spl4),
  inference(superposition,[],[f2105,f501])).
fof(f49186,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl349),
  inference(cnf_transformation,[],[f49186_D])).
fof(f49186_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl349),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl349])])).
fof(f49209,plain,(
  ( ! [X20] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X20))))) ) | $spl6),
  inference(subsumption_resolution,[],[f40230,f587])).
fof(f40230,plain,(
  ( ! [X20] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X20)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(X20)),s('type.realax.real',X20))))) ) | $spl6),
  inference(superposition,[],[f1006,f547])).
fof(f1006,plain,(
  ( ! [X251] : (s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(X251)),s('type.realax.real',X251))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X251))))) ) | $spl6),
  inference(resolution,[],[f587,f550])).
fof(f550,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f331])).
fof(f49208,plain,(
  ( ! [X21] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X21))))) ) | $spl6),
  inference(subsumption_resolution,[],[f40231,f587])).
fof(f40231,plain,(
  ( ! [X21] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X21)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(X21)),s('type.realax.real',X21))))) ) | $spl6),
  inference(superposition,[],[f1006,f549])).
fof(f549,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f330])).
fof(f330,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f133])).
fof(f133,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.real.REAL_LE_SUP','thm.bool.EQ_CLAUSES|split|3')).
fof(f49207,plain,(
  ( ! [X25] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X25))))) ) | $spl6),
  inference(subsumption_resolution,[],[f40235,f1006])).
fof(f40235,plain,(
  ( ! [X25] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X25)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(X25)),s('type.realax.real',X25)))) ) | $spl6),
  inference(superposition,[],[f1006,f572])).
fof(f572,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.real.REAL_LE_SUP',aHL_BOOL_CASES)).
fof(f49188,plain,(
  ( ! [X20] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X20),s('type.realax.real',sK15))))) )),
  inference(subsumption_resolution,[],[f43751,f337])).
fof(f43751,plain,(
  ( ! [X20] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X20),s('type.realax.real',sK15)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X20))))) )),
  inference(superposition,[],[f1182,f547])).
fof(f1182,plain,(
  ( ! [X291] : (s(bool,'T_0') != s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X291))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X291),s('type.realax.real',sK15))))) )),
  inference(resolution,[],[f337,f550])).
fof(f49187,plain,(
  ( ! [X21] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X21),s('type.realax.real',sK15))))) )),
  inference(subsumption_resolution,[],[f43752,f337])).
fof(f43752,plain,(
  ( ! [X21] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X21),s('type.realax.real',sK15)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X21))))) )),
  inference(superposition,[],[f1182,f549])).
fof(f49184,plain,(
  ( ! [X25] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X25),s('type.realax.real',sK15))))) )),
  inference(subsumption_resolution,[],[f43756,f1182])).
fof(f43756,plain,(
  ( ! [X25] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X25),s('type.realax.real',sK15)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X25)))) )),
  inference(superposition,[],[f1182,f572])).
fof(f49143,plain,(
  $false | ($spl1 | $spl6 | $spl130)),
  inference(subsumption_resolution,[],[f49142,f563])).
fof(f49142,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl6 | $spl130)),
  inference(forward_demodulation,[],[f49141,f5648])).
fof(f5648,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))) | $spl1),
  inference(resolution,[],[f583,f547])).
fof(f49141,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ($spl6 | $spl130)),
  inference(duplicate_literal_removal,[],[f49032])).
fof(f49032,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ($spl6 | $spl130)),
  inference(resolution,[],[f1519,f587])).
fof(f45543,plain,(
  $false | $spl347),
  inference(subsumption_resolution,[],[f45437,f575])).
fof(f45437,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK13),s('type.realax.real',sK13)))) | $spl347),
  inference(resolution,[],[f45314,f363])).
fof(f45542,plain,(
  $false | $spl347),
  inference(subsumption_resolution,[],[f45436,f575])).
fof(f45436,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK13),s('type.realax.real',sK13)))) | $spl347),
  inference(resolution,[],[f45314,f366])).
fof(f45312,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16('const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))) | $spl344),
  inference(cnf_transformation,[],[f45312_D])).
fof(f45312_D,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16('const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))) <=> ~$spl344),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl344])])).
fof(f45314,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK13)))) | $spl347),
  inference(cnf_transformation,[],[f45314_D])).
fof(f45314_D,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK13)))) <=> ~$spl347),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl347])])).
fof(f45200,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK13)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16('const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))) | ($spl1 | $spl6)),
  inference(resolution,[],[f5566,f1004])).
fof(f45310,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X16)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X18,X16,sK12)),s('type.realax.real',X18)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X17)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X16),s('type.realax.real',X17))))) ) | $spl342),
  inference(cnf_transformation,[],[f45310_D])).
fof(f45310_D,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X16)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X18,X16,sK12)),s('type.realax.real',X18)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X17)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X16),s('type.realax.real',X17))))) ) <=> ~$spl342),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl342])])).
fof(f45194,plain,(
  ( ! [X19,X17,X18,X16] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X16)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X16),s('type.realax.real',X17)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X17)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X18,X16,sK12)),s('type.realax.real',X18)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X19))))) ) | $spl1),
  inference(resolution,[],[f5566,f343])).
fof(f45309,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X12)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X14,X12,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X13)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X12),s('type.realax.real',X13))))) ) | $spl340),
  inference(cnf_transformation,[],[f45309_D])).
fof(f45309_D,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X12)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X14,X12,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X13)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X12),s('type.realax.real',X13))))) ) <=> ~$spl340),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl340])])).
fof(f45193,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X12)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X12),s('type.realax.real',X13)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X13)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X14,X12,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X15))))) ) | $spl1),
  inference(resolution,[],[f5566,f342])).
fof(f21719,plain,(
  ( ! [X121] : (p(s(bool,X121)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK37(sK12,X121,'type.realax.real')),s('type.realax.real',sK31(sK12)))))) ) | $spl320),
  inference(cnf_transformation,[],[f21719_D])).
fof(f21719_D,plain,(
  ( ! [X121] : (p(s(bool,X121)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK37(sK12,X121,'type.realax.real')),s('type.realax.real',sK31(sK12)))))) ) <=> ~$spl320),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl320])])).
fof(f32777,plain,(
  ( ! [X72,X71] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK37(sK12,X71,'type.realax.real')),s('type.realax.real',sK31(sK12))))) | p(s(bool,X71)) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X72))))) ) | $spl66),
  inference(resolution,[],[f769,f505])).
fof(f505,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK37(X2,X1,X0))))) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f298])).
fof(f298,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK36(X2,X1,X0))))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK37(X2,X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK36,sK37])],[f297])).
fof(f297,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | ? [X5] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(rectify,[],[f296])).
fof(f296,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(flattening,[],[f295])).
fof(f295,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f109])).
fof(f109,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X17] : ! [X10,X22] : ((p(s(bool,X10)) | ? [X7] : p(s(bool,happ(s(fun(X17,bool),X22),s(X17,X7))))) <=> ? [X7] : (p(s(bool,X10)) | p(s(bool,happ(s(fun(X17,bool),X22),s(X17,X7))))))),
  file('thm.real.REAL_LE_SUP','thm.bool.RIGHT_OR_EXISTS_THM')).
fof(f20758,plain,(
  ( ! [X121] : (p(s(bool,X121)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK36(sK12,X121,'type.realax.real')),s('type.realax.real',sK31(sK12)))))) ) | $spl308),
  inference(cnf_transformation,[],[f20758_D])).
fof(f20758_D,plain,(
  ( ! [X121] : (p(s(bool,X121)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK36(sK12,X121,'type.realax.real')),s('type.realax.real',sK31(sK12)))))) ) <=> ~$spl308),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl308])])).
fof(f32776,plain,(
  ( ! [X70,X69] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK36(sK12,X69,'type.realax.real')),s('type.realax.real',sK31(sK12))))) | p(s(bool,X69)) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X70))))) ) | $spl66),
  inference(resolution,[],[f769,f503])).
fof(f503,plain,(
  ( ! [X6,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK36(X2,X1,X0))))) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) )),
  inference(cnf_transformation,[],[f298])).
fof(f8394,plain,(
  ( ! [X95] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X95,sK12)),s('type.realax.real',sK31(sK12)))))) ) | $spl260),
  inference(cnf_transformation,[],[f8394_D])).
fof(f8394_D,plain,(
  ( ! [X95] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X95,sK12)),s('type.realax.real',sK31(sK12)))))) ) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f32764,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X43,sK12)),s('type.realax.real',sK31(sK12))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X44),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X44)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X45))))) ) | $spl66),
  inference(resolution,[],[f769,f358])).
fof(f31837,plain,(
  ( ! [X36] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK23(X36,sK12)),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X36),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) | $spl336),
  inference(cnf_transformation,[],[f31837_D])).
fof(f31837_D,plain,(
  ( ! [X36] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK23(X36,sK12)),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X36),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) <=> ~$spl336),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl336])])).
fof(f32763,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK23(X40,sK12)),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X40),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X41,sK12)),s('type.realax.real',X41)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X42))))) ) | $spl66),
  inference(resolution,[],[f769,f355])).
fof(f32762,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK23(X37,sK12)),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X37),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X38,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X39))))) ) | $spl66),
  inference(resolution,[],[f769,f354])).
fof(f31836,plain,(
  ( ! [X20] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK21(X20,sK12)),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X20),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) | $spl334),
  inference(cnf_transformation,[],[f31836_D])).
fof(f31836_D,plain,(
  ( ! [X20] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK21(X20,sK12)),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X20),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) <=> ~$spl334),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl334])])).
fof(f32758,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK21(X24,sK12)),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X24),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X25,sK12)),s('type.realax.real',X25)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X26))))) ) | $spl66),
  inference(resolution,[],[f769,f349])).
fof(f32757,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK21(X21,sK12)),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X21),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X22,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X23))))) ) | $spl66),
  inference(resolution,[],[f769,f348])).
fof(f31835,plain,(
  ( ! [X8] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK19(X8,sK12)),s('type.realax.real',sK31(sK12))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))),s('type.realax.real',X8))))) ) | $spl332),
  inference(cnf_transformation,[],[f31835_D])).
fof(f31835_D,plain,(
  ( ! [X8] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK19(X8,sK12)),s('type.realax.real',sK31(sK12))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))),s('type.realax.real',X8))))) ) <=> ~$spl332),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl332])])).
fof(f32753,plain,(
  ( ! [X10,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK19(X9,sK12)),s('type.realax.real',sK31(sK12))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))),s('type.realax.real',X9)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X10))))) ) | $spl66),
  inference(resolution,[],[f769,f344])).
fof(f31834,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X4,X5,sK12)),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X6)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X6)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) | $spl330),
  inference(cnf_transformation,[],[f31834_D])).
fof(f31834_D,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X4,X5,sK12)),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X6)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X6)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) <=> ~$spl330),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl330])])).
fof(f32752,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X5,X6,sK12)),s('type.realax.real',sK31(sK12))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X7)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X7)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X8))))) ) | $spl66),
  inference(resolution,[],[f769,f342])).
fof(f32296,plain,(
  ( ! [X62,X63] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X62),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X62)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X63,sK12)))))) ) | $spl58),
  inference(resolution,[],[f763,f358])).
fof(f9179,plain,(
  ( ! [X63] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK29(X63,sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) | $spl292),
  inference(cnf_transformation,[],[f9179_D])).
fof(f9179_D,plain,(
  ( ! [X63] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK29(X63,sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f32297,plain,(
  ( ! [X66,X64,X65] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK29(X64,sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X65,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X66))))) ) | $spl58),
  inference(resolution,[],[f763,f358])).
fof(f769,plain,(
  ( ! [X84] : (~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X84)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X84),s('type.realax.real',sK31(sK12)))))) ) | $spl66),
  inference(cnf_transformation,[],[f769_D])).
fof(f769_D,plain,(
  ( ! [X84] : (~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X84)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X84),s('type.realax.real',sK31(sK12)))))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f32613,plain,(
  ( ! [X2,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X1)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X2))))) ) | $spl294),
  inference(resolution,[],[f9180,f374])).
fof(f5010,plain,(
  ( ! [X88,X89] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X88),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X88)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK29(X89,sK12)),s('type.realax.real',X89))))) ) | $spl50),
  inference(resolution,[],[f758,f374])).
fof(f5849,plain,(
  ( ! [X132,X130,X131] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X130)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X131),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X131)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK29(X132,sK12)),s('type.realax.real',X132))))) ) | $spl8),
  inference(resolution,[],[f588,f374])).
fof(f7380,plain,(
  ( ! [X88,X89] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X88),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X88)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK29(X89,sK12)),s('type.realax.real',X89))))) )),
  inference(resolution,[],[f631,f374])).
fof(f631,plain,(
  p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK38(sK12,'type.realax.real')))))),
  inference(resolution,[],[f336,f508])).
fof(f508,plain,(
  ( ! [X4,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK38(X1,X0)))))) )),
  inference(cnf_transformation,[],[f303])).
fof(f303,plain,(
  ! [X0,X1] : ((p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK38(X1,X0))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK39(X1,X0))))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK38,sK39])],[f302])).
fof(f302,plain,(
  ! [X0,X1] : ((? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f301])).
fof(f301,plain,(
  ! [X0,X1] : ((? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0,X1] : (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(ennf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1] : (~? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(flattening,[],[f112])).
fof(f112,plain,(
  ! [X0] : ! [X1] : (~? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X17] : ! [X10] : (~? [X7] : p(s(bool,happ(s(fun(X17,bool),X10),s(X17,X7)))) <=> ! [X7] : ~p(s(bool,happ(s(fun(X17,bool),X10),s(X17,X7)))))),
  file('thm.real.REAL_LE_SUP','thm.bool.NOT_EXISTS_THM')).
fof(f7588,plain,(
  ( ! [X88,X89] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X88),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X88)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK29(X89,sK12)),s('type.realax.real',X89))))) )),
  inference(resolution,[],[f632,f374])).
fof(f632,plain,(
  p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK39(sK12,'type.realax.real')))))),
  inference(resolution,[],[f336,f509])).
fof(f509,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK39(X1,X0)))))) )),
  inference(cnf_transformation,[],[f303])).
fof(f9465,plain,(
  ( ! [X88,X89] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X88),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X88)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK29(X89,sK12)),s('type.realax.real',X89))))) ) | $spl60),
  inference(resolution,[],[f765,f374])).
fof(f32726,plain,(
  ( ! [X88,X89] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X88),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X88)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK29(X89,sK12)),s('type.realax.real',X89))))) )),
  inference(subsumption_resolution,[],[f12281,f579])).
fof(f12281,plain,(
  ( ! [X88,X89] : (~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X88),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X88)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK29(X89,sK12)),s('type.realax.real',X89))))) )),
  inference(superposition,[],[f374,f715])).
fof(f20645,plain,(
  ( ! [X127,X128,X129] : (p(s(bool,X127)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK36(sK12,X127,'type.realax.real')),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK29(X128,sK12)),s('type.realax.real',X128)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X129))))) )),
  inference(resolution,[],[f633,f374])).
fof(f633,plain,(
  ( ! [X90] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK36(sK12,X90,'type.realax.real'))))) | p(s(bool,X90))) )),
  inference(resolution,[],[f336,f503])).
fof(f20646,plain,(
  ( ! [X132,X130,X131] : (p(s(bool,X130)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X131),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X131)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK29(X132,sK12)),s('type.realax.real',X132))))) )),
  inference(resolution,[],[f633,f374])).
fof(f21328,plain,(
  ( ! [X103,X104] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X103),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X103)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK29(X104,sK12)),s('type.realax.real',X104))))) ) | $spl44),
  inference(resolution,[],[f754,f374])).
fof(f21607,plain,(
  ( ! [X127,X128,X129] : (p(s(bool,X127)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK37(sK12,X127,'type.realax.real')),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK29(X128,sK12)),s('type.realax.real',X128)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X129))))) )),
  inference(resolution,[],[f634,f374])).
fof(f634,plain,(
  ( ! [X91] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK37(sK12,X91,'type.realax.real'))))) | p(s(bool,X91))) )),
  inference(resolution,[],[f336,f505])).
fof(f21608,plain,(
  ( ! [X132,X130,X131] : (p(s(bool,X130)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X131),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X131)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK29(X132,sK12)),s('type.realax.real',X132))))) )),
  inference(resolution,[],[f634,f374])).
fof(f32023,plain,(
  ( ! [X105,X104] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X104),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X104)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK29(X105,sK12)),s('type.realax.real',X105))))) ) | $spl36),
  inference(resolution,[],[f749,f374])).
fof(f32725,plain,(
  ( ! [X101] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK29(X101,sK12)),s('type.realax.real',sK31(sK12)))))) ) | $spl338),
  inference(cnf_transformation,[],[f32725_D])).
fof(f32725_D,plain,(
  ( ! [X101] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK29(X101,sK12)),s('type.realax.real',sK31(sK12)))))) ) <=> ~$spl338),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl338])])).
fof(f32314,plain,(
  ( ! [X103,X101,X102] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK29(X101,sK12)),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK29(X102,sK12)),s('type.realax.real',X102)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X103))))) ) | $spl58),
  inference(resolution,[],[f763,f374])).
fof(f32315,plain,(
  ( ! [X105,X104] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X104),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X104)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK29(X105,sK12)),s('type.realax.real',X105))))) ) | $spl58),
  inference(resolution,[],[f763,f374])).
fof(f32022,plain,(
  ( ! [X103,X101,X102] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X101,sK12)),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK29(X102,sK12)),s('type.realax.real',X102)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X103))))) ) | $spl36),
  inference(resolution,[],[f749,f374])).
fof(f32611,plain,(
  $false | ($spl62 | $spl294)),
  inference(resolution,[],[f9180,f766])).
fof(f9180,plain,(
  ( ! [X83] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK29(X83,sK12)),s('type.realax.real',sK28(sK12)))))) ) | $spl294),
  inference(cnf_transformation,[],[f9180_D])).
fof(f9180_D,plain,(
  ( ! [X83] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK29(X83,sK12)),s('type.realax.real',sK28(sK12)))))) ) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f32308,plain,(
  ( ! [X90,X91,X89] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK29(X89,sK12)),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X90,sK12)),s('type.realax.real',X90)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X91))))) ) | $spl58),
  inference(resolution,[],[f763,f370])).
fof(f32306,plain,(
  ( ! [X85,X86,X84] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK29(X84,sK12)),s('type.realax.real',sK28(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X85,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X86))))) ) | $spl58),
  inference(resolution,[],[f763,f369])).
fof(f32303,plain,(
  ( ! [X80,X81,X79] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK29(X79,sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X80,sK12)),s('type.realax.real',X80)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X81))))) ) | $spl58),
  inference(resolution,[],[f763,f361])).
fof(f32302,plain,(
  ( ! [X78,X77] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X77),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X77)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X78,sK12)),s('type.realax.real',X78))))) ) | $spl58),
  inference(resolution,[],[f763,f361])).
fof(f32299,plain,(
  ( ! [X70,X71,X69] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK29(X69,sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X70,sK12)),s('type.realax.real',X70)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X71))))) ) | $spl58),
  inference(resolution,[],[f763,f359])).
fof(f32298,plain,(
  ( ! [X68,X67] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X67),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X67)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X68,sK12)),s('type.realax.real',X68))))) ) | $spl58),
  inference(resolution,[],[f763,f359])).
fof(f32295,plain,(
  ( ! [X61,X60] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X60),s('type.realax.real',sK23(X60,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X60),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X61,sK12)),s('type.realax.real',X61))))) ) | $spl58),
  inference(resolution,[],[f763,f357])).
fof(f32294,plain,(
  ( ! [X59,X58] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X58),s('type.realax.real',sK23(X58,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X58),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X59,sK12)))))) ) | $spl58),
  inference(resolution,[],[f763,f356])).
fof(f32293,plain,(
  ( ! [X57,X56] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X56,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X56),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X57,sK12)),s('type.realax.real',X57))))) ) | $spl58),
  inference(resolution,[],[f763,f355])).
fof(f32292,plain,(
  ( ! [X54,X55] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X54,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X54),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X55,sK12)))))) ) | $spl58),
  inference(resolution,[],[f763,f354])).
fof(f32291,plain,(
  ( ! [X52,X53,X51] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X51),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X51),s('type.realax.real',X52)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X52)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X53,sK12)),s('type.realax.real',X53))))) ) | $spl58),
  inference(resolution,[],[f763,f353])).
fof(f9178,plain,(
  ( ! [X17,X18] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real',sK29(X18,sK12)))))) ) | $spl290),
  inference(cnf_transformation,[],[f9178_D])).
fof(f9178_D,plain,(
  ( ! [X17,X18] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real',sK29(X18,sK12)))))) ) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f32290,plain,(
  ( ! [X47,X50,X48,X49] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X47),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X47),s('type.realax.real',sK29(X48,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X49,sK12)),s('type.realax.real',X49)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X50))))) ) | $spl58),
  inference(resolution,[],[f763,f353])).
fof(f32289,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X44),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X44),s('type.realax.real',X45)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X45)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X46,sK12)))))) ) | $spl58),
  inference(resolution,[],[f763,f352])).
fof(f32288,plain,(
  ( ! [X43,X41,X42,X40] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X40),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X40),s('type.realax.real',sK29(X41,sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X42,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X43))))) ) | $spl58),
  inference(resolution,[],[f763,f352])).
fof(f32287,plain,(
  ( ! [X39,X38] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X38),s('type.realax.real',sK21(X38,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X38),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X39,sK12)),s('type.realax.real',X39))))) ) | $spl58),
  inference(resolution,[],[f763,f351])).
fof(f32286,plain,(
  ( ! [X37,X36] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X36),s('type.realax.real',sK21(X36,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X36),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X37,sK12)))))) ) | $spl58),
  inference(resolution,[],[f763,f350])).
fof(f32285,plain,(
  ( ! [X35,X34] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X34,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X34),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X35,sK12)),s('type.realax.real',X35))))) ) | $spl58),
  inference(resolution,[],[f763,f349])).
fof(f32284,plain,(
  ( ! [X33,X32] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X32,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X32),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X33,sK12)))))) ) | $spl58),
  inference(resolution,[],[f763,f348])).
fof(f32283,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X29),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X29),s('type.realax.real',X30)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X30)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X31,sK12)),s('type.realax.real',X31))))) ) | $spl58),
  inference(resolution,[],[f763,f347])).
fof(f32282,plain,(
  ( ! [X28,X26,X27,X25] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X25),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X25),s('type.realax.real',sK29(X26,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X27,sK12)),s('type.realax.real',X27)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X28))))) ) | $spl58),
  inference(resolution,[],[f763,f347])).
fof(f32281,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X22),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X22),s('type.realax.real',X23)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X23)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X24,sK12)))))) ) | $spl58),
  inference(resolution,[],[f763,f346])).
fof(f32280,plain,(
  ( ! [X21,X19,X20,X18] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X18),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X18),s('type.realax.real',sK29(X19,sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X20,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X21))))) ) | $spl58),
  inference(resolution,[],[f763,f346])).
fof(f9177,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X13,X11,sK12)),s('type.realax.real',X13)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK29(X12,sK12)))))) ) | $spl288),
  inference(cnf_transformation,[],[f9177_D])).
fof(f9177_D,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X13,X11,sK12)),s('type.realax.real',X13)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK29(X12,sK12)))))) ) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f32277,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X12),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X12),s('type.realax.real',sK29(X13,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X14,X12,sK12)),s('type.realax.real',X14)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X15))))) ) | $spl58),
  inference(resolution,[],[f763,f343])).
fof(f9176,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X6,X4,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',sK29(X5,sK12)))))) ) | $spl286),
  inference(cnf_transformation,[],[f9176_D])).
fof(f9176_D,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X6,X4,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',sK29(X5,sK12)))))) ) <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f32275,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',sK29(X6,sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X7,X5,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X8))))) ) | $spl58),
  inference(resolution,[],[f763,f342])).
fof(f763,plain,(
  ( ! [X80] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X80,sK12)))))) ) | $spl58),
  inference(cnf_transformation,[],[f763_D])).
fof(f763_D,plain,(
  ( ! [X80] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X80,sK12)))))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f32021,plain,(
  ( ! [X99,X100] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X99),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X99)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X100,sK12)))))) ) | $spl36),
  inference(resolution,[],[f749,f373])).
fof(f373,plain,(
  ( ! [X6,X4,X2,X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',sK31(X0))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X2)))) | p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',sK29(X4,X0))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X6))))) )),
  inference(cnf_transformation,[],[f236])).
fof(f8391,plain,(
  ( ! [X73] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X73,sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) | $spl256),
  inference(cnf_transformation,[],[f8391_D])).
fof(f8391_D,plain,(
  ( ! [X73] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X73,sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f32011,plain,(
  ( ! [X80,X81,X79] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X79,sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X80,sK12)),s('type.realax.real',X80)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X81))))) ) | $spl36),
  inference(resolution,[],[f749,f361])).
fof(f32010,plain,(
  ( ! [X78,X77] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X77),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X77)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X78,sK12)),s('type.realax.real',X78))))) ) | $spl36),
  inference(resolution,[],[f749,f361])).
fof(f32007,plain,(
  ( ! [X70,X71,X69] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X69,sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X70,sK12)),s('type.realax.real',X70)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X71))))) ) | $spl36),
  inference(resolution,[],[f749,f359])).
fof(f32006,plain,(
  ( ! [X68,X67] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X67),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X67)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X68,sK12)),s('type.realax.real',X68))))) ) | $spl36),
  inference(resolution,[],[f749,f359])).
fof(f32003,plain,(
  ( ! [X61,X60] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X60),s('type.realax.real',sK23(X60,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X60),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X61,sK12)),s('type.realax.real',X61))))) ) | $spl36),
  inference(resolution,[],[f749,f357])).
fof(f32002,plain,(
  ( ! [X59,X58] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X58),s('type.realax.real',sK23(X58,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X58),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X59,sK12)))))) ) | $spl36),
  inference(resolution,[],[f749,f356])).
fof(f32001,plain,(
  ( ! [X57,X56] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X56,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X56),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X57,sK12)),s('type.realax.real',X57))))) ) | $spl36),
  inference(resolution,[],[f749,f355])).
fof(f32000,plain,(
  ( ! [X54,X55] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X54,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X54),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X55,sK12)))))) ) | $spl36),
  inference(resolution,[],[f749,f354])).
fof(f31999,plain,(
  ( ! [X52,X53,X51] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X51),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X51),s('type.realax.real',X52)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X52)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X53,sK12)),s('type.realax.real',X53))))) ) | $spl36),
  inference(resolution,[],[f749,f353])).
fof(f31998,plain,(
  ( ! [X47,X50,X48,X49] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X47),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X47),s('type.realax.real',sK24(X48,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X49,sK12)),s('type.realax.real',X49)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X50))))) ) | $spl36),
  inference(resolution,[],[f749,f353])).
fof(f31997,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X44),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X44),s('type.realax.real',X45)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X45)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X46,sK12)))))) ) | $spl36),
  inference(resolution,[],[f749,f352])).
fof(f31996,plain,(
  ( ! [X43,X41,X42,X40] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X40),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X40),s('type.realax.real',sK24(X41,sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X42,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X43))))) ) | $spl36),
  inference(resolution,[],[f749,f352])).
fof(f31995,plain,(
  ( ! [X39,X38] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X38),s('type.realax.real',sK21(X38,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X38),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X39,sK12)),s('type.realax.real',X39))))) ) | $spl36),
  inference(resolution,[],[f749,f351])).
fof(f31994,plain,(
  ( ! [X37,X36] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X36),s('type.realax.real',sK21(X36,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X36),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X37,sK12)))))) ) | $spl36),
  inference(resolution,[],[f749,f350])).
fof(f31993,plain,(
  ( ! [X35,X34] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X34,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X34),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X35,sK12)),s('type.realax.real',X35))))) ) | $spl36),
  inference(resolution,[],[f749,f349])).
fof(f31992,plain,(
  ( ! [X33,X32] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X32,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X32),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X33,sK12)))))) ) | $spl36),
  inference(resolution,[],[f749,f348])).
fof(f31991,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X29),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X29),s('type.realax.real',X30)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X30)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X31,sK12)),s('type.realax.real',X31))))) ) | $spl36),
  inference(resolution,[],[f749,f347])).
fof(f31990,plain,(
  ( ! [X28,X26,X27,X25] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X25),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X25),s('type.realax.real',sK24(X26,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X27,sK12)),s('type.realax.real',X27)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X28))))) ) | $spl36),
  inference(resolution,[],[f749,f347])).
fof(f31989,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X22),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X22),s('type.realax.real',X23)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X23)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X24,sK12)))))) ) | $spl36),
  inference(resolution,[],[f749,f346])).
fof(f31988,plain,(
  ( ! [X21,X19,X20,X18] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X18),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X18),s('type.realax.real',sK24(X19,sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X20,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X21))))) ) | $spl36),
  inference(resolution,[],[f749,f346])).
fof(f31985,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X12),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X12),s('type.realax.real',sK24(X13,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X14,X12,sK12)),s('type.realax.real',X14)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X15))))) ) | $spl36),
  inference(resolution,[],[f749,f343])).
fof(f31983,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',sK24(X6,sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X7,X5,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X8))))) ) | $spl36),
  inference(resolution,[],[f749,f342])).
fof(f31852,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X0)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X1))))) ) | $spl36),
  inference(resolution,[],[f8240,f359])).
fof(f31724,plain,(
  ( ! [X70,X71] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK37(sK12,X70,'type.realax.real')),s('type.realax.real',sK31(sK12))))) | p(s(bool,X70)) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X71))))) ) | $spl66),
  inference(resolution,[],[f769,f505])).
fof(f31723,plain,(
  ( ! [X68,X69] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK36(sK12,X68,'type.realax.real')),s('type.realax.real',sK31(sK12))))) | p(s(bool,X68)) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X69))))) ) | $spl66),
  inference(resolution,[],[f769,f503])).
fof(f31710,plain,(
  ( ! [X39,X41,X40] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK23(X39,sK12)),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X39),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X40,sK12)),s('type.realax.real',X40)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X41))))) ) | $spl66),
  inference(resolution,[],[f769,f355])).
fof(f31709,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK23(X36,sK12)),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X36),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X37,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X38))))) ) | $spl66),
  inference(resolution,[],[f769,f354])).
fof(f31705,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK21(X23,sK12)),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X23),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X24,sK12)),s('type.realax.real',X24)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X25))))) ) | $spl66),
  inference(resolution,[],[f769,f349])).
fof(f31704,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK21(X20,sK12)),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X20),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X21,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X22))))) ) | $spl66),
  inference(resolution,[],[f769,f348])).
fof(f31700,plain,(
  ( ! [X8,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK19(X8,sK12)),s('type.realax.real',sK31(sK12))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))),s('type.realax.real',X8)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X9))))) ) | $spl66),
  inference(resolution,[],[f769,f344])).
fof(f31699,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X4,X5,sK12)),s('type.realax.real',sK31(sK12))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X6)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X6)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X7))))) ) | $spl66),
  inference(resolution,[],[f769,f342])).
fof(f21718,plain,(
  ( ! [X105] : (p(s(bool,X105)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK37(sK12,X105,'type.realax.real')),s('type.realax.real',sK28(sK12)))))) ) | $spl318),
  inference(cnf_transformation,[],[f21718_D])).
fof(f21718_D,plain,(
  ( ! [X105] : (p(s(bool,X105)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK37(sK12,X105,'type.realax.real')),s('type.realax.real',sK28(sK12)))))) ) <=> ~$spl318),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl318])])).
fof(f31520,plain,(
  ( ! [X70,X71] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK37(sK12,X70,'type.realax.real')),s('type.realax.real',sK28(sK12))))) | p(s(bool,X70)) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X71))))) ) | $spl56),
  inference(resolution,[],[f762,f505])).
fof(f20757,plain,(
  ( ! [X105] : (p(s(bool,X105)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK36(sK12,X105,'type.realax.real')),s('type.realax.real',sK28(sK12)))))) ) | $spl306),
  inference(cnf_transformation,[],[f20757_D])).
fof(f20757_D,plain,(
  ( ! [X105] : (p(s(bool,X105)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK36(sK12,X105,'type.realax.real')),s('type.realax.real',sK28(sK12)))))) ) <=> ~$spl306),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl306])])).
fof(f31519,plain,(
  ( ! [X68,X69] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK36(sK12,X68,'type.realax.real')),s('type.realax.real',sK28(sK12))))) | p(s(bool,X68)) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X69))))) ) | $spl56),
  inference(resolution,[],[f762,f503])).
fof(f31514,plain,(
  ( ! [X61,X59,X60] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK29(X59,sK12)),s('type.realax.real',sK28(sK12))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X60),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X60)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X61))))) ) | $spl56),
  inference(resolution,[],[f762,f373])).
fof(f31633,plain,(
  ( ! [X36] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK23(X36,sK12)),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X36),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) | $spl328),
  inference(cnf_transformation,[],[f31633_D])).
fof(f31633_D,plain,(
  ( ! [X36] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK23(X36,sK12)),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X36),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) <=> ~$spl328),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl328])])).
fof(f31506,plain,(
  ( ! [X39,X41,X40] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK23(X39,sK12)),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X39),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X40,sK12)),s('type.realax.real',X40)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X41))))) ) | $spl56),
  inference(resolution,[],[f762,f355])).
fof(f31505,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK23(X36,sK12)),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X36),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X37,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X38))))) ) | $spl56),
  inference(resolution,[],[f762,f354])).
fof(f31632,plain,(
  ( ! [X20] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK21(X20,sK12)),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X20),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) | $spl326),
  inference(cnf_transformation,[],[f31632_D])).
fof(f31632_D,plain,(
  ( ! [X20] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK21(X20,sK12)),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X20),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) <=> ~$spl326),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl326])])).
fof(f31501,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK21(X23,sK12)),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X23),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X24,sK12)),s('type.realax.real',X24)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X25))))) ) | $spl56),
  inference(resolution,[],[f762,f349])).
fof(f31500,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK21(X20,sK12)),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X20),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X21,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X22))))) ) | $spl56),
  inference(resolution,[],[f762,f348])).
fof(f31631,plain,(
  ( ! [X8] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK19(X8,sK12)),s('type.realax.real',sK28(sK12))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))),s('type.realax.real',X8))))) ) | $spl324),
  inference(cnf_transformation,[],[f31631_D])).
fof(f31631_D,plain,(
  ( ! [X8] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK19(X8,sK12)),s('type.realax.real',sK28(sK12))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))),s('type.realax.real',X8))))) ) <=> ~$spl324),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl324])])).
fof(f31496,plain,(
  ( ! [X8,X9] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK19(X8,sK12)),s('type.realax.real',sK28(sK12))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))),s('type.realax.real',X8)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X9))))) ) | $spl56),
  inference(resolution,[],[f762,f344])).
fof(f31630,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK18(X4,X5,sK12)),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X6)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X6)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) | $spl322),
  inference(cnf_transformation,[],[f31630_D])).
fof(f31630_D,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK18(X4,X5,sK12)),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X6)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X6)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) <=> ~$spl322),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl322])])).
fof(f31495,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK18(X4,X5,sK12)),s('type.realax.real',sK28(sK12))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X6)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X6)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X7))))) ) | $spl56),
  inference(resolution,[],[f762,f342])).
fof(f22740,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X0)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X1))))) ) | $spl270),
  inference(resolution,[],[f8791,f361])).
fof(f4997,plain,(
  ( ! [X66,X67] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X66),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X66)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X67,sK12)),s('type.realax.real',X67))))) ) | $spl50),
  inference(resolution,[],[f758,f361])).
fof(f3982,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK27(sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | $spl184),
  inference(cnf_transformation,[],[f3982_D])).
fof(f3982_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK27(sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f4998,plain,(
  ( ! [X68,X69] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK27(sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X68,sK12)),s('type.realax.real',X68)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X69))))) ) | $spl50),
  inference(resolution,[],[f758,f361])).
fof(f5836,plain,(
  ( ! [X97,X95,X96] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X95)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X96),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X96)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X97,sK12)),s('type.realax.real',X97))))) ) | $spl8),
  inference(resolution,[],[f588,f361])).
fof(f1519,plain,(
  ( ! [X80] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(X80)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X80))))) ) | $spl130),
  inference(cnf_transformation,[],[f1519_D])).
fof(f1519_D,plain,(
  ( ! [X80] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(X80)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X80))))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f5837,plain,(
  ( ! [X99,X100,X98] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X98)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(X98)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X99,sK12)),s('type.realax.real',X99)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X100))))) ) | $spl8),
  inference(resolution,[],[f588,f361])).
fof(f7367,plain,(
  ( ! [X66,X67] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X66),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X66)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X67,sK12)),s('type.realax.real',X67))))) )),
  inference(resolution,[],[f631,f361])).
fof(f7490,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK38(sK12,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | $spl208),
  inference(cnf_transformation,[],[f7490_D])).
fof(f7490_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK38(sK12,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f7368,plain,(
  ( ! [X68,X69] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK38(sK12,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X68,sK12)),s('type.realax.real',X68)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X69))))) )),
  inference(resolution,[],[f631,f361])).
fof(f7575,plain,(
  ( ! [X66,X67] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X66),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X66)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X67,sK12)),s('type.realax.real',X67))))) )),
  inference(resolution,[],[f632,f361])).
fof(f7698,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK39(sK12,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | $spl220),
  inference(cnf_transformation,[],[f7698_D])).
fof(f7698_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK39(sK12,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f7576,plain,(
  ( ! [X68,X69] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK39(sK12,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X68,sK12)),s('type.realax.real',X68)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X69))))) )),
  inference(resolution,[],[f632,f361])).
fof(f9452,plain,(
  ( ! [X66,X67] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X66),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X66)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X67,sK12)),s('type.realax.real',X67))))) ) | $spl60),
  inference(resolution,[],[f765,f361])).
fof(f5300,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK30(sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | $spl196),
  inference(cnf_transformation,[],[f5300_D])).
fof(f5300_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK30(sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f9453,plain,(
  ( ! [X68,X69] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK30(sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X68,sK12)),s('type.realax.real',X68)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X69))))) ) | $spl60),
  inference(resolution,[],[f765,f361])).
fof(f22852,plain,(
  ( ! [X66,X67] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X66),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X66)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X67,sK12)),s('type.realax.real',X67))))) )),
  inference(subsumption_resolution,[],[f12268,f579])).
fof(f12268,plain,(
  ( ! [X66,X67] : (~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X66),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X66)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X67,sK12)),s('type.realax.real',X67))))) )),
  inference(superposition,[],[f361,f715])).
fof(f20633,plain,(
  ( ! [X97,X95,X96] : (p(s(bool,X95)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X96),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X96)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X97,sK12)),s('type.realax.real',X97))))) )),
  inference(resolution,[],[f633,f361])).
fof(f20756,plain,(
  ( ! [X80] : (p(s(bool,X80)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK36(sK12,X80,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) | $spl304),
  inference(cnf_transformation,[],[f20756_D])).
fof(f20756_D,plain,(
  ( ! [X80] : (p(s(bool,X80)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK36(sK12,X80,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) <=> ~$spl304),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl304])])).
fof(f20634,plain,(
  ( ! [X99,X100,X98] : (p(s(bool,X98)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK36(sK12,X98,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X99,sK12)),s('type.realax.real',X99)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X100))))) )),
  inference(resolution,[],[f633,f361])).
fof(f21315,plain,(
  ( ! [X76,X77] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X76),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X76)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X77,sK12)),s('type.realax.real',X77))))) ) | $spl44),
  inference(resolution,[],[f754,f361])).
fof(f8790,plain,(
  ( ! [X63] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK25(X63,sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) | $spl268),
  inference(cnf_transformation,[],[f8790_D])).
fof(f8790_D,plain,(
  ( ! [X63] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK25(X63,sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f21316,plain,(
  ( ! [X80,X78,X79] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK25(X78,sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X79,sK12)),s('type.realax.real',X79)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X80))))) ) | $spl44),
  inference(resolution,[],[f754,f361])).
fof(f21595,plain,(
  ( ! [X97,X95,X96] : (p(s(bool,X95)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X96),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X96)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X97,sK12)),s('type.realax.real',X97))))) )),
  inference(resolution,[],[f634,f361])).
fof(f21717,plain,(
  ( ! [X80] : (p(s(bool,X80)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK37(sK12,X80,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) | $spl316),
  inference(cnf_transformation,[],[f21717_D])).
fof(f21717_D,plain,(
  ( ! [X80] : (p(s(bool,X80)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK37(sK12,X80,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) <=> ~$spl316),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl316])])).
fof(f21596,plain,(
  ( ! [X99,X100,X98] : (p(s(bool,X98)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK37(sK12,X98,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X99,sK12)),s('type.realax.real',X99)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X100))))) )),
  inference(resolution,[],[f634,f361])).
fof(f22739,plain,(
  $false | ($spl46 | $spl270)),
  inference(resolution,[],[f8791,f755])).
fof(f22571,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X0)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X1))))) ) | $spl260),
  inference(resolution,[],[f8394,f359])).
fof(f4993,plain,(
  ( ! [X59,X58] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X58),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X58)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X59,sK12)),s('type.realax.real',X59))))) ) | $spl50),
  inference(resolution,[],[f758,f359])).
fof(f4994,plain,(
  ( ! [X61,X60] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK27(sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X60,sK12)),s('type.realax.real',X60)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X61))))) ) | $spl50),
  inference(resolution,[],[f758,f359])).
fof(f5832,plain,(
  ( ! [X85,X83,X84] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X83)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X84),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X84)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X85,sK12)),s('type.realax.real',X85))))) ) | $spl8),
  inference(resolution,[],[f588,f359])).
fof(f5833,plain,(
  ( ! [X88,X87,X86] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X86)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(X86)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X87,sK12)),s('type.realax.real',X87)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X88))))) ) | $spl8),
  inference(resolution,[],[f588,f359])).
fof(f7363,plain,(
  ( ! [X59,X58] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X58),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X58)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X59,sK12)),s('type.realax.real',X59))))) )),
  inference(resolution,[],[f631,f359])).
fof(f7364,plain,(
  ( ! [X61,X60] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK38(sK12,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X60,sK12)),s('type.realax.real',X60)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X61))))) )),
  inference(resolution,[],[f631,f359])).
fof(f7571,plain,(
  ( ! [X59,X58] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X58),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X58)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X59,sK12)),s('type.realax.real',X59))))) )),
  inference(resolution,[],[f632,f359])).
fof(f7572,plain,(
  ( ! [X61,X60] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK39(sK12,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X60,sK12)),s('type.realax.real',X60)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X61))))) )),
  inference(resolution,[],[f632,f359])).
fof(f9448,plain,(
  ( ! [X59,X58] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X58),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X58)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X59,sK12)),s('type.realax.real',X59))))) ) | $spl60),
  inference(resolution,[],[f765,f359])).
fof(f9449,plain,(
  ( ! [X61,X60] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK30(sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X60,sK12)),s('type.realax.real',X60)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X61))))) ) | $spl60),
  inference(resolution,[],[f765,f359])).
fof(f22683,plain,(
  ( ! [X59,X58] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X58),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X58)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X59,sK12)),s('type.realax.real',X59))))) )),
  inference(subsumption_resolution,[],[f12264,f579])).
fof(f12264,plain,(
  ( ! [X59,X58] : (~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X58),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X58)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X59,sK12)),s('type.realax.real',X59))))) )),
  inference(superposition,[],[f359,f715])).
fof(f20629,plain,(
  ( ! [X85,X83,X84] : (p(s(bool,X83)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X84),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X84)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X85,sK12)),s('type.realax.real',X85))))) )),
  inference(resolution,[],[f633,f359])).
fof(f20630,plain,(
  ( ! [X88,X87,X86] : (p(s(bool,X86)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK36(sK12,X86,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X87,sK12)),s('type.realax.real',X87)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X88))))) )),
  inference(resolution,[],[f633,f359])).
fof(f21311,plain,(
  ( ! [X66,X67] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X66),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X66)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X67,sK12)),s('type.realax.real',X67))))) ) | $spl44),
  inference(resolution,[],[f754,f359])).
fof(f21312,plain,(
  ( ! [X70,X68,X69] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK25(X68,sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X69,sK12)),s('type.realax.real',X69)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X70))))) ) | $spl44),
  inference(resolution,[],[f754,f359])).
fof(f21591,plain,(
  ( ! [X85,X83,X84] : (p(s(bool,X83)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X84),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X84)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X85,sK12)),s('type.realax.real',X85))))) )),
  inference(resolution,[],[f634,f359])).
fof(f21592,plain,(
  ( ! [X88,X87,X86] : (p(s(bool,X86)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK37(sK12,X86,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X87,sK12)),s('type.realax.real',X87)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X88))))) )),
  inference(resolution,[],[f634,f359])).
fof(f22570,plain,(
  $false | ($spl42 | $spl260)),
  inference(resolution,[],[f8394,f753])).
fof(f22260,plain,(
  ( ! [X6,X5] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X5),s('type.realax.real',sK23(X5,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X5),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X6))))) ) | $spl248),
  inference(resolution,[],[f8187,f357])).
fof(f22259,plain,(
  ( ! [X4,X3] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X3,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X3),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X4))))) ) | $spl248),
  inference(resolution,[],[f8187,f355])).
fof(f22258,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X1)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X2))))) ) | $spl248),
  inference(resolution,[],[f8187,f353])).
fof(f21875,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X1)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X2))))) ) | $spl234),
  inference(resolution,[],[f7988,f347])).
fof(f21876,plain,(
  ( ! [X4,X3] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X3,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X3),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X4))))) ) | $spl234),
  inference(resolution,[],[f7988,f349])).
fof(f21877,plain,(
  ( ! [X6,X5] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X5),s('type.realax.real',sK21(X5,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X5),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X6))))) ) | $spl234),
  inference(resolution,[],[f7988,f351])).
fof(f3980,plain,(
  ( ! [X14] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real',sK27(sK12)))))) ) | $spl182),
  inference(cnf_transformation,[],[f3980_D])).
fof(f3980_D,plain,(
  ( ! [X14] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real',sK27(sK12)))))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f4977,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X20),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X20),s('type.realax.real',sK27(sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X21,sK12)),s('type.realax.real',X21)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X22))))) ) | $spl50),
  inference(resolution,[],[f758,f347])).
fof(f4978,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X23),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X23),s('type.realax.real',X24)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X24)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X25,sK12)),s('type.realax.real',X25))))) ) | $spl50),
  inference(resolution,[],[f758,f347])).
fof(f4980,plain,(
  ( ! [X28,X29] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X28,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X28),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X29,sK12)),s('type.realax.real',X29))))) ) | $spl50),
  inference(resolution,[],[f758,f349])).
fof(f4982,plain,(
  ( ! [X33,X32] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X32),s('type.realax.real',sK21(X32,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X32),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X33,sK12)),s('type.realax.real',X33))))) ) | $spl50),
  inference(resolution,[],[f758,f351])).
fof(f1518,plain,(
  ( ! [X21,X22] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X21)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X22),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X22),s('type.realax.real',sK16(X21)))))) ) | $spl128),
  inference(cnf_transformation,[],[f1518_D])).
fof(f1518_D,plain,(
  ( ! [X21,X22] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X21)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X22),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X22),s('type.realax.real',sK16(X21)))))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f5816,plain,(
  ( ! [X30,X31,X29,X32] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X29)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real',sK16(X29))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X31,sK12)),s('type.realax.real',X31)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X32))))) ) | $spl8),
  inference(resolution,[],[f588,f347])).
fof(f5817,plain,(
  ( ! [X35,X33,X36,X34] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X33)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X34),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X34),s('type.realax.real',X35)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X35)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X36,sK12)),s('type.realax.real',X36))))) ) | $spl8),
  inference(resolution,[],[f588,f347])).
fof(f5819,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X40)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X41,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X41),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X42,sK12)),s('type.realax.real',X42))))) ) | $spl8),
  inference(resolution,[],[f588,f349])).
fof(f5821,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X46)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X47),s('type.realax.real',sK21(X47,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X47),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X48,sK12)),s('type.realax.real',X48))))) ) | $spl8),
  inference(resolution,[],[f588,f351])).
fof(f7488,plain,(
  ( ! [X14] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real',sK38(sK12,'type.realax.real')))))) ) | $spl206),
  inference(cnf_transformation,[],[f7488_D])).
fof(f7488_D,plain,(
  ( ! [X14] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real',sK38(sK12,'type.realax.real')))))) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f7347,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X20),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X20),s('type.realax.real',sK38(sK12,'type.realax.real'))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X21,sK12)),s('type.realax.real',X21)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X22))))) )),
  inference(resolution,[],[f631,f347])).
fof(f7348,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X23),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X23),s('type.realax.real',X24)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X24)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X25,sK12)),s('type.realax.real',X25))))) )),
  inference(resolution,[],[f631,f347])).
fof(f7350,plain,(
  ( ! [X28,X29] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X28,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X28),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X29,sK12)),s('type.realax.real',X29))))) )),
  inference(resolution,[],[f631,f349])).
fof(f7352,plain,(
  ( ! [X33,X32] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X32),s('type.realax.real',sK21(X32,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X32),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X33,sK12)),s('type.realax.real',X33))))) )),
  inference(resolution,[],[f631,f351])).
fof(f7696,plain,(
  ( ! [X14] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real',sK39(sK12,'type.realax.real')))))) ) | $spl218),
  inference(cnf_transformation,[],[f7696_D])).
fof(f7696_D,plain,(
  ( ! [X14] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real',sK39(sK12,'type.realax.real')))))) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f7555,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X20),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X20),s('type.realax.real',sK39(sK12,'type.realax.real'))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X21,sK12)),s('type.realax.real',X21)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X22))))) )),
  inference(resolution,[],[f632,f347])).
fof(f7556,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X23),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X23),s('type.realax.real',X24)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X24)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X25,sK12)),s('type.realax.real',X25))))) )),
  inference(resolution,[],[f632,f347])).
fof(f7558,plain,(
  ( ! [X28,X29] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X28,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X28),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X29,sK12)),s('type.realax.real',X29))))) )),
  inference(resolution,[],[f632,f349])).
fof(f7560,plain,(
  ( ! [X33,X32] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X32),s('type.realax.real',sK21(X32,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X32),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X33,sK12)),s('type.realax.real',X33))))) )),
  inference(resolution,[],[f632,f351])).
fof(f5298,plain,(
  ( ! [X14] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real',sK30(sK12)))))) ) | $spl194),
  inference(cnf_transformation,[],[f5298_D])).
fof(f5298_D,plain,(
  ( ! [X14] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real',sK30(sK12)))))) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f9432,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X20),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X20),s('type.realax.real',sK30(sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X21,sK12)),s('type.realax.real',X21)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X22))))) ) | $spl60),
  inference(resolution,[],[f765,f347])).
fof(f9433,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X23),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X23),s('type.realax.real',X24)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X24)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X25,sK12)),s('type.realax.real',X25))))) ) | $spl60),
  inference(resolution,[],[f765,f347])).
fof(f9435,plain,(
  ( ! [X28,X29] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X28,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X28),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X29,sK12)),s('type.realax.real',X29))))) ) | $spl60),
  inference(resolution,[],[f765,f349])).
fof(f9437,plain,(
  ( ! [X33,X32] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X32),s('type.realax.real',sK21(X32,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X32),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X33,sK12)),s('type.realax.real',X33))))) ) | $spl60),
  inference(resolution,[],[f765,f351])).
fof(f740,plain,(
  ( ! [X14] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real',sK14)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) | $spl18),
  inference(cnf_transformation,[],[f740_D])).
fof(f740_D,plain,(
  ( ! [X14] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real',sK14)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f21992,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X20),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X20),s('type.realax.real',sK14)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X21,sK12)),s('type.realax.real',X21)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X22))))) )),
  inference(subsumption_resolution,[],[f12248,f579])).
fof(f12248,plain,(
  ( ! [X21,X22,X20] : (~p(s(bool,'T_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X20),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X20),s('type.realax.real',sK14)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X21,sK12)),s('type.realax.real',X21)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X22))))) )),
  inference(superposition,[],[f347,f715])).
fof(f21991,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X23),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X23),s('type.realax.real',X24)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X24)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X25,sK12)),s('type.realax.real',X25))))) )),
  inference(subsumption_resolution,[],[f12249,f579])).
fof(f12249,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'T_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X23),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X23),s('type.realax.real',X24)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X24)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X25,sK12)),s('type.realax.real',X25))))) )),
  inference(superposition,[],[f347,f715])).
fof(f21990,plain,(
  ( ! [X28,X29] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X28,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X28),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X29,sK12)),s('type.realax.real',X29))))) )),
  inference(subsumption_resolution,[],[f12251,f579])).
fof(f12251,plain,(
  ( ! [X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X28,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X28),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X29,sK12)),s('type.realax.real',X29))))) )),
  inference(superposition,[],[f349,f715])).
fof(f21989,plain,(
  ( ! [X33,X32] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X32),s('type.realax.real',sK21(X32,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X32),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X33,sK12)),s('type.realax.real',X33))))) )),
  inference(subsumption_resolution,[],[f12253,f579])).
fof(f12253,plain,(
  ( ! [X33,X32] : (~p(s(bool,'T_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X32),s('type.realax.real',sK21(X32,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X32),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X33,sK12)),s('type.realax.real',X33))))) )),
  inference(superposition,[],[f351,f715])).
fof(f20755,plain,(
  ( ! [X21,X22] : (p(s(bool,X21)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X22),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X22),s('type.realax.real',sK36(sK12,X21,'type.realax.real')))))) ) | $spl302),
  inference(cnf_transformation,[],[f20755_D])).
fof(f20755_D,plain,(
  ( ! [X21,X22] : (p(s(bool,X21)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X22),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X22),s('type.realax.real',sK36(sK12,X21,'type.realax.real')))))) ) <=> ~$spl302),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl302])])).
fof(f20613,plain,(
  ( ! [X30,X31,X29,X32] : (p(s(bool,X29)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real',sK36(sK12,X29,'type.realax.real'))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X31,sK12)),s('type.realax.real',X31)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X32))))) )),
  inference(resolution,[],[f633,f347])).
fof(f20614,plain,(
  ( ! [X35,X33,X36,X34] : (p(s(bool,X33)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X34),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X34),s('type.realax.real',X35)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X35)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X36,sK12)),s('type.realax.real',X36))))) )),
  inference(resolution,[],[f633,f347])).
fof(f20616,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,X40)) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X41,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X41),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X42,sK12)),s('type.realax.real',X42))))) )),
  inference(resolution,[],[f633,f349])).
fof(f20618,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,X46)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X47),s('type.realax.real',sK21(X47,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X47),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X48,sK12)),s('type.realax.real',X48))))) )),
  inference(resolution,[],[f633,f351])).
fof(f8789,plain,(
  ( ! [X17,X18] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real',sK25(X18,sK12)))))) ) | $spl266),
  inference(cnf_transformation,[],[f8789_D])).
fof(f8789_D,plain,(
  ( ! [X17,X18] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real',sK25(X18,sK12)))))) ) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f21295,plain,(
  ( ! [X26,X24,X27,X25] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X24),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X24),s('type.realax.real',sK25(X25,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X26,sK12)),s('type.realax.real',X26)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X27))))) ) | $spl44),
  inference(resolution,[],[f754,f347])).
fof(f21296,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X28),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X28),s('type.realax.real',X29)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X29)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X30,sK12)),s('type.realax.real',X30))))) ) | $spl44),
  inference(resolution,[],[f754,f347])).
fof(f21298,plain,(
  ( ! [X33,X34] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X33,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X33),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X34,sK12)),s('type.realax.real',X34))))) ) | $spl44),
  inference(resolution,[],[f754,f349])).
fof(f21300,plain,(
  ( ! [X37,X38] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X37),s('type.realax.real',sK21(X37,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X37),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X38,sK12)),s('type.realax.real',X38))))) ) | $spl44),
  inference(resolution,[],[f754,f351])).
fof(f21716,plain,(
  ( ! [X21,X22] : (p(s(bool,X21)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X22),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X22),s('type.realax.real',sK37(sK12,X21,'type.realax.real')))))) ) | $spl314),
  inference(cnf_transformation,[],[f21716_D])).
fof(f21716_D,plain,(
  ( ! [X21,X22] : (p(s(bool,X21)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X22),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X22),s('type.realax.real',sK37(sK12,X21,'type.realax.real')))))) ) <=> ~$spl314),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl314])])).
fof(f21575,plain,(
  ( ! [X30,X31,X29,X32] : (p(s(bool,X29)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real',sK37(sK12,X29,'type.realax.real'))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X31,sK12)),s('type.realax.real',X31)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X32))))) )),
  inference(resolution,[],[f634,f347])).
fof(f21576,plain,(
  ( ! [X35,X33,X36,X34] : (p(s(bool,X33)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X34),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X34),s('type.realax.real',X35)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X35)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X36,sK12)),s('type.realax.real',X36))))) )),
  inference(resolution,[],[f634,f347])).
fof(f21578,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,X40)) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X41,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X41),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X42,sK12)),s('type.realax.real',X42))))) )),
  inference(resolution,[],[f634,f349])).
fof(f21580,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,X46)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X47),s('type.realax.real',sK21(X47,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X47),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X48,sK12)),s('type.realax.real',X48))))) )),
  inference(resolution,[],[f634,f351])).
fof(f21874,plain,(
  $false | ($spl22 | $spl234)),
  inference(resolution,[],[f7988,f742])).
fof(f21606,plain,(
  ( ! [X125,X126,X124] : (p(s(bool,X124)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X125),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X125)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X126,sK12)))))) )),
  inference(resolution,[],[f634,f373])).
fof(f21605,plain,(
  ( ! [X123,X121,X122] : (p(s(bool,X121)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK37(sK12,X121,'type.realax.real')),s('type.realax.real',sK31(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X122,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X123))))) )),
  inference(resolution,[],[f634,f373])).
fof(f762,plain,(
  ( ! [X74] : (~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X74)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X74),s('type.realax.real',sK28(sK12)))))) ) | $spl56),
  inference(cnf_transformation,[],[f762_D])).
fof(f762_D,plain,(
  ( ! [X74] : (~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X74)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X74),s('type.realax.real',sK28(sK12)))))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f21602,plain,(
  ( ! [X116,X114,X115] : (p(s(bool,X114)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X115),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X115)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X116,sK12)),s('type.realax.real',X116))))) )),
  inference(resolution,[],[f634,f370])).
fof(f21601,plain,(
  ( ! [X111,X112,X113] : (p(s(bool,X111)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK37(sK12,X111,'type.realax.real')),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X112,sK12)),s('type.realax.real',X112)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X113))))) )),
  inference(resolution,[],[f634,f370])).
fof(f21600,plain,(
  ( ! [X109,X110,X108] : (p(s(bool,X108)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X109),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X109)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X110,sK12)))))) )),
  inference(resolution,[],[f634,f369])).
fof(f21599,plain,(
  ( ! [X107,X105,X106] : (p(s(bool,X105)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK37(sK12,X105,'type.realax.real')),s('type.realax.real',sK28(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X106,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X107))))) )),
  inference(resolution,[],[f634,f369])).
fof(f21590,plain,(
  ( ! [X80,X81,X82] : (p(s(bool,X80)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK37(sK12,X80,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X81,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X82))))) )),
  inference(resolution,[],[f634,f358])).
fof(f21589,plain,(
  ( ! [X78,X79,X77] : (p(s(bool,X77)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X78),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X78)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X79,sK12)))))) )),
  inference(resolution,[],[f634,f358])).
fof(f21588,plain,(
  ( ! [X76,X74,X75] : (p(s(bool,X74)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X75),s('type.realax.real',sK23(X75,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X75),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X76,sK12)),s('type.realax.real',X76))))) )),
  inference(resolution,[],[f634,f357])).
fof(f21587,plain,(
  ( ! [X72,X71,X73] : (p(s(bool,X71)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X72),s('type.realax.real',sK23(X72,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X72),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X73,sK12)))))) )),
  inference(resolution,[],[f634,f356])).
fof(f21586,plain,(
  ( ! [X70,X68,X69] : (p(s(bool,X68)) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X69,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X69),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X70,sK12)),s('type.realax.real',X70))))) )),
  inference(resolution,[],[f634,f355])).
fof(f21585,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,X65)) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X66,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X66),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X67,sK12)))))) )),
  inference(resolution,[],[f634,f354])).
fof(f21584,plain,(
  ( ! [X61,X64,X62,X63] : (p(s(bool,X61)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X62),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X62),s('type.realax.real',X63)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X63)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X64,sK12)),s('type.realax.real',X64))))) )),
  inference(resolution,[],[f634,f353])).
fof(f21583,plain,(
  ( ! [X59,X57,X60,X58] : (p(s(bool,X57)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X58),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X58),s('type.realax.real',sK37(sK12,X57,'type.realax.real'))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X59,sK12)),s('type.realax.real',X59)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X60))))) )),
  inference(resolution,[],[f634,f353])).
fof(f21582,plain,(
  ( ! [X54,X56,X55,X53] : (p(s(bool,X53)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X54),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X54),s('type.realax.real',X55)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X55)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X56,sK12)))))) )),
  inference(resolution,[],[f634,f352])).
fof(f21581,plain,(
  ( ! [X52,X50,X51,X49] : (p(s(bool,X49)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X50),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X50),s('type.realax.real',sK37(sK12,X49,'type.realax.real'))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X51,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X52))))) )),
  inference(resolution,[],[f634,f352])).
fof(f21579,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,X43)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X44),s('type.realax.real',sK21(X44,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X44),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X45,sK12)))))) )),
  inference(resolution,[],[f634,f350])).
fof(f21577,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,X37)) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X38,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X38),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X39,sK12)))))) )),
  inference(resolution,[],[f634,f348])).
fof(f21574,plain,(
  ( ! [X28,X26,X27,X25] : (p(s(bool,X25)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X26),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X26),s('type.realax.real',X27)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X27)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X28,sK12)))))) )),
  inference(resolution,[],[f634,f346])).
fof(f21573,plain,(
  ( ! [X24,X23,X21,X22] : (p(s(bool,X21)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X22),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X22),s('type.realax.real',sK37(sK12,X21,'type.realax.real'))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X23,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X24))))) )),
  inference(resolution,[],[f634,f346])).
fof(f1517,plain,(
  ( ! [X20] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))),s('type.realax.real',X20)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK19(X20,sK12)),s('type.realax.real',X20))))) ) | $spl126),
  inference(cnf_transformation,[],[f1517_D])).
fof(f1517_D,plain,(
  ( ! [X20] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))),s('type.realax.real',X20)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK19(X20,sK12)),s('type.realax.real',X20))))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f21572,plain,(
  ( ! [X19,X20] : (p(s(bool,X19)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))),s('type.realax.real',X20)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK19(X20,sK12)),s('type.realax.real',X20))))) )),
  inference(resolution,[],[f634,f345])).
fof(f1516,plain,(
  ( ! [X18] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))),s('type.realax.real',X18)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK19(X18,sK12)))))) ) | $spl124),
  inference(cnf_transformation,[],[f1516_D])).
fof(f1516_D,plain,(
  ( ! [X18] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))),s('type.realax.real',X18)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK19(X18,sK12)))))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f21571,plain,(
  ( ! [X17,X18] : (p(s(bool,X17)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))),s('type.realax.real',X18)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK19(X18,sK12)))))) )),
  inference(resolution,[],[f634,f344])).
fof(f21715,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X15,X14,sK12)),s('type.realax.real',X15)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real',sK37(sK12,X13,'type.realax.real')))))) ) | $spl312),
  inference(cnf_transformation,[],[f21715_D])).
fof(f21715_D,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X15,X14,sK12)),s('type.realax.real',X15)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real',sK37(sK12,X13,'type.realax.real')))))) ) <=> ~$spl312),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl312])])).
fof(f21570,plain,(
  ( ! [X14,X15,X13,X16] : (p(s(bool,X13)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real',sK37(sK12,X13,'type.realax.real'))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X15,X14,sK12)),s('type.realax.real',X15)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X16))))) )),
  inference(resolution,[],[f634,f343])).
fof(f1514,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X10),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X12,X10,sK12)),s('type.realax.real',X12)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X11)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X10),s('type.realax.real',X11))))) ) | $spl120),
  inference(cnf_transformation,[],[f1514_D])).
fof(f1514_D,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X10),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X12,X10,sK12)),s('type.realax.real',X12)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X11)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X10),s('type.realax.real',X11))))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f21569,plain,(
  ( ! [X12,X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X10),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X10),s('type.realax.real',X11)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X11)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X12,X10,sK12)),s('type.realax.real',X12))))) )),
  inference(resolution,[],[f634,f343])).
fof(f21714,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X7,X6,sK12))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',sK37(sK12,X5,'type.realax.real')))))) ) | $spl310),
  inference(cnf_transformation,[],[f21714_D])).
fof(f21714_D,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X7,X6,sK12))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',sK37(sK12,X5,'type.realax.real')))))) ) <=> ~$spl310),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl310])])).
fof(f21568,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,X5)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',sK37(sK12,X5,'type.realax.real'))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X7,X6,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X8))))) )),
  inference(resolution,[],[f634,f342])).
fof(f1511,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X4,X2,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3))))) ) | $spl114),
  inference(cnf_transformation,[],[f1511_D])).
fof(f1511_D,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X4,X2,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3))))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f21567,plain,(
  ( ! [X4,X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X3)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X4,X2,sK12)))))) )),
  inference(resolution,[],[f634,f342])).
fof(f21326,plain,(
  ( ! [X99,X98] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X98),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X98)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X99,sK12)))))) ) | $spl44),
  inference(resolution,[],[f754,f373])).
fof(f8793,plain,(
  ( ! [X95] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK25(X95,sK12)),s('type.realax.real',sK31(sK12)))))) ) | $spl272),
  inference(cnf_transformation,[],[f8793_D])).
fof(f8793_D,plain,(
  ( ! [X95] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK25(X95,sK12)),s('type.realax.real',sK31(sK12)))))) ) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f21325,plain,(
  ( ! [X97,X95,X96] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK25(X95,sK12)),s('type.realax.real',sK31(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X96,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X97))))) ) | $spl44),
  inference(resolution,[],[f754,f373])).
fof(f21322,plain,(
  ( ! [X92,X91] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X91),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X91)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X92,sK12)),s('type.realax.real',X92))))) ) | $spl44),
  inference(resolution,[],[f754,f370])).
fof(f21321,plain,(
  ( ! [X90,X88,X89] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X88,sK12)),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X89,sK12)),s('type.realax.real',X89)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X90))))) ) | $spl44),
  inference(resolution,[],[f754,f370])).
fof(f21320,plain,(
  ( ! [X87,X86] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X86),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X86)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X87,sK12)))))) ) | $spl44),
  inference(resolution,[],[f754,f369])).
fof(f21319,plain,(
  ( ! [X85,X83,X84] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X83,sK12)),s('type.realax.real',sK28(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X84,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X85))))) ) | $spl44),
  inference(resolution,[],[f754,f369])).
fof(f21310,plain,(
  ( ! [X64,X65,X63] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK25(X63,sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X64,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X65))))) ) | $spl44),
  inference(resolution,[],[f754,f358])).
fof(f21309,plain,(
  ( ! [X61,X62] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X61),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X61)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X62,sK12)))))) ) | $spl44),
  inference(resolution,[],[f754,f358])).
fof(f21308,plain,(
  ( ! [X59,X60] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X59),s('type.realax.real',sK23(X59,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X59),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X60,sK12)),s('type.realax.real',X60))))) ) | $spl44),
  inference(resolution,[],[f754,f357])).
fof(f21307,plain,(
  ( ! [X57,X58] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X57),s('type.realax.real',sK23(X57,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X57),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X58,sK12)))))) ) | $spl44),
  inference(resolution,[],[f754,f356])).
fof(f21306,plain,(
  ( ! [X56,X55] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X55,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X55),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X56,sK12)),s('type.realax.real',X56))))) ) | $spl44),
  inference(resolution,[],[f754,f355])).
fof(f21305,plain,(
  ( ! [X54,X53] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X53,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X53),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X54,sK12)))))) ) | $spl44),
  inference(resolution,[],[f754,f354])).
fof(f21304,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X50),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X50),s('type.realax.real',X51)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X51)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X52,sK12)),s('type.realax.real',X52))))) ) | $spl44),
  inference(resolution,[],[f754,f353])).
fof(f21303,plain,(
  ( ! [X47,X48,X46,X49] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X46),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X46),s('type.realax.real',sK25(X47,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X48,sK12)),s('type.realax.real',X48)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X49))))) ) | $spl44),
  inference(resolution,[],[f754,f353])).
fof(f21302,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X43),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X43),s('type.realax.real',X44)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X44)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X45,sK12)))))) ) | $spl44),
  inference(resolution,[],[f754,f352])).
fof(f21301,plain,(
  ( ! [X39,X41,X42,X40] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X39),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X39),s('type.realax.real',sK25(X40,sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X41,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X42))))) ) | $spl44),
  inference(resolution,[],[f754,f352])).
fof(f21299,plain,(
  ( ! [X35,X36] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X35),s('type.realax.real',sK21(X35,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X35),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X36,sK12)))))) ) | $spl44),
  inference(resolution,[],[f754,f350])).
fof(f21297,plain,(
  ( ! [X31,X32] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X31,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X31),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X32,sK12)))))) ) | $spl44),
  inference(resolution,[],[f754,f348])).
fof(f21294,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X21),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X21),s('type.realax.real',X22)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X22)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X23,sK12)))))) ) | $spl44),
  inference(resolution,[],[f754,f346])).
fof(f21293,plain,(
  ( ! [X19,X17,X20,X18] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real',sK25(X18,sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X19,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X20))))) ) | $spl44),
  inference(resolution,[],[f754,f346])).
fof(f8788,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X13,X11,sK12)),s('type.realax.real',X13)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK25(X12,sK12)))))) ) | $spl264),
  inference(cnf_transformation,[],[f8788_D])).
fof(f8788_D,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X13,X11,sK12)),s('type.realax.real',X13)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK25(X12,sK12)))))) ) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f21290,plain,(
  ( ! [X14,X12,X13,X11] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK25(X12,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X13,X11,sK12)),s('type.realax.real',X13)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X14))))) ) | $spl44),
  inference(resolution,[],[f754,f343])).
fof(f8787,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X6,X4,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',sK25(X5,sK12)))))) ) | $spl262),
  inference(cnf_transformation,[],[f8787_D])).
fof(f8787_D,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X6,X4,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',sK25(X5,sK12)))))) ) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f21288,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',sK25(X5,sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X6,X4,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X7))))) ) | $spl44),
  inference(resolution,[],[f754,f342])).
fof(f21043,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X1)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X2))))) ) | $spl28),
  inference(resolution,[],[f8033,f353])).
fof(f8033,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X0,sK12)),s('type.realax.real',sK15))))) ) | $spl28),
  inference(resolution,[],[f745,f337])).
fof(f21044,plain,(
  ( ! [X4,X3] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X3,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X3),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X4))))) ) | $spl28),
  inference(resolution,[],[f8033,f355])).
fof(f21045,plain,(
  ( ! [X6,X5] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X5),s('type.realax.real',sK23(X5,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X5),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X6))))) ) | $spl28),
  inference(resolution,[],[f8033,f357])).
fof(f4985,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X40),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X40),s('type.realax.real',sK27(sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X41,sK12)),s('type.realax.real',X41)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X42))))) ) | $spl50),
  inference(resolution,[],[f758,f353])).
fof(f4986,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X43),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X43),s('type.realax.real',X44)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X44)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X45,sK12)),s('type.realax.real',X45))))) ) | $spl50),
  inference(resolution,[],[f758,f353])).
fof(f4988,plain,(
  ( ! [X48,X49] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X48,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X48),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X49,sK12)),s('type.realax.real',X49))))) ) | $spl50),
  inference(resolution,[],[f758,f355])).
fof(f4990,plain,(
  ( ! [X52,X53] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X52),s('type.realax.real',sK23(X52,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X52),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X53,sK12)),s('type.realax.real',X53))))) ) | $spl50),
  inference(resolution,[],[f758,f357])).
fof(f5824,plain,(
  ( ! [X59,X57,X60,X58] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X57)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X58),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X58),s('type.realax.real',sK16(X57))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X59,sK12)),s('type.realax.real',X59)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X60))))) ) | $spl8),
  inference(resolution,[],[f588,f353])).
fof(f5825,plain,(
  ( ! [X61,X64,X62,X63] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X61)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X62),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X62),s('type.realax.real',X63)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X63)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X64,sK12)),s('type.realax.real',X64))))) ) | $spl8),
  inference(resolution,[],[f588,f353])).
fof(f5827,plain,(
  ( ! [X70,X68,X69] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X68)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X69,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X69),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X70,sK12)),s('type.realax.real',X70))))) ) | $spl8),
  inference(resolution,[],[f588,f355])).
fof(f5829,plain,(
  ( ! [X76,X74,X75] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X74)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X75),s('type.realax.real',sK23(X75,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X75),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X76,sK12)),s('type.realax.real',X76))))) ) | $spl8),
  inference(resolution,[],[f588,f357])).
fof(f7355,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X40),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X40),s('type.realax.real',sK38(sK12,'type.realax.real'))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X41,sK12)),s('type.realax.real',X41)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X42))))) )),
  inference(resolution,[],[f631,f353])).
fof(f7356,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X43),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X43),s('type.realax.real',X44)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X44)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X45,sK12)),s('type.realax.real',X45))))) )),
  inference(resolution,[],[f631,f353])).
fof(f7358,plain,(
  ( ! [X48,X49] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X48,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X48),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X49,sK12)),s('type.realax.real',X49))))) )),
  inference(resolution,[],[f631,f355])).
fof(f7360,plain,(
  ( ! [X52,X53] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X52),s('type.realax.real',sK23(X52,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X52),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X53,sK12)),s('type.realax.real',X53))))) )),
  inference(resolution,[],[f631,f357])).
fof(f7563,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X40),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X40),s('type.realax.real',sK39(sK12,'type.realax.real'))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X41,sK12)),s('type.realax.real',X41)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X42))))) )),
  inference(resolution,[],[f632,f353])).
fof(f7564,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X43),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X43),s('type.realax.real',X44)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X44)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X45,sK12)),s('type.realax.real',X45))))) )),
  inference(resolution,[],[f632,f353])).
fof(f7566,plain,(
  ( ! [X48,X49] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X48,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X48),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X49,sK12)),s('type.realax.real',X49))))) )),
  inference(resolution,[],[f632,f355])).
fof(f7568,plain,(
  ( ! [X52,X53] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X52),s('type.realax.real',sK23(X52,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X52),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X53,sK12)),s('type.realax.real',X53))))) )),
  inference(resolution,[],[f632,f357])).
fof(f9440,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X40),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X40),s('type.realax.real',sK30(sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X41,sK12)),s('type.realax.real',X41)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X42))))) ) | $spl60),
  inference(resolution,[],[f765,f353])).
fof(f9441,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X43),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X43),s('type.realax.real',X44)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X44)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X45,sK12)),s('type.realax.real',X45))))) ) | $spl60),
  inference(resolution,[],[f765,f353])).
fof(f9443,plain,(
  ( ! [X48,X49] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X48,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X48),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X49,sK12)),s('type.realax.real',X49))))) ) | $spl60),
  inference(resolution,[],[f765,f355])).
fof(f9445,plain,(
  ( ! [X52,X53] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X52),s('type.realax.real',sK23(X52,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X52),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X53,sK12)),s('type.realax.real',X53))))) ) | $spl60),
  inference(resolution,[],[f765,f357])).
fof(f21160,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X40),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X40),s('type.realax.real',sK14)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X41,sK12)),s('type.realax.real',X41)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X42))))) )),
  inference(subsumption_resolution,[],[f12256,f579])).
fof(f12256,plain,(
  ( ! [X41,X42,X40] : (~p(s(bool,'T_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X40),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X40),s('type.realax.real',sK14)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X41,sK12)),s('type.realax.real',X41)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X42))))) )),
  inference(superposition,[],[f353,f715])).
fof(f21159,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X43),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X43),s('type.realax.real',X44)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X44)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X45,sK12)),s('type.realax.real',X45))))) )),
  inference(subsumption_resolution,[],[f12257,f579])).
fof(f12257,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,'T_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X43),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X43),s('type.realax.real',X44)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X44)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X45,sK12)),s('type.realax.real',X45))))) )),
  inference(superposition,[],[f353,f715])).
fof(f21158,plain,(
  ( ! [X48,X49] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X48,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X48),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X49,sK12)),s('type.realax.real',X49))))) )),
  inference(subsumption_resolution,[],[f12259,f579])).
fof(f12259,plain,(
  ( ! [X48,X49] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X48,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X48),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X49,sK12)),s('type.realax.real',X49))))) )),
  inference(superposition,[],[f355,f715])).
fof(f21157,plain,(
  ( ! [X52,X53] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X52),s('type.realax.real',sK23(X52,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X52),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X53,sK12)),s('type.realax.real',X53))))) )),
  inference(subsumption_resolution,[],[f12261,f579])).
fof(f12261,plain,(
  ( ! [X52,X53] : (~p(s(bool,'T_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X52),s('type.realax.real',sK23(X52,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X52),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X53,sK12)),s('type.realax.real',X53))))) )),
  inference(superposition,[],[f357,f715])).
fof(f20621,plain,(
  ( ! [X59,X57,X60,X58] : (p(s(bool,X57)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X58),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X58),s('type.realax.real',sK36(sK12,X57,'type.realax.real'))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X59,sK12)),s('type.realax.real',X59)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X60))))) )),
  inference(resolution,[],[f633,f353])).
fof(f20622,plain,(
  ( ! [X61,X64,X62,X63] : (p(s(bool,X61)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X62),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X62),s('type.realax.real',X63)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X63)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X64,sK12)),s('type.realax.real',X64))))) )),
  inference(resolution,[],[f633,f353])).
fof(f20624,plain,(
  ( ! [X70,X68,X69] : (p(s(bool,X68)) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X69,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X69),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X70,sK12)),s('type.realax.real',X70))))) )),
  inference(resolution,[],[f633,f355])).
fof(f20626,plain,(
  ( ! [X76,X74,X75] : (p(s(bool,X74)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X75),s('type.realax.real',sK23(X75,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X75),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X76,sK12)),s('type.realax.real',X76))))) )),
  inference(resolution,[],[f633,f357])).
fof(f7986,plain,(
  ( ! [X39,X40] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X39),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X39),s('type.realax.real',sK20(X40,sK12)))))) ) | $spl230),
  inference(cnf_transformation,[],[f7986_D])).
fof(f7986_D,plain,(
  ( ! [X39,X40] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X39),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X39),s('type.realax.real',sK20(X40,sK12)))))) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f20907,plain,(
  ( ! [X47,X48,X46,X49] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X46),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X46),s('type.realax.real',sK20(X47,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X48,sK12)),s('type.realax.real',X48)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X49))))) ) | $spl16),
  inference(resolution,[],[f739,f353])).
fof(f20908,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X50),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X50),s('type.realax.real',X51)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X51)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X52,sK12)),s('type.realax.real',X52))))) ) | $spl16),
  inference(resolution,[],[f739,f353])).
fof(f20910,plain,(
  ( ! [X56,X55] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X55,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X55),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X56,sK12)),s('type.realax.real',X56))))) ) | $spl16),
  inference(resolution,[],[f739,f355])).
fof(f20912,plain,(
  ( ! [X59,X60] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X59),s('type.realax.real',sK23(X59,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X59),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X60,sK12)),s('type.realax.real',X60))))) ) | $spl16),
  inference(resolution,[],[f739,f357])).
fof(f21042,plain,(
  $false | ($spl28 | $spl30)),
  inference(resolution,[],[f8033,f746])).
fof(f20930,plain,(
  ( ! [X99,X98] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X98),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X98)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X99,sK12)))))) ) | $spl16),
  inference(resolution,[],[f739,f373])).
fof(f7990,plain,(
  ( ! [X95] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK20(X95,sK12)),s('type.realax.real',sK31(sK12)))))) ) | $spl236),
  inference(cnf_transformation,[],[f7990_D])).
fof(f7990_D,plain,(
  ( ! [X95] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK20(X95,sK12)),s('type.realax.real',sK31(sK12)))))) ) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f20929,plain,(
  ( ! [X97,X95,X96] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK20(X95,sK12)),s('type.realax.real',sK31(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X96,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X97))))) ) | $spl16),
  inference(resolution,[],[f739,f373])).
fof(f20926,plain,(
  ( ! [X92,X91] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X91),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X91)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X92,sK12)),s('type.realax.real',X92))))) ) | $spl16),
  inference(resolution,[],[f739,f370])).
fof(f7988,plain,(
  ( ! [X83] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X83,sK12)),s('type.realax.real',sK28(sK12)))))) ) | $spl234),
  inference(cnf_transformation,[],[f7988_D])).
fof(f7988_D,plain,(
  ( ! [X83] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X83,sK12)),s('type.realax.real',sK28(sK12)))))) ) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f20925,plain,(
  ( ! [X90,X88,X89] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X88,sK12)),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X89,sK12)),s('type.realax.real',X89)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X90))))) ) | $spl16),
  inference(resolution,[],[f739,f370])).
fof(f20924,plain,(
  ( ! [X87,X86] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X86),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X86)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X87,sK12)))))) ) | $spl16),
  inference(resolution,[],[f739,f369])).
fof(f20923,plain,(
  ( ! [X85,X83,X84] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X83,sK12)),s('type.realax.real',sK28(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X84,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X85))))) ) | $spl16),
  inference(resolution,[],[f739,f369])).
fof(f754,plain,(
  ( ! [X63] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(X63,sK12)))))) ) | $spl44),
  inference(cnf_transformation,[],[f754_D])).
fof(f754_D,plain,(
  ( ! [X63] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(X63,sK12)))))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f7987,plain,(
  ( ! [X63] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK20(X63,sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) | $spl232),
  inference(cnf_transformation,[],[f7987_D])).
fof(f7987_D,plain,(
  ( ! [X63] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK20(X63,sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f20918,plain,(
  ( ! [X74,X75,X73] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK20(X73,sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(X74,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X75))))) ) | $spl16),
  inference(resolution,[],[f739,f360])).
fof(f20917,plain,(
  ( ! [X72,X71] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X71),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X71)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(X72,sK12)))))) ) | $spl16),
  inference(resolution,[],[f739,f360])).
fof(f20914,plain,(
  ( ! [X64,X65,X63] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK20(X63,sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X64,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X65))))) ) | $spl16),
  inference(resolution,[],[f739,f358])).
fof(f20913,plain,(
  ( ! [X61,X62] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X61),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X61)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X62,sK12)))))) ) | $spl16),
  inference(resolution,[],[f739,f358])).
fof(f20911,plain,(
  ( ! [X57,X58] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X57),s('type.realax.real',sK23(X57,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X57),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X58,sK12)))))) ) | $spl16),
  inference(resolution,[],[f739,f356])).
fof(f20909,plain,(
  ( ! [X54,X53] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X53,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X53),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X54,sK12)))))) ) | $spl16),
  inference(resolution,[],[f739,f354])).
fof(f20906,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X43),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X43),s('type.realax.real',X44)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X44)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X45,sK12)))))) ) | $spl16),
  inference(resolution,[],[f739,f352])).
fof(f20905,plain,(
  ( ! [X39,X41,X42,X40] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X39),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X39),s('type.realax.real',sK20(X40,sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X41,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X42))))) ) | $spl16),
  inference(resolution,[],[f739,f352])).
fof(f7985,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X13,X11,sK12)),s('type.realax.real',X13)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK20(X12,sK12)))))) ) | $spl228),
  inference(cnf_transformation,[],[f7985_D])).
fof(f7985_D,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X13,X11,sK12)),s('type.realax.real',X13)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK20(X12,sK12)))))) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f20894,plain,(
  ( ! [X14,X12,X13,X11] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK20(X12,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X13,X11,sK12)),s('type.realax.real',X13)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X14))))) ) | $spl16),
  inference(resolution,[],[f739,f343])).
fof(f7984,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X6,X4,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',sK20(X5,sK12)))))) ) | $spl226),
  inference(cnf_transformation,[],[f7984_D])).
fof(f7984_D,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X6,X4,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',sK20(X5,sK12)))))) ) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f20892,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',sK20(X5,sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X6,X4,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X7))))) ) | $spl16),
  inference(resolution,[],[f739,f342])).
fof(f1527,plain,(
  ( ! [X138] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK37(sK12,X138,'type.realax.real'))))) | p(s(bool,X138))) ) | $spl142),
  inference(cnf_transformation,[],[f1527_D])).
fof(f1527_D,plain,(
  ( ! [X138] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK37(sK12,X138,'type.realax.real'))))) | p(s(bool,X138))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f20650,plain,(
  ( ! [X138,X137] : (p(s(bool,X137)) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK37(sK12,X138,'type.realax.real'))))) | p(s(bool,X138))) )),
  inference(resolution,[],[f633,f505])).
fof(f20644,plain,(
  ( ! [X125,X126,X124] : (p(s(bool,X124)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X125),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X125)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X126,sK12)))))) )),
  inference(resolution,[],[f633,f373])).
fof(f20643,plain,(
  ( ! [X123,X121,X122] : (p(s(bool,X121)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK36(sK12,X121,'type.realax.real')),s('type.realax.real',sK31(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X122,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X123))))) )),
  inference(resolution,[],[f633,f373])).
fof(f20640,plain,(
  ( ! [X116,X114,X115] : (p(s(bool,X114)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X115),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X115)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X116,sK12)),s('type.realax.real',X116))))) )),
  inference(resolution,[],[f633,f370])).
fof(f20639,plain,(
  ( ! [X111,X112,X113] : (p(s(bool,X111)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK36(sK12,X111,'type.realax.real')),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X112,sK12)),s('type.realax.real',X112)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X113))))) )),
  inference(resolution,[],[f633,f370])).
fof(f20638,plain,(
  ( ! [X109,X110,X108] : (p(s(bool,X108)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X109),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X109)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X110,sK12)))))) )),
  inference(resolution,[],[f633,f369])).
fof(f20637,plain,(
  ( ! [X107,X105,X106] : (p(s(bool,X105)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK36(sK12,X105,'type.realax.real')),s('type.realax.real',sK28(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X106,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X107))))) )),
  inference(resolution,[],[f633,f369])).
fof(f20632,plain,(
  ( ! [X94,X92,X93] : (p(s(bool,X92)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK36(sK12,X92,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(X93,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X94))))) )),
  inference(resolution,[],[f633,f360])).
fof(f20631,plain,(
  ( ! [X90,X91,X89] : (p(s(bool,X89)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X90),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X90)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(X91,sK12)))))) )),
  inference(resolution,[],[f633,f360])).
fof(f20628,plain,(
  ( ! [X80,X81,X82] : (p(s(bool,X80)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK36(sK12,X80,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X81,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X82))))) )),
  inference(resolution,[],[f633,f358])).
fof(f20627,plain,(
  ( ! [X78,X79,X77] : (p(s(bool,X77)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X78),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X78)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X79,sK12)))))) )),
  inference(resolution,[],[f633,f358])).
fof(f20625,plain,(
  ( ! [X72,X71,X73] : (p(s(bool,X71)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X72),s('type.realax.real',sK23(X72,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X72),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X73,sK12)))))) )),
  inference(resolution,[],[f633,f356])).
fof(f20623,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,X65)) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X66,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X66),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X67,sK12)))))) )),
  inference(resolution,[],[f633,f354])).
fof(f20620,plain,(
  ( ! [X54,X56,X55,X53] : (p(s(bool,X53)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X54),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X54),s('type.realax.real',X55)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X55)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X56,sK12)))))) )),
  inference(resolution,[],[f633,f352])).
fof(f20619,plain,(
  ( ! [X52,X50,X51,X49] : (p(s(bool,X49)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X50),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X50),s('type.realax.real',sK36(sK12,X49,'type.realax.real'))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X51,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X52))))) )),
  inference(resolution,[],[f633,f352])).
fof(f20617,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,X43)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X44),s('type.realax.real',sK21(X44,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X44),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X45,sK12)))))) )),
  inference(resolution,[],[f633,f350])).
fof(f20615,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,X37)) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X38,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X38),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X39,sK12)))))) )),
  inference(resolution,[],[f633,f348])).
fof(f20612,plain,(
  ( ! [X28,X26,X27,X25] : (p(s(bool,X25)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X26),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X26),s('type.realax.real',X27)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X27)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X28,sK12)))))) )),
  inference(resolution,[],[f633,f346])).
fof(f20611,plain,(
  ( ! [X24,X23,X21,X22] : (p(s(bool,X21)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X22),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X22),s('type.realax.real',sK36(sK12,X21,'type.realax.real'))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X23,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X24))))) )),
  inference(resolution,[],[f633,f346])).
fof(f20610,plain,(
  ( ! [X19,X20] : (p(s(bool,X19)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))),s('type.realax.real',X20)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK19(X20,sK12)),s('type.realax.real',X20))))) )),
  inference(resolution,[],[f633,f345])).
fof(f20609,plain,(
  ( ! [X17,X18] : (p(s(bool,X17)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))),s('type.realax.real',X18)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK19(X18,sK12)))))) )),
  inference(resolution,[],[f633,f344])).
fof(f20754,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X15,X14,sK12)),s('type.realax.real',X15)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real',sK36(sK12,X13,'type.realax.real')))))) ) | $spl300),
  inference(cnf_transformation,[],[f20754_D])).
fof(f20754_D,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X15,X14,sK12)),s('type.realax.real',X15)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real',sK36(sK12,X13,'type.realax.real')))))) ) <=> ~$spl300),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl300])])).
fof(f20608,plain,(
  ( ! [X14,X15,X13,X16] : (p(s(bool,X13)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real',sK36(sK12,X13,'type.realax.real'))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X15,X14,sK12)),s('type.realax.real',X15)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X16))))) )),
  inference(resolution,[],[f633,f343])).
fof(f20607,plain,(
  ( ! [X12,X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X10),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X10),s('type.realax.real',X11)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X11)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X12,X10,sK12)),s('type.realax.real',X12))))) )),
  inference(resolution,[],[f633,f343])).
fof(f20753,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X7,X6,sK12))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',sK36(sK12,X5,'type.realax.real')))))) ) | $spl298),
  inference(cnf_transformation,[],[f20753_D])).
fof(f20753_D,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X7,X6,sK12))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',sK36(sK12,X5,'type.realax.real')))))) ) <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f20606,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,X5)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',sK36(sK12,X5,'type.realax.real'))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X7,X6,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X8))))) )),
  inference(resolution,[],[f633,f342])).
fof(f20605,plain,(
  ( ! [X4,X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X3)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X4,X2,sK12)))))) )),
  inference(resolution,[],[f633,f342])).
fof(f890,plain,(
  ( ! [X12] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X12,sK13,sK12)),s('type.realax.real',X12))))) ) | $spl74),
  inference(cnf_transformation,[],[f890_D])).
fof(f890_D,plain,(
  ( ! [X12] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X12,sK13,sK12)),s('type.realax.real',X12))))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f889,plain,(
  ( ! [X8] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X8)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X8))))) ) | $spl72),
  inference(cnf_transformation,[],[f889_D])).
fof(f889_D,plain,(
  ( ! [X8] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X8)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X8))))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f19562,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X11)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X11)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X12,sK13,sK12)),s('type.realax.real',X12)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X13))))) ) | $spl1),
  inference(subsumption_resolution,[],[f19331,f517])).
fof(f19331,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,'F_0')) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X11)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X11)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X12,sK13,sK12)),s('type.realax.real',X12)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X13))))) ) | $spl1),
  inference(superposition,[],[f343,f5648])).
fof(f888,plain,(
  ( ! [X9] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X9,sK13,sK12)))))) ) | $spl70),
  inference(cnf_transformation,[],[f888_D])).
fof(f888_D,plain,(
  ( ! [X9] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X9,sK13,sK12)))))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f19561,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X8)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X8)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X9,sK13,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X10))))) ) | $spl1),
  inference(subsumption_resolution,[],[f19330,f541])).
fof(f19330,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,'F_0')) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X8)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X8)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X9,sK13,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X10))))) ) | $spl1),
  inference(superposition,[],[f342,f5648])).
fof(f12521,plain,(
  ( ! [X85,X84] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X84),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X84)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X85,sK12)))))) )),
  inference(subsumption_resolution,[],[f12279,f579])).
fof(f12279,plain,(
  ( ! [X85,X84] : (~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X84),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X84)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X85,sK12)))))) )),
  inference(superposition,[],[f373,f715])).
fof(f12520,plain,(
  ( ! [X78,X79] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X78),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X78)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X79,sK12)),s('type.realax.real',X79))))) )),
  inference(subsumption_resolution,[],[f12275,f579])).
fof(f12275,plain,(
  ( ! [X78,X79] : (~p(s(bool,'T_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X78),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X78)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X79,sK12)),s('type.realax.real',X79))))) )),
  inference(superposition,[],[f370,f715])).
fof(f12519,plain,(
  ( ! [X74,X75] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X74),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X74)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X75,sK12)))))) )),
  inference(subsumption_resolution,[],[f12273,f579])).
fof(f12273,plain,(
  ( ! [X74,X75] : (~p(s(bool,'T_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X74),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X74)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X75,sK12)))))) )),
  inference(superposition,[],[f369,f715])).
fof(f12518,plain,(
  ( ! [X64,X65] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(X64,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X65))))) )),
  inference(subsumption_resolution,[],[f12267,f579])).
fof(f12267,plain,(
  ( ! [X64,X65] : (~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(X64,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X65))))) )),
  inference(superposition,[],[f360,f715])).
fof(f12517,plain,(
  ( ! [X62,X63] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X62),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X62)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(X63,sK12)))))) )),
  inference(subsumption_resolution,[],[f12266,f579])).
fof(f12266,plain,(
  ( ! [X62,X63] : (~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X62),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X62)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(X63,sK12)))))) )),
  inference(superposition,[],[f360,f715])).
fof(f12516,plain,(
  ( ! [X57,X56] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X56,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X57))))) )),
  inference(subsumption_resolution,[],[f12263,f579])).
fof(f12263,plain,(
  ( ! [X57,X56] : (~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X56,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X57))))) )),
  inference(superposition,[],[f358,f715])).
fof(f12515,plain,(
  ( ! [X54,X55] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X54),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X54)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X55,sK12)))))) )),
  inference(subsumption_resolution,[],[f12262,f579])).
fof(f12262,plain,(
  ( ! [X54,X55] : (~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X54),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X54)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X55,sK12)))))) )),
  inference(superposition,[],[f358,f715])).
fof(f12514,plain,(
  ( ! [X50,X51] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X50),s('type.realax.real',sK23(X50,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X50),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X51,sK12)))))) )),
  inference(subsumption_resolution,[],[f12260,f579])).
fof(f12260,plain,(
  ( ! [X50,X51] : (~p(s(bool,'T_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X50),s('type.realax.real',sK23(X50,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X50),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X51,sK12)))))) )),
  inference(superposition,[],[f356,f715])).
fof(f12513,plain,(
  ( ! [X47,X46] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X46,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X46),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X47,sK12)))))) )),
  inference(subsumption_resolution,[],[f12258,f579])).
fof(f12258,plain,(
  ( ! [X47,X46] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X46,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X46),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X47,sK12)))))) )),
  inference(superposition,[],[f354,f715])).
fof(f12512,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X37),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X37),s('type.realax.real',X38)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X38)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X39,sK12)))))) )),
  inference(subsumption_resolution,[],[f12255,f579])).
fof(f12255,plain,(
  ( ! [X39,X37,X38] : (~p(s(bool,'T_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X37),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X37),s('type.realax.real',X38)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X38)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X39,sK12)))))) )),
  inference(superposition,[],[f352,f715])).
fof(f12511,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X34),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X34),s('type.realax.real',sK14)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X35,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X36))))) )),
  inference(subsumption_resolution,[],[f12254,f579])).
fof(f12254,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,'T_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X34),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X34),s('type.realax.real',sK14)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X35,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X36))))) )),
  inference(superposition,[],[f352,f715])).
fof(f12510,plain,(
  ( ! [X30,X31] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real',sK21(X30,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X31,sK12)))))) )),
  inference(subsumption_resolution,[],[f12252,f579])).
fof(f12252,plain,(
  ( ! [X30,X31] : (~p(s(bool,'T_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real',sK21(X30,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X31,sK12)))))) )),
  inference(superposition,[],[f350,f715])).
fof(f12509,plain,(
  ( ! [X26,X27] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X26,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X26),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X27,sK12)))))) )),
  inference(subsumption_resolution,[],[f12250,f579])).
fof(f12250,plain,(
  ( ! [X26,X27] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X26,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X26),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X27,sK12)))))) )),
  inference(superposition,[],[f348,f715])).
fof(f12508,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real',X18)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X18)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X19,sK12)))))) )),
  inference(subsumption_resolution,[],[f12247,f579])).
fof(f12247,plain,(
  ( ! [X19,X17,X18] : (~p(s(bool,'T_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real',X18)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X18)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X19,sK12)))))) )),
  inference(superposition,[],[f346,f715])).
fof(f12507,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real',sK14)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X15,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X16))))) )),
  inference(subsumption_resolution,[],[f12246,f579])).
fof(f12246,plain,(
  ( ! [X14,X15,X16] : (~p(s(bool,'T_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real',sK14)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X15,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X16))))) )),
  inference(superposition,[],[f346,f715])).
fof(f738,plain,(
  ( ! [X10,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X10,X9,sK12)),s('type.realax.real',X10)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',sK14))))) ) | $spl14),
  inference(cnf_transformation,[],[f738_D])).
fof(f738_D,plain,(
  ( ! [X10,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X10,X9,sK12)),s('type.realax.real',X10)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',sK14))))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f12504,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',sK14)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X10,X9,sK12)),s('type.realax.real',X10)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X11))))) )),
  inference(subsumption_resolution,[],[f12243,f579])).
fof(f12243,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',sK14)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X10,X9,sK12)),s('type.realax.real',X10)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X11))))) )),
  inference(superposition,[],[f343,f715])).
fof(f737,plain,(
  ( ! [X4,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X4,X3,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',sK14))))) ) | $spl12),
  inference(cnf_transformation,[],[f737_D])).
fof(f737_D,plain,(
  ( ! [X4,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X4,X3,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',sK14))))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f12502,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',sK14)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X4,X3,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X5))))) )),
  inference(subsumption_resolution,[],[f12241,f579])).
fof(f12241,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',sK14)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X4,X3,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X5))))) )),
  inference(superposition,[],[f342,f715])).
fof(f9463,plain,(
  ( ! [X85,X84] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X84),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X84)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X85,sK12)))))) ) | $spl60),
  inference(resolution,[],[f765,f373])).
fof(f5304,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK30(sK12)),s('type.realax.real',sK31(sK12))))) | $spl200),
  inference(cnf_transformation,[],[f5304_D])).
fof(f5304_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK30(sK12)),s('type.realax.real',sK31(sK12))))) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f9462,plain,(
  ( ! [X83,X82] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK30(sK12)),s('type.realax.real',sK31(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X82,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X83))))) ) | $spl60),
  inference(resolution,[],[f765,f373])).
fof(f9459,plain,(
  ( ! [X78,X79] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X78),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X78)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X79,sK12)),s('type.realax.real',X79))))) ) | $spl60),
  inference(resolution,[],[f765,f370])).
fof(f5302,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK30(sK12)),s('type.realax.real',sK28(sK12))))) | $spl198),
  inference(cnf_transformation,[],[f5302_D])).
fof(f5302_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK30(sK12)),s('type.realax.real',sK28(sK12))))) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f9458,plain,(
  ( ! [X76,X77] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK30(sK12)),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X76,sK12)),s('type.realax.real',X76)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X77))))) ) | $spl60),
  inference(resolution,[],[f765,f370])).
fof(f9457,plain,(
  ( ! [X74,X75] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X74),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X74)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X75,sK12)))))) ) | $spl60),
  inference(resolution,[],[f765,f369])).
fof(f9456,plain,(
  ( ! [X72,X73] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK30(sK12)),s('type.realax.real',sK28(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X72,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X73))))) ) | $spl60),
  inference(resolution,[],[f765,f369])).
fof(f9451,plain,(
  ( ! [X64,X65] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK30(sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(X64,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X65))))) ) | $spl60),
  inference(resolution,[],[f765,f360])).
fof(f9450,plain,(
  ( ! [X62,X63] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X62),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X62)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(X63,sK12)))))) ) | $spl60),
  inference(resolution,[],[f765,f360])).
fof(f9447,plain,(
  ( ! [X57,X56] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK30(sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X56,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X57))))) ) | $spl60),
  inference(resolution,[],[f765,f358])).
fof(f9446,plain,(
  ( ! [X54,X55] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X54),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X54)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X55,sK12)))))) ) | $spl60),
  inference(resolution,[],[f765,f358])).
fof(f9444,plain,(
  ( ! [X50,X51] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X50),s('type.realax.real',sK23(X50,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X50),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X51,sK12)))))) ) | $spl60),
  inference(resolution,[],[f765,f356])).
fof(f9442,plain,(
  ( ! [X47,X46] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X46,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X46),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X47,sK12)))))) ) | $spl60),
  inference(resolution,[],[f765,f354])).
fof(f9439,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X37),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X37),s('type.realax.real',X38)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X38)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X39,sK12)))))) ) | $spl60),
  inference(resolution,[],[f765,f352])).
fof(f9438,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X34),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X34),s('type.realax.real',sK30(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X35,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X36))))) ) | $spl60),
  inference(resolution,[],[f765,f352])).
fof(f9436,plain,(
  ( ! [X30,X31] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real',sK21(X30,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X31,sK12)))))) ) | $spl60),
  inference(resolution,[],[f765,f350])).
fof(f9434,plain,(
  ( ! [X26,X27] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X26,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X26),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X27,sK12)))))) ) | $spl60),
  inference(resolution,[],[f765,f348])).
fof(f9431,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real',X18)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X18)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X19,sK12)))))) ) | $spl60),
  inference(resolution,[],[f765,f346])).
fof(f9430,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real',sK30(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X15,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X16))))) ) | $spl60),
  inference(resolution,[],[f765,f346])).
fof(f5297,plain,(
  ( ! [X10,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X10,X9,sK12)),s('type.realax.real',X10)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',sK30(sK12)))))) ) | $spl192),
  inference(cnf_transformation,[],[f5297_D])).
fof(f5297_D,plain,(
  ( ! [X10,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X10,X9,sK12)),s('type.realax.real',X10)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',sK30(sK12)))))) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f9427,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',sK30(sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X10,X9,sK12)),s('type.realax.real',X10)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X11))))) ) | $spl60),
  inference(resolution,[],[f765,f343])).
fof(f5296,plain,(
  ( ! [X4,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X4,X3,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',sK30(sK12)))))) ) | $spl190),
  inference(cnf_transformation,[],[f5296_D])).
fof(f5296_D,plain,(
  ( ! [X4,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X4,X3,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',sK30(sK12)))))) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f9425,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',sK30(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X4,X3,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X5))))) ) | $spl60),
  inference(resolution,[],[f765,f342])).
fof(f765,plain,(
  p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK30(sK12))))) | $spl60),
  inference(cnf_transformation,[],[f765_D])).
fof(f765_D,plain,(
  p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK30(sK12))))) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f7375,plain,(
  ( ! [X80] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK30(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X80,sK12)))))) )),
  inference(resolution,[],[f631,f371])).
fof(f371,plain,(
  ( ! [X6,X4,X0] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',sK30(X0))))) | p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',sK29(X4,X0))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X6))))) )),
  inference(cnf_transformation,[],[f236])).
fof(f7583,plain,(
  ( ! [X80] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK30(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X80,sK12)))))) )),
  inference(resolution,[],[f632,f371])).
fof(f9064,plain,(
  ( ! [X92,X91] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X91),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X91)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X92,sK12)),s('type.realax.real',X92))))) ) | $spl58),
  inference(resolution,[],[f763,f370])).
fof(f9063,plain,(
  ( ! [X90,X88,X89] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK29(X88,sK12)),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X89,sK12)),s('type.realax.real',X89)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X90))))) ) | $spl58),
  inference(resolution,[],[f763,f370])).
fof(f9062,plain,(
  ( ! [X87,X86] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X86),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X86)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X87,sK12)))))) ) | $spl58),
  inference(resolution,[],[f763,f369])).
fof(f9061,plain,(
  ( ! [X85,X83,X84] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK29(X83,sK12)),s('type.realax.real',sK28(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X84,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X85))))) ) | $spl58),
  inference(resolution,[],[f763,f369])).
fof(f9056,plain,(
  ( ! [X74,X75,X73] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK29(X73,sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(X74,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X75))))) ) | $spl58),
  inference(resolution,[],[f763,f360])).
fof(f9055,plain,(
  ( ! [X72,X71] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X71),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X71)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(X72,sK12)))))) ) | $spl58),
  inference(resolution,[],[f763,f360])).
fof(f9052,plain,(
  ( ! [X64,X65,X63] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK29(X63,sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X64,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X65))))) ) | $spl58),
  inference(resolution,[],[f763,f358])).
fof(f9051,plain,(
  ( ! [X61,X62] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X61),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X61)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X62,sK12)))))) ) | $spl58),
  inference(resolution,[],[f763,f358])).
fof(f9049,plain,(
  ( ! [X57,X58] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X57),s('type.realax.real',sK23(X57,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X57),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X58,sK12)))))) ) | $spl58),
  inference(resolution,[],[f763,f356])).
fof(f9047,plain,(
  ( ! [X54,X53] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X53,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X53),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X54,sK12)))))) ) | $spl58),
  inference(resolution,[],[f763,f354])).
fof(f9044,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X43),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X43),s('type.realax.real',X44)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X44)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X45,sK12)))))) ) | $spl58),
  inference(resolution,[],[f763,f352])).
fof(f9043,plain,(
  ( ! [X39,X41,X42,X40] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X39),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X39),s('type.realax.real',sK29(X40,sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X41,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X42))))) ) | $spl58),
  inference(resolution,[],[f763,f352])).
fof(f9041,plain,(
  ( ! [X35,X36] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X35),s('type.realax.real',sK21(X35,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X35),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X36,sK12)))))) ) | $spl58),
  inference(resolution,[],[f763,f350])).
fof(f9039,plain,(
  ( ! [X31,X32] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X31,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X31),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X32,sK12)))))) ) | $spl58),
  inference(resolution,[],[f763,f348])).
fof(f9036,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X21),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X21),s('type.realax.real',X22)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X22)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X23,sK12)))))) ) | $spl58),
  inference(resolution,[],[f763,f346])).
fof(f9035,plain,(
  ( ! [X19,X17,X20,X18] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real',sK29(X18,sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X19,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X20))))) ) | $spl58),
  inference(resolution,[],[f763,f346])).
fof(f9032,plain,(
  ( ! [X14,X12,X13,X11] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK29(X12,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X13,X11,sK12)),s('type.realax.real',X13)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X14))))) ) | $spl58),
  inference(resolution,[],[f763,f343])).
fof(f9030,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',sK29(X5,sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X6,X4,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X7))))) ) | $spl58),
  inference(resolution,[],[f763,f342])).
fof(f8863,plain,(
  ( ! [X99,X98] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X98),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X98)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X99,sK12)))))) ) | $spl48),
  inference(resolution,[],[f756,f373])).
fof(f8977,plain,(
  ( ! [X95] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X95,sK12)),s('type.realax.real',sK31(sK12)))))) ) | $spl284),
  inference(cnf_transformation,[],[f8977_D])).
fof(f8977_D,plain,(
  ( ! [X95] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X95,sK12)),s('type.realax.real',sK31(sK12)))))) ) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f8862,plain,(
  ( ! [X97,X95,X96] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X95,sK12)),s('type.realax.real',sK31(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X96,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X97))))) ) | $spl48),
  inference(resolution,[],[f756,f373])).
fof(f8859,plain,(
  ( ! [X92,X91] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X91),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X91)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X92,sK12)),s('type.realax.real',X92))))) ) | $spl48),
  inference(resolution,[],[f756,f370])).
fof(f8975,plain,(
  ( ! [X88] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK26(X88,sK12)),s('type.realax.real',sK28(sK12)))))) ) | $spl282),
  inference(cnf_transformation,[],[f8975_D])).
fof(f8975_D,plain,(
  ( ! [X88] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK26(X88,sK12)),s('type.realax.real',sK28(sK12)))))) ) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f8858,plain,(
  ( ! [X90,X88,X89] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK26(X88,sK12)),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X89,sK12)),s('type.realax.real',X89)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X90))))) ) | $spl48),
  inference(resolution,[],[f756,f370])).
fof(f8974,plain,(
  ( ! [X63] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X63,sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) | $spl280),
  inference(cnf_transformation,[],[f8974_D])).
fof(f8974_D,plain,(
  ( ! [X63] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X63,sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f8851,plain,(
  ( ! [X74,X75,X73] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X73,sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(X74,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X75))))) ) | $spl48),
  inference(resolution,[],[f756,f360])).
fof(f8850,plain,(
  ( ! [X72,X71] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X71),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X71)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(X72,sK12)))))) ) | $spl48),
  inference(resolution,[],[f756,f360])).
fof(f8847,plain,(
  ( ! [X64,X65,X63] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X63,sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X64,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X65))))) ) | $spl48),
  inference(resolution,[],[f756,f358])).
fof(f8846,plain,(
  ( ! [X61,X62] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X61),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X61)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X62,sK12)))))) ) | $spl48),
  inference(resolution,[],[f756,f358])).
fof(f8844,plain,(
  ( ! [X57,X58] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X57),s('type.realax.real',sK23(X57,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X57),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X58,sK12)))))) ) | $spl48),
  inference(resolution,[],[f756,f356])).
fof(f8842,plain,(
  ( ! [X54,X53] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X53,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X53),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X54,sK12)))))) ) | $spl48),
  inference(resolution,[],[f756,f354])).
fof(f8839,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X43),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X43),s('type.realax.real',X44)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X44)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X45,sK12)))))) ) | $spl48),
  inference(resolution,[],[f756,f352])).
fof(f8838,plain,(
  ( ! [X39,X41,X42,X40] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X39),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X39),s('type.realax.real',sK26(X40,sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X41,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X42))))) ) | $spl48),
  inference(resolution,[],[f756,f352])).
fof(f8836,plain,(
  ( ! [X35,X36] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X35),s('type.realax.real',sK21(X35,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X35),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X36,sK12)))))) ) | $spl48),
  inference(resolution,[],[f756,f350])).
fof(f8834,plain,(
  ( ! [X31,X32] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X31,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X31),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X32,sK12)))))) ) | $spl48),
  inference(resolution,[],[f756,f348])).
fof(f8831,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X21),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X21),s('type.realax.real',X22)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X22)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X23,sK12)))))) ) | $spl48),
  inference(resolution,[],[f756,f346])).
fof(f8830,plain,(
  ( ! [X19,X17,X20,X18] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real',sK26(X18,sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X19,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X20))))) ) | $spl48),
  inference(resolution,[],[f756,f346])).
fof(f8827,plain,(
  ( ! [X14,X12,X13,X11] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK26(X12,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X13,X11,sK12)),s('type.realax.real',X13)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X14))))) ) | $spl48),
  inference(resolution,[],[f756,f343])).
fof(f8825,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',sK26(X5,sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X6,X4,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X7))))) ) | $spl48),
  inference(resolution,[],[f756,f342])).
fof(f8679,plain,(
  ( ! [X99,X98] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X98),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X98)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X99,sK12)))))) ) | $spl44),
  inference(resolution,[],[f754,f373])).
fof(f8678,plain,(
  ( ! [X97,X95,X96] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK25(X95,sK12)),s('type.realax.real',sK31(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X96,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X97))))) ) | $spl44),
  inference(resolution,[],[f754,f373])).
fof(f8675,plain,(
  ( ! [X92,X91] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X91),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X91)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X92,sK12)),s('type.realax.real',X92))))) ) | $spl44),
  inference(resolution,[],[f754,f370])).
fof(f8674,plain,(
  ( ! [X90,X88,X89] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X88,sK12)),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X89,sK12)),s('type.realax.real',X89)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X90))))) ) | $spl44),
  inference(resolution,[],[f754,f370])).
fof(f8673,plain,(
  ( ! [X87,X86] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X86),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X86)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X87,sK12)))))) ) | $spl44),
  inference(resolution,[],[f754,f369])).
fof(f8672,plain,(
  ( ! [X85,X83,X84] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X83,sK12)),s('type.realax.real',sK28(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X84,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X85))))) ) | $spl44),
  inference(resolution,[],[f754,f369])).
fof(f8663,plain,(
  ( ! [X64,X65,X63] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK25(X63,sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X64,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X65))))) ) | $spl44),
  inference(resolution,[],[f754,f358])).
fof(f8662,plain,(
  ( ! [X61,X62] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X61),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X61)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X62,sK12)))))) ) | $spl44),
  inference(resolution,[],[f754,f358])).
fof(f8660,plain,(
  ( ! [X57,X58] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X57),s('type.realax.real',sK23(X57,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X57),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X58,sK12)))))) ) | $spl44),
  inference(resolution,[],[f754,f356])).
fof(f8658,plain,(
  ( ! [X54,X53] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X53,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X53),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X54,sK12)))))) ) | $spl44),
  inference(resolution,[],[f754,f354])).
fof(f8655,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X43),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X43),s('type.realax.real',X44)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X44)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X45,sK12)))))) ) | $spl44),
  inference(resolution,[],[f754,f352])).
fof(f8654,plain,(
  ( ! [X39,X41,X42,X40] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X39),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X39),s('type.realax.real',sK25(X40,sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X41,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X42))))) ) | $spl44),
  inference(resolution,[],[f754,f352])).
fof(f8652,plain,(
  ( ! [X35,X36] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X35),s('type.realax.real',sK21(X35,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X35),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X36,sK12)))))) ) | $spl44),
  inference(resolution,[],[f754,f350])).
fof(f8650,plain,(
  ( ! [X31,X32] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X31,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X31),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X32,sK12)))))) ) | $spl44),
  inference(resolution,[],[f754,f348])).
fof(f8647,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X21),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X21),s('type.realax.real',X22)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X22)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X23,sK12)))))) ) | $spl44),
  inference(resolution,[],[f754,f346])).
fof(f8646,plain,(
  ( ! [X19,X17,X20,X18] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real',sK25(X18,sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X19,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X20))))) ) | $spl44),
  inference(resolution,[],[f754,f346])).
fof(f8643,plain,(
  ( ! [X14,X12,X13,X11] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK25(X12,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X13,X11,sK12)),s('type.realax.real',X13)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X14))))) ) | $spl44),
  inference(resolution,[],[f754,f343])).
fof(f8641,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',sK25(X5,sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X6,X4,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X7))))) ) | $spl44),
  inference(resolution,[],[f754,f342])).
fof(f8280,plain,(
  ( ! [X99,X98] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X98),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X98)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X99,sK12)))))) ) | $spl36),
  inference(resolution,[],[f749,f373])).
fof(f8279,plain,(
  ( ! [X97,X95,X96] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X95,sK12)),s('type.realax.real',sK31(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X96,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X97))))) ) | $spl36),
  inference(resolution,[],[f749,f373])).
fof(f8276,plain,(
  ( ! [X92,X91] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X91),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X91)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X92,sK12)),s('type.realax.real',X92))))) ) | $spl36),
  inference(resolution,[],[f749,f370])).
fof(f8392,plain,(
  ( ! [X83] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK24(X83,sK12)),s('type.realax.real',sK28(sK12)))))) ) | $spl258),
  inference(cnf_transformation,[],[f8392_D])).
fof(f8392_D,plain,(
  ( ! [X83] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK24(X83,sK12)),s('type.realax.real',sK28(sK12)))))) ) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f8275,plain,(
  ( ! [X90,X88,X89] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK24(X88,sK12)),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X89,sK12)),s('type.realax.real',X89)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X90))))) ) | $spl36),
  inference(resolution,[],[f749,f370])).
fof(f8274,plain,(
  ( ! [X87,X86] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X86),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X86)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X87,sK12)))))) ) | $spl36),
  inference(resolution,[],[f749,f369])).
fof(f8273,plain,(
  ( ! [X85,X83,X84] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK24(X83,sK12)),s('type.realax.real',sK28(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X84,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X85))))) ) | $spl36),
  inference(resolution,[],[f749,f369])).
fof(f8268,plain,(
  ( ! [X74,X75,X73] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X73,sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(X74,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X75))))) ) | $spl36),
  inference(resolution,[],[f749,f360])).
fof(f8267,plain,(
  ( ! [X72,X71] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X71),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X71)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(X72,sK12)))))) ) | $spl36),
  inference(resolution,[],[f749,f360])).
fof(f8261,plain,(
  ( ! [X57,X58] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X57),s('type.realax.real',sK23(X57,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X57),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X58,sK12)))))) ) | $spl36),
  inference(resolution,[],[f749,f356])).
fof(f8259,plain,(
  ( ! [X54,X53] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X53,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X53),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X54,sK12)))))) ) | $spl36),
  inference(resolution,[],[f749,f354])).
fof(f8256,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X43),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X43),s('type.realax.real',X44)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X44)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X45,sK12)))))) ) | $spl36),
  inference(resolution,[],[f749,f352])).
fof(f8255,plain,(
  ( ! [X39,X41,X42,X40] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X39),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X39),s('type.realax.real',sK24(X40,sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X41,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X42))))) ) | $spl36),
  inference(resolution,[],[f749,f352])).
fof(f8253,plain,(
  ( ! [X35,X36] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X35),s('type.realax.real',sK21(X35,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X35),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X36,sK12)))))) ) | $spl36),
  inference(resolution,[],[f749,f350])).
fof(f8251,plain,(
  ( ! [X31,X32] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X31,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X31),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X32,sK12)))))) ) | $spl36),
  inference(resolution,[],[f749,f348])).
fof(f8248,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X21),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X21),s('type.realax.real',X22)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X22)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X23,sK12)))))) ) | $spl36),
  inference(resolution,[],[f749,f346])).
fof(f8247,plain,(
  ( ! [X19,X17,X20,X18] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real',sK24(X18,sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X19,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X20))))) ) | $spl36),
  inference(resolution,[],[f749,f346])).
fof(f8244,plain,(
  ( ! [X14,X12,X13,X11] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK24(X12,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X13,X11,sK12)),s('type.realax.real',X13)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X14))))) ) | $spl36),
  inference(resolution,[],[f749,f343])).
fof(f8242,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',sK24(X5,sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X6,X4,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X7))))) ) | $spl36),
  inference(resolution,[],[f749,f342])).
fof(f8073,plain,(
  ( ! [X99,X98] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X98),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X98)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X99,sK12)))))) ) | $spl28),
  inference(resolution,[],[f745,f373])).
fof(f8187,plain,(
  ( ! [X95] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X95,sK12)),s('type.realax.real',sK31(sK12)))))) ) | $spl248),
  inference(cnf_transformation,[],[f8187_D])).
fof(f8187_D,plain,(
  ( ! [X95] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X95,sK12)),s('type.realax.real',sK31(sK12)))))) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f8072,plain,(
  ( ! [X97,X95,X96] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X95,sK12)),s('type.realax.real',sK31(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X96,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X97))))) ) | $spl28),
  inference(resolution,[],[f745,f373])).
fof(f8069,plain,(
  ( ! [X92,X91] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X91),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X91)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X92,sK12)),s('type.realax.real',X92))))) ) | $spl28),
  inference(resolution,[],[f745,f370])).
fof(f8185,plain,(
  ( ! [X83] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK22(X83,sK12)),s('type.realax.real',sK28(sK12)))))) ) | $spl246),
  inference(cnf_transformation,[],[f8185_D])).
fof(f8185_D,plain,(
  ( ! [X83] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK22(X83,sK12)),s('type.realax.real',sK28(sK12)))))) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f8068,plain,(
  ( ! [X90,X88,X89] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK22(X88,sK12)),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X89,sK12)),s('type.realax.real',X89)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X90))))) ) | $spl28),
  inference(resolution,[],[f745,f370])).
fof(f8067,plain,(
  ( ! [X87,X86] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X86),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X86)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X87,sK12)))))) ) | $spl28),
  inference(resolution,[],[f745,f369])).
fof(f8066,plain,(
  ( ! [X85,X83,X84] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK22(X83,sK12)),s('type.realax.real',sK28(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X84,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X85))))) ) | $spl28),
  inference(resolution,[],[f745,f369])).
fof(f8184,plain,(
  ( ! [X63] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X63,sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) | $spl244),
  inference(cnf_transformation,[],[f8184_D])).
fof(f8184_D,plain,(
  ( ! [X63] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X63,sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f8061,plain,(
  ( ! [X74,X75,X73] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X73,sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(X74,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X75))))) ) | $spl28),
  inference(resolution,[],[f745,f360])).
fof(f8060,plain,(
  ( ! [X72,X71] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X71),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X71)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(X72,sK12)))))) ) | $spl28),
  inference(resolution,[],[f745,f360])).
fof(f8057,plain,(
  ( ! [X64,X65,X63] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X63,sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X64,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X65))))) ) | $spl28),
  inference(resolution,[],[f745,f358])).
fof(f8056,plain,(
  ( ! [X61,X62] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X61),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X61)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X62,sK12)))))) ) | $spl28),
  inference(resolution,[],[f745,f358])).
fof(f8046,plain,(
  ( ! [X35,X36] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X35),s('type.realax.real',sK21(X35,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X35),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X36,sK12)))))) ) | $spl28),
  inference(resolution,[],[f745,f350])).
fof(f8044,plain,(
  ( ! [X31,X32] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X31,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X31),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X32,sK12)))))) ) | $spl28),
  inference(resolution,[],[f745,f348])).
fof(f8041,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X21),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X21),s('type.realax.real',X22)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X22)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X23,sK12)))))) ) | $spl28),
  inference(resolution,[],[f745,f346])).
fof(f8183,plain,(
  ( ! [X17,X18] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real',sK22(X18,sK12)))))) ) | $spl242),
  inference(cnf_transformation,[],[f8183_D])).
fof(f8183_D,plain,(
  ( ! [X17,X18] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real',sK22(X18,sK12)))))) ) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f8040,plain,(
  ( ! [X19,X17,X20,X18] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real',sK22(X18,sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X19,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X20))))) ) | $spl28),
  inference(resolution,[],[f745,f346])).
fof(f8182,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X13,X11,sK12)),s('type.realax.real',X13)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK22(X12,sK12)))))) ) | $spl240),
  inference(cnf_transformation,[],[f8182_D])).
fof(f8182_D,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X13,X11,sK12)),s('type.realax.real',X13)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK22(X12,sK12)))))) ) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f8037,plain,(
  ( ! [X14,X12,X13,X11] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK22(X12,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X13,X11,sK12)),s('type.realax.real',X13)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X14))))) ) | $spl28),
  inference(resolution,[],[f745,f343])).
fof(f8181,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X6,X4,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',sK22(X5,sK12)))))) ) | $spl238),
  inference(cnf_transformation,[],[f8181_D])).
fof(f8181_D,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X6,X4,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',sK22(X5,sK12)))))) ) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f8035,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',sK22(X5,sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X6,X4,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X7))))) ) | $spl28),
  inference(resolution,[],[f745,f342])).
fof(f7876,plain,(
  ( ! [X99,X98] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X98),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X98)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X99,sK12)))))) ) | $spl16),
  inference(resolution,[],[f739,f373])).
fof(f7875,plain,(
  ( ! [X97,X95,X96] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK20(X95,sK12)),s('type.realax.real',sK31(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X96,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X97))))) ) | $spl16),
  inference(resolution,[],[f739,f373])).
fof(f7872,plain,(
  ( ! [X92,X91] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X91),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X91)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X92,sK12)),s('type.realax.real',X92))))) ) | $spl16),
  inference(resolution,[],[f739,f370])).
fof(f7871,plain,(
  ( ! [X90,X88,X89] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X88,sK12)),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X89,sK12)),s('type.realax.real',X89)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X90))))) ) | $spl16),
  inference(resolution,[],[f739,f370])).
fof(f7870,plain,(
  ( ! [X87,X86] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X86),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X86)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X87,sK12)))))) ) | $spl16),
  inference(resolution,[],[f739,f369])).
fof(f7869,plain,(
  ( ! [X85,X83,X84] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X83,sK12)),s('type.realax.real',sK28(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X84,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X85))))) ) | $spl16),
  inference(resolution,[],[f739,f369])).
fof(f7864,plain,(
  ( ! [X74,X75,X73] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK20(X73,sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(X74,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X75))))) ) | $spl16),
  inference(resolution,[],[f739,f360])).
fof(f7863,plain,(
  ( ! [X72,X71] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X71),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X71)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(X72,sK12)))))) ) | $spl16),
  inference(resolution,[],[f739,f360])).
fof(f7860,plain,(
  ( ! [X64,X65,X63] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK20(X63,sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X64,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X65))))) ) | $spl16),
  inference(resolution,[],[f739,f358])).
fof(f7859,plain,(
  ( ! [X61,X62] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X61),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X61)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X62,sK12)))))) ) | $spl16),
  inference(resolution,[],[f739,f358])).
fof(f7857,plain,(
  ( ! [X57,X58] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X57),s('type.realax.real',sK23(X57,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X57),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X58,sK12)))))) ) | $spl16),
  inference(resolution,[],[f739,f356])).
fof(f7855,plain,(
  ( ! [X54,X53] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X53,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X53),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X54,sK12)))))) ) | $spl16),
  inference(resolution,[],[f739,f354])).
fof(f7852,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X43),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X43),s('type.realax.real',X44)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X44)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X45,sK12)))))) ) | $spl16),
  inference(resolution,[],[f739,f352])).
fof(f7851,plain,(
  ( ! [X39,X41,X42,X40] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X39),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X39),s('type.realax.real',sK20(X40,sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X41,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X42))))) ) | $spl16),
  inference(resolution,[],[f739,f352])).
fof(f7840,plain,(
  ( ! [X14,X12,X13,X11] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',sK20(X12,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X13,X11,sK12)),s('type.realax.real',X13)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X14))))) ) | $spl16),
  inference(resolution,[],[f739,f343])).
fof(f7838,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',sK20(X5,sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X6,X4,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X7))))) ) | $spl16),
  inference(resolution,[],[f739,f342])).
fof(f7586,plain,(
  ( ! [X85,X84] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X84),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X84)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X85,sK12)))))) )),
  inference(resolution,[],[f632,f373])).
fof(f7703,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK39(sK12,'type.realax.real')),s('type.realax.real',sK31(sK12))))) | $spl224),
  inference(cnf_transformation,[],[f7703_D])).
fof(f7703_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK39(sK12,'type.realax.real')),s('type.realax.real',sK31(sK12))))) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f7585,plain,(
  ( ! [X83,X82] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK39(sK12,'type.realax.real')),s('type.realax.real',sK31(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X82,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X83))))) )),
  inference(resolution,[],[f632,f373])).
fof(f7582,plain,(
  ( ! [X78,X79] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X78),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X78)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X79,sK12)),s('type.realax.real',X79))))) )),
  inference(resolution,[],[f632,f370])).
fof(f7700,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK39(sK12,'type.realax.real')),s('type.realax.real',sK28(sK12))))) | $spl222),
  inference(cnf_transformation,[],[f7700_D])).
fof(f7700_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK39(sK12,'type.realax.real')),s('type.realax.real',sK28(sK12))))) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f7581,plain,(
  ( ! [X76,X77] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK39(sK12,'type.realax.real')),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X76,sK12)),s('type.realax.real',X76)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X77))))) )),
  inference(resolution,[],[f632,f370])).
fof(f7580,plain,(
  ( ! [X74,X75] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X74),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X74)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X75,sK12)))))) )),
  inference(resolution,[],[f632,f369])).
fof(f7579,plain,(
  ( ! [X72,X73] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK39(sK12,'type.realax.real')),s('type.realax.real',sK28(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X72,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X73))))) )),
  inference(resolution,[],[f632,f369])).
fof(f7574,plain,(
  ( ! [X64,X65] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK39(sK12,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(X64,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X65))))) )),
  inference(resolution,[],[f632,f360])).
fof(f7573,plain,(
  ( ! [X62,X63] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X62),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X62)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(X63,sK12)))))) )),
  inference(resolution,[],[f632,f360])).
fof(f7570,plain,(
  ( ! [X57,X56] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK39(sK12,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X56,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X57))))) )),
  inference(resolution,[],[f632,f358])).
fof(f7569,plain,(
  ( ! [X54,X55] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X54),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X54)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X55,sK12)))))) )),
  inference(resolution,[],[f632,f358])).
fof(f7567,plain,(
  ( ! [X50,X51] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X50),s('type.realax.real',sK23(X50,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X50),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X51,sK12)))))) )),
  inference(resolution,[],[f632,f356])).
fof(f7565,plain,(
  ( ! [X47,X46] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X46,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X46),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X47,sK12)))))) )),
  inference(resolution,[],[f632,f354])).
fof(f7562,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X37),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X37),s('type.realax.real',X38)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X38)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X39,sK12)))))) )),
  inference(resolution,[],[f632,f352])).
fof(f7561,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X34),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X34),s('type.realax.real',sK39(sK12,'type.realax.real'))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X35,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X36))))) )),
  inference(resolution,[],[f632,f352])).
fof(f7559,plain,(
  ( ! [X30,X31] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real',sK21(X30,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X31,sK12)))))) )),
  inference(resolution,[],[f632,f350])).
fof(f7557,plain,(
  ( ! [X26,X27] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X26,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X26),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X27,sK12)))))) )),
  inference(resolution,[],[f632,f348])).
fof(f7554,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real',X18)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X18)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X19,sK12)))))) )),
  inference(resolution,[],[f632,f346])).
fof(f7553,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real',sK39(sK12,'type.realax.real'))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X15,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X16))))) )),
  inference(resolution,[],[f632,f346])).
fof(f7695,plain,(
  ( ! [X10,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X10,X9,sK12)),s('type.realax.real',X10)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',sK39(sK12,'type.realax.real')))))) ) | $spl216),
  inference(cnf_transformation,[],[f7695_D])).
fof(f7695_D,plain,(
  ( ! [X10,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X10,X9,sK12)),s('type.realax.real',X10)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',sK39(sK12,'type.realax.real')))))) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f7550,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',sK39(sK12,'type.realax.real'))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X10,X9,sK12)),s('type.realax.real',X10)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X11))))) )),
  inference(resolution,[],[f632,f343])).
fof(f7694,plain,(
  ( ! [X4,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X4,X3,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',sK39(sK12,'type.realax.real')))))) ) | $spl214),
  inference(cnf_transformation,[],[f7694_D])).
fof(f7694_D,plain,(
  ( ! [X4,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X4,X3,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',sK39(sK12,'type.realax.real')))))) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f7548,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',sK39(sK12,'type.realax.real'))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X4,X3,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X5))))) )),
  inference(resolution,[],[f632,f342])).
fof(f7378,plain,(
  ( ! [X85,X84] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X84),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X84)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X85,sK12)))))) )),
  inference(resolution,[],[f631,f373])).
fof(f7495,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK38(sK12,'type.realax.real')),s('type.realax.real',sK31(sK12))))) | $spl212),
  inference(cnf_transformation,[],[f7495_D])).
fof(f7495_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK38(sK12,'type.realax.real')),s('type.realax.real',sK31(sK12))))) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f7377,plain,(
  ( ! [X83,X82] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK38(sK12,'type.realax.real')),s('type.realax.real',sK31(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X82,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X83))))) )),
  inference(resolution,[],[f631,f373])).
fof(f7374,plain,(
  ( ! [X78,X79] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X78),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X78)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X79,sK12)),s('type.realax.real',X79))))) )),
  inference(resolution,[],[f631,f370])).
fof(f7492,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK38(sK12,'type.realax.real')),s('type.realax.real',sK28(sK12))))) | $spl210),
  inference(cnf_transformation,[],[f7492_D])).
fof(f7492_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK38(sK12,'type.realax.real')),s('type.realax.real',sK28(sK12))))) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f7373,plain,(
  ( ! [X76,X77] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK38(sK12,'type.realax.real')),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X76,sK12)),s('type.realax.real',X76)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X77))))) )),
  inference(resolution,[],[f631,f370])).
fof(f7372,plain,(
  ( ! [X74,X75] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X74),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X74)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X75,sK12)))))) )),
  inference(resolution,[],[f631,f369])).
fof(f7371,plain,(
  ( ! [X72,X73] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK38(sK12,'type.realax.real')),s('type.realax.real',sK28(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X72,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X73))))) )),
  inference(resolution,[],[f631,f369])).
fof(f7366,plain,(
  ( ! [X64,X65] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK38(sK12,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(X64,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X65))))) )),
  inference(resolution,[],[f631,f360])).
fof(f7365,plain,(
  ( ! [X62,X63] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X62),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X62)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(X63,sK12)))))) )),
  inference(resolution,[],[f631,f360])).
fof(f7362,plain,(
  ( ! [X57,X56] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK38(sK12,'type.realax.real')),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X56,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X57))))) )),
  inference(resolution,[],[f631,f358])).
fof(f7361,plain,(
  ( ! [X54,X55] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X54),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X54)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X55,sK12)))))) )),
  inference(resolution,[],[f631,f358])).
fof(f7359,plain,(
  ( ! [X50,X51] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X50),s('type.realax.real',sK23(X50,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X50),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X51,sK12)))))) )),
  inference(resolution,[],[f631,f356])).
fof(f7357,plain,(
  ( ! [X47,X46] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X46,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X46),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X47,sK12)))))) )),
  inference(resolution,[],[f631,f354])).
fof(f7354,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X37),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X37),s('type.realax.real',X38)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X38)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X39,sK12)))))) )),
  inference(resolution,[],[f631,f352])).
fof(f7353,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X34),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X34),s('type.realax.real',sK38(sK12,'type.realax.real'))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X35,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X36))))) )),
  inference(resolution,[],[f631,f352])).
fof(f7351,plain,(
  ( ! [X30,X31] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real',sK21(X30,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X31,sK12)))))) )),
  inference(resolution,[],[f631,f350])).
fof(f7349,plain,(
  ( ! [X26,X27] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X26,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X26),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X27,sK12)))))) )),
  inference(resolution,[],[f631,f348])).
fof(f7346,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real',X18)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X18)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X19,sK12)))))) )),
  inference(resolution,[],[f631,f346])).
fof(f7345,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real',sK38(sK12,'type.realax.real'))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X15,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X16))))) )),
  inference(resolution,[],[f631,f346])).
fof(f7487,plain,(
  ( ! [X10,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X10,X9,sK12)),s('type.realax.real',X10)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',sK38(sK12,'type.realax.real')))))) ) | $spl204),
  inference(cnf_transformation,[],[f7487_D])).
fof(f7487_D,plain,(
  ( ! [X10,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X10,X9,sK12)),s('type.realax.real',X10)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',sK38(sK12,'type.realax.real')))))) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f7342,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',sK38(sK12,'type.realax.real'))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X10,X9,sK12)),s('type.realax.real',X10)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X11))))) )),
  inference(resolution,[],[f631,f343])).
fof(f7486,plain,(
  ( ! [X4,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X4,X3,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',sK38(sK12,'type.realax.real')))))) ) | $spl202),
  inference(cnf_transformation,[],[f7486_D])).
fof(f7486_D,plain,(
  ( ! [X4,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X4,X3,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',sK38(sK12,'type.realax.real')))))) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f7340,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',sK38(sK12,'type.realax.real'))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X4,X3,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X5))))) )),
  inference(resolution,[],[f631,f342])).
fof(f6166,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X0,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X1))))) ) | $spl61),
  inference(resolution,[],[f764,f371])).
fof(f764,plain,(
  ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK30(sK12))))) | $spl61),
  inference(cnf_transformation,[],[f764_D])).
fof(f764_D,plain,(
  ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK30(sK12))))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f5853,plain,(
  ( ! [X138,X137] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X137)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK37(sK12,X138,'type.realax.real'))))) | p(s(bool,X138))) ) | $spl8),
  inference(resolution,[],[f588,f505])).
fof(f1526,plain,(
  ( ! [X136] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK36(sK12,X136,'type.realax.real'))))) | p(s(bool,X136))) ) | $spl140),
  inference(cnf_transformation,[],[f1526_D])).
fof(f1526_D,plain,(
  ( ! [X136] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK36(sK12,X136,'type.realax.real'))))) | p(s(bool,X136))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f5852,plain,(
  ( ! [X136,X135] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X135)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK36(sK12,X136,'type.realax.real'))))) | p(s(bool,X136))) ) | $spl8),
  inference(resolution,[],[f588,f503])).
fof(f1525,plain,(
  p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK39(sK12,'type.realax.real'))))) | $spl138),
  inference(cnf_transformation,[],[f1525_D])).
fof(f1525_D,plain,(
  p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK39(sK12,'type.realax.real'))))) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f5851,plain,(
  ( ! [X134] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X134)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK39(sK12,'type.realax.real')))))) ) | $spl8),
  inference(resolution,[],[f588,f509])).
fof(f1523,plain,(
  p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK38(sK12,'type.realax.real'))))) | $spl136),
  inference(cnf_transformation,[],[f1523_D])).
fof(f1523_D,plain,(
  p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK38(sK12,'type.realax.real'))))) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f5850,plain,(
  ( ! [X133] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X133)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK38(sK12,'type.realax.real')))))) ) | $spl8),
  inference(resolution,[],[f588,f508])).
fof(f5847,plain,(
  ( ! [X125,X126,X124] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X124)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X125),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X125)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X126,sK12)))))) ) | $spl8),
  inference(resolution,[],[f588,f373])).
fof(f1521,plain,(
  ( ! [X121] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(X121)),s('type.realax.real',sK31(sK12))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X121))))) ) | $spl134),
  inference(cnf_transformation,[],[f1521_D])).
fof(f1521_D,plain,(
  ( ! [X121] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(X121)),s('type.realax.real',sK31(sK12))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X121))))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f5846,plain,(
  ( ! [X123,X121,X122] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X121)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(X121)),s('type.realax.real',sK31(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X122,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X123))))) ) | $spl8),
  inference(resolution,[],[f588,f373])).
fof(f5844,plain,(
  ( ! [X118,X117] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X117)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK30(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X118,sK12)))))) ) | $spl8),
  inference(resolution,[],[f588,f371])).
fof(f5843,plain,(
  ( ! [X116,X114,X115] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X114)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X115),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X115)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X116,sK12)),s('type.realax.real',X116))))) ) | $spl8),
  inference(resolution,[],[f588,f370])).
fof(f1520,plain,(
  ( ! [X105] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK16(X105)),s('type.realax.real',sK28(sK12))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X105))))) ) | $spl132),
  inference(cnf_transformation,[],[f1520_D])).
fof(f1520_D,plain,(
  ( ! [X105] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK16(X105)),s('type.realax.real',sK28(sK12))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X105))))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f5842,plain,(
  ( ! [X111,X112,X113] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X111)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK16(X111)),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X112,sK12)),s('type.realax.real',X112)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X113))))) ) | $spl8),
  inference(resolution,[],[f588,f370])).
fof(f5841,plain,(
  ( ! [X109,X110,X108] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X108)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X109),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X109)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X110,sK12)))))) ) | $spl8),
  inference(resolution,[],[f588,f369])).
fof(f5840,plain,(
  ( ! [X107,X105,X106] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X105)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK16(X105)),s('type.realax.real',sK28(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X106,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X107))))) ) | $spl8),
  inference(resolution,[],[f588,f369])).
fof(f5835,plain,(
  ( ! [X94,X92,X93] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X92)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(X92)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(X93,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X94))))) ) | $spl8),
  inference(resolution,[],[f588,f360])).
fof(f5834,plain,(
  ( ! [X90,X91,X89] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X89)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X90),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X90)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(X91,sK12)))))) ) | $spl8),
  inference(resolution,[],[f588,f360])).
fof(f5831,plain,(
  ( ! [X80,X81,X82] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X80)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(X80)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X81,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X82))))) ) | $spl8),
  inference(resolution,[],[f588,f358])).
fof(f5830,plain,(
  ( ! [X78,X79,X77] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X77)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X78),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X78)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X79,sK12)))))) ) | $spl8),
  inference(resolution,[],[f588,f358])).
fof(f5828,plain,(
  ( ! [X72,X71,X73] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X71)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X72),s('type.realax.real',sK23(X72,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X72),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X73,sK12)))))) ) | $spl8),
  inference(resolution,[],[f588,f356])).
fof(f5826,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X65)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X66,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X66),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X67,sK12)))))) ) | $spl8),
  inference(resolution,[],[f588,f354])).
fof(f5823,plain,(
  ( ! [X54,X56,X55,X53] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X53)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X54),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X54),s('type.realax.real',X55)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X55)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X56,sK12)))))) ) | $spl8),
  inference(resolution,[],[f588,f352])).
fof(f5822,plain,(
  ( ! [X52,X50,X51,X49] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X49)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X50),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X50),s('type.realax.real',sK16(X49))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X51,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X52))))) ) | $spl8),
  inference(resolution,[],[f588,f352])).
fof(f5820,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X43)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X44),s('type.realax.real',sK21(X44,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X44),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X45,sK12)))))) ) | $spl8),
  inference(resolution,[],[f588,f350])).
fof(f5818,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X37)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X38,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X38),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X39,sK12)))))) ) | $spl8),
  inference(resolution,[],[f588,f348])).
fof(f5815,plain,(
  ( ! [X28,X26,X27,X25] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X25)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X26),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X26),s('type.realax.real',X27)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X27)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X28,sK12)))))) ) | $spl8),
  inference(resolution,[],[f588,f346])).
fof(f5814,plain,(
  ( ! [X24,X23,X21,X22] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X21)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X22),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X22),s('type.realax.real',sK16(X21))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X23,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X24))))) ) | $spl8),
  inference(resolution,[],[f588,f346])).
fof(f5813,plain,(
  ( ! [X19,X20] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X19)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))),s('type.realax.real',X20)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK19(X20,sK12)),s('type.realax.real',X20))))) ) | $spl8),
  inference(resolution,[],[f588,f345])).
fof(f5812,plain,(
  ( ! [X17,X18] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X17)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))),s('type.realax.real',X18)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK19(X18,sK12)))))) ) | $spl8),
  inference(resolution,[],[f588,f344])).
fof(f1515,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X13)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X15,X14,sK12)),s('type.realax.real',X15)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real',sK16(X13)))))) ) | $spl122),
  inference(cnf_transformation,[],[f1515_D])).
fof(f1515_D,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X13)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X15,X14,sK12)),s('type.realax.real',X15)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real',sK16(X13)))))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f5811,plain,(
  ( ! [X14,X15,X13,X16] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X13)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real',sK16(X13))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X15,X14,sK12)),s('type.realax.real',X15)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X16))))) ) | $spl8),
  inference(resolution,[],[f588,f343])).
fof(f5810,plain,(
  ( ! [X12,X10,X11,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X9)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X10),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X10),s('type.realax.real',X11)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X11)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X12,X10,sK12)),s('type.realax.real',X12))))) ) | $spl8),
  inference(resolution,[],[f588,f343])).
fof(f1513,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X5)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X7,X6,sK12))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',sK16(X5)))))) ) | $spl118),
  inference(cnf_transformation,[],[f1513_D])).
fof(f1513_D,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X5)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X7,X6,sK12))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',sK16(X5)))))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f5809,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X5)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',sK16(X5))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X7,X6,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X8))))) ) | $spl8),
  inference(resolution,[],[f588,f342])).
fof(f5808,plain,(
  ( ! [X4,X2,X3,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X1)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X3)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X4,X2,sK12)))))) ) | $spl8),
  inference(resolution,[],[f588,f342])).
fof(f5563,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X11)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X11)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X12,sK13,sK12)),s('type.realax.real',X12)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X13))))) ) | $spl1),
  inference(resolution,[],[f583,f343])).
fof(f5562,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X8)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X8)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X9,sK13,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X10))))) ) | $spl1),
  inference(resolution,[],[f583,f342])).
fof(f887,plain,(
  ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK13)))) | $spl69),
  inference(cnf_transformation,[],[f887_D])).
fof(f887_D,plain,(
  ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK13)))) <=> ~$spl69),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl69])])).
fof(f5560,plain,(
  ( ! [X4,X5] : (~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK13)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X4,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X5))))) ) | $spl1),
  inference(resolution,[],[f583,f358])).
fof(f5558,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK13)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(X0,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X1))))) ) | $spl1),
  inference(resolution,[],[f583,f360])).
fof(f5438,plain,(
  $false | ($spl1 | $spl116)),
  inference(subsumption_resolution,[],[f583,f1512])).
fof(f5309,plain,(
  $false | ($spl3 | $spl116)),
  inference(resolution,[],[f1512,f585])).
fof(f5188,plain,(
  ( ! [X85,X84] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X84),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X84)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X85,sK12)))))) ) | $spl60),
  inference(resolution,[],[f765,f373])).
fof(f5187,plain,(
  ( ! [X83,X82] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK30(sK12)),s('type.realax.real',sK31(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X82,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X83))))) ) | $spl60),
  inference(resolution,[],[f765,f373])).
fof(f5184,plain,(
  ( ! [X78,X79] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X78),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X78)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X79,sK12)),s('type.realax.real',X79))))) ) | $spl60),
  inference(resolution,[],[f765,f370])).
fof(f5183,plain,(
  ( ! [X76,X77] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK30(sK12)),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X76,sK12)),s('type.realax.real',X76)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X77))))) ) | $spl60),
  inference(resolution,[],[f765,f370])).
fof(f5182,plain,(
  ( ! [X74,X75] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X74),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X74)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X75,sK12)))))) ) | $spl60),
  inference(resolution,[],[f765,f369])).
fof(f5181,plain,(
  ( ! [X72,X73] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK30(sK12)),s('type.realax.real',sK28(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X72,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X73))))) ) | $spl60),
  inference(resolution,[],[f765,f369])).
fof(f5176,plain,(
  ( ! [X64,X65] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK30(sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(X64,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X65))))) ) | $spl60),
  inference(resolution,[],[f765,f360])).
fof(f5175,plain,(
  ( ! [X62,X63] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X62),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X62)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(X63,sK12)))))) ) | $spl60),
  inference(resolution,[],[f765,f360])).
fof(f5172,plain,(
  ( ! [X57,X56] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK30(sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X56,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X57))))) ) | $spl60),
  inference(resolution,[],[f765,f358])).
fof(f5171,plain,(
  ( ! [X54,X55] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X54),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X54)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X55,sK12)))))) ) | $spl60),
  inference(resolution,[],[f765,f358])).
fof(f5169,plain,(
  ( ! [X50,X51] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X50),s('type.realax.real',sK23(X50,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X50),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X51,sK12)))))) ) | $spl60),
  inference(resolution,[],[f765,f356])).
fof(f5167,plain,(
  ( ! [X47,X46] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X46,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X46),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X47,sK12)))))) ) | $spl60),
  inference(resolution,[],[f765,f354])).
fof(f5164,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X37),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X37),s('type.realax.real',X38)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X38)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X39,sK12)))))) ) | $spl60),
  inference(resolution,[],[f765,f352])).
fof(f5163,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X34),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X34),s('type.realax.real',sK30(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X35,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X36))))) ) | $spl60),
  inference(resolution,[],[f765,f352])).
fof(f5161,plain,(
  ( ! [X30,X31] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real',sK21(X30,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X31,sK12)))))) ) | $spl60),
  inference(resolution,[],[f765,f350])).
fof(f5159,plain,(
  ( ! [X26,X27] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X26,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X26),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X27,sK12)))))) ) | $spl60),
  inference(resolution,[],[f765,f348])).
fof(f5156,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real',X18)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X18)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X19,sK12)))))) ) | $spl60),
  inference(resolution,[],[f765,f346])).
fof(f5155,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real',sK30(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X15,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X16))))) ) | $spl60),
  inference(resolution,[],[f765,f346])).
fof(f5152,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',sK30(sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X10,X9,sK12)),s('type.realax.real',X10)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X11))))) ) | $spl60),
  inference(resolution,[],[f765,f343])).
fof(f5150,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',sK30(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X4,X3,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X5))))) ) | $spl60),
  inference(resolution,[],[f765,f342])).
fof(f5008,plain,(
  ( ! [X85,X84] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X84),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X84)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X85,sK12)))))) ) | $spl50),
  inference(resolution,[],[f758,f373])).
fof(f3986,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK27(sK12)),s('type.realax.real',sK31(sK12))))) | $spl188),
  inference(cnf_transformation,[],[f3986_D])).
fof(f3986_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK27(sK12)),s('type.realax.real',sK31(sK12))))) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f5007,plain,(
  ( ! [X83,X82] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK27(sK12)),s('type.realax.real',sK31(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X82,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X83))))) ) | $spl50),
  inference(resolution,[],[f758,f373])).
fof(f5005,plain,(
  ( ! [X80] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK30(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X80,sK12)))))) ) | $spl50),
  inference(resolution,[],[f758,f371])).
fof(f5004,plain,(
  ( ! [X78,X79] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X78),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X78)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X79,sK12)),s('type.realax.real',X79))))) ) | $spl50),
  inference(resolution,[],[f758,f370])).
fof(f3984,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK27(sK12)),s('type.realax.real',sK28(sK12))))) | $spl186),
  inference(cnf_transformation,[],[f3984_D])).
fof(f3984_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK27(sK12)),s('type.realax.real',sK28(sK12))))) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f5003,plain,(
  ( ! [X76,X77] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK27(sK12)),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X76,sK12)),s('type.realax.real',X76)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X77))))) ) | $spl50),
  inference(resolution,[],[f758,f370])).
fof(f5002,plain,(
  ( ! [X74,X75] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X74),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X74)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X75,sK12)))))) ) | $spl50),
  inference(resolution,[],[f758,f369])).
fof(f5001,plain,(
  ( ! [X72,X73] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK27(sK12)),s('type.realax.real',sK28(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X72,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X73))))) ) | $spl50),
  inference(resolution,[],[f758,f369])).
fof(f4996,plain,(
  ( ! [X64,X65] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK27(sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(X64,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X65))))) ) | $spl50),
  inference(resolution,[],[f758,f360])).
fof(f4995,plain,(
  ( ! [X62,X63] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X62),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X62)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(X63,sK12)))))) ) | $spl50),
  inference(resolution,[],[f758,f360])).
fof(f4992,plain,(
  ( ! [X57,X56] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK27(sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X56,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X57))))) ) | $spl50),
  inference(resolution,[],[f758,f358])).
fof(f4991,plain,(
  ( ! [X54,X55] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X54),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X54)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X55,sK12)))))) ) | $spl50),
  inference(resolution,[],[f758,f358])).
fof(f4989,plain,(
  ( ! [X50,X51] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X50),s('type.realax.real',sK23(X50,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X50),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X51,sK12)))))) ) | $spl50),
  inference(resolution,[],[f758,f356])).
fof(f4987,plain,(
  ( ! [X47,X46] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X46,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X46),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X47,sK12)))))) ) | $spl50),
  inference(resolution,[],[f758,f354])).
fof(f4984,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X37),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X37),s('type.realax.real',X38)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X38)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X39,sK12)))))) ) | $spl50),
  inference(resolution,[],[f758,f352])).
fof(f4983,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X34),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X34),s('type.realax.real',sK27(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X35,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X36))))) ) | $spl50),
  inference(resolution,[],[f758,f352])).
fof(f4981,plain,(
  ( ! [X30,X31] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real',sK21(X30,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X31,sK12)))))) ) | $spl50),
  inference(resolution,[],[f758,f350])).
fof(f4979,plain,(
  ( ! [X26,X27] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X26,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X26),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X27,sK12)))))) ) | $spl50),
  inference(resolution,[],[f758,f348])).
fof(f4976,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real',X18)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X18)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X19,sK12)))))) ) | $spl50),
  inference(resolution,[],[f758,f346])).
fof(f4975,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real',sK27(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X15,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X16))))) ) | $spl50),
  inference(resolution,[],[f758,f346])).
fof(f3979,plain,(
  ( ! [X10,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X10,X9,sK12)),s('type.realax.real',X10)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',sK27(sK12)))))) ) | $spl180),
  inference(cnf_transformation,[],[f3979_D])).
fof(f3979_D,plain,(
  ( ! [X10,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X10,X9,sK12)),s('type.realax.real',X10)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',sK27(sK12)))))) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f4972,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',sK27(sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X10,X9,sK12)),s('type.realax.real',X10)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X11))))) ) | $spl50),
  inference(resolution,[],[f758,f343])).
fof(f3978,plain,(
  ( ! [X4,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X4,X3,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',sK27(sK12)))))) ) | $spl178),
  inference(cnf_transformation,[],[f3978_D])).
fof(f3978_D,plain,(
  ( ! [X4,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X4,X3,sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',sK27(sK12)))))) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f4970,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',sK27(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X4,X3,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X5))))) ) | $spl50),
  inference(resolution,[],[f758,f342])).
fof(f4724,plain,(
  ( ! [X2,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X1)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X2))))) ) | $spl160),
  inference(resolution,[],[f2039,f370])).
fof(f758,plain,(
  p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK27(sK12))))) | $spl50),
  inference(cnf_transformation,[],[f758_D])).
fof(f758_D,plain,(
  p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK27(sK12))))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f4723,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK27(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X0))))) ) | $spl160),
  inference(resolution,[],[f2039,f368])).
fof(f368,plain,(
  ( ! [X0,X12,X10] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',sK27(X0))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X10,X0)),s('type.realax.real',X10)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X12))))) )),
  inference(cnf_transformation,[],[f236])).
fof(f4000,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK27(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X0))))) ) | $spl92),
  inference(resolution,[],[f1213,f368])).
fof(f4001,plain,(
  ( ! [X2,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X1)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X2))))) ) | $spl92),
  inference(resolution,[],[f1213,f370])).
fof(f3999,plain,(
  $false | ($spl52 | $spl92)),
  inference(resolution,[],[f1213,f759])).
fof(f3870,plain,(
  ( ! [X85,X84] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X84),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X84)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X85,sK12)))))) ) | $spl50),
  inference(resolution,[],[f758,f373])).
fof(f3869,plain,(
  ( ! [X83,X82] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK27(sK12)),s('type.realax.real',sK31(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X82,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X83))))) ) | $spl50),
  inference(resolution,[],[f758,f373])).
fof(f3867,plain,(
  ( ! [X80] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK30(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X80,sK12)))))) ) | $spl50),
  inference(resolution,[],[f758,f371])).
fof(f3864,plain,(
  ( ! [X74,X75] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X74),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X74)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X75,sK12)))))) ) | $spl50),
  inference(resolution,[],[f758,f369])).
fof(f3863,plain,(
  ( ! [X72,X73] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK27(sK12)),s('type.realax.real',sK28(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X72,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X73))))) ) | $spl50),
  inference(resolution,[],[f758,f369])).
fof(f3860,plain,(
  ( ! [X68,X69] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK27(sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X68,sK12)),s('type.realax.real',X68)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X69))))) ) | $spl50),
  inference(resolution,[],[f758,f361])).
fof(f3859,plain,(
  ( ! [X66,X67] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X66),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X66)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X67,sK12)),s('type.realax.real',X67))))) ) | $spl50),
  inference(resolution,[],[f758,f361])).
fof(f3858,plain,(
  ( ! [X64,X65] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK27(sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(X64,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X65))))) ) | $spl50),
  inference(resolution,[],[f758,f360])).
fof(f3857,plain,(
  ( ! [X62,X63] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X62),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X62)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(X63,sK12)))))) ) | $spl50),
  inference(resolution,[],[f758,f360])).
fof(f3856,plain,(
  ( ! [X61,X60] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK27(sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X60,sK12)),s('type.realax.real',X60)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X61))))) ) | $spl50),
  inference(resolution,[],[f758,f359])).
fof(f3855,plain,(
  ( ! [X59,X58] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X58),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X58)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X59,sK12)),s('type.realax.real',X59))))) ) | $spl50),
  inference(resolution,[],[f758,f359])).
fof(f3854,plain,(
  ( ! [X57,X56] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK27(sK12)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X56,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X57))))) ) | $spl50),
  inference(resolution,[],[f758,f358])).
fof(f3853,plain,(
  ( ! [X54,X55] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X54),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X54)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X55,sK12)))))) ) | $spl50),
  inference(resolution,[],[f758,f358])).
fof(f3851,plain,(
  ( ! [X50,X51] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X50),s('type.realax.real',sK23(X50,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X50),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X51,sK12)))))) ) | $spl50),
  inference(resolution,[],[f758,f356])).
fof(f3849,plain,(
  ( ! [X47,X46] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X46,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X46),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X47,sK12)))))) ) | $spl50),
  inference(resolution,[],[f758,f354])).
fof(f3846,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X37),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X37),s('type.realax.real',X38)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X38)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X39,sK12)))))) ) | $spl50),
  inference(resolution,[],[f758,f352])).
fof(f3845,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X34),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X34),s('type.realax.real',sK27(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X35,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X36))))) ) | $spl50),
  inference(resolution,[],[f758,f352])).
fof(f3843,plain,(
  ( ! [X30,X31] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real',sK21(X30,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X31,sK12)))))) ) | $spl50),
  inference(resolution,[],[f758,f350])).
fof(f3841,plain,(
  ( ! [X26,X27] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X26,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X26),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X27,sK12)))))) ) | $spl50),
  inference(resolution,[],[f758,f348])).
fof(f3838,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real',X18)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X18)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X19,sK12)))))) ) | $spl50),
  inference(resolution,[],[f758,f346])).
fof(f3837,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real',sK27(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X15,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X16))))) ) | $spl50),
  inference(resolution,[],[f758,f346])).
fof(f3834,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',sK27(sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X10,X9,sK12)),s('type.realax.real',X10)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X11))))) ) | $spl50),
  inference(resolution,[],[f758,f343])).
fof(f3832,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',sK27(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X4,X3,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X5))))) ) | $spl50),
  inference(resolution,[],[f758,f342])).
fof(f2430,plain,(
  $false | $spl10),
  inference(resolution,[],[f736,f336])).
fof(f2051,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK39(sK12,'type.realax.real')),s('type.realax.real',sK17)))) | $spl176),
  inference(cnf_transformation,[],[f2051_D])).
fof(f2051_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK39(sK12,'type.realax.real')),s('type.realax.real',sK17)))) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f1924,plain,(
  ( ! [X69] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK39(sK12,'type.realax.real')),s('type.realax.real',sK17)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X69))))) ) | $spl4),
  inference(resolution,[],[f586,f509])).
fof(f2049,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK38(sK12,'type.realax.real')),s('type.realax.real',sK17)))) | $spl174),
  inference(cnf_transformation,[],[f2049_D])).
fof(f2049_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK38(sK12,'type.realax.real')),s('type.realax.real',sK17)))) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f1923,plain,(
  ( ! [X68] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK38(sK12,'type.realax.real')),s('type.realax.real',sK17)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X68))))) ) | $spl4),
  inference(resolution,[],[f586,f508])).
fof(f2047,plain,(
  ( ! [X66] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK37(sK12,X66,'type.realax.real')),s('type.realax.real',sK17)))) | p(s(bool,X66))) ) | $spl172),
  inference(cnf_transformation,[],[f2047_D])).
fof(f2047_D,plain,(
  ( ! [X66] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK37(sK12,X66,'type.realax.real')),s('type.realax.real',sK17)))) | p(s(bool,X66))) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f1922,plain,(
  ( ! [X66,X67] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK37(sK12,X66,'type.realax.real')),s('type.realax.real',sK17)))) | p(s(bool,X66)) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X67))))) ) | $spl4),
  inference(resolution,[],[f586,f505])).
fof(f2046,plain,(
  ( ! [X64] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK36(sK12,X64,'type.realax.real')),s('type.realax.real',sK17)))) | p(s(bool,X64))) ) | $spl170),
  inference(cnf_transformation,[],[f2046_D])).
fof(f2046_D,plain,(
  ( ! [X64] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK36(sK12,X64,'type.realax.real')),s('type.realax.real',sK17)))) | p(s(bool,X64))) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f1921,plain,(
  ( ! [X64,X65] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK36(sK12,X64,'type.realax.real')),s('type.realax.real',sK17)))) | p(s(bool,X64)) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X65))))) ) | $spl4),
  inference(resolution,[],[f586,f503])).
fof(f1223,plain,(
  ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK41(sK12,'type.realax.real'))))) | $spl105),
  inference(cnf_transformation,[],[f1223_D])).
fof(f1223_D,plain,(
  ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK41(sK12,'type.realax.real'))))) <=> ~$spl105),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl105])])).
fof(f2045,plain,(
  ( ! [X62] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X62),s('type.realax.real',sK17))))) ) | $spl168),
  inference(cnf_transformation,[],[f2045_D])).
fof(f2045_D,plain,(
  ( ! [X62] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X62),s('type.realax.real',sK17))))) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f1920,plain,(
  ( ! [X63] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X63),s('type.realax.real',sK17)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK41(sK12,'type.realax.real')))))) ) | $spl4),
  inference(resolution,[],[f586,f511])).
fof(f511,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK41(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f306])).
fof(f306,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK40(X1,X0)))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK41(X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK40,sK41])],[f305])).
fof(f305,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f304])).
fof(f304,plain,(
  ! [X0,X1] : ((! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0,X1] : (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(ennf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0,X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(flattening,[],[f114])).
fof(f114,plain,(
  ! [X0] : ! [X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X17] : ! [X10] : (~! [X7] : p(s(bool,happ(s(fun(X17,bool),X10),s(X17,X7)))) <=> ? [X7] : ~p(s(bool,happ(s(fun(X17,bool),X10),s(X17,X7)))))),
  file('thm.real.REAL_LE_SUP','thm.bool.NOT_FORALL_THM')).
fof(f1220,plain,(
  ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK40(sK12,'type.realax.real'))))) | $spl101),
  inference(cnf_transformation,[],[f1220_D])).
fof(f1220_D,plain,(
  ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK40(sK12,'type.realax.real'))))) <=> ~$spl101),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl101])])).
fof(f1919,plain,(
  ( ! [X62] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X62),s('type.realax.real',sK17)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK40(sK12,'type.realax.real')))))) ) | $spl4),
  inference(resolution,[],[f586,f510])).
fof(f510,plain,(
  ( ! [X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK40(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f306])).
fof(f1918,plain,(
  ( ! [X61,X60] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK30(sK12)),s('type.realax.real',sK17)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK29(X60,sK12)),s('type.realax.real',X60)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X61))))) ) | $spl4),
  inference(resolution,[],[f586,f372])).
fof(f1917,plain,(
  ( ! [X59,X58] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK30(sK12)),s('type.realax.real',sK17)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X58,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X59))))) ) | $spl4),
  inference(resolution,[],[f586,f371])).
fof(f2042,plain,(
  ( ! [X53] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK29(X53,sK12)),s('type.realax.real',sK17))))) ) | $spl164),
  inference(cnf_transformation,[],[f2042_D])).
fof(f2042_D,plain,(
  ( ! [X53] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK29(X53,sK12)),s('type.realax.real',sK17))))) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f1916,plain,(
  ( ! [X57,X56,X55] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK29(X55,sK12)),s('type.realax.real',sK17)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X56),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X56)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X57))))) ) | $spl4),
  inference(resolution,[],[f586,f373])).
fof(f1915,plain,(
  ( ! [X54,X53] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK29(X53,sK12)),s('type.realax.real',sK17)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK30(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X54))))) ) | $spl4),
  inference(resolution,[],[f586,f371])).
fof(f2041,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK27(sK12)),s('type.realax.real',sK17)))) | $spl162),
  inference(cnf_transformation,[],[f2041_D])).
fof(f2041_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK27(sK12)),s('type.realax.real',sK17)))) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f1914,plain,(
  ( ! [X52,X51] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK27(sK12)),s('type.realax.real',sK17)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X51,sK12)),s('type.realax.real',X51)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X52))))) ) | $spl4),
  inference(resolution,[],[f586,f368])).
fof(f1913,plain,(
  ( ! [X50,X49] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK27(sK12)),s('type.realax.real',sK17)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X49,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X50))))) ) | $spl4),
  inference(resolution,[],[f586,f367])).
fof(f367,plain,(
  ( ! [X0,X12,X10] : (p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',sK27(X0))))) | p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',sK26(X10,X0))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X0),s('type.realax.real',X12))))) )),
  inference(cnf_transformation,[],[f236])).
fof(f2039,plain,(
  ( ! [X44] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X44,sK12)),s('type.realax.real',sK17))))) ) | $spl160),
  inference(cnf_transformation,[],[f2039_D])).
fof(f2039_D,plain,(
  ( ! [X44] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X44,sK12)),s('type.realax.real',sK17))))) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f1912,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X46,sK12)),s('type.realax.real',sK17)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X47),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X47)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X48))))) ) | $spl4),
  inference(resolution,[],[f586,f369])).
fof(f1911,plain,(
  ( ! [X45,X44] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X44,sK12)),s('type.realax.real',sK17)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK27(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X45))))) ) | $spl4),
  inference(resolution,[],[f586,f367])).
fof(f1910,plain,(
  ( ! [X43,X41,X42] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK25(X41,sK12)),s('type.realax.real',sK17)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X42),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X42)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X43))))) ) | $spl4),
  inference(resolution,[],[f586,f360])).
fof(f1909,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X38,sK12)),s('type.realax.real',sK17)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X39),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X39)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X40))))) ) | $spl4),
  inference(resolution,[],[f586,f358])).
fof(f1908,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK23(X35,sK12)),s('type.realax.real',sK17)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X35),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X36,sK12)),s('type.realax.real',X36)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X37))))) ) | $spl4),
  inference(resolution,[],[f586,f355])).
fof(f1907,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK23(X32,sK12)),s('type.realax.real',sK17)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X32),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X33,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X34))))) ) | $spl4),
  inference(resolution,[],[f586,f354])).
fof(f1906,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X29,sK12)),s('type.realax.real',sK17)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real',sK23(X30,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X31))))) ) | $spl4),
  inference(resolution,[],[f586,f356])).
fof(f1905,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X26,sK12)),s('type.realax.real',sK17)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X27,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X27),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X28))))) ) | $spl4),
  inference(resolution,[],[f586,f354])).
fof(f1904,plain,(
  ( ! [X24,X23,X25,X22] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X22,sK12)),s('type.realax.real',sK17)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X23),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X23),s('type.realax.real',X24)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X24)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X25))))) ) | $spl4),
  inference(resolution,[],[f586,f352])).
fof(f1903,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK21(X19,sK12)),s('type.realax.real',sK17)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X19),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X20,sK12)),s('type.realax.real',X20)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X21))))) ) | $spl4),
  inference(resolution,[],[f586,f349])).
fof(f1902,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK21(X16,sK12)),s('type.realax.real',sK17)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X16),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X17,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X18))))) ) | $spl4),
  inference(resolution,[],[f586,f348])).
fof(f1901,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK20(X13,sK12)),s('type.realax.real',sK17)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real',sK21(X14,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X15))))) ) | $spl4),
  inference(resolution,[],[f586,f350])).
fof(f1900,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK20(X10,sK12)),s('type.realax.real',sK17)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X11,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X11),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X12))))) ) | $spl4),
  inference(resolution,[],[f586,f348])).
fof(f1899,plain,(
  ( ! [X6,X8,X7,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK20(X6,sK12)),s('type.realax.real',sK17)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X7),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X7),s('type.realax.real',X8)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X8)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X9))))) ) | $spl4),
  inference(resolution,[],[f586,f346])).
fof(f1898,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK19(X4,sK12)),s('type.realax.real',sK17)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))),s('type.realax.real',X4)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X5))))) ) | $spl4),
  inference(resolution,[],[f586,f344])).
fof(f1897,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X0,X1,sK12)),s('type.realax.real',sK17)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X2)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X3))))) ) | $spl4),
  inference(resolution,[],[f586,f342])).
fof(f1409,plain,(
  ( ! [X138,X137] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X137)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK37(sK12,X138,'type.realax.real'))))) | p(s(bool,X138))) ) | $spl8),
  inference(resolution,[],[f588,f505])).
fof(f1408,plain,(
  ( ! [X136,X135] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X135)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK36(sK12,X136,'type.realax.real'))))) | p(s(bool,X136))) ) | $spl8),
  inference(resolution,[],[f588,f503])).
fof(f1407,plain,(
  ( ! [X134] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X134)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK39(sK12,'type.realax.real')))))) ) | $spl8),
  inference(resolution,[],[f588,f509])).
fof(f1406,plain,(
  ( ! [X133] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X133)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK38(sK12,'type.realax.real')))))) ) | $spl8),
  inference(resolution,[],[f588,f508])).
fof(f1405,plain,(
  ( ! [X132,X130,X131] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X130)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X131),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X131)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK29(X132,sK12)),s('type.realax.real',X132))))) ) | $spl8),
  inference(resolution,[],[f588,f374])).
fof(f1404,plain,(
  ( ! [X127,X128,X129] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X127)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(X127)),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK29(X128,sK12)),s('type.realax.real',X128)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X129))))) ) | $spl8),
  inference(resolution,[],[f588,f374])).
fof(f1403,plain,(
  ( ! [X125,X126,X124] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X124)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X125),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X125)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X126,sK12)))))) ) | $spl8),
  inference(resolution,[],[f588,f373])).
fof(f1402,plain,(
  ( ! [X123,X121,X122] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X121)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(X121)),s('type.realax.real',sK31(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X122,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X123))))) ) | $spl8),
  inference(resolution,[],[f588,f373])).
fof(f1401,plain,(
  ( ! [X120,X119] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X119)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK30(sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK29(X120,sK12)),s('type.realax.real',X120))))) ) | $spl8),
  inference(resolution,[],[f588,f372])).
fof(f1400,plain,(
  ( ! [X118,X117] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X117)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK30(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X118,sK12)))))) ) | $spl8),
  inference(resolution,[],[f588,f371])).
fof(f1399,plain,(
  ( ! [X116,X114,X115] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X114)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X115),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X115)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X116,sK12)),s('type.realax.real',X116))))) ) | $spl8),
  inference(resolution,[],[f588,f370])).
fof(f1398,plain,(
  ( ! [X111,X112,X113] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X111)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK16(X111)),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X112,sK12)),s('type.realax.real',X112)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X113))))) ) | $spl8),
  inference(resolution,[],[f588,f370])).
fof(f1397,plain,(
  ( ! [X109,X110,X108] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X108)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X109),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X109)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X110,sK12)))))) ) | $spl8),
  inference(resolution,[],[f588,f369])).
fof(f1396,plain,(
  ( ! [X107,X105,X106] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X105)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK16(X105)),s('type.realax.real',sK28(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X106,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X107))))) ) | $spl8),
  inference(resolution,[],[f588,f369])).
fof(f1395,plain,(
  ( ! [X103,X104] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X103)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK27(sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X104,sK12)),s('type.realax.real',X104))))) ) | $spl8),
  inference(resolution,[],[f588,f368])).
fof(f1394,plain,(
  ( ! [X101,X102] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X101)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK27(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X102,sK12)))))) ) | $spl8),
  inference(resolution,[],[f588,f367])).
fof(f1393,plain,(
  ( ! [X99,X100,X98] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X98)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(X98)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X99,sK12)),s('type.realax.real',X99)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X100))))) ) | $spl8),
  inference(resolution,[],[f588,f361])).
fof(f1392,plain,(
  ( ! [X97,X95,X96] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X95)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X96),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X96)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X97,sK12)),s('type.realax.real',X97))))) ) | $spl8),
  inference(resolution,[],[f588,f361])).
fof(f1391,plain,(
  ( ! [X94,X92,X93] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X92)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(X92)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(X93,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X94))))) ) | $spl8),
  inference(resolution,[],[f588,f360])).
fof(f1390,plain,(
  ( ! [X90,X91,X89] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X89)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X90),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X90)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(X91,sK12)))))) ) | $spl8),
  inference(resolution,[],[f588,f360])).
fof(f1389,plain,(
  ( ! [X88,X87,X86] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X86)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(X86)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X87,sK12)),s('type.realax.real',X87)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X88))))) ) | $spl8),
  inference(resolution,[],[f588,f359])).
fof(f1388,plain,(
  ( ! [X85,X83,X84] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X83)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X84),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X84)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X85,sK12)),s('type.realax.real',X85))))) ) | $spl8),
  inference(resolution,[],[f588,f359])).
fof(f1387,plain,(
  ( ! [X80,X81,X82] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X80)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(X80)),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X81,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X82))))) ) | $spl8),
  inference(resolution,[],[f588,f358])).
fof(f1386,plain,(
  ( ! [X78,X79,X77] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X77)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X78),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X78)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X79,sK12)))))) ) | $spl8),
  inference(resolution,[],[f588,f358])).
fof(f1385,plain,(
  ( ! [X76,X74,X75] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X74)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X75),s('type.realax.real',sK23(X75,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X75),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X76,sK12)),s('type.realax.real',X76))))) ) | $spl8),
  inference(resolution,[],[f588,f357])).
fof(f1384,plain,(
  ( ! [X72,X71,X73] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X71)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X72),s('type.realax.real',sK23(X72,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X72),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X73,sK12)))))) ) | $spl8),
  inference(resolution,[],[f588,f356])).
fof(f1383,plain,(
  ( ! [X70,X68,X69] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X68)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X69,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X69),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X70,sK12)),s('type.realax.real',X70))))) ) | $spl8),
  inference(resolution,[],[f588,f355])).
fof(f1382,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X65)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X66,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X66),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X67,sK12)))))) ) | $spl8),
  inference(resolution,[],[f588,f354])).
fof(f1381,plain,(
  ( ! [X61,X64,X62,X63] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X61)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X62),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X62),s('type.realax.real',X63)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X63)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X64,sK12)),s('type.realax.real',X64))))) ) | $spl8),
  inference(resolution,[],[f588,f353])).
fof(f1380,plain,(
  ( ! [X59,X57,X60,X58] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X57)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X58),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X58),s('type.realax.real',sK16(X57))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X59,sK12)),s('type.realax.real',X59)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X60))))) ) | $spl8),
  inference(resolution,[],[f588,f353])).
fof(f1379,plain,(
  ( ! [X54,X56,X55,X53] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X53)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X54),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X54),s('type.realax.real',X55)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X55)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X56,sK12)))))) ) | $spl8),
  inference(resolution,[],[f588,f352])).
fof(f1378,plain,(
  ( ! [X52,X50,X51,X49] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X49)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X50),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X50),s('type.realax.real',sK16(X49))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X51,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X52))))) ) | $spl8),
  inference(resolution,[],[f588,f352])).
fof(f1377,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X46)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X47),s('type.realax.real',sK21(X47,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X47),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X48,sK12)),s('type.realax.real',X48))))) ) | $spl8),
  inference(resolution,[],[f588,f351])).
fof(f1376,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X43)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X44),s('type.realax.real',sK21(X44,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X44),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X45,sK12)))))) ) | $spl8),
  inference(resolution,[],[f588,f350])).
fof(f1375,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X40)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X41,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X41),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X42,sK12)),s('type.realax.real',X42))))) ) | $spl8),
  inference(resolution,[],[f588,f349])).
fof(f1374,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X37)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X38,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X38),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X39,sK12)))))) ) | $spl8),
  inference(resolution,[],[f588,f348])).
fof(f1373,plain,(
  ( ! [X35,X33,X36,X34] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X33)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X34),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X34),s('type.realax.real',X35)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X35)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X36,sK12)),s('type.realax.real',X36))))) ) | $spl8),
  inference(resolution,[],[f588,f347])).
fof(f1372,plain,(
  ( ! [X30,X31,X29,X32] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X29)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real',sK16(X29))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X31,sK12)),s('type.realax.real',X31)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X32))))) ) | $spl8),
  inference(resolution,[],[f588,f347])).
fof(f1371,plain,(
  ( ! [X28,X26,X27,X25] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X25)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X26),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X26),s('type.realax.real',X27)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X27)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X28,sK12)))))) ) | $spl8),
  inference(resolution,[],[f588,f346])).
fof(f1370,plain,(
  ( ! [X24,X23,X21,X22] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X21)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X22),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X22),s('type.realax.real',sK16(X21))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X23,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X24))))) ) | $spl8),
  inference(resolution,[],[f588,f346])).
fof(f1369,plain,(
  ( ! [X19,X20] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X19)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))),s('type.realax.real',X20)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK19(X20,sK12)),s('type.realax.real',X20))))) ) | $spl8),
  inference(resolution,[],[f588,f345])).
fof(f1368,plain,(
  ( ! [X17,X18] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X17)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))),s('type.realax.real',X18)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK19(X18,sK12)))))) ) | $spl8),
  inference(resolution,[],[f588,f344])).
fof(f1367,plain,(
  ( ! [X14,X15,X13,X16] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X13)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real',sK16(X13))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X15,X14,sK12)),s('type.realax.real',X15)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X16))))) ) | $spl8),
  inference(resolution,[],[f588,f343])).
fof(f1366,plain,(
  ( ! [X12,X10,X11,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X9)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X10),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X10),s('type.realax.real',X11)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X11)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X12,X10,sK12)),s('type.realax.real',X12))))) ) | $spl8),
  inference(resolution,[],[f588,f343])).
fof(f1365,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X5)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',sK16(X5))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X7,X6,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X8))))) ) | $spl8),
  inference(resolution,[],[f588,f342])).
fof(f1364,plain,(
  ( ! [X4,X2,X3,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X1)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X3)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X4,X2,sK12)))))) ) | $spl8),
  inference(resolution,[],[f588,f342])).
fof(f1229,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK39(sK12,'type.realax.real')),s('type.realax.real',sK15)))) | $spl112),
  inference(cnf_transformation,[],[f1229_D])).
fof(f1229_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK39(sK12,'type.realax.real')),s('type.realax.real',sK15)))) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f1098,plain,(
  ( ! [X69] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK39(sK12,'type.realax.real')),s('type.realax.real',sK15)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X69))))) )),
  inference(resolution,[],[f337,f509])).
fof(f1227,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK38(sK12,'type.realax.real')),s('type.realax.real',sK15)))) | $spl110),
  inference(cnf_transformation,[],[f1227_D])).
fof(f1227_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK38(sK12,'type.realax.real')),s('type.realax.real',sK15)))) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f1097,plain,(
  ( ! [X68] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK38(sK12,'type.realax.real')),s('type.realax.real',sK15)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X68))))) )),
  inference(resolution,[],[f337,f508])).
fof(f1225,plain,(
  ( ! [X66] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK37(sK12,X66,'type.realax.real')),s('type.realax.real',sK15)))) | p(s(bool,X66))) ) | $spl108),
  inference(cnf_transformation,[],[f1225_D])).
fof(f1225_D,plain,(
  ( ! [X66] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK37(sK12,X66,'type.realax.real')),s('type.realax.real',sK15)))) | p(s(bool,X66))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f1096,plain,(
  ( ! [X66,X67] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK37(sK12,X66,'type.realax.real')),s('type.realax.real',sK15)))) | p(s(bool,X66)) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X67))))) )),
  inference(resolution,[],[f337,f505])).
fof(f1224,plain,(
  ( ! [X64] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK36(sK12,X64,'type.realax.real')),s('type.realax.real',sK15)))) | p(s(bool,X64))) ) | $spl106),
  inference(cnf_transformation,[],[f1224_D])).
fof(f1224_D,plain,(
  ( ! [X64] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK36(sK12,X64,'type.realax.real')),s('type.realax.real',sK15)))) | p(s(bool,X64))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f1095,plain,(
  ( ! [X64,X65] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK36(sK12,X64,'type.realax.real')),s('type.realax.real',sK15)))) | p(s(bool,X64)) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X65))))) )),
  inference(resolution,[],[f337,f503])).
fof(f1094,plain,(
  ( ! [X63] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X63),s('type.realax.real',sK15)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK41(sK12,'type.realax.real')))))) )),
  inference(resolution,[],[f337,f511])).
fof(f1093,plain,(
  ( ! [X62] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X62),s('type.realax.real',sK15)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK40(sK12,'type.realax.real')))))) )),
  inference(resolution,[],[f337,f510])).
fof(f1218,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK30(sK12)),s('type.realax.real',sK15)))) | $spl98),
  inference(cnf_transformation,[],[f1218_D])).
fof(f1218_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK30(sK12)),s('type.realax.real',sK15)))) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f1092,plain,(
  ( ! [X61,X60] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK30(sK12)),s('type.realax.real',sK15)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK29(X60,sK12)),s('type.realax.real',X60)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X61))))) )),
  inference(resolution,[],[f337,f372])).
fof(f1091,plain,(
  ( ! [X59,X58] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK30(sK12)),s('type.realax.real',sK15)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X58,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X59))))) )),
  inference(resolution,[],[f337,f371])).
fof(f1216,plain,(
  ( ! [X53] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK29(X53,sK12)),s('type.realax.real',sK15))))) ) | $spl96),
  inference(cnf_transformation,[],[f1216_D])).
fof(f1216_D,plain,(
  ( ! [X53] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK29(X53,sK12)),s('type.realax.real',sK15))))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f1090,plain,(
  ( ! [X57,X56,X55] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK29(X55,sK12)),s('type.realax.real',sK15)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X56),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X56)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X57))))) )),
  inference(resolution,[],[f337,f373])).
fof(f1089,plain,(
  ( ! [X54,X53] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK29(X53,sK12)),s('type.realax.real',sK15)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK30(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X54))))) )),
  inference(resolution,[],[f337,f371])).
fof(f1215,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK27(sK12)),s('type.realax.real',sK15)))) | $spl94),
  inference(cnf_transformation,[],[f1215_D])).
fof(f1215_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK27(sK12)),s('type.realax.real',sK15)))) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f1088,plain,(
  ( ! [X52,X51] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK27(sK12)),s('type.realax.real',sK15)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X51,sK12)),s('type.realax.real',X51)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X52))))) )),
  inference(resolution,[],[f337,f368])).
fof(f1087,plain,(
  ( ! [X50,X49] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK27(sK12)),s('type.realax.real',sK15)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X49,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X50))))) )),
  inference(resolution,[],[f337,f367])).
fof(f1213,plain,(
  ( ! [X44] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X44,sK12)),s('type.realax.real',sK15))))) ) | $spl92),
  inference(cnf_transformation,[],[f1213_D])).
fof(f1213_D,plain,(
  ( ! [X44] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X44,sK12)),s('type.realax.real',sK15))))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f1086,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X46,sK12)),s('type.realax.real',sK15)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X47),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X47)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X48))))) )),
  inference(resolution,[],[f337,f369])).
fof(f1085,plain,(
  ( ! [X45,X44] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X44,sK12)),s('type.realax.real',sK15)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK27(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X45))))) )),
  inference(resolution,[],[f337,f367])).
fof(f1212,plain,(
  ( ! [X41] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK25(X41,sK12)),s('type.realax.real',sK15))))) ) | $spl90),
  inference(cnf_transformation,[],[f1212_D])).
fof(f1212_D,plain,(
  ( ! [X41] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK25(X41,sK12)),s('type.realax.real',sK15))))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f1084,plain,(
  ( ! [X43,X41,X42] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK25(X41,sK12)),s('type.realax.real',sK15)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X42),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X42)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X43))))) )),
  inference(resolution,[],[f337,f360])).
fof(f1083,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X38,sK12)),s('type.realax.real',sK15)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X39),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X39)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X40))))) )),
  inference(resolution,[],[f337,f358])).
fof(f1210,plain,(
  ( ! [X32] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK23(X32,sK12)),s('type.realax.real',sK15)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X32),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) | $spl86),
  inference(cnf_transformation,[],[f1210_D])).
fof(f1210_D,plain,(
  ( ! [X32] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK23(X32,sK12)),s('type.realax.real',sK15)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X32),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f1082,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK23(X35,sK12)),s('type.realax.real',sK15)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X35),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X36,sK12)),s('type.realax.real',X36)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X37))))) )),
  inference(resolution,[],[f337,f355])).
fof(f1081,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK23(X32,sK12)),s('type.realax.real',sK15)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X32),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X33,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X34))))) )),
  inference(resolution,[],[f337,f354])).
fof(f1209,plain,(
  ( ! [X22] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X22,sK12)),s('type.realax.real',sK15))))) ) | $spl84),
  inference(cnf_transformation,[],[f1209_D])).
fof(f1209_D,plain,(
  ( ! [X22] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X22,sK12)),s('type.realax.real',sK15))))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f1080,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X29,sK12)),s('type.realax.real',sK15)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real',sK23(X30,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X31))))) )),
  inference(resolution,[],[f337,f356])).
fof(f1079,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X26,sK12)),s('type.realax.real',sK15)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X27,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X27),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X28))))) )),
  inference(resolution,[],[f337,f354])).
fof(f1078,plain,(
  ( ! [X24,X23,X25,X22] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X22,sK12)),s('type.realax.real',sK15)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X23),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X23),s('type.realax.real',X24)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X24)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X25))))) )),
  inference(resolution,[],[f337,f352])).
fof(f1208,plain,(
  ( ! [X16] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK21(X16,sK12)),s('type.realax.real',sK15)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X16),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) | $spl82),
  inference(cnf_transformation,[],[f1208_D])).
fof(f1208_D,plain,(
  ( ! [X16] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK21(X16,sK12)),s('type.realax.real',sK15)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X16),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f1077,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK21(X19,sK12)),s('type.realax.real',sK15)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X19),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X20,sK12)),s('type.realax.real',X20)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X21))))) )),
  inference(resolution,[],[f337,f349])).
fof(f1076,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK21(X16,sK12)),s('type.realax.real',sK15)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X16),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X17,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X18))))) )),
  inference(resolution,[],[f337,f348])).
fof(f1207,plain,(
  ( ! [X6] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK20(X6,sK12)),s('type.realax.real',sK15))))) ) | $spl80),
  inference(cnf_transformation,[],[f1207_D])).
fof(f1207_D,plain,(
  ( ! [X6] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK20(X6,sK12)),s('type.realax.real',sK15))))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f1075,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK20(X13,sK12)),s('type.realax.real',sK15)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real',sK21(X14,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X15))))) )),
  inference(resolution,[],[f337,f350])).
fof(f1074,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK20(X10,sK12)),s('type.realax.real',sK15)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X11,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X11),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X12))))) )),
  inference(resolution,[],[f337,f348])).
fof(f1073,plain,(
  ( ! [X6,X8,X7,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK20(X6,sK12)),s('type.realax.real',sK15)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X7),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X7),s('type.realax.real',X8)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X8)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X9))))) )),
  inference(resolution,[],[f337,f346])).
fof(f1206,plain,(
  ( ! [X4] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK19(X4,sK12)),s('type.realax.real',sK15)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))),s('type.realax.real',X4))))) ) | $spl78),
  inference(cnf_transformation,[],[f1206_D])).
fof(f1206_D,plain,(
  ( ! [X4] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK19(X4,sK12)),s('type.realax.real',sK15)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))),s('type.realax.real',X4))))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f1072,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK19(X4,sK12)),s('type.realax.real',sK15)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))),s('type.realax.real',X4)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X5))))) )),
  inference(resolution,[],[f337,f344])).
fof(f1205,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X0,X1,sK12)),s('type.realax.real',sK15)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) | $spl76),
  inference(cnf_transformation,[],[f1205_D])).
fof(f1205_D,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X0,X1,sK12)),s('type.realax.real',sK15)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f1071,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X0,X1,sK12)),s('type.realax.real',sK15)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X2)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X3))))) )),
  inference(resolution,[],[f337,f342])).
fof(f779,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X11)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X11)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X12,sK13,sK12)),s('type.realax.real',X12)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X13))))) ) | $spl1),
  inference(resolution,[],[f583,f343])).
fof(f778,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X8)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X8)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X9,sK13,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X10))))) ) | $spl1),
  inference(resolution,[],[f583,f342])).
fof(f777,plain,(
  ( ! [X6,X7] : (~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK13)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X6,sK12)),s('type.realax.real',X6)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X7))))) ) | $spl1),
  inference(resolution,[],[f583,f359])).
fof(f776,plain,(
  ( ! [X4,X5] : (~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK13)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X4,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X5))))) ) | $spl1),
  inference(resolution,[],[f583,f358])).
fof(f775,plain,(
  ( ! [X2,X3] : (~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK13)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X2,sK12)),s('type.realax.real',X2)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X3))))) ) | $spl1),
  inference(resolution,[],[f583,f361])).
fof(f774,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK13)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(X0,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X1))))) ) | $spl1),
  inference(resolution,[],[f583,f360])).
fof(f630,plain,(
  ( ! [X88,X89] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X88),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X88)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK29(X89,sK12)),s('type.realax.real',X89))))) )),
  inference(resolution,[],[f336,f374])).
fof(f768,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK31(sK12))))) | $spl64),
  inference(cnf_transformation,[],[f768_D])).
fof(f768_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK31(sK12))))) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f629,plain,(
  ( ! [X87,X86] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK29(X86,sK12)),s('type.realax.real',X86)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X87))))) )),
  inference(resolution,[],[f336,f374])).
fof(f628,plain,(
  ( ! [X85,X84] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X84),s('type.realax.real',sK31(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X84)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X85,sK12)))))) )),
  inference(resolution,[],[f336,f373])).
fof(f627,plain,(
  ( ! [X83,X82] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real',sK31(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X82,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X83))))) )),
  inference(resolution,[],[f336,f373])).
fof(f626,plain,(
  ( ! [X81] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK30(sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK29(X81,sK12)),s('type.realax.real',X81))))) )),
  inference(resolution,[],[f336,f372])).
fof(f625,plain,(
  ( ! [X80] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK30(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK29(X80,sK12)))))) )),
  inference(resolution,[],[f336,f371])).
fof(f624,plain,(
  ( ! [X78,X79] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X78),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X78)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X79,sK12)),s('type.realax.real',X79))))) )),
  inference(resolution,[],[f336,f370])).
fof(f761,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK14),s('type.realax.real',sK28(sK12))))) | $spl54),
  inference(cnf_transformation,[],[f761_D])).
fof(f761_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK14),s('type.realax.real',sK28(sK12))))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f623,plain,(
  ( ! [X76,X77] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK14),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X76,sK12)),s('type.realax.real',X76)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X77))))) )),
  inference(resolution,[],[f336,f370])).
fof(f622,plain,(
  ( ! [X74,X75] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X74),s('type.realax.real',sK28(sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X74)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X75,sK12)))))) )),
  inference(resolution,[],[f336,f369])).
fof(f621,plain,(
  ( ! [X72,X73] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK14),s('type.realax.real',sK28(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X72,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X73))))) )),
  inference(resolution,[],[f336,f369])).
fof(f620,plain,(
  ( ! [X71] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK27(sK12))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK26(X71,sK12)),s('type.realax.real',X71))))) )),
  inference(resolution,[],[f336,f368])).
fof(f619,plain,(
  ( ! [X70] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK27(sK12))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK26(X70,sK12)))))) )),
  inference(resolution,[],[f336,f367])).
fof(f618,plain,(
  ( ! [X68,X69] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X68,sK12)),s('type.realax.real',X68)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X69))))) )),
  inference(resolution,[],[f336,f361])).
fof(f617,plain,(
  ( ! [X66,X67] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X66),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X66)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK25(X67,sK12)),s('type.realax.real',X67))))) )),
  inference(resolution,[],[f336,f361])).
fof(f616,plain,(
  ( ! [X64,X65] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(X64,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X65))))) )),
  inference(resolution,[],[f336,f360])).
fof(f615,plain,(
  ( ! [X62,X63] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X62),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X62)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK25(X63,sK12)))))) )),
  inference(resolution,[],[f336,f360])).
fof(f614,plain,(
  ( ! [X61,X60] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X60,sK12)),s('type.realax.real',X60)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X61))))) )),
  inference(resolution,[],[f336,f359])).
fof(f613,plain,(
  ( ! [X59,X58] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X58),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X58)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK24(X59,sK12)),s('type.realax.real',X59))))) )),
  inference(resolution,[],[f336,f359])).
fof(f612,plain,(
  ( ! [X57,X56] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X56,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X57))))) )),
  inference(resolution,[],[f336,f358])).
fof(f611,plain,(
  ( ! [X54,X55] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X54),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X54)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK24(X55,sK12)))))) )),
  inference(resolution,[],[f336,f358])).
fof(f610,plain,(
  ( ! [X52,X53] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X52),s('type.realax.real',sK23(X52,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X52),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X53,sK12)),s('type.realax.real',X53))))) )),
  inference(resolution,[],[f336,f357])).
fof(f609,plain,(
  ( ! [X50,X51] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X50),s('type.realax.real',sK23(X50,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X50),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X51,sK12)))))) )),
  inference(resolution,[],[f336,f356])).
fof(f608,plain,(
  ( ! [X48,X49] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X48,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X48),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X49,sK12)),s('type.realax.real',X49))))) )),
  inference(resolution,[],[f336,f355])).
fof(f607,plain,(
  ( ! [X47,X46] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK23(X46,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X46),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X47,sK12)))))) )),
  inference(resolution,[],[f336,f354])).
fof(f606,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X43),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X43),s('type.realax.real',X44)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X44)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X45,sK12)),s('type.realax.real',X45))))) )),
  inference(resolution,[],[f336,f353])).
fof(f605,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X40),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X40),s('type.realax.real',sK14)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK22(X41,sK12)),s('type.realax.real',X41)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X42))))) )),
  inference(resolution,[],[f336,f353])).
fof(f604,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X37),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X37),s('type.realax.real',X38)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X38)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X39,sK12)))))) )),
  inference(resolution,[],[f336,f352])).
fof(f603,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X34),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X34),s('type.realax.real',sK14)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK22(X35,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X36))))) )),
  inference(resolution,[],[f336,f352])).
fof(f602,plain,(
  ( ! [X33,X32] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X32),s('type.realax.real',sK21(X32,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X32),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X33,sK12)),s('type.realax.real',X33))))) )),
  inference(resolution,[],[f336,f351])).
fof(f601,plain,(
  ( ! [X30,X31] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real',sK21(X30,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X30),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X31,sK12)))))) )),
  inference(resolution,[],[f336,f350])).
fof(f600,plain,(
  ( ! [X28,X29] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X28,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X28),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X29,sK12)),s('type.realax.real',X29))))) )),
  inference(resolution,[],[f336,f349])).
fof(f599,plain,(
  ( ! [X26,X27] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK21(X26,sK12))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X26),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X27,sK12)))))) )),
  inference(resolution,[],[f336,f348])).
fof(f598,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X23),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X23),s('type.realax.real',X24)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X24)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X25,sK12)),s('type.realax.real',X25))))) )),
  inference(resolution,[],[f336,f347])).
fof(f597,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X20),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X20),s('type.realax.real',sK14)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK20(X21,sK12)),s('type.realax.real',X21)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X22))))) )),
  inference(resolution,[],[f336,f347])).
fof(f596,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X17),s('type.realax.real',X18)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X18)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X19,sK12)))))) )),
  inference(resolution,[],[f336,f346])).
fof(f595,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real',sK14)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK20(X15,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X16))))) )),
  inference(resolution,[],[f336,f346])).
fof(f592,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',sK14)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK18(X10,X9,sK12)),s('type.realax.real',X10)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X11))))) )),
  inference(resolution,[],[f336,f343])).
fof(f590,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',sK14)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK18(X4,X3,sK12))))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X5))))) )),
  inference(resolution,[],[f336,f342])).
fof(f588,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK16(X4))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X4))))) ) | $spl8),
  inference(cnf_transformation,[],[f588_D])).
fof(f588_D,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK16(X4))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X4))))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f338,plain,(
  ( ! [X4] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X4)))) | p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',sK16(X4))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) )),
  inference(cnf_transformation,[],[f218])).
fof(f587,plain,(
  ( ! [X4] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(X4)),s('type.realax.real',X4)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X4))))) ) | $spl6),
  inference(cnf_transformation,[],[f587_D])).
fof(f587_D,plain,(
  ( ! [X4] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(X4)),s('type.realax.real',X4)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X4))))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f339,plain,(
  ( ! [X4] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',X4)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(X4)),s('type.realax.real',X4)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) )),
  inference(cnf_transformation,[],[f218])).
fof(f583,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) | $spl1),
  inference(cnf_transformation,[],[f583_D])).
fof(f583_D,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12)))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f586,plain,(
  ( ! [X3] : (~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X3)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',sK17))))) ) | $spl4),
  inference(cnf_transformation,[],[f586_D])).
fof(f586_D,plain,(
  ( ! [X3] : (~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X3)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',sK17))))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f340,plain,(
  ( ! [X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',sK17)))) | ~p(s(bool,happ(s(fun('type.realax.real',bool),sK12),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))) )),
  inference(cnf_transformation,[],[f218])).
fof(f585,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK17)))) | $spl3),
  inference(cnf_transformation,[],[f585_D])).
fof(f585_D,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK17)))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f341,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real',sK17)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK13),s('type.realax.real','const.real.sup_1'(s(fun('type.realax.real',bool),sK12))))))),
  inference(cnf_transformation,[],[f218])).
% SZS output end Proof for thm.real
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 2499
% Passive clauses: 25186
% Generated clauses: 87721
% Final active clauses: 1513
% Final passive clauses: 11894
% Input formulas: 69
% Initial clauses: 240
% 
% Selected by SInE selection: 60
% SInE iterations: 1
% 
% Duplicate literals: 2219
% Trivial inequalities: 276
% Fw subsumption resolutions: 11693
% Fw demodulations: 9018
% Bw demodulations: 66
% 
% Simple tautologies: 4346
% Equational tautologies: 1
% Forward subsumptions: 32721
% 
% Binary resolution: 45344
% Forward superposition: 4870
% Backward superposition: 9100
% Equality resolution: 2
% 
% Split clauses: 1061
% Split components: 2527
% SAT solver clauses: 1127
% SAT solver unit clauses: 13
% SAT solver binary clauses: 164
% 
% Sat splits: 1061
% Sat splitting refutations: 57
% 
% Memory used [KB]: 23155
% Time elapsed: 1.547 s
% ------------------------------
% ------------------------------
% Success in time 13.29 s
