Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.real
% (71959)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1155
% Passive clauses: 6456
% Generated clauses: 32858
% Final active clauses: 1027
% Final passive clauses: 4720
% Input formulas: 117
% Initial clauses: 310
% 
% Duplicate literals: 371
% Trivial inequalities: 76
% Fw subsumption resolutions: 2653
% Fw demodulations: 3210
% Bw demodulations: 82
% 
% Simple tautologies: 273
% Equational tautologies: 53
% Forward subsumptions: 18654
% Fw demodulations to eq. taut.: 3
% 
% Binary resolution: 8277
% Factoring: 1325
% Forward superposition: 6780
% Backward superposition: 9409
% Self superposition: 53
% Equality factoring: 26
% Equality resolution: 16
% 
% Split clauses: 1066
% Split components: 2134
% SAT solver clauses: 1528
% SAT solver unit clauses: 16
% SAT solver binary clauses: 108
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 2
% 
% Sat splits: 1066
% Sat splitting refutations: 25
% 
% TWLsolver clauses: 1522
% TWLsolver calls for satisfiability: 508
% 
% Memory used [KB]: 18038
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.real
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.real
% SZS output start Proof for thm.real
fof(f6886,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f6885,f991,f440,f777,f775,f6795,f6839,f589,f868,f6618,f6627,f6602,f6612,f439,f776,f774,f587,f866])).
fof(f866,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))) | s('type.realax.real',X0) = s('type.realax.real',X1)) ) | $spl130),
  inference(cnf_transformation,[],[f866_D])).
fof(f866_D,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))) | s('type.realax.real',X0) = s('type.realax.real',X1)) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f587,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X0) = s('type.realax.real',X1) | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))) )),
  inference(cnf_transformation,[],[f365])).
fof(f365,plain,(
  ! [X0,X1] : ((s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))) | s('type.realax.real',X0) = s('type.realax.real',X1)) & (s('type.realax.real',X0) != s('type.realax.real',X1) | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))))),
  inference(nnf_transformation,[],[f220])).
fof(f220,plain,(
  ! [X0,X1] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))) <=> s('type.realax.real',X0) = s('type.realax.real',X1))),
  inference(rectify,[],[f103])).
fof(f103,axiom,(
  ! [X13,X14] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X13))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X14))) <=> s('type.realax.real',X13) = s('type.realax.real',X14))),
  file('thm.real.eq_ints|split|3','thm.real.REAL_EQ_NEG')).
fof(f774,plain,(
  s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15))))) | $spl0),
  inference(cnf_transformation,[],[f774_D])).
fof(f774_D,plain,(
  s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f776,plain,(
  s('type.num.num',sK14) = s('type.num.num',sK15) | $spl2),
  inference(cnf_transformation,[],[f776_D])).
fof(f776_D,plain,(
  s('type.num.num',sK14) = s('type.num.num',sK15) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f439,plain,(
  s('type.num.num',sK14) = s('type.num.num',sK15) | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15)))))),
  inference(cnf_transformation,[],[f316])).
fof(f316,plain,(
  (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15))))) | s('type.num.num',sK14) = s('type.num.num',sK15)) & (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15))))) | s('type.num.num',sK14) != s('type.num.num',sK15))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14,sK15])],[f315])).
fof(f315,plain,(
  ? [X0,X1] : ((s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))))) | s('type.num.num',X0) != s('type.num.num',X1)))),
  inference(nnf_transformation,[],[f256])).
fof(f256,plain,(
  ? [X0,X1] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))))) <~> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(ennf_transformation,[],[f119])).
fof(f119,plain,(
  ~! [X0,X1] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X6))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X7))))) <=> s('type.num.num',X6) = s('type.num.num',X7))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X6))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X7))))) <=> s('type.num.num',X6) = s('type.num.num',X7))),
  file('thm.real.eq_ints|split|3','thm.real.eq_ints|split|3')).
fof(f6612,plain,(
  ( ! [X1] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15)))) ) | $spl1740),
  inference(cnf_transformation,[],[f6612_D])).
fof(f6612_D,plain,(
  ( ! [X1] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15)))) ) <=> ~$spl1740),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1740])])).
fof(f6602,plain,(
  ( ! [X1] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15)))) ) | ($spl0 | $spl130)),
  inference(superposition,[],[f866,f774])).
fof(f6627,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15))) | $spl1746),
  inference(cnf_transformation,[],[f6627_D])).
fof(f6627_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15))) <=> ~$spl1746),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1746])])).
fof(f6618,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15))) | $spl1740),
  inference(equality_resolution,[],[f6612])).
fof(f868,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) ) | $spl134),
  inference(cnf_transformation,[],[f868_D])).
fof(f868_D,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f589,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num',X1) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f366])).
fof(f366,plain,(
  ! [X0,X1] : ((s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.num.num',X0) != s('type.num.num',X1) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f102])).
fof(f102,axiom,(
  ! [X7,X6] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X6))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X7))) <=> s('type.num.num',X6) = s('type.num.num',X7))),
  file('thm.real.eq_ints|split|3','thm.real.REAL_INJ')).
fof(f6839,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))) | s('type.num.num',sK15) = s('type.num.num',X0)) ) | $spl1814),
  inference(cnf_transformation,[],[f6839_D])).
fof(f6839_D,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))) | s('type.num.num',sK15) = s('type.num.num',X0)) ) <=> ~$spl1814),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1814])])).
fof(f6795,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))) | s('type.num.num',sK15) = s('type.num.num',X0)) ) | ($spl134 | $spl1746)),
  inference(superposition,[],[f868,f6627])).
fof(f775,plain,(
  s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15))))) | $spl1),
  inference(cnf_transformation,[],[f775_D])).
fof(f775_D,plain,(
  s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f777,plain,(
  s('type.num.num',sK14) != s('type.num.num',sK15) | $spl3),
  inference(cnf_transformation,[],[f777_D])).
fof(f777_D,plain,(
  s('type.num.num',sK14) != s('type.num.num',sK15) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f440,plain,(
  s('type.num.num',sK14) != s('type.num.num',sK15) | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK15)))))),
  inference(cnf_transformation,[],[f316])).
fof(f991,plain,(
  $false | ($spl1 | $spl2)),
  inference(trivial_inequality_removal,[],[f990])).
fof(f990,plain,(
  s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',sK14))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f775,f776])).
fof(f6885,plain,(
  $false | ($spl3 | $spl1814)),
  inference(subsumption_resolution,[],[f6869,f777])).
fof(f6869,plain,(
  s('type.num.num',sK14) = s('type.num.num',sK15) | $spl1814),
  inference(equality_resolution,[],[f6839])).
% SZS output end Proof for thm.real
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 351
% Passive clauses: 1021
% Generated clauses: 6567
% Final active clauses: 317
% Final passive clauses: 467
% Input formulas: 117
% Initial clauses: 310
% 
% Duplicate literals: 86
% Trivial inequalities: 42
% Fw subsumption resolutions: 455
% Fw demodulations: 409
% Bw demodulations: 223
% Condensations: 56
% Global subsumptions: 24
% 
% Simple tautologies: 215
% Equational tautologies: 9
% Forward subsumptions: 3061
% Fw demodulations to eq. taut.: 3
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 1188
% Factoring: 110
% Forward superposition: 1267
% Backward superposition: 1259
% Self superposition: 26
% Equality factoring: 2
% Equality resolution: 15
% 
% Split clauses: 7
% Split components: 14
% SAT solver clauses: 1451
% SAT solver unit clauses: 363
% SAT solver binary clauses: 114
% SAT solver learnt clauses: 50
% SAT solver learnt literals: 15
% 
% Sat splits: 7
% Sat splitting refutations: 25
% 
% TWLsolver clauses: 1353
% TWLsolver calls for satisfiability: 2149
% 
% Memory used [KB]: 5245
% Time elapsed: 0.135 s
% ------------------------------
% ------------------------------
% Success in time 0.752 s
