Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.quotient
% (70347)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 137
% Passive clauses: 3006
% Generated clauses: 16386
% Final active clauses: 117
% Final passive clauses: 2709
% Input formulas: 25
% Initial clauses: 42
% 
% Duplicate literals: 617
% Trivial inequalities: 172
% Fw subsumption resolutions: 1492
% Fw demodulations: 2928
% Bw demodulations: 18
% 
% Simple tautologies: 313
% Equational tautologies: 21
% Forward subsumptions: 7153
% 
% Binary resolution: 144
% Forward superposition: 5657
% Backward superposition: 4949
% Self superposition: 156
% Equality factoring: 45
% Equality resolution: 25
% 
% Split clauses: 364
% Split components: 736
% SAT solver clauses: 739
% SAT solver unit clauses: 11
% SAT solver binary clauses: 64
% SAT solver learnt clauses: 1
% 
% Sat splits: 364
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 737
% TWLsolver calls for satisfiability: 133
% 
% Memory used [KB]: 14200
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.quotient
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.quotient
% SZS output start Proof for thm.quotient
fof(f752,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f751,f87,f135,f169,f172,f85,f131,f92,f139,f741,f742,f89,f137,f157,f158,f119,f150,f123,f154,f187,f192,f96,f141,f165,f167,f160,f161,f117,f148,f118,f149,f120,f151,f162,f164,f86,f129,f116,f147])).
fof(f147,plain,(
  ( ! [X2,X0,X1] : (s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) = s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) ) | $spl20),
  inference(cnf_transformation,[],[f147_D])).
fof(f147_D,plain,(
  ( ! [X2,X0,X1] : (s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) = s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f116,plain,(
  ( ! [X2,X0,X1] : (s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) = s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) )),
  inference(cnf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0,X1,X2] : s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) = s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))),
  inference(flattening,[],[f43])).
fof(f43,plain,(
  ! [X0] : ! [X1,X2] : s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) = s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X9] : ! [X19,X10] : s(bool,happ(s(fun(X9,bool),X19),s(X9,X10))) = s(bool,'const.bool.IN_2'(s(X9,X10),s(fun(X9,bool),X19)))),
  file('thm.quotient.quotient_45_prove_ltb','thm.pred_set.SPECIFICATION')).
fof(f129,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK3,happ(s(fun(sK2,sK3),sK6),s(sK2,sK4))),s(fun(sK3,bool),'const.quotient.respects_1'(s(fun(sK3,fun(sK3,bool)),sK10)))))) | $spl1),
  inference(cnf_transformation,[],[f129_D])).
fof(f129_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK3,happ(s(fun(sK2,sK3),sK6),s(sK2,sK4))),s(fun(sK3,bool),'const.quotient.respects_1'(s(fun(sK3,fun(sK3,bool)),sK10)))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f86,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK3,happ(s(fun(sK2,sK3),sK6),s(sK2,sK4))),s(fun(sK3,bool),'const.quotient.respects_1'(s(fun(sK3,fun(sK3,bool)),sK10))))))),
  inference(cnf_transformation,[],[f67])).
fof(f67,plain,(
  p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK9),s(fun(sK1,sK0),sK7),s(fun(sK0,sK1),sK5)))) & p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK3,fun(sK3,bool)),sK10),s(fun(sK3,sK2),sK8),s(fun(sK2,sK3),sK6)))) & ~p(s(bool,'const.bool.IN_2'(s(sK3,happ(s(fun(sK2,sK3),sK6),s(sK2,sK4))),s(fun(sK3,bool),'const.quotient.respects_1'(s(fun(sK3,fun(sK3,bool)),sK10))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK9,sK10])],[f63])).
fof(f63,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X9),s(fun(X1,X0),X7),s(fun(X0,X1),X5)))) & p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X3,fun(X3,bool)),X10),s(fun(X3,X2),X8),s(fun(X2,X3),X6)))) & ~p(s(bool,'const.bool.IN_2'(s(X3,happ(s(fun(X2,X3),X6),s(X2,X4))),s(fun(X3,bool),'const.quotient.respects_1'(s(fun(X3,fun(X3,bool)),X10)))))))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X9),s(fun(X1,X0),X7),s(fun(X0,X1),X5)))) & (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X3,fun(X3,bool)),X10),s(fun(X3,X2),X8),s(fun(X2,X3),X6)))) & ~p(s(bool,'const.bool.IN_2'(s(X3,happ(s(fun(X2,X3),X6),s(X2,X4))),s(fun(X3,bool),'const.quotient.respects_1'(s(fun(X3,fun(X3,bool)),X10))))))))),
  inference(ennf_transformation,[],[f28])).
fof(f28,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X9),s(fun(X1,X0),X7),s(fun(X0,X1),X5)))) => (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X3,fun(X3,bool)),X10),s(fun(X3,X2),X8),s(fun(X2,X3),X6)))) => p(s(bool,'const.bool.IN_2'(s(X3,happ(s(fun(X2,X3),X6),s(X2,X4))),s(fun(X3,bool),'const.quotient.respects_1'(s(fun(X3,fun(X3,bool)),X10))))))))),
  inference(flattening,[],[f27])).
fof(f27,plain,(
  ~! [X0,X1,X2,X3] : ! [X4,X5,X6,X7,X8,X9,X10] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X9),s(fun(X1,X0),X7),s(fun(X0,X1),X5)))) => (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X3,fun(X3,bool)),X10),s(fun(X3,X2),X8),s(fun(X2,X3),X6)))) => p(s(bool,'const.bool.IN_2'(s(X3,happ(s(fun(X2,X3),X6),s(X2,X4))),s(fun(X3,bool),'const.quotient.respects_1'(s(fun(X3,fun(X3,bool)),X10))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8,X9] : ! [X10,X11,X12,X13,X14,X15,X16] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X7,fun(X7,bool)),X15),s(fun(X7,X6),X13),s(fun(X6,X7),X11)))) => (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X9,fun(X9,bool)),X16),s(fun(X9,X8),X14),s(fun(X8,X9),X12)))) => p(s(bool,'const.bool.IN_2'(s(X9,happ(s(fun(X8,X9),X12),s(X8,X10))),s(fun(X9,bool),'const.quotient.respects_1'(s(fun(X9,fun(X9,bool)),X16))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8,X9] : ! [X10,X11,X12,X13,X14,X15,X16] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X7,fun(X7,bool)),X15),s(fun(X7,X6),X13),s(fun(X6,X7),X11)))) => (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X9,fun(X9,bool)),X16),s(fun(X9,X8),X14),s(fun(X8,X9),X12)))) => p(s(bool,'const.bool.IN_2'(s(X9,happ(s(fun(X8,X9),X12),s(X8,X10))),s(fun(X9,bool),'const.quotient.respects_1'(s(fun(X9,fun(X9,bool)),X16))))))))),
  file('thm.quotient.quotient_45_prove_ltb','thm.quotient.quotient_45_prove_ltb')).
fof(f164,plain,(
  ~p(s(bool,happ(s(fun(sK3,bool),'const.quotient.respects_1'(s(fun(sK3,fun(sK3,bool)),sK10))),s(sK3,happ(s(fun(sK2,sK3),sK6),s(sK2,sK4)))))) | $spl45),
  inference(cnf_transformation,[],[f164_D])).
fof(f164_D,plain,(
  ~p(s(bool,happ(s(fun(sK3,bool),'const.quotient.respects_1'(s(fun(sK3,fun(sK3,bool)),sK10))),s(sK3,happ(s(fun(sK2,sK3),sK6),s(sK2,sK4)))))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f162,plain,(
  ~p(s(bool,happ(s(fun(sK3,bool),'const.quotient.respects_1'(s(fun(sK3,fun(sK3,bool)),sK10))),s(sK3,happ(s(fun(sK2,sK3),sK6),s(sK2,sK4)))))) | ($spl1 | $spl20)),
  inference(forward_demodulation,[],[f129,f147])).
fof(f151,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.combin.W_2'(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,fun(X1,X0)),fun(X1,X0)),'const.combin.W_0'),s(fun(X1,fun(X1,X0)),X2))),s(X1,X3)))) ) | $spl28),
  inference(cnf_transformation,[],[f151_D])).
fof(f151_D,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.combin.W_2'(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,fun(X1,X0)),fun(X1,X0)),'const.combin.W_0'),s(fun(X1,fun(X1,X0)),X2))),s(X1,X3)))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f120,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.combin.W_2'(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,fun(X1,X0)),fun(X1,X0)),'const.combin.W_0'),s(fun(X1,fun(X1,X0)),X2))),s(X1,X3)))) )),
  inference(cnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.combin.W_2'(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,fun(X1,X0)),fun(X1,X0)),'const.combin.W_0'),s(fun(X1,fun(X1,X0)),X2))),s(X1,X3)))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X0,'const.combin.W_2'(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,fun(X1,X0)),fun(X1,X0)),'const.combin.W_0'),s(fun(X1,fun(X1,X0)),X2))),s(X1,X3)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X7,X9] : ! [X26,X27] : s(X7,'const.combin.W_2'(s(fun(X9,fun(X9,X7)),X26),s(X9,X27))) = s(X7,happ(s(fun(X9,X7),happ(s(fun(fun(X9,fun(X9,X7)),fun(X9,X7)),'const.combin.W_0'),s(fun(X9,fun(X9,X7)),X26))),s(X9,X27)))),
  file('thm.quotient.quotient_45_prove_ltb','ARITY_0_2_const.combin.W')).
fof(f149,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,X1),'const.quotient.respects_1'(s(fun(X0,fun(X0,X1)),X2))) = s(fun(X0,X1),happ(s(fun(fun(X0,fun(X0,X1)),fun(X0,X1)),'const.quotient.respects_0'),s(fun(X0,fun(X0,X1)),X2)))) ) | $spl24),
  inference(cnf_transformation,[],[f149_D])).
fof(f149_D,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,X1),'const.quotient.respects_1'(s(fun(X0,fun(X0,X1)),X2))) = s(fun(X0,X1),happ(s(fun(fun(X0,fun(X0,X1)),fun(X0,X1)),'const.quotient.respects_0'),s(fun(X0,fun(X0,X1)),X2)))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f118,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,X1),'const.quotient.respects_1'(s(fun(X0,fun(X0,X1)),X2))) = s(fun(X0,X1),happ(s(fun(fun(X0,fun(X0,X1)),fun(X0,X1)),'const.quotient.respects_0'),s(fun(X0,fun(X0,X1)),X2)))) )),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1,X2] : s(fun(X0,X1),'const.quotient.respects_1'(s(fun(X0,fun(X0,X1)),X2))) = s(fun(X0,X1),happ(s(fun(fun(X0,fun(X0,X1)),fun(X0,X1)),'const.quotient.respects_0'),s(fun(X0,fun(X0,X1)),X2)))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0,X1] : ! [X2] : s(fun(X0,X1),'const.quotient.respects_1'(s(fun(X0,fun(X0,X1)),X2))) = s(fun(X0,X1),happ(s(fun(fun(X0,fun(X0,X1)),fun(X0,X1)),'const.quotient.respects_0'),s(fun(X0,fun(X0,X1)),X2)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X9,X7] : ! [X26] : s(fun(X9,X7),'const.quotient.respects_1'(s(fun(X9,fun(X9,X7)),X26))) = s(fun(X9,X7),happ(s(fun(fun(X9,fun(X9,X7)),fun(X9,X7)),'const.quotient.respects_0'),s(fun(X9,fun(X9,X7)),X26)))),
  file('thm.quotient.quotient_45_prove_ltb','ARITY_0_1_const.quotient.respects')).
fof(f148,plain,(
  ( ! [X0,X1] : (s(fun(fun(X0,fun(X0,X1)),fun(X0,X1)),'const.quotient.respects_0') = s(fun(fun(X0,fun(X0,X1)),fun(X0,X1)),'const.combin.W_0')) ) | $spl22),
  inference(cnf_transformation,[],[f148_D])).
fof(f148_D,plain,(
  ( ! [X0,X1] : (s(fun(fun(X0,fun(X0,X1)),fun(X0,X1)),'const.quotient.respects_0') = s(fun(fun(X0,fun(X0,X1)),fun(X0,X1)),'const.combin.W_0')) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f117,plain,(
  ( ! [X0,X1] : (s(fun(fun(X0,fun(X0,X1)),fun(X0,X1)),'const.quotient.respects_0') = s(fun(fun(X0,fun(X0,X1)),fun(X0,X1)),'const.combin.W_0')) )),
  inference(cnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0,X1] : s(fun(fun(X0,fun(X0,X1)),fun(X0,X1)),'const.quotient.respects_0') = s(fun(fun(X0,fun(X0,X1)),fun(X0,X1)),'const.combin.W_0')),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X9,X7] : s(fun(fun(X9,fun(X9,X7)),fun(X9,X7)),'const.quotient.respects_0') = s(fun(fun(X9,fun(X9,X7)),fun(X9,X7)),'const.combin.W_0')),
  file('thm.quotient.quotient_45_prove_ltb','thm.quotient.respects_def')).
fof(f161,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.combin.W_2'(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),'const.quotient.respects_1'(s(fun(X1,fun(X1,X0)),X2))),s(X1,X3)))) ) | $spl42),
  inference(cnf_transformation,[],[f161_D])).
fof(f161_D,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.combin.W_2'(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),'const.quotient.respects_1'(s(fun(X1,fun(X1,X0)),X2))),s(X1,X3)))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f160,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.combin.W_2'(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),'const.quotient.respects_1'(s(fun(X1,fun(X1,X0)),X2))),s(X1,X3)))) ) | ($spl22 | $spl24 | $spl28)),
  inference(forward_demodulation,[],[f159,f149])).
fof(f159,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.combin.W_2'(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(fun(X1,fun(X1,X0)),fun(X1,X0)),'const.quotient.respects_0'),s(fun(X1,fun(X1,X0)),X2))),s(X1,X3)))) ) | ($spl22 | $spl28)),
  inference(backward_demodulation,[],[f148,f151])).
fof(f167,plain,(
  ~p(s(bool,'const.combin.W_2'(s(fun(sK3,fun(sK3,bool)),sK10),s(sK3,happ(s(fun(sK2,sK3),sK6),s(sK2,sK4)))))) | $spl47),
  inference(cnf_transformation,[],[f167_D])).
fof(f167_D,plain,(
  ~p(s(bool,'const.combin.W_2'(s(fun(sK3,fun(sK3,bool)),sK10),s(sK3,happ(s(fun(sK2,sK3),sK6),s(sK2,sK4)))))) <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f165,plain,(
  ~p(s(bool,'const.combin.W_2'(s(fun(sK3,fun(sK3,bool)),sK10),s(sK3,happ(s(fun(sK2,sK3),sK6),s(sK2,sK4)))))) | ($spl42 | $spl45)),
  inference(backward_demodulation,[],[f161,f164])).
fof(f141,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl16),
  inference(cnf_transformation,[],[f141_D])).
fof(f141_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f96,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f32])).
fof(f32,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X17] : (s(bool,'F_0') = s(bool,X17) <=> ~p(s(bool,X17)))),
  file('thm.quotient.quotient_45_prove_ltb','thm.bool.EQ_CLAUSES|split|2')).
fof(f192,plain,(
  s(bool,'F_0') = s(bool,'const.combin.W_2'(s(fun(sK3,fun(sK3,bool)),sK10),s(sK3,happ(s(fun(sK2,sK3),sK6),s(sK2,sK4))))) | $spl52),
  inference(cnf_transformation,[],[f192_D])).
fof(f192_D,plain,(
  s(bool,'F_0') = s(bool,'const.combin.W_2'(s(fun(sK3,fun(sK3,bool)),sK10),s(sK3,happ(s(fun(sK2,sK3),sK6),s(sK2,sK4))))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f187,plain,(
  s(bool,'F_0') = s(bool,'const.combin.W_2'(s(fun(sK3,fun(sK3,bool)),sK10),s(sK3,happ(s(fun(sK2,sK3),sK6),s(sK2,sK4))))) | ($spl16 | $spl47)),
  inference(resolution,[],[f141,f167])).
fof(f154,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5)))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),X2),s(fun(X0,X1),X3),s(fun(X1,X0),X4))))) ) | $spl34),
  inference(cnf_transformation,[],[f154_D])).
fof(f154_D,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5)))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),X2),s(fun(X0,X1),X3),s(fun(X1,X0),X4))))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f123,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5)))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),X2),s(fun(X0,X1),X3),s(fun(X1,X0),X4))))) )),
  inference(cnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),X2),s(fun(X0,X1),X3),s(fun(X1,X0),X4)))) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5)))))))),
  inference(ennf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),X2),s(fun(X0,X1),X3),s(fun(X1,X0),X4)))) => ! [X5] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5)))))))),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),X2),s(fun(X0,X1),X3),s(fun(X1,X0),X4)))) => ! [X5] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5))))),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5)))))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X9,X7] : ! [X20,X21,X22] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X9,fun(X9,bool)),X20),s(fun(X9,X7),X21),s(fun(X7,X9),X22)))) => ! [X23] : p(s(bool,happ(s(fun(X9,bool),happ(s(fun(X9,fun(X9,bool)),X20),s(X9,happ(s(fun(X7,X9),X22),s(X7,X23))))),s(X9,happ(s(fun(X7,X9),X22),s(X7,X23)))))))),
  file('thm.quotient.quotient_45_prove_ltb','thm.quotient.QUOTIENT_REP_REFL')).
fof(f150,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.combin.W_2'(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))),s(X1,X3)))) ) | $spl26),
  inference(cnf_transformation,[],[f150_D])).
fof(f150_D,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.combin.W_2'(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))),s(X1,X3)))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f119,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.combin.W_2'(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))),s(X1,X3)))) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.combin.W_2'(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))),s(X1,X3)))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X0,'const.combin.W_2'(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),happ(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))),s(X1,X3)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X7,X9] : ! [X18,X10] : s(X7,'const.combin.W_2'(s(fun(X9,fun(X9,X7)),X18),s(X9,X10))) = s(X7,happ(s(fun(X9,X7),happ(s(fun(X9,fun(X9,X7)),X18),s(X9,X10))),s(X9,X10)))),
  file('thm.quotient.quotient_45_prove_ltb','thm.combin.W_THM')).
fof(f158,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.combin.W_2'(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5)))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),X2),s(fun(X0,X1),X3),s(fun(X1,X0),X4))))) ) | $spl40),
  inference(cnf_transformation,[],[f158_D])).
fof(f158_D,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.combin.W_2'(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5)))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),X2),s(fun(X0,X1),X3),s(fun(X1,X0),X4))))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f157,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.combin.W_2'(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(X1,X0),X4),s(X1,X5)))))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X0,fun(X0,bool)),X2),s(fun(X0,X1),X3),s(fun(X1,X0),X4))))) ) | ($spl26 | $spl34)),
  inference(backward_demodulation,[],[f150,f154])).
fof(f137,plain,(
  ~p(s(bool,'F_0')) | $spl9),
  inference(cnf_transformation,[],[f137_D])).
fof(f137_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f89,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f29])).
fof(f29,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.quotient.quotient_45_prove_ltb',aHL_FALSITY)).
fof(f742,plain,(
  ( ! [X0] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK3,fun(sK3,bool)),sK10),s(fun(sK3,sK2),X0),s(fun(sK2,sK3),sK6))))) ) | $spl156),
  inference(cnf_transformation,[],[f742_D])).
fof(f742_D,plain,(
  ( ! [X0] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK3,fun(sK3,bool)),sK10),s(fun(sK3,sK2),X0),s(fun(sK2,sK3),sK6))))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f741,plain,(
  ( ! [X0] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK3,fun(sK3,bool)),sK10),s(fun(sK3,sK2),X0),s(fun(sK2,sK3),sK6))))) ) | ($spl9 | $spl40 | $spl52)),
  inference(subsumption_resolution,[],[f724,f137])).
fof(f724,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK3,fun(sK3,bool)),sK10),s(fun(sK3,sK2),X0),s(fun(sK2,sK3),sK6))))) ) | ($spl40 | $spl52)),
  inference(superposition,[],[f158,f192])).
fof(f139,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(cnf_transformation,[],[f139_D])).
fof(f139_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f92,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X17] : (s(bool,'T_0') = s(bool,X17) <=> p(s(bool,X17)))),
  file('thm.quotient.quotient_45_prove_ltb','thm.bool.EQ_CLAUSES|split|1')).
fof(f131,plain,(
  p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK3,fun(sK3,bool)),sK10),s(fun(sK3,sK2),sK8),s(fun(sK2,sK3),sK6)))) | $spl2),
  inference(cnf_transformation,[],[f131_D])).
fof(f131_D,plain,(
  p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK3,fun(sK3,bool)),sK10),s(fun(sK3,sK2),sK8),s(fun(sK2,sK3),sK6)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f85,plain,(
  p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK3,fun(sK3,bool)),sK10),s(fun(sK3,sK2),sK8),s(fun(sK2,sK3),sK6))))),
  inference(cnf_transformation,[],[f67])).
fof(f172,plain,(
  s(bool,'T_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK3,fun(sK3,bool)),sK10),s(fun(sK3,sK2),sK8),s(fun(sK2,sK3),sK6))) | $spl48),
  inference(cnf_transformation,[],[f172_D])).
fof(f172_D,plain,(
  s(bool,'T_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK3,fun(sK3,bool)),sK10),s(fun(sK3,sK2),sK8),s(fun(sK2,sK3),sK6))) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f169,plain,(
  s(bool,'T_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK3,fun(sK3,bool)),sK10),s(fun(sK3,sK2),sK8),s(fun(sK2,sK3),sK6))) | ($spl2 | $spl12)),
  inference(resolution,[],[f139,f131])).
fof(f135,plain,(
  p(s(bool,'T_0')) | $spl6),
  inference(cnf_transformation,[],[f135_D])).
fof(f135_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f87,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.quotient.quotient_45_prove_ltb','thm.bool.TRUTH')).
fof(f751,plain,(
  $false | ($spl6 | $spl48 | $spl156)),
  inference(subsumption_resolution,[],[f748,f135])).
fof(f748,plain,(
  ~p(s(bool,'T_0')) | ($spl48 | $spl156)),
  inference(superposition,[],[f742,f172])).
% SZS output end Proof for thm.quotient
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 42
% Passive clauses: 94
% Generated clauses: 684
% Final active clauses: 29
% Final passive clauses: 33
% Input formulas: 25
% Initial clauses: 42
% 
% Duplicate literals: 6
% Trivial inequalities: 14
% Fw subsumption resolutions: 25
% Fw demodulations: 105
% Bw demodulations: 23
% Condensations: 6
% Global subsumptions: 4
% 
% Simple tautologies: 19
% Equational tautologies: 5
% Forward subsumptions: 294
% 
% Binary resolution: 15
% Forward superposition: 217
% Backward superposition: 97
% Self superposition: 12
% Equality factoring: 2
% Equality resolution: 6
% 
% Split clauses: 4
% Split components: 8
% SAT solver clauses: 157
% SAT solver unit clauses: 56
% SAT solver binary clauses: 22
% SAT solver learnt clauses: 9
% 
% Sat splits: 4
% Sat splitting refutations: 8
% 
% TWLsolver clauses: 140
% TWLsolver calls for satisfiability: 127
% 
% Memory used [KB]: 1535
% Time elapsed: 0.027 s
% ------------------------------
% ------------------------------
% Success in time 0.642 s
