Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.integer
% (51584)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1273
% Passive clauses: 8661
% Generated clauses: 35221
% Final active clauses: 1029
% Final passive clauses: 6793
% Input formulas: 261
% Initial clauses: 558
% 
% Duplicate literals: 299
% Trivial inequalities: 97
% Fw subsumption resolutions: 2182
% Fw demodulations: 4199
% Bw demodulations: 236
% 
% Simple tautologies: 357
% Equational tautologies: 36
% Forward subsumptions: 19116
% 
% Binary resolution: 9563
% Factoring: 1109
% Forward superposition: 9298
% Backward superposition: 7332
% Self superposition: 45
% Equality factoring: 2
% Equality resolution: 23
% 
% Split clauses: 498
% Split components: 998
% SAT solver clauses: 750
% SAT solver unit clauses: 17
% SAT solver binary clauses: 84
% 
% Sat splits: 498
% Sat splitting refutations: 8
% 
% TWLsolver clauses: 747
% TWLsolver calls for satisfiability: 491
% 
% Memory used [KB]: 15351
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.integer
% (51599)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 698
% Passive clauses: 4488
% Generated clauses: 24177
% Final active clauses: 601
% Final passive clauses: 3411
% Input formulas: 261
% Initial clauses: 558
% 
% Duplicate literals: 109
% Trivial inequalities: 41
% Fw subsumption resolutions: 897
% Fw demodulations: 3274
% Bw demodulations: 444
% Condensations: 124
% Global subsumptions: 24
% 
% Simple tautologies: 330
% Equational tautologies: 17
% Forward subsumptions: 9733
% Fw demodulations to eq. taut.: 10
% 
% Binary resolution: 3436
% Factoring: 172
% Forward superposition: 6298
% Backward superposition: 3988
% Self superposition: 63
% Equality factoring: 2
% Equality resolution: 20
% 
% Split clauses: 50
% Split components: 100
% SAT solver clauses: 5617
% SAT solver unit clauses: 652
% SAT solver binary clauses: 179
% SAT solver learnt clauses: 49
% SAT solver learnt literals: 14
% 
% Sat splits: 50
% Sat splitting refutations: 28
% 
% TWLsolver clauses: 5520
% TWLsolver calls for satisfiability: 4339
% 
% Memory used [KB]: 22771
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.integer
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.integer
% SZS output start Proof for thm.integer
fof(f41244,plain,(
  $false),
  inference(subsumption_resolution,[],[f41230,f1525])).
fof(f1525,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_|48|_0')))),
  inference(forward_demodulation,[],[f951,f945])).
fof(f945,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),
  inference(cnf_transformation,[],[f184])).
fof(f184,axiom,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),
  file('thm.integer.INT_MOD_|49|','thm.integer.INT_|48|')).
fof(f951,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(cnf_transformation,[],[f222])).
fof(f222,axiom,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.integer.INT_MOD_|49|','thm.integer.INT_NEG_|48|')).
fof(f41230,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_|48|_0')))),
  inference(backward_demodulation,[],[f41229,f7133])).
fof(f7133,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_|49|_0')))))),
  inference(unit_resulting_resolution,[],[f1527,f1620])).
fof(f1620,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',X0)) )),
  inference(forward_demodulation,[],[f1619,f945])).
fof(f1619,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',X0) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))) )),
  inference(forward_demodulation,[],[f1071,f945])).
fof(f1071,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f737])).
fof(f737,plain,(
  ! [X0] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))),
  inference(nnf_transformation,[],[f365])).
fof(f365,plain,(
  ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) <=> s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f218])).
fof(f218,axiom,(
  ! [X11] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X11))) <=> s('type.integer.int',X11) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.integer.INT_MOD_|49|','thm.integer.INT_NEG_EQ|48|')).
fof(f1527,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_|49|_0')))),
  inference(backward_demodulation,[],[f1526,f1524])).
fof(f1524,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(backward_demodulation,[],[f950,f1523])).
fof(f1523,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(backward_demodulation,[],[f945,f941])).
fof(f941,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(cnf_transformation,[],[f710])).
fof(f710,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',sK15),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK15])],[f602])).
fof(f602,plain,(
  ? [X0] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(ennf_transformation,[],[f263])).
fof(f263,plain,(
  ~! [X0] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X6),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X6),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  file('thm.integer.INT_MOD_|49|','thm.integer.INT_MOD_|49|')).
fof(f950,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f162])).
fof(f162,axiom,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  file('thm.integer.INT_MOD_|49|','thm.numeral.numeral_distrib|split|13')).
fof(f1526,plain,(
  s('type.integer.int','const.integer.int_|49|_0') = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f953,f950])).
fof(f953,plain,(
  s('type.integer.int','const.integer.int_|49|_0') = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(cnf_transformation,[],[f11])).
fof(f11,axiom,(
  s('type.integer.int','const.integer.int_|49|_0') = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  file('thm.integer.INT_MOD_|49|','thm.integer.INT_|49|')).
fof(f41229,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_|49|_0')))) )),
  inference(forward_demodulation,[],[f41061,f1574])).
fof(f1574,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_|49|_0')))) )),
  inference(forward_demodulation,[],[f1573,f1526])).
fof(f1573,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(forward_demodulation,[],[f1011,f950])).
fof(f1011,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(cnf_transformation,[],[f315])).
fof(f315,plain,(
  ! [X0] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X11] : s('type.integer.int',X11) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X11),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  file('thm.integer.INT_MOD_|49|','thm.integer.INT_MUL_RID')).
fof(f41061,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_|49|_0'))),s('type.integer.int','const.integer.int_|49|_0')))) )),
  inference(unit_resulting_resolution,[],[f6782,f1599])).
fof(f1599,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))),s('type.integer.int',X0))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',X0)) )),
  inference(forward_demodulation,[],[f1598,f945])).
fof(f1598,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))),s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f1046,f945])).
fof(f1046,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))),s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f623])).
fof(f623,plain,(
  ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ! [X1] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))),s('type.integer.int',X0))))),
  inference(ennf_transformation,[],[f351])).
fof(f351,plain,(
  ! [X0] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) => ! [X1] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))),s('type.integer.int',X0))))),
  inference(flattening,[],[f350])).
fof(f350,plain,(
  ! [X0] : (~s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) => ! [X1] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X0))),s('type.integer.int',X0))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X16] : (~s('type.integer.int',X16) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) => ! [X8] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X8),s('type.integer.int',X16))),s('type.integer.int',X16))))),
  file('thm.integer.INT_MOD_|49|','thm.integer.INT_MOD_COMMON_FACTOR')).
fof(f6782,plain,(
  s('type.integer.int','const.integer.int_|49|_0') != s('type.integer.int','const.integer.int_|48|_0')),
  inference(subsumption_resolution,[],[f6712,f942])).
fof(f942,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f9])).
fof(f9,axiom,(
  p(s(bool,'T_0'))),
  file('thm.integer.INT_MOD_|49|','thm.bool.TRUTH')).
fof(f6712,plain,(
  ~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_|49|_0') != s('type.integer.int','const.integer.int_|48|_0')),
  inference(superposition,[],[f1246,f1873])).
fof(f1873,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|49|_0')))),
  inference(unit_resulting_resolution,[],[f1530,f1058])).
fof(f1058,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f730])).
fof(f730,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f356])).
fof(f356,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X18] : (s(bool,'T_0') = s(bool,X18) <=> p(s(bool,X18)))),
  file('thm.integer.INT_MOD_|49|','thm.bool.EQ_CLAUSES|split|1')).
fof(f1530,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|49|_0'))))),
  inference(forward_demodulation,[],[f1529,f945])).
fof(f1529,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|49|_0'))))),
  inference(forward_demodulation,[],[f1528,f1526])).
fof(f1528,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),
  inference(forward_demodulation,[],[f954,f950])).
fof(f954,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(cnf_transformation,[],[f30])).
fof(f30,axiom,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  file('thm.integer.INT_MOD_|49|','thm.integer.INT_LT_|48||49|')).
fof(f1246,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) != s('type.integer.int',X1)) )),
  inference(cnf_transformation,[],[f831])).
fof(f831,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & s('type.integer.int',X0) != s('type.integer.int',X1))) & (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(flattening,[],[f830])).
fof(f830,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & s('type.integer.int',X0) != s('type.integer.int',X1))) & ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))),
  inference(nnf_transformation,[],[f482])).
fof(f482,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & s('type.integer.int',X0) != s('type.integer.int',X1)))),
  inference(flattening,[],[f481])).
fof(f481,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) <=> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & ~s('type.integer.int',X0) = s('type.integer.int',X1)))),
  inference(rectify,[],[f183])).
fof(f183,axiom,(
  ! [X11,X15] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X11),s('type.integer.int',X15)))) <=> (p(s(bool,'const.integer.int_le_2'(s('type.integer.int',X11),s('type.integer.int',X15)))) & ~s('type.integer.int',X11) = s('type.integer.int',X15)))),
  file('thm.integer.INT_MOD_|49|','thm.integer.INT_LT_LE')).
% SZS output end Proof for thm.integer
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 896
% Passive clauses: 7521
% Generated clauses: 41532
% Final active clauses: 677
% Final passive clauses: 4065
% Input formulas: 261
% Initial clauses: 558
% 
% Duplicate literals: 8
% Trivial inequalities: 124
% Fw subsumption resolutions: 785
% Fw demodulations: 7905
% Bw demodulations: 2006
% Condensations: 76
% Global subsumptions: 1
% 
% Simple tautologies: 130
% Equational tautologies: 62
% Forward subsumptions: 23522
% Fw demodulations to eq. taut.: 496
% Bw demodulations to eq. taut.: 29
% 
% Unit resulting resolution: 13202
% Factoring: 1
% Forward superposition: 8553
% Backward superposition: 6910
% Self superposition: 74
% Equality factoring: 2
% Equality resolution: 28
% 
% Split clauses: 24
% Split components: 48
% SAT solver clauses: 8284
% SAT solver unit clauses: 3092
% SAT solver binary clauses: 3266
% 
% Sat splits: 24
% Sat splitting refutations: 712
% 
% Memory used [KB]: 22259
% Time elapsed: 1.001 s
% ------------------------------
% ------------------------------
% Success in time 2.32 s
