Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.ringNorm
% (89280)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 66
% Passive clauses: 2179
% Generated clauses: 12315
% Final active clauses: 64
% Final passive clauses: 2060
% Input formulas: 20
% Initial clauses: 21
% 
% Duplicate literals: 752
% Trivial inequalities: 109
% Fw subsumption resolutions: 959
% Fw demodulations: 617
% 
% Simple tautologies: 245
% Equational tautologies: 95
% Forward subsumptions: 5872
% 
% Binary resolution: 25
% Forward superposition: 4850
% Backward superposition: 4284
% Self superposition: 356
% Equality factoring: 262
% Equality resolution: 6
% 
% Split clauses: 403
% Split components: 816
% SAT solver clauses: 1293
% SAT solver unit clauses: 444
% SAT solver binary clauses: 409
% 
% Sat splits: 403
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1289
% TWLsolver calls for satisfiability: 45
% 
% Memory used [KB]: 10234
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ringNorm
% (89290)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 100
% Passive clauses: 1152
% Generated clauses: 16512
% Final active clauses: 88
% Final passive clauses: 641
% Input formulas: 20
% Initial clauses: 21
% 
% Duplicate literals: 771
% Trivial inequalities: 80
% Fw subsumption resolutions: 709
% Fw demodulations: 1920
% Bw demodulations: 313
% Condensations: 997
% 
% Simple tautologies: 183
% Equational tautologies: 21
% Forward subsumptions: 8237
% Bw demodulations to eq. taut.: 106
% 
% Binary resolution: 17
% Forward superposition: 5638
% Backward superposition: 4975
% Self superposition: 187
% Equality factoring: 61
% Equality resolution: 9
% 
% Split clauses: 119
% Split components: 243
% SAT solver clauses: 1754
% SAT solver unit clauses: 36
% SAT solver binary clauses: 24
% 
% Sat splits: 119
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 1754
% TWLsolver calls for satisfiability: 116
% 
% Memory used [KB]: 17270
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.ringNorm
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ringNorm
% SZS output start Proof for thm.ringNorm
fof(f33931,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f33930])).
fof(f33930,plain,(
  s(sK0,'const.quote.varmap_find_2'(s('type.quote.index',sK3),s('type.quote.varmap'(sK0),sK2))) != s(sK0,'const.quote.varmap_find_2'(s('type.quote.index',sK3),s('type.quote.varmap'(sK0),sK2)))),
  inference(superposition,[],[f62,f7306])).
fof(f7306,plain,(
  ( ! [X6,X8,X7,X5] : (s(X5,'const.quote.varmap_find_2'(s('type.quote.index',X8),s('type.quote.varmap'(X5),X7))) = s(X5,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(X5),X6),s('type.quote.varmap'(X5),X7),s('type.quote.index',X8),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))) )),
  inference(superposition,[],[f1933,f70])).
fof(f70,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.quote.varmap_find_2'(s('type.quote.index',X3),s('type.quote.varmap'(X0),X2))) = s(X0,'const.canonical.ivl_aux_4'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.quote.varmap_find_2'(s('type.quote.index',X3),s('type.quote.varmap'(X0),X2))) = s(X0,'const.canonical.ivl_aux_4'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))),
  inference(flattening,[],[f31])).
fof(f31,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.quote.varmap_find_2'(s('type.quote.index',X3),s('type.quote.varmap'(X0),X2))) = s(X0,'const.canonical.ivl_aux_4'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X11,X8,X9] : s(X6,'const.quote.varmap_find_2'(s('type.quote.index',X9),s('type.quote.varmap'(X6),X8))) = s(X6,'const.canonical.ivl_aux_4'(s('type.semi_ring.semi_ring'(X6),X11),s('type.quote.varmap'(X6),X8),s('type.quote.index',X9),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))),
  file('thm.ringNorm.ivl_aux_def|split|0','thm.canonical.ivl_aux_def|split|0')).
fof(f1933,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,'const.canonical.ivl_aux_4'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),X4)))) )),
  inference(forward_demodulation,[],[f1853,f80])).
fof(f80,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0)),happ(s(fun('type.quote.varmap'(X0),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0))),'const.ringNorm.r_ivl_aux_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.quote.index',X3))),s('type.list.list'('type.quote.index'),X4)))) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0)),happ(s(fun('type.quote.varmap'(X0),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0))),'const.ringNorm.r_ivl_aux_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.quote.index',X3))),s('type.list.list'('type.quote.index'),X4)))),
  inference(flattening,[],[f51])).
fof(f51,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(X0,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0)),happ(s(fun('type.quote.varmap'(X0),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0))),'const.ringNorm.r_ivl_aux_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.quote.index',X3))),s('type.list.list'('type.quote.index'),X4)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6] : ! [X18,X19,X20,X21] : s(X6,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(X6),X18),s('type.quote.varmap'(X6),X19),s('type.quote.index',X20),s('type.list.list'('type.quote.index'),X21))) = s(X6,happ(s(fun('type.list.list'('type.quote.index'),X6),happ(s(fun('type.quote.index',fun('type.list.list'('type.quote.index'),X6)),happ(s(fun('type.quote.varmap'(X6),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X6))),'const.ringNorm.r_ivl_aux_1'(s('type.ring.ring'(X6),X18))),s('type.quote.varmap'(X6),X19))),s('type.quote.index',X20))),s('type.list.list'('type.quote.index'),X21)))),
  file('thm.ringNorm.ivl_aux_def|split|0','ARITY_1_4_const.ringNorm.r_ivl_aux')).
fof(f1853,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.canonical.ivl_aux_4'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0)),happ(s(fun('type.quote.varmap'(X0),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0))),'const.ringNorm.r_ivl_aux_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.quote.index',X3))),s('type.list.list'('type.quote.index'),X4)))) )),
  inference(superposition,[],[f79,f66])).
fof(f66,plain,(
  ( ! [X0,X1] : (s(fun('type.quote.varmap'(X0),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0))),'const.ringNorm.r_ivl_aux_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.quote.varmap'(X0),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0))),'const.canonical.ivl_aux_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f26])).
fof(f26,plain,(
  ! [X0,X1] : s(fun('type.quote.varmap'(X0),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0))),'const.ringNorm.r_ivl_aux_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.quote.varmap'(X0),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0))),'const.canonical.ivl_aux_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f25])).
fof(f25,plain,(
  ! [X0] : ! [X1] : s(fun('type.quote.varmap'(X0),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0))),'const.ringNorm.r_ivl_aux_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.quote.varmap'(X0),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0))),'const.canonical.ivl_aux_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : ! [X7] : s(fun('type.quote.varmap'(X6),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X6))),'const.ringNorm.r_ivl_aux_1'(s('type.ring.ring'(X6),X7))) = s(fun('type.quote.varmap'(X6),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X6))),'const.canonical.ivl_aux_1'(s('type.semi_ring.semi_ring'(X6),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X6),X7)))))),
  file('thm.ringNorm.ivl_aux_def|split|0','thm.ringNorm.r_ivl_aux_def')).
fof(f79,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.canonical.ivl_aux_4'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0)),happ(s(fun('type.quote.varmap'(X0),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0))),'const.canonical.ivl_aux_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.quote.index',X3))),s('type.list.list'('type.quote.index'),X4)))) )),
  inference(cnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,'const.canonical.ivl_aux_4'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0)),happ(s(fun('type.quote.varmap'(X0),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0))),'const.canonical.ivl_aux_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.quote.index',X3))),s('type.list.list'('type.quote.index'),X4)))),
  inference(flattening,[],[f49])).
fof(f49,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(X0,'const.canonical.ivl_aux_4'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0)),happ(s(fun('type.quote.varmap'(X0),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0))),'const.canonical.ivl_aux_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.quote.index',X3))),s('type.list.list'('type.quote.index'),X4)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X18,X19,X20,X21] : s(X6,'const.canonical.ivl_aux_4'(s('type.semi_ring.semi_ring'(X6),X18),s('type.quote.varmap'(X6),X19),s('type.quote.index',X20),s('type.list.list'('type.quote.index'),X21))) = s(X6,happ(s(fun('type.list.list'('type.quote.index'),X6),happ(s(fun('type.quote.index',fun('type.list.list'('type.quote.index'),X6)),happ(s(fun('type.quote.varmap'(X6),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X6))),'const.canonical.ivl_aux_1'(s('type.semi_ring.semi_ring'(X6),X18))),s('type.quote.varmap'(X6),X19))),s('type.quote.index',X20))),s('type.list.list'('type.quote.index'),X21)))),
  file('thm.ringNorm.ivl_aux_def|split|0','ARITY_1_4_const.canonical.ivl_aux')).
fof(f62,plain,(
  s(sK0,'const.quote.varmap_find_2'(s('type.quote.index',sK3),s('type.quote.varmap'(sK0),sK2))) != s(sK0,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(sK0),sK1),s('type.quote.varmap'(sK0),sK2),s('type.quote.index',sK3),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))),
  inference(cnf_transformation,[],[f59])).
fof(f59,plain,(
  s(sK0,'const.quote.varmap_find_2'(s('type.quote.index',sK3),s('type.quote.varmap'(sK0),sK2))) != s(sK0,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(sK0),sK1),s('type.quote.varmap'(sK0),sK2),s('type.quote.index',sK3),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f57])).
fof(f57,plain,(
  ? [X0,X1,X2,X3] : s(X0,'const.quote.varmap_find_2'(s('type.quote.index',X3),s('type.quote.varmap'(X0),X2))) != s(X0,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))),
  inference(ennf_transformation,[],[f23])).
fof(f23,plain,(
  ~! [X0,X1,X2,X3] : s(X0,'const.quote.varmap_find_2'(s('type.quote.index',X3),s('type.quote.varmap'(X0),X2))) = s(X0,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))),
  inference(flattening,[],[f22])).
fof(f22,plain,(
  ~! [X0] : ! [X1,X2,X3] : s(X0,'const.quote.varmap_find_2'(s('type.quote.index',X3),s('type.quote.varmap'(X0),X2))) = s(X0,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : s(X6,'const.quote.varmap_find_2'(s('type.quote.index',X9),s('type.quote.varmap'(X6),X8))) = s(X6,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(X6),X7),s('type.quote.varmap'(X6),X8),s('type.quote.index',X9),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : s(X6,'const.quote.varmap_find_2'(s('type.quote.index',X9),s('type.quote.varmap'(X6),X8))) = s(X6,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(X6),X7),s('type.quote.varmap'(X6),X8),s('type.quote.index',X9),s('type.list.list'('type.quote.index'),'const.list.NIL_0')))),
  file('thm.ringNorm.ivl_aux_def|split|0','thm.ringNorm.ivl_aux_def|split|0')).
% SZS output end Proof for thm.ringNorm
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 149
% Passive clauses: 1826
% Generated clauses: 34688
% Final active clauses: 96
% Final passive clauses: 715
% Input formulas: 20
% Initial clauses: 21
% 
% Duplicate literals: 1846
% Trivial inequalities: 329
% Fw subsumption resolutions: 1761
% Fw demodulations: 5001
% Bw demodulations: 664
% Condensations: 2998
% Global subsumptions: 5
% 
% Simple tautologies: 299
% Equational tautologies: 133
% Forward subsumptions: 20410
% Fw demodulations to eq. taut.: 2
% Bw demodulations to eq. taut.: 186
% 
% Unit resulting resolution: 332
% Forward superposition: 10338
% Backward superposition: 9778
% Self superposition: 361
% Equality factoring: 136
% Equality resolution: 17
% 
% Split clauses: 112
% Split components: 232
% SAT solver clauses: 2077
% SAT solver unit clauses: 38
% SAT solver binary clauses: 208
% 
% Sat splits: 112
% Sat splitting refutations: 20
% 
% Memory used [KB]: 25585
% Time elapsed: 1.617 s
% ------------------------------
% ------------------------------
% Success in time 2.987 s
