Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.int_arith
% (58739)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 324
% Passive clauses: 3886
% Generated clauses: 24333
% Final active clauses: 296
% Final passive clauses: 3316
% Input formulas: 37
% Initial clauses: 123
% 
% Duplicate literals: 836
% Trivial inequalities: 264
% Fw subsumption resolutions: 2908
% Fw demodulations: 3812
% Bw demodulations: 24
% 
% Simple tautologies: 356
% Equational tautologies: 48
% Forward subsumptions: 11406
% 
% Binary resolution: 1574
% Factoring: 13
% Forward superposition: 8302
% Backward superposition: 6170
% Self superposition: 94
% Equality factoring: 62
% Equality resolution: 22
% 
% Split clauses: 556
% Split components: 1127
% SAT solver clauses: 1042
% SAT solver unit clauses: 8
% SAT solver binary clauses: 107
% 
% Sat splits: 556
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1037
% TWLsolver calls for satisfiability: 200
% 
% Memory used [KB]: 13176
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.int_arith
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.int_arith
% SZS output start Proof for thm.int_arith
fof(f2119,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f2109,f180,f308,f2084,f2085,f181,f316,f251,f353,f405,f411,f179,f310,f195,f324,f408,f415,f178,f312])).
fof(f312,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK9)))) | $spl4),
  inference(cnf_transformation,[],[f312_D])).
fof(f312_D,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK9)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f178,plain,(
  p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK9))))),
  inference(cnf_transformation,[],[f124])).
fof(f124,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK8)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK9)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK9),s('type.integer.int',sK8)))) & s('type.integer.int',sK9) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK7),s('type.integer.int',sK8))),s('type.integer.int',sK9))),s('type.integer.int',sK8)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7,sK8,sK9])],[f94])).
fof(f94,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X2)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int',X1)))) & s('type.integer.int',X2) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int',X2))),s('type.integer.int',X1))))),
  inference(flattening,[],[f93])).
fof(f93,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X2)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int',X1))))) & s('type.integer.int',X2) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int',X2))),s('type.integer.int',X1))))),
  inference(ennf_transformation,[],[f39])).
fof(f39,plain,(
  ~! [X0,X1,X2] : ((p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X2)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int',X1))))) => s('type.integer.int',X2) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int',X2))),s('type.integer.int',X1))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : ((p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X7)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X8)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X8),s('type.integer.int',X7))))) => s('type.integer.int',X8) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X6),s('type.integer.int',X7))),s('type.integer.int',X8))),s('type.integer.int',X7))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : ((p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X7)))) & p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X8)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X8),s('type.integer.int',X7))))) => s('type.integer.int',X8) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X6),s('type.integer.int',X7))),s('type.integer.int',X8))),s('type.integer.int',X7))))),
  file('thm.int_arith.int_arith_63_prove_ltb','thm.int_arith.int_arith_63_prove_ltb')).
fof(f415,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK9))) | $spl148),
  inference(cnf_transformation,[],[f415_D])).
fof(f415_D,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK9))) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f408,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK9))) | ($spl4 | $spl14)),
  inference(resolution,[],[f324,f312])).
fof(f324,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(cnf_transformation,[],[f324_D])).
fof(f324_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f195,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.int_arith.int_arith_63_prove_ltb','thm.bool.EQ_CLAUSES|split|0')).
fof(f310,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK9),s('type.integer.int',sK8)))) | $spl2),
  inference(cnf_transformation,[],[f310_D])).
fof(f310_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK9),s('type.integer.int',sK8)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f179,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK9),s('type.integer.int',sK8))))),
  inference(cnf_transformation,[],[f124])).
fof(f411,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK9),s('type.integer.int',sK8))) | $spl144),
  inference(cnf_transformation,[],[f411_D])).
fof(f411_D,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK9),s('type.integer.int',sK8))) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f405,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK9),s('type.integer.int',sK8))) | ($spl2 | $spl14)),
  inference(resolution,[],[f324,f310])).
fof(f353,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X2)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X2))),s('type.integer.int',X0))),s('type.integer.int',X2)))) ) | $spl54),
  inference(cnf_transformation,[],[f353_D])).
fof(f353_D,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X2)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X2))),s('type.integer.int',X0))),s('type.integer.int',X2)))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f251,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X2))),s('type.integer.int',X0))),s('type.integer.int',X2))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X2)))) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X2)))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X2))),s('type.integer.int',X0))),s('type.integer.int',X2))))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X2))))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X2))),s('type.integer.int',X0))),s('type.integer.int',X2))))),
  inference(ennf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X2))))) => s('type.integer.int',X0) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X2))),s('type.integer.int',X0))),s('type.integer.int',X2))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X8,X7,X9] : ((p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X8)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X8),s('type.integer.int',X9))))) => s('type.integer.int',X8) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X7),s('type.integer.int',X9))),s('type.integer.int',X8))),s('type.integer.int',X9))))),
  file('thm.int_arith.int_arith_63_prove_ltb','thm.integer.integer_287_prove_ltb')).
fof(f316,plain,(
  p(s(bool,'T_0')) | $spl8),
  inference(cnf_transformation,[],[f316_D])).
fof(f316_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f181,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.int_arith.int_arith_63_prove_ltb',aHL_TRUTH)).
fof(f2085,plain,(
  ( ! [X0] : (s('type.integer.int',sK9) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',sK8))),s('type.integer.int',sK9))),s('type.integer.int',sK8)))) ) | $spl512),
  inference(cnf_transformation,[],[f2085_D])).
fof(f2085_D,plain,(
  ( ! [X0] : (s('type.integer.int',sK9) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',sK8))),s('type.integer.int',sK9))),s('type.integer.int',sK8)))) ) <=> ~$spl512),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl512])])).
fof(f2084,plain,(
  ( ! [X0] : (s('type.integer.int',sK9) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',sK8))),s('type.integer.int',sK9))),s('type.integer.int',sK8)))) ) | ($spl8 | $spl54 | $spl144 | $spl148)),
  inference(subsumption_resolution,[],[f2083,f316])).
fof(f2083,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.integer.int',sK9) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',sK8))),s('type.integer.int',sK9))),s('type.integer.int',sK8)))) ) | ($spl8 | $spl54 | $spl144 | $spl148)),
  inference(forward_demodulation,[],[f2082,f415])).
fof(f2082,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK9)))) | s('type.integer.int',sK9) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',sK8))),s('type.integer.int',sK9))),s('type.integer.int',sK8)))) ) | ($spl8 | $spl54 | $spl144)),
  inference(subsumption_resolution,[],[f2041,f316])).
fof(f2041,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.integer.int_le_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK9)))) | s('type.integer.int',sK9) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int',sK8))),s('type.integer.int',sK9))),s('type.integer.int',sK8)))) ) | ($spl54 | $spl144)),
  inference(superposition,[],[f353,f411])).
fof(f308,plain,(
  s('type.integer.int',sK9) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK7),s('type.integer.int',sK8))),s('type.integer.int',sK9))),s('type.integer.int',sK8))) | $spl1),
  inference(cnf_transformation,[],[f308_D])).
fof(f308_D,plain,(
  s('type.integer.int',sK9) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK7),s('type.integer.int',sK8))),s('type.integer.int',sK9))),s('type.integer.int',sK8))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f180,plain,(
  s('type.integer.int',sK9) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',sK7),s('type.integer.int',sK8))),s('type.integer.int',sK9))),s('type.integer.int',sK8)))),
  inference(cnf_transformation,[],[f124])).
fof(f2109,plain,(
  $false | ($spl1 | $spl512)),
  inference(trivial_inequality_removal,[],[f2102])).
fof(f2102,plain,(
  s('type.integer.int',sK9) != s('type.integer.int',sK9) | ($spl1 | $spl512)),
  inference(superposition,[],[f308,f2085])).
% SZS output end Proof for thm.int_arith
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 126
% Passive clauses: 269
% Generated clauses: 1931
% Final active clauses: 108
% Final passive clauses: 119
% Input formulas: 37
% Initial clauses: 123
% 
% Duplicate literals: 24
% Trivial inequalities: 9
% Fw subsumption resolutions: 104
% Fw demodulations: 291
% Bw demodulations: 35
% Condensations: 35
% Global subsumptions: 13
% 
% Simple tautologies: 91
% Equational tautologies: 7
% Forward subsumptions: 747
% 
% Binary resolution: 283
% Factoring: 2
% Forward superposition: 494
% Backward superposition: 226
% Self superposition: 9
% Equality factoring: 2
% Equality resolution: 3
% 
% Split clauses: 18
% Split components: 36
% SAT solver clauses: 493
% SAT solver unit clauses: 149
% SAT solver binary clauses: 53
% SAT solver learnt clauses: 27
% SAT solver learnt literals: 10
% 
% Sat splits: 18
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 440
% TWLsolver calls for satisfiability: 940
% 
% Memory used [KB]: 2174
% Time elapsed: 0.048 s
% ------------------------------
% ------------------------------
% Success in time 0.661 s
