Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.numpair
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.numpair
% SZS output start Proof for thm.numpair
fof(f13067,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f13031])).
fof(f13031,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0')),
  inference(superposition,[],[f65,f12977])).
fof(f12977,plain,(
  ( ! [X2] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numpair.nmap_2'(s(fun('type.num.num','type.num.num'),X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f2439,f88])).
fof(f88,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.numpair.nlistrec_3'(s(X0,X1),s(fun('type.num.num',fun('type.num.num',fun(X0,X0))),X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.numpair.nlistrec_3'(s(X0,X1),s(fun('type.num.num',fun('type.num.num',fun(X0,X0))),X2),s('type.num.num','const.num.|48|_0')))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.numpair.nlistrec_3'(s(X0,X1),s(fun('type.num.num',fun('type.num.num',fun(X0,X0))),X2),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X8] : ! [X12,X6] : s(X8,X12) = s(X8,'const.numpair.nlistrec_3'(s(X8,X12),s(fun('type.num.num',fun('type.num.num',fun(X8,X8))),X6),s('type.num.num','const.num.|48|_0')))),
  file('thm.numpair.nmap_thm|split|0','thm.numpair.nlistrec_thm|split|0')).
fof(f2439,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.numpair.nmap_2'(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))) = s('type.num.num','const.numpair.nlistrec_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',fun('type.num.num',fun('type.num.num','type.num.num'))),'_dst_x00x10_2'(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.numpair.ncons_0'),s(fun('type.num.num','type.num.num'),X0))),s('type.num.num',X1)))) )),
  inference(forward_demodulation,[],[f2357,f91])).
fof(f91,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.numpair.nmap_2'(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.numpair.nmap_1'(s(fun('type.num.num','type.num.num'),X0))),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0,X1] : s('type.num.num','const.numpair.nmap_2'(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.numpair.nmap_1'(s(fun('type.num.num','type.num.num'),X0))),s('type.num.num',X1)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X13,X14] : s('type.num.num','const.numpair.nmap_2'(s(fun('type.num.num','type.num.num'),X13),s('type.num.num',X14))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.numpair.nmap_1'(s(fun('type.num.num','type.num.num'),X13))),s('type.num.num',X14)))),
  file('thm.numpair.nmap_thm|split|0','ARITY_1_2_const.numpair.nmap')).
fof(f2357,plain,(
  ( ! [X0,X1] : (s('type.num.num',happ(s(fun('type.num.num','type.num.num'),'const.numpair.nmap_1'(s(fun('type.num.num','type.num.num'),X0))),s('type.num.num',X1))) = s('type.num.num','const.numpair.nlistrec_3'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',fun('type.num.num',fun('type.num.num','type.num.num'))),'_dst_x00x10_2'(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.numpair.ncons_0'),s(fun('type.num.num','type.num.num'),X0))),s('type.num.num',X1)))) )),
  inference(superposition,[],[f89,f69])).
fof(f69,plain,(
  ( ! [X0] : (s(fun('type.num.num','type.num.num'),'const.numpair.nmap_1'(s(fun('type.num.num','type.num.num'),X0))) = s(fun('type.num.num','type.num.num'),'const.numpair.nlistrec_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',fun('type.num.num',fun('type.num.num','type.num.num'))),'_dst_x00x10_2'(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.numpair.ncons_0'),s(fun('type.num.num','type.num.num'),X0)))))) )),
  inference(cnf_transformation,[],[f25])).
fof(f25,plain,(
  ! [X0] : s(fun('type.num.num','type.num.num'),'const.numpair.nmap_1'(s(fun('type.num.num','type.num.num'),X0))) = s(fun('type.num.num','type.num.num'),'const.numpair.nlistrec_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',fun('type.num.num',fun('type.num.num','type.num.num'))),'_dst_x00x10_2'(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.numpair.ncons_0'),s(fun('type.num.num','type.num.num'),X0)))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X6] : s(fun('type.num.num','type.num.num'),'const.numpair.nmap_1'(s(fun('type.num.num','type.num.num'),X6))) = s(fun('type.num.num','type.num.num'),'const.numpair.nlistrec_2'(s('type.num.num','const.num.|48|_0'),s(fun('type.num.num',fun('type.num.num',fun('type.num.num','type.num.num'))),'_dst_x00x10_2'(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.numpair.ncons_0'),s(fun('type.num.num','type.num.num'),X6)))))),
  file('thm.numpair.nmap_thm|split|0','thm.numpair.nmap_def')).
fof(f89,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.numpair.nlistrec_3'(s(X0,X1),s(fun('type.num.num',fun('type.num.num',fun(X0,X0))),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),'const.numpair.nlistrec_2'(s(X0,X1),s(fun('type.num.num',fun('type.num.num',fun(X0,X0))),X2))),s('type.num.num',X3)))) )),
  inference(cnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.numpair.nlistrec_3'(s(X0,X1),s(fun('type.num.num',fun('type.num.num',fun(X0,X0))),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),'const.numpair.nlistrec_2'(s(X0,X1),s(fun('type.num.num',fun('type.num.num',fun(X0,X0))),X2))),s('type.num.num',X3)))),
  inference(flattening,[],[f35])).
fof(f35,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.numpair.nlistrec_3'(s(X0,X1),s(fun('type.num.num',fun('type.num.num',fun(X0,X0))),X2),s('type.num.num',X3))) = s(X0,happ(s(fun('type.num.num',X0),'const.numpair.nlistrec_2'(s(X0,X1),s(fun('type.num.num',fun('type.num.num',fun(X0,X0))),X2))),s('type.num.num',X3)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X8] : ! [X13,X14,X15] : s(X8,'const.numpair.nlistrec_3'(s(X8,X13),s(fun('type.num.num',fun('type.num.num',fun(X8,X8))),X14),s('type.num.num',X15))) = s(X8,happ(s(fun('type.num.num',X8),'const.numpair.nlistrec_2'(s(X8,X13),s(fun('type.num.num',fun('type.num.num',fun(X8,X8))),X14))),s('type.num.num',X15)))),
  file('thm.numpair.nmap_thm|split|0','ARITY_2_3_const.numpair.nlistrec')).
fof(f65,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.nmap_2'(s(fun('type.num.num','type.num.num'),sK0),s('type.num.num','const.num.|48|_0')))),
  inference(cnf_transformation,[],[f50])).
fof(f50,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.nmap_2'(s(fun('type.num.num','type.num.num'),sK0),s('type.num.num','const.num.|48|_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f48])).
fof(f48,plain,(
  ? [X0] : s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numpair.nmap_2'(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.|48|_0')))),
  inference(ennf_transformation,[],[f23])).
fof(f23,plain,(
  ~! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numpair.nmap_2'(s(fun('type.num.num','type.num.num'),X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numpair.nmap_2'(s(fun('type.num.num','type.num.num'),X6),s('type.num.num','const.num.|48|_0')))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numpair.nmap_2'(s(fun('type.num.num','type.num.num'),X6),s('type.num.num','const.num.|48|_0')))),
  file('thm.numpair.nmap_thm|split|0','thm.numpair.nmap_thm|split|0')).
% SZS output end Proof for thm.numpair
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 68
% Passive clauses: 1863
% Generated clauses: 12610
% Final active clauses: 66
% Final passive clauses: 1745
% Input formulas: 21
% Initial clauses: 34
% 
% Duplicate literals: 890
% Trivial inequalities: 166
% Fw subsumption resolutions: 1541
% Fw demodulations: 486
% 
% Simple tautologies: 275
% Equational tautologies: 102
% Forward subsumptions: 5615
% 
% Binary resolution: 33
% Forward superposition: 5264
% Backward superposition: 3437
% Self superposition: 368
% Equality factoring: 283
% Equality resolution: 19
% 
% Split clauses: 578
% Split components: 1166
% SAT solver clauses: 1691
% SAT solver unit clauses: 268
% SAT solver binary clauses: 503
% 
% Sat splits: 578
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1683
% TWLsolver calls for satisfiability: 62
% 
% Memory used [KB]: 6524
% Time elapsed: 0.544 s
% ------------------------------
% ------------------------------
% Success in time 0.553 s
