Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.integer
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.integer
% SZS output start Proof for thm.integer
fof(f7474,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f2397,f626,f428,f433,f363,f431,f437,f281,f432,f6781,f6782,f4841,f7378,f4940])).
fof(f4940,plain,(
  $false | ($spl1 | $spl2 | $spl26)),
  inference(subsumption_resolution,[],[f4867,f433])).
fof(f4867,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl26)),
  inference(backward_demodulation,[],[f4841,f432])).
fof(f7378,plain,(
  $false | $spl34),
  inference(subsumption_resolution,[],[f7351,f6782])).
fof(f7351,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK0),s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f378,f1774])).
fof(f1774,plain,(
  s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(trivial_inequality_removal,[],[f1773])).
fof(f1773,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f1772,f316])).
fof(f316,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.integer.INT_LT_REDUCE|split|4',aHL_BOOL_CASES)).
fof(f1772,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f269,f283])).
fof(f283,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7] : s('type.num.num',X7) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7)))),
  file('thm.integer.INT_LT_REDUCE|split|4','thm.arithmetic.NUMERAL_DEF')).
fof(f269,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(cnf_transformation,[],[f216])).
fof(f216,plain,(
  s(bool,'F_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f207])).
fof(f207,plain,(
  ? [X0] : s(bool,'F_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(ennf_transformation,[],[f103])).
fof(f103,plain,(
  ~! [X0] : s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.integer.INT_LT_REDUCE|split|4','thm.integer.INT_LT_REDUCE|split|4')).
fof(f378,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f174])).
fof(f174,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X6,X8] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X8))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X6))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X8)))))),
  file('thm.integer.INT_LT_REDUCE|split|4','thm.integer.INT_LT_CALCULATE|split|0')).
fof(f4841,plain,(
  ( ! [X15] : (s(bool,'T_0') = s(bool,X15)) ) | $spl26),
  inference(cnf_transformation,[],[f4841_D])).
fof(f4841_D,plain,(
  ( ! [X15] : (s(bool,'T_0') = s(bool,X15)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f6782,plain,(
  ( ! [X26] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X26),s('type.num.num','const.num.|48|_0')))) ) | $spl34),
  inference(cnf_transformation,[],[f6782_D])).
fof(f6782_D,plain,(
  ( ! [X26] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X26),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f6781,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X26),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X25)) ) | ($spl2 | $spl8 | $spl13)),
  inference(forward_demodulation,[],[f6759,f287])).
fof(f287,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f86])).
fof(f86,axiom,(
  ! [X6] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)))),
  file('thm.integer.INT_LT_REDUCE|split|4','thm.numeral.numeral_distrib|split|26')).
fof(f6759,plain,(
  ( ! [X26,X27,X25] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X26),s('type.num.num','const.num.|48|_0'))) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X27))) | s(bool,'T_0') = s(bool,X25)) ) | ($spl2 | $spl8 | $spl13)),
  inference(duplicate_literal_removal,[],[f6701])).
fof(f6701,plain,(
  ( ! [X26,X27,X25] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X26),s('type.num.num','const.num.|48|_0'))) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X27))) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) ) | ($spl2 | $spl8 | $spl13)),
  inference(superposition,[],[f956,f4672])).
fof(f4672,plain,(
  ( ! [X17,X18] : (s(bool,X18) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X17),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X18)) ) | $spl13),
  inference(subsumption_resolution,[],[f4370,f626])).
fof(f4370,plain,(
  ( ! [X17,X18] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X18) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X17),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f597,f293])).
fof(f293,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f80])).
fof(f80,axiom,(
  ! [X6] : s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))),
  file('thm.integer.INT_LT_REDUCE|split|4','thm.numeral.numeral_distrib|split|20')).
fof(f597,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f316,f316])).
fof(f956,plain,(
  ( ! [X6,X7] : (s(bool,X7) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6))) | s(bool,'T_0') = s(bool,X7)) ) | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f949,f433])).
fof(f949,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,X7) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6))) | s(bool,'T_0') = s(bool,X7)) ) | $spl8),
  inference(superposition,[],[f603,f287])).
fof(f603,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(superposition,[],[f437,f316])).
fof(f432,plain,(
  ~p(s(bool,'F_0')) | $spl1),
  inference(cnf_transformation,[],[f432_D])).
fof(f432_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f281,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f106])).
fof(f106,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.integer.INT_LT_REDUCE|split|4',aHL_FALSITY)).
fof(f437,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl8),
  inference(cnf_transformation,[],[f437_D])).
fof(f437_D,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f431,plain,(
  p(s(bool,'F_0')) | $spl0),
  inference(cnf_transformation,[],[f431_D])).
fof(f431_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f363,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | ~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f247])).
fof(f247,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X0)) & (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f246])).
fof(f246,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X0)) & (s(bool,'F_0') != s(bool,X0) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f208])).
fof(f208,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X0))),
  inference(ennf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X0))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X11] : ((p(s(bool,X11)) => p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X11))),
  file('thm.integer.INT_LT_REDUCE|split|4','thm.bool.IMP_F_EQ_F')).
fof(f433,plain,(
  p(s(bool,'T_0')) | $spl2),
  inference(cnf_transformation,[],[f433_D])).
fof(f433_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f428,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f366])).
fof(f366,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f249])).
fof(f249,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f164])).
fof(f164,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f163])).
fof(f163,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X10] : ! [X7] : (s(X10,X7) = s(X10,X7) <=> p(s(bool,'T_0')))),
  file('thm.integer.INT_LT_REDUCE|split|4','thm.bool.REFL_CLAUSE')).
fof(f626,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl13),
  inference(cnf_transformation,[],[f626_D])).
fof(f626_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f2397,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X0)) ) | ($spl2 | $spl8)),
  inference(forward_demodulation,[],[f2396,f286])).
fof(f286,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f89])).
fof(f89,axiom,(
  ! [X6] : s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))),
  file('thm.integer.INT_LT_REDUCE|split|4','thm.numeral.numeral_distrib|split|29')).
fof(f2396,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl2 | $spl8)),
  inference(duplicate_literal_removal,[],[f2361])).
fof(f2361,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | ($spl2 | $spl8)),
  inference(superposition,[],[f957,f316])).
fof(f957,plain,(
  ( ! [X14,X15] : (s(bool,X15) != s(bool,'const.arithmetic.>=_2'(s('type.num.num',X14),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X15)) ) | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f953,f433])).
fof(f953,plain,(
  ( ! [X14,X15] : (~p(s(bool,'T_0')) | s(bool,X15) != s(bool,'const.arithmetic.>=_2'(s('type.num.num',X14),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X15)) ) | $spl8),
  inference(superposition,[],[f603,f286])).
% SZS output end Proof for thm.integer
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 150
% Passive clauses: 1104
% Generated clauses: 7341
% Final active clauses: 104
% Final passive clauses: 741
% Input formulas: 101
% Initial clauses: 159
% 
% Duplicate literals: 147
% Trivial inequalities: 130
% Fw subsumption resolutions: 344
% Fw demodulations: 1479
% Bw demodulations: 38
% 
% Simple tautologies: 109
% Equational tautologies: 22
% Forward subsumptions: 3876
% Fw demodulations to eq. taut.: 5
% 
% Binary resolution: 65
% Forward superposition: 2966
% Backward superposition: 1663
% Self superposition: 59
% Equality factoring: 32
% Equality resolution: 29
% 
% Split clauses: 178
% Split components: 360
% SAT solver clauses: 234
% SAT solver unit clauses: 14
% SAT solver binary clauses: 106
% SAT solver learnt clauses: 1
% 
% Sat splits: 178
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 215
% TWLsolver calls for satisfiability: 23
% 
% Memory used [KB]: 3198
% Time elapsed: 0.149 s
% ------------------------------
% ------------------------------
% Success in time 0.153 s
