Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.extreal
% (101650)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 663
% Passive clauses: 7492
% Generated clauses: 21396
% Final active clauses: 657
% Final passive clauses: 6529
% Input formulas: 69
% Initial clauses: 197
% 
% Duplicate literals: 426
% Trivial inequalities: 108
% Fw subsumption resolutions: 1557
% Fw demodulations: 2279
% Bw demodulations: 4
% 
% Simple tautologies: 317
% Equational tautologies: 49
% Forward subsumptions: 8135
% 
% Binary resolution: 6069
% Factoring: 1117
% Forward superposition: 3840
% Backward superposition: 5625
% Self superposition: 87
% Equality factoring: 26
% Equality resolution: 22
% 
% Split clauses: 197
% Split components: 394
% SAT solver clauses: 241
% SAT solver unit clauses: 6
% SAT solver binary clauses: 41
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 197
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 240
% TWLsolver calls for satisfiability: 36
% 
% Memory used [KB]: 16630
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.extreal
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.extreal
% SZS output start Proof for thm.extreal
fof(f5988,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f5960,f656,f291,f497,f495,f653,f654,f347,f529,f348,f530,f5778,f5779,f290,f496,f494,f373,f545])).
fof(f545,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun('type.realax.real',X0),X4),s('type.realax.real',X1))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))),s(X0,X2),s(X0,X3),s(fun('type.realax.real',X0),X4)))) ) | $spl76),
  inference(cnf_transformation,[],[f545_D])).
fof(f545_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun('type.realax.real',X0),X4),s('type.realax.real',X1))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))),s(X0,X2),s(X0,X3),s(fun('type.realax.real',X0),X4)))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f373,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun('type.realax.real',X0),X4),s('type.realax.real',X1))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))),s(X0,X2),s(X0,X3),s(fun('type.realax.real',X0),X4)))) )),
  inference(cnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,happ(s(fun('type.realax.real',X0),X4),s('type.realax.real',X1))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))),s(X0,X2),s(X0,X3),s(fun('type.realax.real',X0),X4)))),
  inference(flattening,[],[f139])).
fof(f139,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(X0,happ(s(fun('type.realax.real',X0),X4),s('type.realax.real',X1))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))),s(X0,X2),s(X0,X3),s(fun('type.realax.real',X0),X4)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X10] : ! [X6,X38,X39,X23] : s(X10,happ(s(fun('type.realax.real',X10),X23),s('type.realax.real',X6))) = s(X10,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6))),s(X10,X38),s(X10,X39),s(fun('type.realax.real',X10),X23)))),
  file('thm.extreal.extreal_|49||49|','thm.extreal.extreal_case_def|split|2')).
fof(f494,plain,(
  s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK11))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK12))) | $spl0),
  inference(cnf_transformation,[],[f494_D])).
fof(f494_D,plain,(
  s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK11))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK12))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f496,plain,(
  s('type.realax.real',sK11) = s('type.realax.real',sK12) | $spl2),
  inference(cnf_transformation,[],[f496_D])).
fof(f496_D,plain,(
  s('type.realax.real',sK11) = s('type.realax.real',sK12) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f290,plain,(
  s('type.realax.real',sK11) = s('type.realax.real',sK12) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK11))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK12)))),
  inference(cnf_transformation,[],[f212])).
fof(f212,plain,(
  (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK11))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK12))) | s('type.realax.real',sK11) = s('type.realax.real',sK12)) & (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK11))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK12))) | s('type.realax.real',sK11) != s('type.realax.real',sK12))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK11,sK12])],[f211])).
fof(f211,plain,(
  ? [X0,X1] : ((s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))) | s('type.realax.real',X0) = s('type.realax.real',X1)) & (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))) | s('type.realax.real',X0) != s('type.realax.real',X1)))),
  inference(nnf_transformation,[],[f175])).
fof(f175,plain,(
  ? [X0,X1] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))) <~> s('type.realax.real',X0) = s('type.realax.real',X1))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))) <=> s('type.realax.real',X0) = s('type.realax.real',X1))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X7))) <=> s('type.realax.real',X6) = s('type.realax.real',X7))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X7))) <=> s('type.realax.real',X6) = s('type.realax.real',X7))),
  file('thm.extreal.extreal_|49||49|','thm.extreal.extreal_|49||49|')).
fof(f5779,plain,(
  ( ! [X151,X150] : (s(X150,happ(s(fun('type.realax.real',X150),X151),s('type.realax.real',sK11))) = s(X150,happ(s(fun('type.realax.real',X150),X151),s('type.realax.real',sK12)))) ) | $spl2202),
  inference(cnf_transformation,[],[f5779_D])).
fof(f5779_D,plain,(
  ( ! [X151,X150] : (s(X150,happ(s(fun('type.realax.real',X150),X151),s('type.realax.real',sK11))) = s(X150,happ(s(fun('type.realax.real',X150),X151),s('type.realax.real',sK12)))) ) <=> ~$spl2202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2202])])).
fof(f5778,plain,(
  ( ! [X151,X150] : (s(X150,happ(s(fun('type.realax.real',X150),X151),s('type.realax.real',sK11))) = s(X150,happ(s(fun('type.realax.real',X150),X151),s('type.realax.real',sK12)))) ) | ($spl0 | $spl76)),
  inference(forward_demodulation,[],[f5484,f545])).
fof(f5484,plain,(
  ( ! [X152,X151,X153,X150] : (s(X150,happ(s(fun('type.realax.real',X150),X151),s('type.realax.real',sK12))) = s(X150,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK11))),s(X150,X152),s(X150,X153),s(fun('type.realax.real',X150),X151)))) ) | ($spl0 | $spl76)),
  inference(superposition,[],[f545,f494])).
fof(f530,plain,(
  ( ! [X0,X1] : (s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) ) | $spl46),
  inference(cnf_transformation,[],[f530_D])).
fof(f530_D,plain,(
  ( ! [X0,X1] : (s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f348,plain,(
  ( ! [X0,X1] : (s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0,X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0] : ! [X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X10] : ! [X20] : s(X10,'const.combin.I_1'(s(X10,X20))) = s(X10,happ(s(fun(X10,X10),'const.combin.I_0'),s(X10,X20)))),
  file('thm.extreal.extreal_|49||49|','ARITY_0_1_const.combin.I')).
fof(f529,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))) ) | $spl44),
  inference(cnf_transformation,[],[f529_D])).
fof(f529_D,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f347,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(flattening,[],[f108])).
fof(f108,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X10] : ! [X13] : s(X10,X13) = s(X10,'const.combin.I_1'(s(X10,X13)))),
  file('thm.extreal.extreal_|49||49|','thm.combin.I_THM')).
fof(f654,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) ) | $spl276),
  inference(cnf_transformation,[],[f654_D])).
fof(f654_D,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) ) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f653,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) ) | ($spl44 | $spl46)),
  inference(backward_demodulation,[],[f529,f530])).
fof(f495,plain,(
  s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK11))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK12))) | $spl1),
  inference(cnf_transformation,[],[f495_D])).
fof(f495_D,plain,(
  s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK11))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK12))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f497,plain,(
  s('type.realax.real',sK11) != s('type.realax.real',sK12) | $spl3),
  inference(cnf_transformation,[],[f497_D])).
fof(f497_D,plain,(
  s('type.realax.real',sK11) != s('type.realax.real',sK12) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f291,plain,(
  s('type.realax.real',sK11) != s('type.realax.real',sK12) | s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK11))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK12)))),
  inference(cnf_transformation,[],[f212])).
fof(f656,plain,(
  $false | ($spl1 | $spl2)),
  inference(trivial_inequality_removal,[],[f655])).
fof(f655,plain,(
  s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK11))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK11))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f495,f496])).
fof(f5960,plain,(
  $false | ($spl3 | $spl276 | $spl2202)),
  inference(subsumption_resolution,[],[f5959,f497])).
fof(f5959,plain,(
  s('type.realax.real',sK11) = s('type.realax.real',sK12) | ($spl276 | $spl2202)),
  inference(forward_demodulation,[],[f5845,f654])).
fof(f5845,plain,(
  s('type.realax.real',sK12) = s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.combin.I_0'),s('type.realax.real',sK11))) | ($spl276 | $spl2202)),
  inference(superposition,[],[f5779,f654])).
% SZS output end Proof for thm.extreal
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 243
% Passive clauses: 1126
% Generated clauses: 5617
% Final active clauses: 239
% Final passive clauses: 879
% Input formulas: 69
% Initial clauses: 197
% 
% Duplicate literals: 96
% Trivial inequalities: 4
% Fw subsumption resolutions: 198
% Fw demodulations: 549
% Bw demodulations: 5
% Condensations: 83
% Global subsumptions: 8
% 
% Simple tautologies: 171
% Equational tautologies: 11
% Forward subsumptions: 2101
% 
% Binary resolution: 955
% Factoring: 148
% Forward superposition: 888
% Backward superposition: 1307
% Self superposition: 24
% Equality factoring: 2
% Equality resolution: 3
% 
% Split clauses: 8
% Split components: 16
% SAT solver clauses: 1414
% SAT solver unit clauses: 215
% SAT solver binary clauses: 55
% SAT solver learnt clauses: 18
% SAT solver learnt literals: 1
% 
% Sat splits: 8
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 1379
% TWLsolver calls for satisfiability: 1953
% 
% Memory used [KB]: 8315
% Time elapsed: 0.210 s
% ------------------------------
% ------------------------------
% Success in time 0.822 s
