Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.list
% (67168)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 200
% Passive clauses: 3244
% Generated clauses: 22526
% Final active clauses: 187
% Final passive clauses: 2751
% Input formulas: 24
% Initial clauses: 38
% 
% Duplicate literals: 728
% Trivial inequalities: 303
% Fw subsumption resolutions: 2169
% Fw demodulations: 4347
% Bw demodulations: 9
% 
% Simple tautologies: 367
% Equational tautologies: 23
% Forward subsumptions: 10467
% 
% Binary resolution: 268
% Forward superposition: 9441
% Backward superposition: 5216
% Self superposition: 86
% Equality factoring: 36
% Equality resolution: 43
% 
% Split clauses: 478
% Split components: 974
% SAT solver clauses: 1109
% SAT solver unit clauses: 59
% SAT solver binary clauses: 344
% 
% Sat splits: 478
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 1096
% TWLsolver calls for satisfiability: 210
% 
% Memory used [KB]: 15351
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.list
% (67182)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 185
% Passive clauses: 1201
% Generated clauses: 20231
% Final active clauses: 109
% Final passive clauses: 685
% Input formulas: 24
% Initial clauses: 38
% 
% Duplicate literals: 565
% Trivial inequalities: 189
% Fw subsumption resolutions: 1792
% Fw demodulations: 2886
% Bw demodulations: 307
% Condensations: 1577
% Global subsumptions: 4
% 
% Simple tautologies: 174
% Equational tautologies: 25
% Forward subsumptions: 8661
% Bw demodulations to eq. taut.: 5
% 
% Binary resolution: 201
% Forward superposition: 7393
% Backward superposition: 2834
% Self superposition: 94
% Equality factoring: 38
% Equality resolution: 33
% 
% Split clauses: 435
% Split components: 891
% SAT solver clauses: 3167
% SAT solver unit clauses: 62
% SAT solver binary clauses: 58
% SAT solver learnt clauses: 12
% SAT solver learnt literals: 5
% 
% Sat splits: 435
% Sat splitting refutations: 26
% 
% TWLsolver clauses: 3147
% TWLsolver calls for satisfiability: 332
% 
% Memory used [KB]: 13432
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.list
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.list
% SZS output start Proof for thm.list
fof(f6110,plain,(
  $false),
  inference(subsumption_resolution,[],[f6109,f102])).
fof(f102,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.list.LIST_TO_SET_2'(s('type.list.list'(X0),'const.list.NIL_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0,X1] : s(bool,'F_0') = s(bool,'const.list.LIST_TO_SET_2'(s('type.list.list'(X0),'const.list.NIL_0'),s(X0,X1)))),
  inference(flattening,[],[f40])).
fof(f40,plain,(
  ! [X0] : ! [X1] : s(bool,'F_0') = s(bool,'const.list.LIST_TO_SET_2'(s('type.list.list'(X0),'const.list.NIL_0'),s(X0,X1)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X9] : s(bool,'F_0') = s(bool,'const.list.LIST_TO_SET_2'(s('type.list.list'(X6),'const.list.NIL_0'),s(X6,X9)))),
  file('thm.list.LIST_TO_SET|split|0','thm.list.LIST_TO_SET_DEF|split|0')).
fof(f6109,plain,(
  s(bool,'F_0') != s(bool,'const.list.LIST_TO_SET_2'(s('type.list.list'(sK0),'const.list.NIL_0'),s(sK0,sK2('const.pred_set.EMPTY_0','const.list.LIST_TO_SET_1'(s('type.list.list'(sK0),'const.list.NIL_0')),sK0,bool))))),
  inference(forward_demodulation,[],[f6108,f122])).
fof(f122,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(X0,X1)))) )),
  inference(forward_demodulation,[],[f103,f101])).
fof(f101,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.pred_set.EMPTY_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0,X1] : s(bool,'F_0') = s(bool,'const.pred_set.EMPTY_1'(s(X0,X1)))),
  inference(flattening,[],[f38])).
fof(f38,plain,(
  ! [X0] : ! [X1] : s(bool,'F_0') = s(bool,'const.pred_set.EMPTY_1'(s(X0,X1)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6] : ! [X9] : s(bool,'F_0') = s(bool,'const.pred_set.EMPTY_1'(s(X6,X9)))),
  file('thm.list.LIST_TO_SET|split|0','thm.pred_set.EMPTY_applied')).
fof(f103,plain,(
  ( ! [X0,X1] : (s(bool,'const.pred_set.EMPTY_1'(s(X0,X1))) = s(bool,happ(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0,X1] : s(bool,'const.pred_set.EMPTY_1'(s(X0,X1))) = s(bool,happ(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(X0,X1)))),
  inference(flattening,[],[f42])).
fof(f42,plain,(
  ! [X0] : ! [X1] : s(bool,'const.pred_set.EMPTY_1'(s(X0,X1))) = s(bool,happ(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(X0,X1)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : ! [X16] : s(bool,'const.pred_set.EMPTY_1'(s(X6,X16))) = s(bool,happ(s(fun(X6,bool),'const.pred_set.EMPTY_0'),s(X6,X16)))),
  file('thm.list.LIST_TO_SET|split|0','ARITY_0_1_const.pred_set.EMPTY')).
fof(f6108,plain,(
  s(bool,'const.list.LIST_TO_SET_2'(s('type.list.list'(sK0),'const.list.NIL_0'),s(sK0,sK2('const.pred_set.EMPTY_0','const.list.LIST_TO_SET_1'(s('type.list.list'(sK0),'const.list.NIL_0')),sK0,bool)))) != s(bool,happ(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK2('const.pred_set.EMPTY_0','const.list.LIST_TO_SET_1'(s('type.list.list'(sK0),'const.list.NIL_0')),sK0,bool))))),
  inference(forward_demodulation,[],[f6061,f109])).
fof(f109,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.list.LIST_TO_SET_2'(s('type.list.list'(X0),X1),s(X0,X2))) = s(bool,happ(s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1))),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0,X1,X2] : s(bool,'const.list.LIST_TO_SET_2'(s('type.list.list'(X0),X1),s(X0,X2))) = s(bool,happ(s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1))),s(X0,X2)))),
  inference(flattening,[],[f49])).
fof(f49,plain,(
  ! [X0] : ! [X1,X2] : s(bool,'const.list.LIST_TO_SET_2'(s('type.list.list'(X0),X1),s(X0,X2))) = s(bool,happ(s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1))),s(X0,X2)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : ! [X16,X17] : s(bool,'const.list.LIST_TO_SET_2'(s('type.list.list'(X6),X16),s(X6,X17))) = s(bool,happ(s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X16))),s(X6,X17)))),
  file('thm.list.LIST_TO_SET|split|0','ARITY_1_2_const.list.LIST_TO_SET')).
fof(f6061,plain,(
  s(bool,happ(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(sK0,sK2('const.pred_set.EMPTY_0','const.list.LIST_TO_SET_1'(s('type.list.list'(sK0),'const.list.NIL_0')),sK0,bool)))) != s(bool,happ(s(fun(sK0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK0),'const.list.NIL_0'))),s(sK0,sK2('const.pred_set.EMPTY_0','const.list.LIST_TO_SET_1'(s('type.list.list'(sK0),'const.list.NIL_0')),sK0,bool))))),
  inference(unit_resulting_resolution,[],[f82,f119])).
fof(f119,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK2(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK2(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)) )),
  inference(cnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK2(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK2(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f80])).
fof(f80,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(rectify,[],[f79])).
fof(f79,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(nnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0,X1,X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(flattening,[],[f59])).
fof(f59,plain,(
  ! [X0,X1] : ! [X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X10,X6] : ! [X11,X12] : (s(fun(X6,X10),X11) = s(fun(X6,X10),X12) <=> ! [X9] : s(X10,happ(s(fun(X6,X10),X11),s(X6,X9))) = s(X10,happ(s(fun(X6,X10),X12),s(X6,X9))))),
  file('thm.list.LIST_TO_SET|split|0','thm.bool.FUN_EQ_THM')).
fof(f82,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK0),'const.list.NIL_0')))),
  inference(cnf_transformation,[],[f63])).
fof(f63,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK0),'const.list.NIL_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f61])).
fof(f61,plain,(
  ? [X0] : s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(ennf_transformation,[],[f26])).
fof(f26,plain,(
  ~! [X0] : s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),'const.list.NIL_0')))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),'const.list.NIL_0')))),
  file('thm.list.LIST_TO_SET|split|0','thm.list.LIST_TO_SET|split|0')).
% SZS output end Proof for thm.list
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 80
% Passive clauses: 259
% Generated clauses: 6068
% Final active clauses: 59
% Final passive clauses: 113
% Input formulas: 24
% Initial clauses: 38
% 
% Duplicate literals: 273
% Trivial inequalities: 42
% Fw subsumption resolutions: 144
% Fw demodulations: 1184
% Bw demodulations: 83
% Condensations: 321
% Global subsumptions: 1
% 
% Simple tautologies: 54
% Equational tautologies: 82
% Forward subsumptions: 3629
% Bw demodulations to eq. taut.: 2
% 
% Unit resulting resolution: 487
% Forward superposition: 1921
% Backward superposition: 1368
% Self superposition: 69
% Equality factoring: 38
% Equality resolution: 7
% 
% Split clauses: 13
% Split components: 26
% SAT solver clauses: 307
% SAT solver unit clauses: 58
% SAT solver binary clauses: 83
% 
% Sat splits: 13
% Sat splitting refutations: 19
% 
% Memory used [KB]: 3837
% Time elapsed: 0.165 s
% ------------------------------
% ------------------------------
% Success in time 1.45 s
