Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.list
% (45470)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1110
% Passive clauses: 6671
% Generated clauses: 39374
% Final active clauses: 1016
% Final passive clauses: 5185
% Input formulas: 261
% Initial clauses: 561
% 
% Duplicate literals: 203
% Trivial inequalities: 84
% Fw subsumption resolutions: 2445
% Fw demodulations: 3586
% Bw demodulations: 91
% 
% Simple tautologies: 283
% Equational tautologies: 23
% Forward subsumptions: 25646
% 
% Binary resolution: 6240
% Factoring: 838
% Forward superposition: 10384
% Backward superposition: 14686
% Self superposition: 53
% Equality factoring: 2
% Equality resolution: 9
% 
% Split clauses: 386
% Split components: 774
% SAT solver clauses: 627
% SAT solver unit clauses: 14
% SAT solver binary clauses: 96
% 
% Sat splits: 386
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 623
% TWLsolver calls for satisfiability: 375
% 
% Memory used [KB]: 19957
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.list
% (45482)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 581
% Passive clauses: 3423
% Generated clauses: 25726
% Final active clauses: 557
% Final passive clauses: 2710
% Input formulas: 261
% Initial clauses: 561
% 
% Duplicate literals: 89
% Trivial inequalities: 24
% Fw subsumption resolutions: 852
% Fw demodulations: 3214
% Bw demodulations: 147
% Condensations: 45
% Global subsumptions: 25
% 
% Simple tautologies: 276
% Equational tautologies: 12
% Forward subsumptions: 13807
% Fw demodulations to eq. taut.: 14
% 
% Binary resolution: 1984
% Factoring: 86
% Forward superposition: 7583
% Backward superposition: 7584
% Self superposition: 73
% Equality factoring: 2
% Equality resolution: 8
% 
% Split clauses: 44
% Split components: 88
% SAT solver clauses: 4342
% SAT solver unit clauses: 645
% SAT solver binary clauses: 170
% SAT solver learnt clauses: 50
% SAT solver learnt literals: 14
% 
% Sat splits: 44
% Sat splitting refutations: 8
% 
% TWLsolver clauses: 4242
% TWLsolver calls for satisfiability: 3478
% 
% Memory used [KB]: 23922
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.list
% (45500)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1271
% Passive clauses: 10788
% Generated clauses: 146564
% Final active clauses: 954
% Final passive clauses: 5569
% Input formulas: 261
% Initial clauses: 561
% 
% Duplicate literals: 10
% Trivial inequalities: 58
% Fw subsumption resolutions: 1536
% Fw demodulations: 59761
% Bw demodulations: 2680
% Condensations: 29
% 
% Simple tautologies: 144
% Equational tautologies: 60
% Forward subsumptions: 62206
% Fw demodulations to eq. taut.: 10765
% Bw demodulations to eq. taut.: 29
% 
% Unit resulting resolution: 35440
% Factoring: 1
% Forward superposition: 23421
% Backward superposition: 20802
% Self superposition: 174
% Equality factoring: 2
% Equality resolution: 21
% 
% Split clauses: 14
% Split components: 28
% SAT solver clauses: 12011
% SAT solver unit clauses: 2642
% SAT solver binary clauses: 6177
% 
% Sat splits: 14
% Sat splitting refutations: 1193
% 
% Memory used [KB]: 239740
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.list
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.list
% SZS output start Proof for thm.list
fof(f34179,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f1084,f1087,f1086,f1088,f711,f1092,f1085,f1107,f1108,f1127,f1110,f1131,f1146,f1148,f1150,f1152,f1153,f1257,f1363,f1365,f1367,f1369,f1370,f1474,f1482,f1747,f1695,f1673,f1819,f1820,f1821,f1822,f1903,f1902,f1901,f1908,f2004,f2003,f2009,f2030,f2414,f2416,f2418,f2420,f2421,f2525,f2595,f2597,f2598,f2599,f2602,f2700,f2772,f2774,f2776,f2778,f2779,f2883,f2572,f3159,f3298,f3300,f3302,f3304,f3305,f3409,f3539,f3541,f3542,f3543,f3546,f3644,f3737,f3739,f3741,f3743,f3744,f3848,f3923,f3925,f3927,f3929,f3930,f4034,f4042,f1538,f4108,f4282,f4284,f4316,f4318,f4319,f4320,f4323,f4421,f4483,f4512,f4514,f4515,f4516,f4519,f4617,f5172,f5174,f5175,f5176,f5179,f5277,f5360,f5362,f5364,f5366,f5367,f5471,f5546,f5569,f5593,f5616,f5618,f5620,f5622,f5623,f5727,f5806,f5830,f5853,f5877,f5902,f5955,f5984,f5986,f5988,f5990,f5991,f6095,f7370,f7372,f7374,f7376,f7377,f7481,f7583,f7585,f7587,f7589,f7590,f7694,f1525,f7769,f4120,f7792,f5581,f7813,f7815,f7817,f7819,f7820,f7924,f3464,f8003,f5606,f8027,f5818,f8050,f5842,f8072,f5866,f8095,f5891,f8147,f5916,f8171,f5969,f8199,f8201,f8203,f8205,f8206,f8310,f1320,f8430,f8432,f8434,f8436,f8437,f8541,f8666,f8715,f8717,f8718,f8719,f8722,f8820,f2756,f10045,f10046,f10176,f10177,f10227,f10228,f10256,f10257,f10268,f8667,f10394,f10395,f10462,f10463,f10532,f10533,f10574,f10581,f10575,f10583,f10576,f10585,f10622,f10623,f10632,f10639,f10633,f10641,f10634,f10643,f10852,f10775,f10774,f10773,f10941,f10942,f10943,f10944,f10945,f10946,f10951,f10952,f10975,f11002,f11028,f11030,f11031,f11032,f11035,f11133,f3712,f11221,f11227,f11251,f11280,f11309,f11339,f11368,f11397,f11428,f11458,f11468,f11489,f11516,f11905,f11920,f11932,f11947,f11961,f11974,f12186,f12136,f12108,f12107,f12106,f12275,f12276,f12277,f12278,f12279,f12280,f12285,f12286,f12300,f12313,f12322,f11226,f13072,f13107,f13183,f13217,f13297,f13330,f13363,f13396,f13445,f13447,f14028,f15137,f15639,f15641,f15642,f15643,f15646,f15744,f15818,f19410,f19382,f19381,f19380,f19552,f19553,f19554,f19555,f19556,f19557,f19562,f19563,f19737,f19774,f19776,f19778,f19780,f19782,f19784,f19786,f19788,f19790,f19792,f19794,f19796,f19798,f19800,f19804,f19807,f10346,f10347,f19828,f19829,f19831,f19806,f7557,f7560,f20405,f20418,f20448,f20461,f1089,f20482,f20483,f20503,f20504,f20519,f20521,f20522,f20523,f20526,f20624,f20748,f20750,f20751,f20752,f20755,f20853,f20919,f19863,f19862,f1121,f20922,f20923,f1124,f20924,f19858,f20925,f20921,f20926,f1129,f19846,f1126,f19845,f1123,f1122,f21037,f21039,f21040,f21041,f21044,f21142,f20909,f21359,f21361,f21362,f21363,f21366,f21464,f22185,f22163,f22308,f22309,f22310,f22311,f22338,f22339,f22757,f22758,f7559,f22773,f22775,f22776,f22777,f22780,f22878,f5337,f23014,f10976,f23157,f11003,f23258,f11252,f23288,f11281,f23336,f11310,f23384,f11340,f23452,f11369,f23511,f11398,f23581,f11429,f23701,f11459,f23810,f23831,f1091,f23834,f11490,f23925,f11517,f23949,f23951,f23952,f23953,f23956,f24054,f21532,f24199,f24198,f24205,f24224,f24225,f24226,f24227,f24228,f24229,f1090,f24564,f24647,f13298,f25429,f13331,f25527,f13364,f25568,f13397,f25583,f10642,f25584,f25591,f25586,f10638,f26073,f14030,f27299,f19801,f27389,f19799,f27420,f19797,f27460,f19795,f27501,f19793,f27540,f19791,f27575,f19789,f27736,f19787,f27809,f19785,f27946,f19783,f28061,f19781,f28210,f19779,f28249,f19777,f28291,f19775,f1536,f28300,f2238,f2236,f28683,f28685,f28687,f28689,f28690,f28794,f28857,f28971,f29522,f29524,f29525,f29526,f29529,f29627,f29715,f29747,f29749,f29750,f29751,f29754,f29852,f29959,f29991,f29993,f29994,f29995,f29998,f30096,f30203,f31964,f31967,f31968,f31969,f31966,f32012,f32014,f32016,f32018,f32019,f32123,f7528,f33620,f33622,f33623,f33624,f33627,f33725,f33794,f33978,f33980,f33982,f1319,f33984,f1322,f33985,f1321,f34089,f34150,f1323,f34171])).
fof(f34171,plain,(
  $false | ($spl0 | $spl6 | $spl42)),
  inference(subsumption_resolution,[],[f34170,f1017])).
fof(f1017,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f382])).
fof(f382,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.list.FILTER_NEQ_ID',aHL_FALSITY)).
fof(f34170,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl6 | $spl42)),
  inference(forward_demodulation,[],[f34169,f24330])).
fof(f24330,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14))) | $spl42),
  inference(resolution,[],[f24197,f953])).
fof(f953,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f639])).
fof(f639,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X0)) & (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f638])).
fof(f638,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X0)) & (s(bool,'F_0') != s(bool,X0) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f471])).
fof(f471,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X0))),
  inference(ennf_transformation,[],[f348])).
fof(f348,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X0))),
  inference(rectify,[],[f158])).
fof(f158,axiom,(
  ! [X12] : ((p(s(bool,X12)) => p(s(bool,'F_0'))) <=> s(bool,'F_0') = s(bool,X12))),
  file('thm.list.FILTER_NEQ_ID','thm.bool.IMP_F_EQ_F')).
fof(f24197,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)))) | $spl42),
  inference(equality_resolution,[],[f2572])).
fof(f34169,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)))) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f34113,f1080])).
fof(f1080,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f946])).
fof(f946,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f634])).
fof(f634,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f343])).
fof(f343,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f342])).
fof(f342,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6] : ! [X9] : (s(X6,X9) = s(X6,X9) <=> p(s(bool,'T_0')))),
  file('thm.list.FILTER_NEQ_ID','thm.bool.REFL_CLAUSE')).
fof(f34113,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)))) | ($spl0 | $spl6)),
  inference(superposition,[],[f11184,f5274])).
fof(f5274,plain,(
  s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | $spl0),
  inference(resolution,[],[f4976,f1003])).
fof(f1003,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f670])).
fof(f670,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f367])).
fof(f367,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.list.FILTER_NEQ_ID','thm.bool.EQ_CLAUSES|split|0')).
fof(f4976,plain,(
  p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) | $spl0),
  inference(trivial_inequality_removal,[],[f4796])).
fof(f4796,plain,(
  s('type.list.list'(sK13),sK15) != s('type.list.list'(sK13),sK15) | p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) | $spl0),
  inference(superposition,[],[f1066,f1086])).
fof(f1066,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | s('type.list.list'(X0),X2) != s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) )),
  inference(cnf_transformation,[],[f707])).
fof(f707,plain,(
  ! [X0,X1,X2] : ((s('type.list.list'(X0),X2) != s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))) & (~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))))),
  inference(nnf_transformation,[],[f424])).
fof(f424,plain,(
  ! [X0,X1,X2] : (s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) <=> p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))))),
  inference(flattening,[],[f423])).
fof(f423,plain,(
  ! [X0] : ! [X1,X2] : (s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) <=> p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6] : ! [X7,X8] : (s('type.list.list'(X6),X8) = s('type.list.list'(X6),'const.list.FILTER_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X8))) <=> p(s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X8)))))),
  file('thm.list.FILTER_NEQ_ID','thm.list.FILTER_EQ_ID')).
fof(f11184,plain,(
  ( ! [X12] : (~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X12),s('type.list.list'(sK13),sK15)))) | p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),X12))))) ) | $spl6),
  inference(forward_demodulation,[],[f11018,f760])).
fof(f760,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))) = s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f296])).
fof(f296,plain,(
  ! [X0,X1,X2] : s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))) = s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))),
  inference(flattening,[],[f295])).
fof(f295,plain,(
  ! [X0] : ! [X1,X2] : s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))) = s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))),
  inference(rectify,[],[f109])).
fof(f109,axiom,(
  ! [X6] : ! [X5,X48] : s(bool,'const.bool.IN_2'(s(X6,X5),s(fun(X6,bool),X48))) = s(bool,happ(s(fun(X6,bool),X48),s(X6,X5)))),
  file('thm.list.FILTER_NEQ_ID','thm.bool.IN_DEF')).
fof(f11018,plain,(
  ( ! [X12] : (p(s(bool,happ(s(fun(sK13,bool),X12),s(sK13,sK16)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X12),s('type.list.list'(sK13),sK15))))) ) | $spl6),
  inference(resolution,[],[f1092,f1052])).
fof(f1052,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))) )),
  inference(cnf_transformation,[],[f703])).
fof(f703,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))))) & ((p(s(bool,'const.bool.IN_2'(s(X0,sK58(X2,X1,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK58(X2,X1,X0)))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK58])],[f702])).
fof(f702,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))))),
  inference(rectify,[],[f701])).
fof(f701,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))))),
  inference(nnf_transformation,[],[f486])).
fof(f486,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) <=> ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))))),
  inference(ennf_transformation,[],[f405])).
fof(f405,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))))),
  inference(flattening,[],[f404])).
fof(f404,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X6] : ! [X7,X8] : (p(s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X8)))) <=> ! [X13] : (p(s(bool,'const.bool.IN_2'(s(X6,X13),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X8)))))) => p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X13))))))),
  file('thm.list.FILTER_NEQ_ID','thm.list.EVERY_MEM')).
fof(f1323,plain,(
  ( ! [X61] : (p(s(bool,X61))) ) | $spl26),
  inference(cnf_transformation,[],[f1323_D])).
fof(f1323_D,plain,(
  ( ! [X61] : (p(s(bool,X61))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f34150,plain,(
  ( ! [X22,X20] : (p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),X20)))) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X20),s('type.list.list'(sK13),sK15)))) ) | $spl246),
  inference(cnf_transformation,[],[f34150_D])).
fof(f34150_D,plain,(
  ( ! [X22,X20] : (p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),X20)))) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X20),s('type.list.list'(sK13),sK15)))) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f34089,plain,(
  ( ! [X313,X315,X314] : (p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),X313)))) | p(s(bool,X314)) | ~p(s(bool,X315)) | s(bool,X315) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X313),s('type.list.list'(sK13),sK15)))) ) | $spl6),
  inference(resolution,[],[f11184,f1075])).
fof(f1075,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f787])).
fof(f787,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f565])).
fof(f565,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f564])).
fof(f564,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f448])).
fof(f448,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f447])).
fof(f447,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f305])).
fof(f305,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f125])).
fof(f125,axiom,(
  ! [X9,X52,X22,X53] : ((s(bool,X9) = s(bool,X52) & (p(s(bool,X52)) => s(bool,X22) = s(bool,X53))) => ((p(s(bool,X9)) => p(s(bool,X22))) <=> (p(s(bool,X52)) => p(s(bool,X53)))))),
  file('thm.list.FILTER_NEQ_ID','thm.bool.IMP_CONG')).
fof(f1321,plain,(
  ( ! [X47,X45] : (s(bool,X45) != s(bool,X47) | ~p(s(bool,X45))) ) | $spl22),
  inference(cnf_transformation,[],[f1321_D])).
fof(f1321_D,plain,(
  ( ! [X47,X45] : (s(bool,X45) != s(bool,X47) | ~p(s(bool,X45))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f33985,plain,(
  ( ! [X50,X48,X51,X49] : (p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),X48)))) | ~p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,X50) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X48),s('type.list.list'(sK13),sK15))) | s(bool,X49) != s(bool,X51)) ) | $spl6),
  inference(resolution,[],[f11184,f790])).
fof(f790,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f565])).
fof(f1322,plain,(
  ( ! [X63] : (~p(s(bool,X63))) ) | $spl24),
  inference(cnf_transformation,[],[f1322_D])).
fof(f1322_D,plain,(
  ( ! [X63] : (~p(s(bool,X63))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f33984,plain,(
  ( ! [X47,X45,X46,X44] : (p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),X44)))) | p(s(bool,X45)) | ~p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,X46) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X44),s('type.list.list'(sK13),sK15)))) ) | $spl6),
  inference(resolution,[],[f11184,f789])).
fof(f789,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f565])).
fof(f1319,plain,(
  ( ! [X39,X37] : (s(bool,X37) != s(bool,X39) | p(s(bool,X37))) ) | $spl18),
  inference(cnf_transformation,[],[f1319_D])).
fof(f1319_D,plain,(
  ( ! [X39,X37] : (s(bool,X37) != s(bool,X39) | p(s(bool,X37))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f33982,plain,(
  ( ! [X39,X37,X38,X36] : (p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),X36)))) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,X39) | s(bool,X38) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X36),s('type.list.list'(sK13),sK15)))) ) | $spl6),
  inference(resolution,[],[f11184,f788])).
fof(f788,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f565])).
fof(f33980,plain,(
  ( ! [X30,X28,X31,X29] : (p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),X28)))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X30) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X28),s('type.list.list'(sK13),sK15))) | s(bool,X29) != s(bool,X31)) ) | $spl6),
  inference(resolution,[],[f11184,f786])).
fof(f786,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f565])).
fof(f33978,plain,(
  ( ! [X23,X21,X22,X20] : (p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),X20)))) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X20),s('type.list.list'(sK13),sK15)))) ) | $spl6),
  inference(resolution,[],[f11184,f784])).
fof(f784,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f565])).
fof(f33794,plain,(
  ( ! [X36,X34] : (~p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),X34)))) | p(s(bool,X36)) | s(bool,X36) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X34),s('type.list.list'(sK13),sK15)))) ) | $spl244),
  inference(cnf_transformation,[],[f33794_D])).
fof(f33794_D,plain,(
  ( ! [X36,X34] : (~p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),X34)))) | p(s(bool,X36)) | s(bool,X36) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X34),s('type.list.list'(sK13),sK15)))) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f33725,plain,(
  ( ! [X313,X315,X314] : (~p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),X313)))) | p(s(bool,X314)) | p(s(bool,X315)) | s(bool,X315) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X313),s('type.list.list'(sK13),sK15)))) ) | $spl6),
  inference(resolution,[],[f11183,f1075])).
fof(f11183,plain,(
  ( ! [X11] : (p(s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X11),s('type.list.list'(sK13),sK15)))) | ~p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),X11))))) ) | $spl6),
  inference(forward_demodulation,[],[f11017,f760])).
fof(f11017,plain,(
  ( ! [X11] : (p(s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X11),s('type.list.list'(sK13),sK15)))) | ~p(s(bool,happ(s(fun(sK13,bool),X11),s(sK13,sK16))))) ) | $spl6),
  inference(resolution,[],[f1092,f1051])).
fof(f1051,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2))))))) )),
  inference(cnf_transformation,[],[f700])).
fof(f700,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | (p(s(bool,'const.bool.IN_2'(s(X0,sK57(X2,X1,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK57(X2,X1,X0))))))) & (! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK57])],[f699])).
fof(f699,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | ? [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))))) & (! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))))),
  inference(rectify,[],[f698])).
fof(f698,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))) & (! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))))),
  inference(nnf_transformation,[],[f403])).
fof(f403,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) <=> ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))))),
  inference(flattening,[],[f402])).
fof(f402,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) <=> ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6] : ! [X7,X8] : (p(s(bool,'const.list.EXISTS_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X8)))) <=> ? [X13] : (p(s(bool,'const.bool.IN_2'(s(X6,X13),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X8)))))) & p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X13))))))),
  file('thm.list.FILTER_NEQ_ID','thm.list.EXISTS_MEM')).
fof(f33627,plain,(
  ( ! [X64,X62,X65,X63] : (~p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),X62)))) | p(s(bool,X63)) | ~p(s(bool,X64)) | s(bool,X63) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X62),s('type.list.list'(sK13),sK15))) | s(bool,X64) != s(bool,X65)) ) | $spl6),
  inference(resolution,[],[f11183,f790])).
fof(f33624,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),X50)))) | p(s(bool,X51)) | ~p(s(bool,X52)) | p(s(bool,X53)) | s(bool,X53) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X50),s('type.list.list'(sK13),sK15)))) ) | $spl6),
  inference(resolution,[],[f11183,f789])).
fof(f33623,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),X46)))) | p(s(bool,X47)) | p(s(bool,X48)) | s(bool,X47) != s(bool,X49) | s(bool,X48) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X46),s('type.list.list'(sK13),sK15)))) ) | $spl6),
  inference(resolution,[],[f11183,f788])).
fof(f33622,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),X42)))) | p(s(bool,X43)) | ~p(s(bool,X44)) | s(bool,X43) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X42),s('type.list.list'(sK13),sK15))) | s(bool,X44) != s(bool,X45)) ) | $spl6),
  inference(resolution,[],[f11183,f786])).
fof(f33620,plain,(
  ( ! [X37,X35,X36,X34] : (~p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),X34)))) | p(s(bool,X35)) | p(s(bool,X36)) | s(bool,X35) != s(bool,X37) | s(bool,X36) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X34),s('type.list.list'(sK13),sK15)))) ) | $spl6),
  inference(resolution,[],[f11183,f784])).
fof(f7528,plain,(
  ( ! [X14] : (~p(s(bool,X14)) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) ) | $spl72),
  inference(cnf_transformation,[],[f7528_D])).
fof(f7528_D,plain,(
  ( ! [X14] : (~p(s(bool,X14)) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f32123,plain,(
  ( ! [X196,X195] : (p(s(bool,X195)) | ~p(s(bool,X196)) | s(bool,X196) != s(bool,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) ) | $spl95),
  inference(resolution,[],[f31960,f1075])).
fof(f31960,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) | $spl95),
  inference(trivial_inequality_removal,[],[f31955])).
fof(f31955,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) | $spl95),
  inference(superposition,[],[f10583,f1005])).
fof(f1005,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f671])).
fof(f671,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f368])).
fof(f368,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.list.FILTER_NEQ_ID','thm.bool.EQ_CLAUSES|split|1')).
fof(f32019,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X35) != s(bool,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),sK14))) | s(bool,X34) != s(bool,X36)) ) | $spl95),
  inference(resolution,[],[f31960,f790])).
fof(f32018,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) ) | $spl95),
  inference(resolution,[],[f31960,f789])).
fof(f32016,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) ) | $spl95),
  inference(resolution,[],[f31960,f788])).
fof(f32014,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X20) != s(bool,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),sK14))) | s(bool,X19) != s(bool,X21)) ) | $spl95),
  inference(resolution,[],[f31960,f786])).
fof(f32012,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) ) | $spl95),
  inference(resolution,[],[f31960,f784])).
fof(f31966,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) | $spl243),
  inference(cnf_transformation,[],[f31966_D])).
fof(f31966_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) <=> ~$spl243),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl243])])).
fof(f31969,plain,(
  ( ! [X15] : (~p(s(bool,X15)) | ~p(s(bool,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),sK14))))) ) | $spl95),
  inference(subsumption_resolution,[],[f31949,f1003])).
fof(f31949,plain,(
  ( ! [X15] : (s(bool,'T_0') != s(bool,X15) | ~p(s(bool,X15)) | ~p(s(bool,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),sK14))))) ) | $spl95),
  inference(superposition,[],[f10583,f941])).
fof(f941,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f469])).
fof(f469,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f468])).
fof(f468,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f339])).
fof(f339,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X19,X20] : ((p(s(bool,X19)) => p(s(bool,X20))) => ((p(s(bool,X20)) => p(s(bool,X19))) => s(bool,X19) = s(bool,X20)))),
  file('thm.list.FILTER_NEQ_ID','thm.bool.IMP_ANTISYM_AX')).
fof(f31968,plain,(
  ( ! [X14] : (~p(s(bool,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X14))) ) | $spl95),
  inference(subsumption_resolution,[],[f31948,f1003])).
fof(f31948,plain,(
  ( ! [X14] : (s(bool,'T_0') != s(bool,X14) | ~p(s(bool,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X14))) ) | $spl95),
  inference(superposition,[],[f10583,f941])).
fof(f31967,plain,(
  ( ! [X9] : (~p(s(bool,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X9))) ) | $spl95),
  inference(subsumption_resolution,[],[f31943,f1003])).
fof(f31943,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,X9) | ~p(s(bool,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X9))) ) | $spl95),
  inference(superposition,[],[f10583,f932])).
fof(f932,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f629])).
fof(f629,plain,(
  ! [X0,X1] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X0)) & p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X1)))) & (((~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X1)))) | s(bool,X0) = s(bool,X1)))),
  inference(flattening,[],[f628])).
fof(f628,plain,(
  ! [X0,X1] : ((s(bool,X0) != s(bool,X1) | ((p(s(bool,X0)) & p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X1))))) & (((~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X1)))) | s(bool,X0) = s(bool,X1)))),
  inference(nnf_transformation,[],[f331])).
fof(f331,plain,(
  ! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> ((p(s(bool,X0)) & p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X1)))))),
  inference(flattening,[],[f330])).
fof(f330,plain,(
  ! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> ((p(s(bool,X0)) & p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X1)))))),
  inference(rectify,[],[f220])).
fof(f220,axiom,(
  ! [X19,X20] : (s(bool,X19) = s(bool,X20) <=> ((p(s(bool,X19)) & p(s(bool,X20))) | (~p(s(bool,X19)) & ~p(s(bool,X20)))))),
  file('thm.list.FILTER_NEQ_ID','thm.bool.EQ_EXPAND')).
fof(f31964,plain,(
  ( ! [X8] : (~p(s(bool,X8)) | ~p(s(bool,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),sK14))))) ) | $spl95),
  inference(subsumption_resolution,[],[f31942,f1003])).
fof(f31942,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,X8) | ~p(s(bool,X8)) | ~p(s(bool,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),sK14))))) ) | $spl95),
  inference(superposition,[],[f10583,f932])).
fof(f30203,plain,(
  ( ! [X61,X59] : (p(s(bool,X59)) | s(bool,X59) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),X61),s('type.list.list'(sK13),sK15)))))))) ) | $spl240),
  inference(cnf_transformation,[],[f30203_D])).
fof(f30203_D,plain,(
  ( ! [X61,X59] : (p(s(bool,X59)) | s(bool,X59) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),X61),s('type.list.list'(sK13),sK15)))))))) ) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f30096,plain,(
  ( ! [X329,X328,X330] : (p(s(bool,X328)) | p(s(bool,X329)) | s(bool,X329) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),X330),s('type.list.list'(sK13),sK15)))))))) ) | $spl6),
  inference(resolution,[],[f11013,f1075])).
fof(f11013,plain,(
  ( ! [X8] : (p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),X8),s('type.list.list'(sK13),sK15))))))))) ) | $spl6),
  inference(resolution,[],[f1092,f1037])).
fof(f1037,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3)))))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X3))))))) )),
  inference(cnf_transformation,[],[f686])).
fof(f686,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3)))))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X3))))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X3))))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3))))))))))),
  inference(flattening,[],[f685])).
fof(f685,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3)))))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X3)))))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X3))))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3))))))))))),
  inference(nnf_transformation,[],[f393])).
fof(f393,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3)))))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X3))))))))),
  inference(flattening,[],[f392])).
fof(f392,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3)))))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X3))))))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X13,X15,X16] : (p(s(bool,'const.bool.IN_2'(s(X6,X13),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),'const.list.APPEND_2'(s('type.list.list'(X6),X15),s('type.list.list'(X6),X16)))))))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X13),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X15)))))) | p(s(bool,'const.bool.IN_2'(s(X6,X13),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X16))))))))),
  file('thm.list.FILTER_NEQ_ID','thm.list.MEM_APPEND')).
fof(f29998,plain,(
  ( ! [X88,X87,X85,X86] : (p(s(bool,X85)) | ~p(s(bool,X86)) | s(bool,X85) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),X87),s('type.list.list'(sK13),sK15))))))) | s(bool,X86) != s(bool,X88)) ) | $spl6),
  inference(resolution,[],[f11013,f790])).
fof(f29995,plain,(
  ( ! [X76,X74,X77,X75] : (p(s(bool,X74)) | ~p(s(bool,X75)) | p(s(bool,X76)) | s(bool,X76) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),X77),s('type.list.list'(sK13),sK15)))))))) ) | $spl6),
  inference(resolution,[],[f11013,f789])).
fof(f29994,plain,(
  ( ! [X70,X72,X71,X73] : (p(s(bool,X70)) | p(s(bool,X71)) | s(bool,X70) != s(bool,X72) | s(bool,X71) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),X73),s('type.list.list'(sK13),sK15)))))))) ) | $spl6),
  inference(resolution,[],[f11013,f788])).
fof(f29993,plain,(
  ( ! [X68,X66,X69,X67] : (p(s(bool,X66)) | ~p(s(bool,X67)) | s(bool,X66) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),X68),s('type.list.list'(sK13),sK15))))))) | s(bool,X67) != s(bool,X69)) ) | $spl6),
  inference(resolution,[],[f11013,f786])).
fof(f29991,plain,(
  ( ! [X61,X59,X60,X58] : (p(s(bool,X58)) | p(s(bool,X59)) | s(bool,X58) != s(bool,X60) | s(bool,X59) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),X61),s('type.list.list'(sK13),sK15)))))))) ) | $spl6),
  inference(resolution,[],[f11013,f784])).
fof(f29959,plain,(
  ( ! [X61,X59] : (p(s(bool,X59)) | s(bool,X59) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),X61)))))))) ) | $spl238),
  inference(cnf_transformation,[],[f29959_D])).
fof(f29959_D,plain,(
  ( ! [X61,X59] : (p(s(bool,X59)) | s(bool,X59) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),X61)))))))) ) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f29852,plain,(
  ( ! [X329,X328,X330] : (p(s(bool,X328)) | p(s(bool,X329)) | s(bool,X329) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),X330)))))))) ) | $spl6),
  inference(resolution,[],[f11012,f1075])).
fof(f11012,plain,(
  ( ! [X7] : (p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),X7))))))))) ) | $spl6),
  inference(resolution,[],[f1092,f1036])).
fof(f1036,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.APPEND_2'(s('type.list.list'(X0),X2),s('type.list.list'(X0),X3)))))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2))))))) )),
  inference(cnf_transformation,[],[f686])).
fof(f29754,plain,(
  ( ! [X88,X87,X85,X86] : (p(s(bool,X85)) | ~p(s(bool,X86)) | s(bool,X85) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),X87))))))) | s(bool,X86) != s(bool,X88)) ) | $spl6),
  inference(resolution,[],[f11012,f790])).
fof(f29751,plain,(
  ( ! [X76,X74,X77,X75] : (p(s(bool,X74)) | ~p(s(bool,X75)) | p(s(bool,X76)) | s(bool,X76) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),X77)))))))) ) | $spl6),
  inference(resolution,[],[f11012,f789])).
fof(f29750,plain,(
  ( ! [X70,X72,X71,X73] : (p(s(bool,X70)) | p(s(bool,X71)) | s(bool,X70) != s(bool,X72) | s(bool,X71) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),X73)))))))) ) | $spl6),
  inference(resolution,[],[f11012,f788])).
fof(f29749,plain,(
  ( ! [X68,X66,X69,X67] : (p(s(bool,X66)) | ~p(s(bool,X67)) | s(bool,X66) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),X68))))))) | s(bool,X67) != s(bool,X69)) ) | $spl6),
  inference(resolution,[],[f11012,f786])).
fof(f29747,plain,(
  ( ! [X61,X59,X60,X58] : (p(s(bool,X58)) | p(s(bool,X59)) | s(bool,X58) != s(bool,X60) | s(bool,X59) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.APPEND_2'(s('type.list.list'(sK13),sK15),s('type.list.list'(sK13),X61)))))))) ) | $spl6),
  inference(resolution,[],[f11012,f784])).
fof(f29715,plain,(
  ( ! [X60,X58] : (p(s(bool,X58)) | s(bool,X58) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,X60),s('type.list.list'(sK13),sK15)))))))) ) | $spl236),
  inference(cnf_transformation,[],[f29715_D])).
fof(f29715_D,plain,(
  ( ! [X60,X58] : (p(s(bool,X58)) | s(bool,X58) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,X60),s('type.list.list'(sK13),sK15)))))))) ) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f29627,plain,(
  ( ! [X327,X329,X328] : (p(s(bool,X327)) | p(s(bool,X328)) | s(bool,X328) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,X329),s('type.list.list'(sK13),sK15)))))))) ) | $spl6),
  inference(resolution,[],[f11185,f1075])).
fof(f11185,plain,(
  ( ! [X13] : (p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,X13),s('type.list.list'(sK13),sK15))))))))) ) | $spl6),
  inference(forward_demodulation,[],[f11019,f1056])).
fof(f1056,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X1))))) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f409])).
fof(f409,plain,(
  ! [X0,X1,X2] : s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X1))))) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1)))))),
  inference(flattening,[],[f408])).
fof(f408,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X1))))) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1)))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X14] : ! [X12,X11] : s(fun(X14,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X14),'const.list.CONS_2'(s(X14,X11),s('type.list.list'(X14),X12))))) = s(fun(X14,bool),'const.pred_set.INSERT_2'(s(X14,X11),s(fun(X14,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X14),X12)))))),
  file('thm.list.FILTER_NEQ_ID','thm.list.LIST_TO_SET|split|1')).
fof(f11019,plain,(
  ( ! [X13] : (p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.INSERT_2'(s(sK13,X13),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))))))) ) | $spl6),
  inference(resolution,[],[f1092,f1040])).
fof(f1040,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f688])).
fof(f688,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) & ((s(X0,X1) != s(X0,X2) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))))))),
  inference(flattening,[],[f687])).
fof(f687,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))))) & ((s(X0,X1) != s(X0,X2) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))))))),
  inference(nnf_transformation,[],[f395])).
fof(f395,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) <=> (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))))),
  inference(flattening,[],[f394])).
fof(f394,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) <=> (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X6] : ! [X9,X22,X34] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X22),s(fun(X6,bool),X34)))))) <=> (s(X6,X9) = s(X6,X22) | p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X34))))))),
  file('thm.list.FILTER_NEQ_ID','thm.pred_set.IN_INSERT')).
fof(f29529,plain,(
  ( ! [X87,X85,X86,X84] : (p(s(bool,X84)) | ~p(s(bool,X85)) | s(bool,X84) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,X86),s('type.list.list'(sK13),sK15))))))) | s(bool,X85) != s(bool,X87)) ) | $spl6),
  inference(resolution,[],[f11185,f790])).
fof(f29526,plain,(
  ( ! [X76,X74,X75,X73] : (p(s(bool,X73)) | ~p(s(bool,X74)) | p(s(bool,X75)) | s(bool,X75) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,X76),s('type.list.list'(sK13),sK15)))))))) ) | $spl6),
  inference(resolution,[],[f11185,f789])).
fof(f29525,plain,(
  ( ! [X70,X72,X71,X69] : (p(s(bool,X69)) | p(s(bool,X70)) | s(bool,X69) != s(bool,X71) | s(bool,X70) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,X72),s('type.list.list'(sK13),sK15)))))))) ) | $spl6),
  inference(resolution,[],[f11185,f788])).
fof(f29524,plain,(
  ( ! [X68,X66,X67,X65] : (p(s(bool,X65)) | ~p(s(bool,X66)) | s(bool,X65) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,X67),s('type.list.list'(sK13),sK15))))))) | s(bool,X66) != s(bool,X68)) ) | $spl6),
  inference(resolution,[],[f11185,f786])).
fof(f29522,plain,(
  ( ! [X59,X57,X60,X58] : (p(s(bool,X57)) | p(s(bool,X58)) | s(bool,X57) != s(bool,X59) | s(bool,X58) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,X60),s('type.list.list'(sK13),sK15)))))))) ) | $spl6),
  inference(resolution,[],[f11185,f784])).
fof(f28971,plain,(
  $false | ($spl0 | $spl93)),
  inference(subsumption_resolution,[],[f10581,f5274])).
fof(f28857,plain,(
  ( ! [X13] : (~p(s(bool,X13)) | s(bool,X13) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(sK13,bool),sK14))),s('type.list.list'(sK13),sK15)))) ) | $spl234),
  inference(cnf_transformation,[],[f28857_D])).
fof(f28857_D,plain,(
  ( ! [X13] : (~p(s(bool,X13)) | s(bool,X13) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(sK13,bool),sK14))),s('type.list.list'(sK13),sK15)))) ) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f28794,plain,(
  ( ! [X194,X195] : (p(s(bool,X194)) | ~p(s(bool,X195)) | s(bool,X195) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(sK13,bool),sK14))),s('type.list.list'(sK13),sK15)))) ) | $spl9),
  inference(resolution,[],[f1122,f1075])).
fof(f28690,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(sK13,bool),sK14))),s('type.list.list'(sK13),sK15))) | s(bool,X33) != s(bool,X35)) ) | $spl9),
  inference(resolution,[],[f1122,f790])).
fof(f28689,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(sK13,bool),sK14))),s('type.list.list'(sK13),sK15)))) ) | $spl9),
  inference(resolution,[],[f1122,f789])).
fof(f28687,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(sK13,bool),sK14))),s('type.list.list'(sK13),sK15)))) ) | $spl9),
  inference(resolution,[],[f1122,f788])).
fof(f28685,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(sK13,bool),sK14))),s('type.list.list'(sK13),sK15))) | s(bool,X18) != s(bool,X20)) ) | $spl9),
  inference(resolution,[],[f1122,f786])).
fof(f28683,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(sK13,bool),sK14))),s('type.list.list'(sK13),sK15)))) ) | $spl9),
  inference(resolution,[],[f1122,f784])).
fof(f2236,plain,(
  s('type.list.list'(sK13),sK60(sK15,sK13)) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK60(sK15,sK13),sK13)),s('type.list.list'(sK13),sK60(sK60(sK15,sK13),sK13)))) | $spl38),
  inference(cnf_transformation,[],[f2236_D])).
fof(f2236_D,plain,(
  s('type.list.list'(sK13),sK60(sK15,sK13)) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK60(sK15,sK13),sK13)),s('type.list.list'(sK13),sK60(sK60(sK15,sK13),sK13)))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f2238,plain,(
  s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),'const.list.NIL_0'))) | $spl40),
  inference(cnf_transformation,[],[f2238_D])).
fof(f2238_D,plain,(
  s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),'const.list.NIL_0'))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f28300,plain,(
  s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),'const.list.NIL_0'))) | s('type.list.list'(sK13),sK60(sK15,sK13)) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK60(sK15,sK13),sK13)),s('type.list.list'(sK13),sK60(sK60(sK15,sK13),sK13)))) | $spl6),
  inference(superposition,[],[f11204,f1074])).
fof(f1074,plain,(
  ( ! [X0,X1] : (s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK59(X1,X0)),s('type.list.list'(X0),sK60(X1,X0)))) | s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1)) )),
  inference(cnf_transformation,[],[f710])).
fof(f710,plain,(
  ! [X0,X1] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1) | s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK59(X1,X0)),s('type.list.list'(X0),sK60(X1,X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK59,sK60])],[f434])).
fof(f434,plain,(
  ! [X0,X1] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1) | ? [X2,X3] : s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))),
  inference(flattening,[],[f433])).
fof(f433,plain,(
  ! [X0] : ! [X1] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1) | ? [X2,X3] : s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))),
  inference(rectify,[],[f151])).
fof(f151,axiom,(
  ! [X6] : ! [X8] : (s('type.list.list'(X6),'const.list.NIL_0') = s('type.list.list'(X6),X8) | ? [X11,X12] : s('type.list.list'(X6),X8) = s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X11),s('type.list.list'(X6),X12))))),
  file('thm.list.FILTER_NEQ_ID','thm.list.list_CASES')).
fof(f11204,plain,(
  s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13)))) | $spl6),
  inference(subsumption_resolution,[],[f11203,f1055])).
fof(f1055,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(cnf_transformation,[],[f407])).
fof(f407,plain,(
  ! [X0,X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(flattening,[],[f406])).
fof(f406,plain,(
  ! [X0] : ! [X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X6] : ! [X9] : ~p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),'const.pred_set.EMPTY_0'))))),
  file('thm.list.FILTER_NEQ_ID','thm.pred_set.NOT_IN_EMPTY')).
fof(f11203,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13)))) | $spl6),
  inference(forward_demodulation,[],[f11141,f1058])).
fof(f1058,plain,(
  ( ! [X0] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f412])).
fof(f412,plain,(
  ! [X0] : s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X6] : s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),'const.list.NIL_0')))),
  file('thm.list.FILTER_NEQ_ID','thm.list.LIST_TO_SET|split|0')).
fof(f11141,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13)))) | $spl6),
  inference(superposition,[],[f1092,f1074])).
fof(f1536,plain,(
  s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13)))) | $spl30),
  inference(cnf_transformation,[],[f1536_D])).
fof(f1536_D,plain,(
  s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13)))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f19775,plain,(
  ( ! [X0,X1] : (~sP10(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X1) | p(s(bool,X0))) ) | $spl156),
  inference(cnf_transformation,[],[f19775_D])).
fof(f19775_D,plain,(
  ( ! [X0,X1] : (~sP10(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X1) | p(s(bool,X0))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f28291,plain,(
  ( ! [X0,X1] : (~sP10(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X1) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f28272,f1058])).
fof(f28272,plain,(
  ( ! [X0,X1] : (~sP10(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f11096,f1074])).
fof(f11096,plain,(
  ( ! [X182,X183] : (~sP10(X182,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X183) | p(s(bool,X182))) ) | $spl6),
  inference(resolution,[],[f1092,f866])).
fof(f866,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f603])).
fof(f603,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ((~p(s(bool,X2)) | p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0))))) & ((p(s(bool,X2)) & ~p(s(bool,X0))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f602])).
fof(f602,plain,(
  ! [X2,X1,X0] : ((~sP10(X2,X1,X0) | ((~p(s(bool,X0)) | p(s(bool,X2))) & (~p(s(bool,X1)) | p(s(bool,X2))))) & ((p(s(bool,X0)) & ~p(s(bool,X2))) | (p(s(bool,X1)) & ~p(s(bool,X2))) | sP10(X2,X1,X0)))),
  inference(flattening,[],[f601])).
fof(f601,plain,(
  ! [X2,X1,X0] : ((~sP10(X2,X1,X0) | ((~p(s(bool,X0)) | p(s(bool,X2))) & (~p(s(bool,X1)) | p(s(bool,X2))))) & (((p(s(bool,X0)) & ~p(s(bool,X2))) | (p(s(bool,X1)) & ~p(s(bool,X2)))) | sP10(X2,X1,X0)))),
  inference(nnf_transformation,[],[f501])).
fof(f501,plain,(
  ! [X2,X1,X0] : (sP10(X2,X1,X0) <=> ((~p(s(bool,X0)) | p(s(bool,X2))) & (~p(s(bool,X1)) | p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f19777,plain,(
  ( ! [X0,X1] : (~sP10(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X0))) ) | $spl158),
  inference(cnf_transformation,[],[f19777_D])).
fof(f19777_D,plain,(
  ( ! [X0,X1] : (~sP10(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X0))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f28249,plain,(
  ( ! [X0,X1] : (~sP10(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f28230,f1058])).
fof(f28230,plain,(
  ( ! [X0,X1] : (~sP10(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f11095,f1074])).
fof(f11095,plain,(
  ( ! [X180,X181] : (~sP10(X180,X181,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | p(s(bool,X180))) ) | $spl6),
  inference(resolution,[],[f1092,f865])).
fof(f865,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X2)) | ~sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f603])).
fof(f19779,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,X0))) ) | $spl160),
  inference(cnf_transformation,[],[f19779_D])).
fof(f19779_D,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,X0))) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f28210,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f28185,f1058])).
fof(f28185,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f11066,f1074])).
fof(f11066,plain,(
  ( ! [X123,X122] : (~sP5('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X122,X123) | p(s(bool,X122))) ) | $spl6),
  inference(resolution,[],[f1092,f827])).
fof(f827,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f584])).
fof(f584,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP4(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP4(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f583])).
fof(f583,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | (sP4(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP4(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP5(X2,X0,X1)))),
  inference(flattening,[],[f582])).
fof(f582,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | (sP4(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP4(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP5(X2,X0,X1)))),
  inference(nnf_transformation,[],[f494])).
fof(f494,plain,(
  ! [X2,X0,X1] : (sP5(X2,X0,X1) <=> (sP4(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f19781,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,X1))) ) | $spl162),
  inference(cnf_transformation,[],[f19781_D])).
fof(f19781_D,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,X1))) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f28061,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,X1)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f28036,f1058])).
fof(f28036,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1) | p(s(bool,X1)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f11065,f1074])).
fof(f11065,plain,(
  ( ! [X121,X120] : (~sP5('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X121,X120) | p(s(bool,X120))) ) | $spl6),
  inference(resolution,[],[f1092,f826])).
fof(f826,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X2)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f584])).
fof(f19783,plain,(
  ( ! [X0,X1] : (sP3('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,X0))) ) | $spl164),
  inference(cnf_transformation,[],[f19783_D])).
fof(f19783_D,plain,(
  ( ! [X0,X1] : (sP3('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,X0))) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f27946,plain,(
  ( ! [X0,X1] : (sP3('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f27922,f1058])).
fof(f27922,plain,(
  ( ! [X0,X1] : (sP3('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f11061,f1074])).
fof(f11061,plain,(
  ( ! [X112,X113] : (sP3('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X112,X113) | ~p(s(bool,X112))) ) | $spl6),
  inference(resolution,[],[f1092,f821])).
fof(f821,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f581])).
fof(f581,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP3(X2,X0,X1)) & (~sP3(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f580])).
fof(f580,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP3(X2,X0,X1)) & (~sP3(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f492])).
fof(f492,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP3(X2,X0,X1))),
  inference(definition_folding,[],[f309,f491,f490])).
fof(f490,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f491,plain,(
  ! [X2,X0,X1] : (sP3(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP2(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f309,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f308])).
fof(f308,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f80])).
fof(f80,axiom,(
  ! [X42,X44,X45] : ((p(s(bool,X45)) <=> (p(s(bool,X44)) | p(s(bool,X42)))) <=> ((p(s(bool,X45)) | ~p(s(bool,X44))) & (p(s(bool,X45)) | ~p(s(bool,X42))) & (p(s(bool,X44)) | p(s(bool,X42)) | ~p(s(bool,X45)))))),
  file('thm.list.FILTER_NEQ_ID','thm.sat.sat_25_prove_ltb')).
fof(f19785,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X1) | ~p(s(bool,X0))) ) | $spl166),
  inference(cnf_transformation,[],[f19785_D])).
fof(f19785_D,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X1) | ~p(s(bool,X0))) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f27809,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X1) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f27793,f1058])).
fof(f27793,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f11060,f1074])).
fof(f11060,plain,(
  ( ! [X111,X110] : (sP3(X110,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X111) | ~p(s(bool,X110))) ) | $spl6),
  inference(resolution,[],[f1092,f821])).
fof(f19787,plain,(
  ( ! [X0,X1] : (sP3('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,X1))) ) | $spl168),
  inference(cnf_transformation,[],[f19787_D])).
fof(f19787_D,plain,(
  ( ! [X0,X1] : (sP3('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,X1))) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f27736,plain,(
  ( ! [X0,X1] : (sP3('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,X1)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f27712,f1058])).
fof(f27712,plain,(
  ( ! [X0,X1] : (sP3('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1) | ~p(s(bool,X1)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f11059,f1074])).
fof(f11059,plain,(
  ( ! [X109,X108] : (sP3('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X108,X109) | ~p(s(bool,X109))) ) | $spl6),
  inference(resolution,[],[f1092,f820])).
fof(f820,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f581])).
fof(f19789,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X0))) ) | $spl170),
  inference(cnf_transformation,[],[f19789_D])).
fof(f19789_D,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X0))) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f27575,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f27560,f1058])).
fof(f27560,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f11058,f1074])).
fof(f11058,plain,(
  ( ! [X107,X106] : (sP3(X106,X107,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | ~p(s(bool,X106))) ) | $spl6),
  inference(resolution,[],[f1092,f820])).
fof(f19791,plain,(
  ( ! [X0,X1] : (~sP3(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X1) | p(s(bool,X0))) ) | $spl172),
  inference(cnf_transformation,[],[f19791_D])).
fof(f19791_D,plain,(
  ( ! [X0,X1] : (~sP3(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X1) | p(s(bool,X0))) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f27540,plain,(
  ( ! [X0,X1] : (~sP3(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X1) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f27520,f1058])).
fof(f27520,plain,(
  ( ! [X0,X1] : (~sP3(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f11052,f1074])).
fof(f11052,plain,(
  ( ! [X94,X95] : (~sP3(X94,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X95) | p(s(bool,X94))) ) | $spl6),
  inference(resolution,[],[f1092,f809])).
fof(f809,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f576])).
fof(f576,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP2(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP2(X0,X1,X2) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f575])).
fof(f575,plain,(
  ! [X2,X0,X1] : ((~sP3(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP2(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP2(X2,X0,X1) | sP3(X2,X0,X1)))),
  inference(flattening,[],[f574])).
fof(f574,plain,(
  ! [X2,X0,X1] : ((~sP3(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP2(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP2(X2,X0,X1)) | sP3(X2,X0,X1)))),
  inference(nnf_transformation,[],[f491])).
fof(f19793,plain,(
  ( ! [X0,X1] : (~sP3(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X0))) ) | $spl174),
  inference(cnf_transformation,[],[f19793_D])).
fof(f19793_D,plain,(
  ( ! [X0,X1] : (~sP3(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X0))) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f27501,plain,(
  ( ! [X0,X1] : (~sP3(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f27481,f1058])).
fof(f27481,plain,(
  ( ! [X0,X1] : (~sP3(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f11051,f1074])).
fof(f11051,plain,(
  ( ! [X92,X93] : (~sP3(X92,X93,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | p(s(bool,X92))) ) | $spl6),
  inference(resolution,[],[f1092,f808])).
fof(f808,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X0)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f576])).
fof(f19795,plain,(
  ( ! [X0,X1] : (sP1('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,X0))) ) | $spl176),
  inference(cnf_transformation,[],[f19795_D])).
fof(f19795_D,plain,(
  ( ! [X0,X1] : (sP1('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,X0))) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f27460,plain,(
  ( ! [X0,X1] : (sP1('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f27438,f1058])).
fof(f27438,plain,(
  ( ! [X0,X1] : (sP1('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f11047,f1074])).
fof(f11047,plain,(
  ( ! [X85,X84] : (sP1('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X84,X85) | ~p(s(bool,X84))) ) | $spl6),
  inference(resolution,[],[f1092,f804])).
fof(f804,plain,(
  ( ! [X2,X0,X1] : (sP1(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f573])).
fof(f573,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP1(X2,X0,X1)) & (~sP1(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f572])).
fof(f572,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP1(X2,X0,X1)) & (~sP1(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f489])).
fof(f489,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP1(X2,X0,X1))),
  inference(definition_folding,[],[f449,f488,f487])).
fof(f487,plain,(
  ! [X2,X0,X1] : (sP0(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f488,plain,(
  ! [X2,X0,X1] : (sP1(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP0(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f449,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f307])).
fof(f307,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f306])).
fof(f306,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f78])).
fof(f78,axiom,(
  ! [X42,X44,X45] : ((p(s(bool,X45)) <=> (p(s(bool,X44)) => p(s(bool,X42)))) <=> ((p(s(bool,X45)) | p(s(bool,X44))) & (p(s(bool,X45)) | ~p(s(bool,X42))) & (~p(s(bool,X44)) | p(s(bool,X42)) | ~p(s(bool,X45)))))),
  file('thm.list.FILTER_NEQ_ID','thm.sat.sat_27_prove_ltb')).
fof(f19797,plain,(
  ( ! [X0,X1] : (sP1(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X1) | ~p(s(bool,X0))) ) | $spl178),
  inference(cnf_transformation,[],[f19797_D])).
fof(f19797_D,plain,(
  ( ! [X0,X1] : (sP1(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X1) | ~p(s(bool,X0))) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f27420,plain,(
  ( ! [X0,X1] : (sP1(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X1) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f27405,f1058])).
fof(f27405,plain,(
  ( ! [X0,X1] : (sP1(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f11046,f1074])).
fof(f11046,plain,(
  ( ! [X83,X82] : (sP1(X82,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X83) | ~p(s(bool,X82))) ) | $spl6),
  inference(resolution,[],[f1092,f804])).
fof(f19799,plain,(
  ( ! [X0,X1] : (sP1('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,X1))) ) | $spl180),
  inference(cnf_transformation,[],[f19799_D])).
fof(f19799_D,plain,(
  ( ! [X0,X1] : (sP1('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,X1))) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f27389,plain,(
  ( ! [X0,X1] : (sP1('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,X1)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f27368,f1058])).
fof(f27368,plain,(
  ( ! [X0,X1] : (sP1('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1) | p(s(bool,X1)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f11045,f1074])).
fof(f11045,plain,(
  ( ! [X80,X81] : (sP1('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X80,X81) | p(s(bool,X81))) ) | $spl6),
  inference(resolution,[],[f1092,f803])).
fof(f803,plain,(
  ( ! [X2,X0,X1] : (sP1(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f573])).
fof(f19801,plain,(
  ( ! [X0,X1] : (~sP1(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X1) | p(s(bool,X0))) ) | $spl182),
  inference(cnf_transformation,[],[f19801_D])).
fof(f19801_D,plain,(
  ( ! [X0,X1] : (~sP1(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X1) | p(s(bool,X0))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f27299,plain,(
  ( ! [X0,X1] : (~sP1(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X1) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f27279,f1058])).
fof(f27279,plain,(
  ( ! [X0,X1] : (~sP1(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f11036,f1074])).
fof(f11036,plain,(
  ( ! [X62,X63] : (~sP1(X62,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X63) | p(s(bool,X62))) ) | $spl6),
  inference(resolution,[],[f1092,f792])).
fof(f792,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | ~sP1(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f568])).
fof(f568,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP0(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP0(X0,X1,X2) | sP1(X0,X1,X2)))),
  inference(rectify,[],[f567])).
fof(f567,plain,(
  ! [X2,X0,X1] : ((~sP1(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP0(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP0(X2,X0,X1) | sP1(X2,X0,X1)))),
  inference(flattening,[],[f566])).
fof(f566,plain,(
  ! [X2,X0,X1] : ((~sP1(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP0(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP0(X2,X0,X1)) | sP1(X2,X0,X1)))),
  inference(nnf_transformation,[],[f488])).
fof(f14030,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | $spl152),
  inference(cnf_transformation,[],[f14030_D])).
fof(f14030_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f26073,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13)))) | $spl6),
  inference(forward_demodulation,[],[f25625,f1058])).
fof(f25625,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13)))) | $spl6),
  inference(superposition,[],[f11130,f1074])).
fof(f11130,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | $spl6),
  inference(resolution,[],[f1092,f1003])).
fof(f10638,plain,(
  s(bool,'F_0') != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | $spl101),
  inference(cnf_transformation,[],[f10638_D])).
fof(f10638_D,plain,(
  s(bool,'F_0') != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) <=> ~$spl101),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl101])])).
fof(f25586,plain,(
  ( ! [X18] : (p(s(bool,X18)) | s(bool,'F_0') != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f23530,f5236])).
fof(f5236,plain,(
  ( ! [X161,X162] : (sP6('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),X161,X162)) ) | $spl0),
  inference(resolution,[],[f4976,f860])).
fof(f860,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f599])).
fof(f599,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f598])).
fof(f598,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f496])).
fof(f496,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f23530,plain,(
  ( ! [X4,X5] : (~sP6(X5,'F_0',X4) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f23525])).
fof(f23525,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | ~sP6(X5,'F_0',X4) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5)) ) | $spl6),
  inference(resolution,[],[f23469,f864])).
fof(f864,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | s(bool,X0) != s(bool,X1) | ~sP9(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f600])).
fof(f600,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP9(X2,X1,X0)) & (~sP9(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f500])).
fof(f500,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP9(X2,X1,X0))),
  inference(definition_folding,[],[f313,f499,f498,f497,f496])).
fof(f497,plain,(
  ! [X1,X0,X2] : (sP7(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f498,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f499,plain,(
  ! [X2,X1,X0] : (sP9(X2,X1,X0) <=> (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f313,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f312])).
fof(f312,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f74])).
fof(f74,axiom,(
  ! [X42,X44,X45] : ((p(s(bool,X45)) <=> s(bool,X42) = s(bool,X44)) <=> ((p(s(bool,X45)) | p(s(bool,X44)) | p(s(bool,X42))) & (p(s(bool,X45)) | ~p(s(bool,X42)) | ~p(s(bool,X44))) & (p(s(bool,X44)) | ~p(s(bool,X42)) | ~p(s(bool,X45))) & (p(s(bool,X42)) | ~p(s(bool,X44)) | ~p(s(bool,X45)))))),
  file('thm.list.FILTER_NEQ_ID','thm.sat.sat_2|49|_prove_ltb')).
fof(f23469,plain,(
  ( ! [X6,X7] : (sP9(X6,'F_0',X7) | p(s(bool,X6)) | ~sP6(X7,'F_0',X6)) ) | $spl6),
  inference(subsumption_resolution,[],[f23467,f852])).
fof(f852,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f594])).
fof(f594,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f593])).
fof(f593,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f592])).
fof(f592,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f498])).
fof(f23467,plain,(
  ( ! [X6,X7] : (sP9(X6,'F_0',X7) | p(s(bool,X6)) | ~sP8(X6,X7,'F_0') | ~sP6(X7,'F_0',X6)) ) | $spl6),
  inference(resolution,[],[f23463,f848])).
fof(f848,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X0)) | ~sP8(X0,X2,X1) | ~sP7(X1,X2,X0) | ~sP6(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f591])).
fof(f591,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | (sP6(X2,X1,X0) & sP7(X1,X2,X0) & sP8(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X2,X1,X0) | ~sP7(X1,X2,X0) | ~sP8(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f590])).
fof(f590,plain,(
  ! [X2,X1,X0] : ((~sP9(X2,X1,X0) | (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | ~sP7(X1,X0,X2) | ~sP8(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP9(X2,X1,X0)))),
  inference(flattening,[],[f589])).
fof(f589,plain,(
  ! [X2,X1,X0] : ((~sP9(X2,X1,X0) | (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | ~sP7(X1,X0,X2) | ~sP8(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP9(X2,X1,X0)))),
  inference(nnf_transformation,[],[f499])).
fof(f23463,plain,(
  ( ! [X4,X5] : (sP7('F_0',X4,X5)) ) | $spl6),
  inference(subsumption_resolution,[],[f23459,f854])).
fof(f854,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f597])).
fof(f597,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f596])).
fof(f596,plain,(
  ! [X1,X0,X2] : ((~sP7(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP7(X1,X0,X2)))),
  inference(flattening,[],[f595])).
fof(f595,plain,(
  ! [X1,X0,X2] : ((~sP7(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP7(X1,X0,X2)))),
  inference(nnf_transformation,[],[f497])).
fof(f23459,plain,(
  ( ! [X4,X5] : (sP7('F_0',X4,X5) | p(s(bool,X5))) ) | $spl6),
  inference(resolution,[],[f23427,f852])).
fof(f23427,plain,(
  ( ! [X6,X7] : (~sP8(X7,X6,'F_0') | sP7('F_0',X6,X7)) ) | $spl6),
  inference(subsumption_resolution,[],[f23423,f855])).
fof(f855,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f597])).
fof(f23423,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~sP8(X7,X6,'F_0') | sP7('F_0',X6,X7)) ) | $spl6),
  inference(resolution,[],[f23417,f843])).
fof(f843,plain,(
  ( ! [X2,X0,X1] : (sP7(X1,X2,X0) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f591])).
fof(f23417,plain,(
  ( ! [X2,X3] : (sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0')) ) | $spl6),
  inference(subsumption_resolution,[],[f23416,f1003])).
fof(f23416,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0')) ) | $spl6),
  inference(subsumption_resolution,[],[f23413,f23403])).
fof(f23403,plain,(
  ( ! [X4,X5] : (sP6(X4,'F_0',X5) | s(bool,'T_0') != s(bool,X4)) ) | $spl6),
  inference(subsumption_resolution,[],[f23398,f858])).
fof(f858,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f599])).
fof(f23398,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | sP6(X4,'F_0',X5)) ) | $spl6),
  inference(resolution,[],[f20235,f842])).
fof(f842,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X1,X0) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f591])).
fof(f20235,plain,(
  ( ! [X6,X7] : (sP9(X7,'F_0',X6) | s(bool,'T_0') != s(bool,X6) | p(s(bool,X7))) ) | $spl6),
  inference(superposition,[],[f13574,f861])).
fof(f861,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f600])).
fof(f13574,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl6),
  inference(backward_demodulation,[],[f11130,f11128])).
fof(f11128,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | $spl6),
  inference(resolution,[],[f1092,f998])).
fof(f998,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f668])).
fof(f668,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f364])).
fof(f364,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f363])).
fof(f363,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.list.FILTER_NEQ_ID','thm.bool.EQ_CLAUSES|split|2')).
fof(f23413,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0') | ~sP6(X2,'F_0',X3)) ) | $spl6),
  inference(resolution,[],[f23404,f846])).
fof(f846,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X2)) | ~sP8(X0,X2,X1) | ~sP7(X1,X2,X0) | ~sP6(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f591])).
fof(f23404,plain,(
  ( ! [X6,X7] : (sP7('F_0',X6,X7) | s(bool,'T_0') != s(bool,X6)) ) | $spl6),
  inference(subsumption_resolution,[],[f23399,f854])).
fof(f23399,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | p(s(bool,X7)) | sP7('F_0',X6,X7)) ) | $spl6),
  inference(resolution,[],[f20235,f843])).
fof(f25591,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),sK14))) | $spl233),
  inference(cnf_transformation,[],[f25591_D])).
fof(f25591_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),sK14))) <=> ~$spl233),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl233])])).
fof(f25584,plain,(
  ( ! [X16] : (p(s(bool,X16)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) ) | ($spl0 | $spl6 | $spl11)),
  inference(resolution,[],[f23530,f20812])).
fof(f20812,plain,(
  ( ! [X161,X162] : (sP6('const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)),X161,X162)) ) | ($spl0 | $spl11)),
  inference(resolution,[],[f4982,f860])).
fof(f4982,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) | ($spl0 | $spl11)),
  inference(forward_demodulation,[],[f4981,f760])).
fof(f4981,plain,(
  p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK53(sK15,sK14,sK13))))) | ($spl0 | $spl11)),
  inference(subsumption_resolution,[],[f4779,f1124])).
fof(f4779,plain,(
  s('type.list.list'(sK13),'const.list.NIL_0') = s('type.list.list'(sK13),sK15) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK53(sK15,sK14,sK13))))) | $spl0),
  inference(superposition,[],[f1086,f1042])).
fof(f1042,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK53(X2,X1,X0))))) | s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) )),
  inference(cnf_transformation,[],[f691])).
fof(f691,plain,(
  ! [X0,X1,X2] : ((s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) | (p(s(bool,'const.bool.IN_2'(s(X0,sK53(X2,X1,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK53(X2,X1,X0))))))) & (! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK53])],[f690])).
fof(f690,plain,(
  ! [X0,X1,X2] : ((s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) | ? [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))))) & (! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))))),
  inference(rectify,[],[f689])).
fof(f689,plain,(
  ! [X0,X1,X2] : ((s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) | ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))) & (! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))))),
  inference(nnf_transformation,[],[f397])).
fof(f397,plain,(
  ! [X0,X1,X2] : (s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) <=> ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))))),
  inference(flattening,[],[f396])).
fof(f396,plain,(
  ! [X0] : ! [X1,X2] : (~s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) <=> ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : ! [X7,X8] : (~s('type.list.list'(X6),'const.list.NIL_0') = s('type.list.list'(X6),'const.list.FILTER_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X8))) <=> ? [X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X8)))))) & p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X9))))))),
  file('thm.list.FILTER_NEQ_ID','thm.list.FILTER_NEQ_NIL')).
fof(f10642,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | $spl105),
  inference(cnf_transformation,[],[f10642_D])).
fof(f10642_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) <=> ~$spl105),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl105])])).
fof(f25583,plain,(
  ( ! [X15] : (p(s(bool,X15)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl6),
  inference(resolution,[],[f23530,f3066])).
fof(f3066,plain,(
  ( ! [X175,X174] : (sP6('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X174,X175)) ) | $spl6),
  inference(resolution,[],[f1092,f860])).
fof(f13397,plain,(
  ( ! [X0] : (sP11(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X0))) ) | $spl148),
  inference(cnf_transformation,[],[f13397_D])).
fof(f13397_D,plain,(
  ( ! [X0] : (sP11(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X0))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f25568,plain,(
  ( ! [X0] : (sP11(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f25550,f1058])).
fof(f25550,plain,(
  ( ! [X0] : (sP11(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f11105,f1074])).
fof(f11105,plain,(
  ( ! [X197] : (sP11(X197,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | p(s(bool,X197))) ) | $spl6),
  inference(resolution,[],[f1092,f900])).
fof(f900,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f616])).
fof(f616,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP11(X0,X1)))),
  inference(rectify,[],[f615])).
fof(f615,plain,(
  ! [X1,X0] : ((~sP11(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP11(X1,X0)))),
  inference(flattening,[],[f614])).
fof(f614,plain,(
  ! [X1,X0] : ((~sP11(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP11(X1,X0)))),
  inference(nnf_transformation,[],[f503])).
fof(f503,plain,(
  ! [X1,X0] : (sP11(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f13364,plain,(
  ( ! [X0] : (sP11('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0) | p(s(bool,X0))) ) | $spl146),
  inference(cnf_transformation,[],[f13364_D])).
fof(f13364_D,plain,(
  ( ! [X0] : (sP11('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0) | p(s(bool,X0))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f25527,plain,(
  ( ! [X0] : (sP11('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f25509,f1058])).
fof(f25509,plain,(
  ( ! [X0] : (sP11('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f11104,f1074])).
fof(f11104,plain,(
  ( ! [X196] : (sP11('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X196) | p(s(bool,X196))) ) | $spl6),
  inference(resolution,[],[f1092,f897])).
fof(f897,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f616])).
fof(f13331,plain,(
  ( ! [X0] : (~sP11(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X0))) ) | $spl144),
  inference(cnf_transformation,[],[f13331_D])).
fof(f13331_D,plain,(
  ( ! [X0] : (~sP11(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X0))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f25429,plain,(
  ( ! [X0] : (~sP11(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f25411,f1058])).
fof(f25411,plain,(
  ( ! [X0] : (~sP11(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f11103,f1074])).
fof(f11103,plain,(
  ( ! [X195] : (~sP11(X195,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | ~p(s(bool,X195))) ) | $spl6),
  inference(resolution,[],[f1092,f896])).
fof(f896,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP11(X0,X1)) )),
  inference(cnf_transformation,[],[f616])).
fof(f13298,plain,(
  ( ! [X0] : (~sP11('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0) | ~p(s(bool,X0))) ) | $spl142),
  inference(cnf_transformation,[],[f13298_D])).
fof(f13298_D,plain,(
  ( ! [X0] : (~sP11('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0) | ~p(s(bool,X0))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f24647,plain,(
  ( ! [X0] : (~sP11('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f24630,f1058])).
fof(f24630,plain,(
  ( ! [X0] : (~sP11('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f11102,f1074])).
fof(f11102,plain,(
  ( ! [X194] : (~sP11('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X194) | ~p(s(bool,X194))) ) | $spl6),
  inference(resolution,[],[f1092,f896])).
fof(f24564,plain,(
  $false | ($spl42 | $spl44)),
  inference(equality_resolution,[],[f24499])).
fof(f24499,plain,(
  ( ! [X26] : (s(bool,X26) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)))) ) | ($spl42 | $spl44)),
  inference(subsumption_resolution,[],[f2756,f2572])).
fof(f1090,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)))) | $spl5),
  inference(cnf_transformation,[],[f1090_D])).
fof(f1090_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f24229,plain,(
  ( ! [X33] : (~p(s(bool,X33)) | ~p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14))))) ) | $spl42),
  inference(subsumption_resolution,[],[f24192,f1003])).
fof(f24192,plain,(
  ( ! [X33] : (s(bool,'T_0') != s(bool,X33) | ~p(s(bool,X33)) | ~p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14))))) ) | $spl42),
  inference(superposition,[],[f2572,f1005])).
fof(f24228,plain,(
  ( ! [X32] : (~p(s(bool,X32)) | ~p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14))))) ) | $spl42),
  inference(subsumption_resolution,[],[f24191,f1003])).
fof(f24191,plain,(
  ( ! [X32] : (s(bool,'T_0') != s(bool,X32) | ~p(s(bool,X32)) | ~p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14))))) ) | $spl42),
  inference(superposition,[],[f2572,f1003])).
fof(f24227,plain,(
  ( ! [X26,X27] : (~p(s(bool,X27)) | ~p(s(bool,X26)) | ~p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14))))) ) | $spl42),
  inference(subsumption_resolution,[],[f24186,f932])).
fof(f24186,plain,(
  ( ! [X26,X27] : (s(bool,X26) != s(bool,X27) | ~p(s(bool,X27)) | ~p(s(bool,X26)) | ~p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14))))) ) | $spl42),
  inference(superposition,[],[f2572,f941])).
fof(f24226,plain,(
  ( ! [X24,X25] : (~p(s(bool,X25)) | ~p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X24))) ) | $spl42),
  inference(subsumption_resolution,[],[f24185,f932])).
fof(f24185,plain,(
  ( ! [X24,X25] : (s(bool,X24) != s(bool,X25) | ~p(s(bool,X25)) | ~p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X24))) ) | $spl42),
  inference(superposition,[],[f2572,f941])).
fof(f24225,plain,(
  ( ! [X14,X15] : (~p(s(bool,X15)) | ~p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X14))) ) | $spl42),
  inference(subsumption_resolution,[],[f24180,f932])).
fof(f24180,plain,(
  ( ! [X14,X15] : (s(bool,X14) != s(bool,X15) | ~p(s(bool,X15)) | ~p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X14))) ) | $spl42),
  inference(superposition,[],[f2572,f932])).
fof(f24224,plain,(
  ( ! [X12,X13] : (~p(s(bool,X13)) | ~p(s(bool,X12)) | ~p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14))))) ) | $spl42),
  inference(subsumption_resolution,[],[f24179,f932])).
fof(f24179,plain,(
  ( ! [X12,X13] : (s(bool,X12) != s(bool,X13) | ~p(s(bool,X13)) | ~p(s(bool,X12)) | ~p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14))))) ) | $spl42),
  inference(superposition,[],[f2572,f932])).
fof(f24205,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14))) | $spl231),
  inference(cnf_transformation,[],[f24205_D])).
fof(f24205_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14))) <=> ~$spl231),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl231])])).
fof(f24198,plain,(
  ( ! [X33] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14))) | ~p(s(bool,X33))) ) | $spl42),
  inference(duplicate_literal_removal,[],[f24150])).
fof(f24150,plain,(
  ( ! [X33] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14))) | ~p(s(bool,X33)) | ~p(s(bool,X33))) ) | $spl42),
  inference(superposition,[],[f2572,f1005])).
fof(f24199,plain,(
  ( ! [X32] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14))) | ~p(s(bool,X32))) ) | $spl42),
  inference(duplicate_literal_removal,[],[f24149])).
fof(f24149,plain,(
  ( ! [X32] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14))) | ~p(s(bool,X32)) | ~p(s(bool,X32))) ) | $spl42),
  inference(superposition,[],[f2572,f1003])).
fof(f21532,plain,(
  ( ! [X39] : (p(s(bool,X39)) | s(bool,X39) != s(bool,'const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl210),
  inference(cnf_transformation,[],[f21532_D])).
fof(f21532_D,plain,(
  ( ! [X39] : (p(s(bool,X39)) | s(bool,X39) != s(bool,'const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f24054,plain,(
  ( ! [X212,X213] : (p(s(bool,X212)) | p(s(bool,X213)) | s(bool,X213) != s(bool,'const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl12),
  inference(resolution,[],[f1126,f1075])).
fof(f23956,plain,(
  ( ! [X61,X59,X60] : (p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X59) != s(bool,'const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X60) != s(bool,X61)) ) | $spl12),
  inference(resolution,[],[f1126,f790])).
fof(f23953,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | p(s(bool,X52)) | s(bool,X52) != s(bool,'const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl12),
  inference(resolution,[],[f1126,f789])).
fof(f23952,plain,(
  ( ! [X47,X48,X49] : (p(s(bool,X47)) | p(s(bool,X48)) | s(bool,X47) != s(bool,X49) | s(bool,X48) != s(bool,'const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl12),
  inference(resolution,[],[f1126,f788])).
fof(f23951,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,X44)) | ~p(s(bool,X45)) | s(bool,X44) != s(bool,'const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X45) != s(bool,X46)) ) | $spl12),
  inference(resolution,[],[f1126,f786])).
fof(f23949,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X39) != s(bool,'const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl12),
  inference(resolution,[],[f1126,f784])).
fof(f11517,plain,(
  ( ! [X0,X1] : (sP10('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1)) ) | $spl130),
  inference(cnf_transformation,[],[f11517_D])).
fof(f11517_D,plain,(
  ( ! [X0,X1] : (sP10('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1)) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f23925,plain,(
  ( ! [X0,X1] : (sP10('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f23912,f1058])).
fof(f23912,plain,(
  ( ! [X0,X1] : (sP10('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f3073,f1074])).
fof(f3073,plain,(
  ( ! [X189,X188] : (sP10('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X188,X189)) ) | $spl6),
  inference(resolution,[],[f1092,f872])).
fof(f872,plain,(
  ( ! [X2,X0,X1] : (sP10(X2,X1,X0) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f605])).
fof(f605,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X0)) | p(s(bool,X1))) & ~p(s(bool,X2))) | sP10(X2,X1,X0)) & (~sP10(X2,X1,X0) | (~p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,X2))))),
  inference(flattening,[],[f604])).
fof(f604,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X0)) | p(s(bool,X1))) & ~p(s(bool,X2))) | sP10(X2,X1,X0)) & (~sP10(X2,X1,X0) | ((~p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,X2)))))),
  inference(nnf_transformation,[],[f502])).
fof(f502,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,X2))) <=> sP10(X2,X1,X0))),
  inference(definition_folding,[],[f450,f501])).
fof(f450,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,X2))) <=> ((~p(s(bool,X0)) | p(s(bool,X2))) & (~p(s(bool,X1)) | p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f314])).
fof(f314,plain,(
  ! [X0,X1,X2] : (((p(s(bool,X0)) | p(s(bool,X1))) => p(s(bool,X2))) <=> ((p(s(bool,X0)) => p(s(bool,X2))) & (p(s(bool,X1)) => p(s(bool,X2)))))),
  inference(rectify,[],[f128])).
fof(f128,axiom,(
  ! [X7,X31,X56] : (((p(s(bool,X7)) | p(s(bool,X31))) => p(s(bool,X56))) <=> ((p(s(bool,X7)) => p(s(bool,X56))) & (p(s(bool,X31)) => p(s(bool,X56)))))),
  file('thm.list.FILTER_NEQ_ID','thm.bool.DISJ_IMP_THM')).
fof(f11490,plain,(
  ( ! [X0,X1] : (sP6('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1)) ) | $spl128),
  inference(cnf_transformation,[],[f11490_D])).
fof(f11490_D,plain,(
  ( ! [X0,X1] : (sP6('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1)) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f23834,plain,(
  ( ! [X0,X1] : (sP6('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f23818,f1058])).
fof(f23818,plain,(
  ( ! [X0,X1] : (sP6('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f3066,f1074])).
fof(f1091,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) | $spl7),
  inference(cnf_transformation,[],[f1091_D])).
fof(f1091_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f23831,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | ~sP7(X3,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X2) | ~sP8(X2,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X3)) ) | $spl228),
  inference(cnf_transformation,[],[f23831_D])).
fof(f23831_D,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | ~sP7(X3,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X2) | ~sP8(X2,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X3)) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f23810,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | ~p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) | ~sP8(X2,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X3) | ~sP7(X3,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X2)) ) | $spl6),
  inference(resolution,[],[f3066,f846])).
fof(f11459,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X1)) ) | $spl126),
  inference(cnf_transformation,[],[f11459_D])).
fof(f11459_D,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X1)) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f23701,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f23683,f1058])).
fof(f23683,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f3065,f1074])).
fof(f3065,plain,(
  ( ! [X173,X172] : (sP6(X172,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X173)) ) | $spl6),
  inference(resolution,[],[f1092,f859])).
fof(f859,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f599])).
fof(f11429,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl124),
  inference(cnf_transformation,[],[f11429_D])).
fof(f11429_D,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f23581,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f23563,f1058])).
fof(f23563,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f3064,f1074])).
fof(f3064,plain,(
  ( ! [X171,X170] : (sP6(X170,X171,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl6),
  inference(resolution,[],[f1092,f858])).
fof(f11398,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl122),
  inference(cnf_transformation,[],[f11398_D])).
fof(f11398_D,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f23511,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f23495,f1058])).
fof(f23495,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f3063,f1074])).
fof(f3063,plain,(
  ( ! [X169,X168] : (sP7(X168,X169,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl6),
  inference(resolution,[],[f1092,f854])).
fof(f11369,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl120),
  inference(cnf_transformation,[],[f11369_D])).
fof(f11369_D,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f23452,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f23436,f1058])).
fof(f23436,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f3060,f1074])).
fof(f3060,plain,(
  ( ! [X163,X162] : (sP8(X162,X163,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl6),
  inference(resolution,[],[f1092,f850])).
fof(f850,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f594])).
fof(f11340,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl118),
  inference(cnf_transformation,[],[f11340_D])).
fof(f11340_D,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f23384,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f23368,f1058])).
fof(f23368,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f3047,f1074])).
fof(f3047,plain,(
  ( ! [X136,X137] : (sP4(X136,X137,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl6),
  inference(resolution,[],[f1092,f833])).
fof(f833,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f586])).
fof(f586,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(flattening,[],[f585])).
fof(f585,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(nnf_transformation,[],[f493])).
fof(f493,plain,(
  ! [X0,X1,X2] : (sP4(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f11310,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X1)) ) | $spl116),
  inference(cnf_transformation,[],[f11310_D])).
fof(f11310_D,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X1)) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f23336,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f23320,f1058])).
fof(f23320,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f3031,f1074])).
fof(f3031,plain,(
  ( ! [X105,X104] : (sP2(X104,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X105)) ) | $spl6),
  inference(resolution,[],[f1092,f817])).
fof(f817,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f579])).
fof(f579,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f578])).
fof(f578,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f577])).
fof(f577,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f490])).
fof(f11281,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl114),
  inference(cnf_transformation,[],[f11281_D])).
fof(f11281_D,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f23288,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f23272,f1058])).
fof(f23272,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f3030,f1074])).
fof(f3030,plain,(
  ( ! [X103,X102] : (sP2(X102,X103,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl6),
  inference(resolution,[],[f1092,f816])).
fof(f816,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f579])).
fof(f11252,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X1)) ) | $spl112),
  inference(cnf_transformation,[],[f11252_D])).
fof(f11252_D,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X1)) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f23258,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f23242,f1058])).
fof(f23242,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f3017,f1074])).
fof(f3017,plain,(
  ( ! [X76,X77] : (sP0(X76,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X77)) ) | $spl6),
  inference(resolution,[],[f1092,f800])).
fof(f800,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f571])).
fof(f571,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(rectify,[],[f570])).
fof(f570,plain,(
  ! [X2,X0,X1] : ((~sP0(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP0(X2,X0,X1)))),
  inference(flattening,[],[f569])).
fof(f569,plain,(
  ! [X2,X0,X1] : ((~sP0(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP0(X2,X0,X1)))),
  inference(nnf_transformation,[],[f487])).
fof(f11003,plain,(
  ( ! [X0] : (sP12('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0)) ) | $spl108),
  inference(cnf_transformation,[],[f11003_D])).
fof(f11003_D,plain,(
  ( ! [X0] : (sP12('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0)) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f23157,plain,(
  ( ! [X0] : (sP12('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f23142,f1058])).
fof(f23142,plain,(
  ( ! [X0] : (sP12('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f3085,f1074])).
fof(f3085,plain,(
  ( ! [X203] : (sP12('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X203)) ) | $spl6),
  inference(resolution,[],[f1092,f913])).
fof(f913,plain,(
  ( ! [X0,X1] : (sP12(X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f621])).
fof(f621,plain,(
  ! [X0,X1] : ((~sP12(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP12(X0,X1)))),
  inference(flattening,[],[f620])).
fof(f620,plain,(
  ! [X0,X1] : ((~sP12(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP12(X0,X1)))),
  inference(nnf_transformation,[],[f505])).
fof(f505,plain,(
  ! [X0,X1] : (sP12(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP12])])).
fof(f10976,plain,(
  ( ! [X0] : (sP12(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl106),
  inference(cnf_transformation,[],[f10976_D])).
fof(f10976_D,plain,(
  ( ! [X0] : (sP12(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f23014,plain,(
  ( ! [X0] : (sP12(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f22999,f1058])).
fof(f22999,plain,(
  ( ! [X0] : (sP12(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f3084,f1074])).
fof(f3084,plain,(
  ( ! [X202] : (sP12(X202,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl6),
  inference(resolution,[],[f1092,f912])).
fof(f912,plain,(
  ( ! [X0,X1] : (sP12(X0,X1) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f621])).
fof(f5337,plain,(
  ( ! [X26] : (p(s(bool,X26)) | s(bool,X26) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl54),
  inference(cnf_transformation,[],[f5337_D])).
fof(f5337_D,plain,(
  ( ! [X26] : (p(s(bool,X26)) | s(bool,X26) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f22878,plain,(
  ( ! [X200,X199] : (p(s(bool,X199)) | p(s(bool,X200)) | s(bool,X200) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl0),
  inference(resolution,[],[f4976,f1075])).
fof(f22780,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,X46) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | s(bool,X47) != s(bool,X48)) ) | $spl0),
  inference(resolution,[],[f4976,f790])).
fof(f22777,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,X37)) | ~p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X39) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl0),
  inference(resolution,[],[f4976,f789])).
fof(f22776,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X34) != s(bool,X36) | s(bool,X35) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl0),
  inference(resolution,[],[f4976,f788])).
fof(f22775,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | s(bool,X32) != s(bool,X33)) ) | $spl0),
  inference(resolution,[],[f4976,f786])).
fof(f22773,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl0),
  inference(resolution,[],[f4976,f784])).
fof(f7559,plain,(
  p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) | $spl74),
  inference(cnf_transformation,[],[f7559_D])).
fof(f7559_D,plain,(
  p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f22758,plain,(
  ( ! [X1] : (p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) | p(s(bool,X1))) ) | $spl0),
  inference(subsumption_resolution,[],[f22743,f968])).
fof(f968,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f649])).
fof(f649,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f648])).
fof(f648,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f353])).
fof(f353,plain,(
  ! [X0] : ((p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X12] : ((p(s(bool,'F_0')) | p(s(bool,X12))) <=> p(s(bool,X12)))),
  file('thm.list.FILTER_NEQ_ID','thm.bool.OR_CLAUSES|split|2')).
fof(f22743,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) | p(s(bool,X1))) ) | $spl0),
  inference(resolution,[],[f5255,f911])).
fof(f911,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP12(X0,X1)) )),
  inference(cnf_transformation,[],[f621])).
fof(f5255,plain,(
  ( ! [X190] : (sP12('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),X190)) ) | $spl0),
  inference(resolution,[],[f4976,f913])).
fof(f22757,plain,(
  ( ! [X0] : (p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) | p(s(bool,X0))) ) | $spl0),
  inference(subsumption_resolution,[],[f22742,f968])).
fof(f22742,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) | p(s(bool,X0))) ) | $spl0),
  inference(resolution,[],[f5255,f915])).
fof(f915,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP12(X0,X1)) )),
  inference(cnf_transformation,[],[f623])).
fof(f623,plain,(
  ! [X0,X1] : ((~sP12(X0,X1) | (p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP12(X0,X1)))),
  inference(flattening,[],[f622])).
fof(f622,plain,(
  ! [X0,X1] : ((~sP12(X0,X1) | ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0')))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP12(X0,X1)))),
  inference(nnf_transformation,[],[f506])).
fof(f506,plain,(
  ! [X0,X1] : (sP12(X0,X1) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(definition_folding,[],[f458,f505])).
fof(f458,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f457])).
fof(f457,plain,(
  ! [X0,X1] : (((p(s(bool,X1)) | p(s(bool,X0))) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(ennf_transformation,[],[f326])).
fof(f326,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(flattening,[],[f325])).
fof(f325,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(rectify,[],[f73])).
fof(f73,axiom,(
  ! [X14,X6] : ((~(p(s(bool,X6)) | p(s(bool,X14))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X6)) => p(s(bool,'F_0'))) => (~p(s(bool,X14)) => p(s(bool,'F_0')))))),
  file('thm.list.FILTER_NEQ_ID','thm.sat.sat_9_prove_ltb')).
fof(f22339,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))))) ) | $spl0),
  inference(subsumption_resolution,[],[f22324,f968])).
fof(f22324,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))))) ) | $spl0),
  inference(resolution,[],[f5254,f911])).
fof(f5254,plain,(
  ( ! [X189] : (sP12(X189,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl0),
  inference(resolution,[],[f4976,f912])).
fof(f22338,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))))) ) | $spl0),
  inference(subsumption_resolution,[],[f22323,f968])).
fof(f22323,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))))) ) | $spl0),
  inference(resolution,[],[f5254,f915])).
fof(f22311,plain,(
  $false | ($spl6 | $spl10)),
  inference(subsumption_resolution,[],[f22165,f22025])).
fof(f22025,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl6 | $spl10)),
  inference(subsumption_resolution,[],[f22024,f1055])).
fof(f22024,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl6 | $spl10)),
  inference(subsumption_resolution,[],[f22012,f968])).
fof(f22012,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl6 | $spl10)),
  inference(resolution,[],[f21783,f915])).
fof(f21783,plain,(
  ( ! [X202] : (sP12(X202,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | ($spl6 | $spl10)),
  inference(forward_demodulation,[],[f21555,f1058])).
fof(f21555,plain,(
  ( ! [X202] : (sP12(X202,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))))) ) | ($spl6 | $spl10)),
  inference(backward_demodulation,[],[f1123,f3084])).
fof(f22165,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl10)),
  inference(resolution,[],[f22025,f1021])).
fof(f1021,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f673])).
fof(f673,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f384])).
fof(f384,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f21])).
fof(f21,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.list.FILTER_NEQ_ID','thm.bool.NOT_CLAUSES|split|2')).
fof(f22310,plain,(
  $false | ($spl6 | $spl10)),
  inference(subsumption_resolution,[],[f22164,f22025])).
fof(f22164,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl10)),
  inference(resolution,[],[f22025,f1019])).
fof(f1019,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f672])).
fof(f672,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f383])).
fof(f383,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f20])).
fof(f20,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.list.FILTER_NEQ_ID','thm.bool.NOT_CLAUSES|split|1')).
fof(f22309,plain,(
  $false | ($spl6 | $spl10)),
  inference(subsumption_resolution,[],[f22161,f22025])).
fof(f22161,plain,(
  ( ! [X326] : (~p(s(bool,X326))) ) | ($spl6 | $spl10)),
  inference(resolution,[],[f22025,f1013])).
fof(f1013,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f481])).
fof(f481,plain,(
  ! [X0] : ((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f376])).
fof(f376,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(flattening,[],[f375])).
fof(f375,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(rectify,[],[f79])).
fof(f79,axiom,(
  ! [X12] : ((p(s(bool,X12)) => p(s(bool,'F_0'))) => ~p(s(bool,X12)))),
  file('thm.list.FILTER_NEQ_ID','thm.bool.IMP_F')).
fof(f22308,plain,(
  $false | ($spl6 | $spl10)),
  inference(subsumption_resolution,[],[f22154,f22025])).
fof(f22154,plain,(
  ( ! [X319] : (~p(s(bool,X319))) ) | ($spl6 | $spl10)),
  inference(resolution,[],[f22025,f951])).
fof(f951,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f637])).
fof(f637,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f636])).
fof(f636,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f470])).
fof(f470,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f347])).
fof(f347,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f346])).
fof(f346,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X12] : ((p(s(bool,X12)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X12)))),
  file('thm.list.FILTER_NEQ_ID','thm.bool.IMP_CLAUSES|split|4')).
fof(f22163,plain,(
  $false | ($spl6 | $spl10)),
  inference(resolution,[],[f22025,f1017])).
fof(f22185,plain,(
  $false | ($spl6 | $spl10)),
  inference(resolution,[],[f22025,f1055])).
fof(f21464,plain,(
  ( ! [X212,X213] : (p(s(bool,X212)) | p(s(bool,X213)) | s(bool,X213) != s(bool,'const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl12),
  inference(resolution,[],[f1126,f1075])).
fof(f21366,plain,(
  ( ! [X61,X59,X60] : (p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X59) != s(bool,'const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X60) != s(bool,X61)) ) | $spl12),
  inference(resolution,[],[f1126,f790])).
fof(f21363,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | p(s(bool,X52)) | s(bool,X52) != s(bool,'const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl12),
  inference(resolution,[],[f1126,f789])).
fof(f21362,plain,(
  ( ! [X47,X48,X49] : (p(s(bool,X47)) | p(s(bool,X48)) | s(bool,X47) != s(bool,X49) | s(bool,X48) != s(bool,'const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl12),
  inference(resolution,[],[f1126,f788])).
fof(f21361,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,X44)) | ~p(s(bool,X45)) | s(bool,X44) != s(bool,'const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X45) != s(bool,X46)) ) | $spl12),
  inference(resolution,[],[f1126,f786])).
fof(f21359,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X39) != s(bool,'const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl12),
  inference(resolution,[],[f1126,f784])).
fof(f20909,plain,(
  ( ! [X26] : (p(s(bool,X26)) | s(bool,X26) != s(bool,'const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) ) | $spl204),
  inference(cnf_transformation,[],[f20909_D])).
fof(f20909_D,plain,(
  ( ! [X26] : (p(s(bool,X26)) | s(bool,X26) != s(bool,'const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f21142,plain,(
  ( ! [X200,X199] : (p(s(bool,X199)) | p(s(bool,X200)) | s(bool,X200) != s(bool,'const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) ) | $spl14),
  inference(resolution,[],[f1129,f1075])).
fof(f21044,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,X46) != s(bool,'const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),sK14))) | s(bool,X47) != s(bool,X48)) ) | $spl14),
  inference(resolution,[],[f1129,f790])).
fof(f21041,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,X37)) | ~p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X39) != s(bool,'const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) ) | $spl14),
  inference(resolution,[],[f1129,f789])).
fof(f21040,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X34) != s(bool,X36) | s(bool,X35) != s(bool,'const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) ) | $spl14),
  inference(resolution,[],[f1129,f788])).
fof(f21039,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),sK14))) | s(bool,X32) != s(bool,X33)) ) | $spl14),
  inference(resolution,[],[f1129,f786])).
fof(f21037,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) ) | $spl14),
  inference(resolution,[],[f1129,f784])).
fof(f1122,plain,(
  ~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(sK13,bool),sK14))),s('type.list.list'(sK13),sK15)))) | $spl9),
  inference(cnf_transformation,[],[f1122_D])).
fof(f1122_D,plain,(
  ~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(sK13,bool),sK14))),s('type.list.list'(sK13),sK15)))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f1123,plain,(
  s('type.list.list'(sK13),'const.list.NIL_0') = s('type.list.list'(sK13),sK15) | $spl10),
  inference(cnf_transformation,[],[f1123_D])).
fof(f1123_D,plain,(
  s('type.list.list'(sK13),'const.list.NIL_0') = s('type.list.list'(sK13),sK15) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f19845,plain,(
  s('type.list.list'(sK13),'const.list.NIL_0') = s('type.list.list'(sK13),sK15) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(sK13,bool),sK14))),s('type.list.list'(sK13),sK15)))) | $spl0),
  inference(superposition,[],[f1086,f1065])).
fof(f1065,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(X0,bool),X1))),s('type.list.list'(X0),X2))))) )),
  inference(cnf_transformation,[],[f706])).
fof(f706,plain,(
  ! [X0,X1,X2] : ((s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(X0,bool),X1))),s('type.list.list'(X0),X2))))) & (~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(X0,bool),X1))),s('type.list.list'(X0),X2)))) | s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))))),
  inference(nnf_transformation,[],[f422])).
fof(f422,plain,(
  ! [X0,X1,X2] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) <=> p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(X0,bool),X1))),s('type.list.list'(X0),X2)))))),
  inference(flattening,[],[f421])).
fof(f421,plain,(
  ! [X0] : ! [X1,X2] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) <=> p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(X0,bool),X1))),s('type.list.list'(X0),X2)))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : ! [X7,X8] : (s('type.list.list'(X6),'const.list.NIL_0') = s('type.list.list'(X6),'const.list.FILTER_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X8))) <=> p(s(bool,'const.list.EVERY_2'(s(fun(X6,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(X6,bool),X7))),s('type.list.list'(X6),X8)))))),
  file('thm.list.FILTER_NEQ_ID','thm.list.FILTER_EQ_NIL')).
fof(f1126,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) | $spl12),
  inference(cnf_transformation,[],[f1126_D])).
fof(f1126_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f19846,plain,(
  s('type.list.list'(sK13),'const.list.NIL_0') = s('type.list.list'(sK13),sK15) | p(s(bool,'const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) | $spl0),
  inference(superposition,[],[f1086,f1041])).
fof(f1041,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK53(X2,X1,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) )),
  inference(cnf_transformation,[],[f691])).
fof(f1129,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) | $spl14),
  inference(cnf_transformation,[],[f1129_D])).
fof(f1129_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f20926,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) | s('type.list.list'(sK13),'const.list.NIL_0') = s('type.list.list'(sK13),sK15) | $spl0),
  inference(forward_demodulation,[],[f19847,f760])).
fof(f19847,plain,(
  s('type.list.list'(sK13),'const.list.NIL_0') = s('type.list.list'(sK13),sK15) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK53(sK15,sK14,sK13))))) | $spl0),
  inference(superposition,[],[f1086,f1042])).
fof(f20921,plain,(
  s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)),s('type.list.list'(sK13),sK60('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)))) | $spl206),
  inference(cnf_transformation,[],[f20921_D])).
fof(f20921_D,plain,(
  s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)),s('type.list.list'(sK13),sK60('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)))) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f20925,plain,(
  s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)),s('type.list.list'(sK13),sK60('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)))) | s('type.list.list'(sK13),'const.list.NIL_0') = s('type.list.list'(sK13),sK15) | $spl0),
  inference(forward_demodulation,[],[f19848,f1086])).
fof(f19848,plain,(
  s('type.list.list'(sK13),'const.list.NIL_0') = s('type.list.list'(sK13),sK15) | s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)),s('type.list.list'(sK13),sK60('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)))) | $spl0),
  inference(superposition,[],[f1086,f1074])).
fof(f19858,plain,(
  s('type.list.list'(sK13),'const.list.NIL_0') = s('type.list.list'(sK13),sK15) | p(s(bool,'const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) | $spl0),
  inference(superposition,[],[f1041,f1086])).
fof(f20924,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) | s('type.list.list'(sK13),'const.list.NIL_0') = s('type.list.list'(sK13),sK15) | $spl0),
  inference(forward_demodulation,[],[f19859,f760])).
fof(f19859,plain,(
  s('type.list.list'(sK13),'const.list.NIL_0') = s('type.list.list'(sK13),sK15) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK53(sK15,sK14,sK13))))) | $spl0),
  inference(superposition,[],[f1042,f1086])).
fof(f1124,plain,(
  s('type.list.list'(sK13),'const.list.NIL_0') != s('type.list.list'(sK13),sK15) | $spl11),
  inference(cnf_transformation,[],[f1124_D])).
fof(f1124_D,plain,(
  s('type.list.list'(sK13),'const.list.NIL_0') != s('type.list.list'(sK13),sK15) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f20923,plain,(
  ( ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(sK13,X3),s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.bool.IN_2'(s(sK13,X3),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl208),
  inference(cnf_transformation,[],[f20923_D])).
fof(f20923_D,plain,(
  ( ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(sK13,X3),s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.bool.IN_2'(s(sK13,X3),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))))) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f20922,plain,(
  ( ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(sK13,X3),s(fun(sK13,bool),sK14)))) | s('type.list.list'(sK13),'const.list.NIL_0') != s('type.list.list'(sK13),sK15) | ~p(s(bool,'const.bool.IN_2'(s(sK13,X3),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl0),
  inference(forward_demodulation,[],[f19860,f760])).
fof(f19860,plain,(
  ( ! [X3] : (s('type.list.list'(sK13),'const.list.NIL_0') != s('type.list.list'(sK13),sK15) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X3)))) | ~p(s(bool,'const.bool.IN_2'(s(sK13,X3),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl0),
  inference(superposition,[],[f1043,f1086])).
fof(f1043,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2))))))) )),
  inference(cnf_transformation,[],[f691])).
fof(f1121,plain,(
  p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(sK13,bool),sK14))),s('type.list.list'(sK13),sK15)))) | $spl8),
  inference(cnf_transformation,[],[f1121_D])).
fof(f1121_D,plain,(
  p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(sK13,bool),sK14))),s('type.list.list'(sK13),sK15)))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f19862,plain,(
  s('type.list.list'(sK13),'const.list.NIL_0') != s('type.list.list'(sK13),sK15) | p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(sK13,bool),sK14))),s('type.list.list'(sK13),sK15)))) | $spl0),
  inference(superposition,[],[f1064,f1086])).
fof(f1064,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(X0,bool),X1))),s('type.list.list'(X0),X2)))) | s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) )),
  inference(cnf_transformation,[],[f706])).
fof(f19863,plain,(
  s('type.list.list'(sK13),'const.list.NIL_0') = s('type.list.list'(sK13),sK15) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(sK13,bool),sK14))),s('type.list.list'(sK13),sK15)))) | $spl0),
  inference(superposition,[],[f1065,f1086])).
fof(f20919,plain,(
  s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)),s('type.list.list'(sK13),sK60('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)))) | s('type.list.list'(sK13),'const.list.NIL_0') = s('type.list.list'(sK13),sK15) | $spl0),
  inference(forward_demodulation,[],[f19936,f1086])).
fof(f19936,plain,(
  s('type.list.list'(sK13),'const.list.NIL_0') = s('type.list.list'(sK13),sK15) | s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)),s('type.list.list'(sK13),sK60('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)))) | $spl0),
  inference(superposition,[],[f1074,f1086])).
fof(f20853,plain,(
  ( ! [X200,X199] : (p(s(bool,X199)) | p(s(bool,X200)) | s(bool,X200) != s(bool,'const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) ) | ($spl0 | $spl11)),
  inference(resolution,[],[f4982,f1075])).
fof(f20755,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,X46) != s(bool,'const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),sK14))) | s(bool,X47) != s(bool,X48)) ) | ($spl0 | $spl11)),
  inference(resolution,[],[f4982,f790])).
fof(f20752,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,X37)) | ~p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X39) != s(bool,'const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) ) | ($spl0 | $spl11)),
  inference(resolution,[],[f4982,f789])).
fof(f20751,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X34) != s(bool,X36) | s(bool,X35) != s(bool,'const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) ) | ($spl0 | $spl11)),
  inference(resolution,[],[f4982,f788])).
fof(f20750,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),sK14))) | s(bool,X32) != s(bool,X33)) ) | ($spl0 | $spl11)),
  inference(resolution,[],[f4982,f786])).
fof(f20748,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) ) | ($spl0 | $spl11)),
  inference(resolution,[],[f4982,f784])).
fof(f20624,plain,(
  ( ! [X200,X199] : (p(s(bool,X199)) | p(s(bool,X200)) | s(bool,X200) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl0),
  inference(resolution,[],[f4976,f1075])).
fof(f20526,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,X46) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | s(bool,X47) != s(bool,X48)) ) | $spl0),
  inference(resolution,[],[f4976,f790])).
fof(f20523,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,X37)) | ~p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X39) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl0),
  inference(resolution,[],[f4976,f789])).
fof(f20522,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X34) != s(bool,X36) | s(bool,X35) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl0),
  inference(resolution,[],[f4976,f788])).
fof(f20521,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | s(bool,X32) != s(bool,X33)) ) | $spl0),
  inference(resolution,[],[f4976,f786])).
fof(f20519,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl0),
  inference(resolution,[],[f4976,f784])).
fof(f20504,plain,(
  ( ! [X1] : (p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) | p(s(bool,X1))) ) | $spl0),
  inference(subsumption_resolution,[],[f20489,f968])).
fof(f20489,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) | p(s(bool,X1))) ) | $spl0),
  inference(resolution,[],[f5255,f911])).
fof(f20503,plain,(
  ( ! [X0] : (p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) | p(s(bool,X0))) ) | $spl0),
  inference(subsumption_resolution,[],[f20488,f968])).
fof(f20488,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) | p(s(bool,X0))) ) | $spl0),
  inference(resolution,[],[f5255,f915])).
fof(f20483,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))))) ) | $spl0),
  inference(subsumption_resolution,[],[f20468,f968])).
fof(f20468,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))))) ) | $spl0),
  inference(resolution,[],[f5254,f911])).
fof(f20482,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))))) ) | $spl0),
  inference(subsumption_resolution,[],[f20467,f968])).
fof(f20467,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))))) ) | $spl0),
  inference(resolution,[],[f5254,f915])).
fof(f1089,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)))) | $spl4),
  inference(cnf_transformation,[],[f1089_D])).
fof(f1089_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f20461,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)),X5) | ~sP6(X5,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)),X4) | ~sP8(X4,X5,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)))) ) | $spl202),
  inference(cnf_transformation,[],[f20461_D])).
fof(f20461_D,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)),X5) | ~sP6(X5,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)),X4) | ~sP8(X4,X5,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)))) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f20448,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)),X5) | p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)))) | ~sP8(X4,X5,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14))) | ~sP6(X5,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)),X4)) ) | $spl5),
  inference(resolution,[],[f2475,f847])).
fof(f847,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X1)) | ~sP8(X0,X2,X1) | ~sP7(X1,X2,X0) | ~sP6(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f591])).
fof(f2475,plain,(
  ( ! [X144,X143] : (sP7('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)),X143,X144)) ) | $spl5),
  inference(resolution,[],[f1090,f856])).
fof(f856,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f597])).
fof(f20418,plain,(
  ( ! [X6,X7] : (sP9('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)),X6,X7) | ~sP6(X7,X6,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14))) | ~sP7(X6,X7,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)))) ) | $spl198),
  inference(cnf_transformation,[],[f20418_D])).
fof(f20418_D,plain,(
  ( ! [X6,X7] : (sP9('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)),X6,X7) | ~sP6(X7,X6,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14))) | ~sP7(X6,X7,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)))) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f20405,plain,(
  ( ! [X6,X7] : (sP9('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)),X6,X7) | p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)))) | ~sP7(X6,X7,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14))) | ~sP6(X7,X6,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)))) ) | $spl5),
  inference(resolution,[],[f2472,f848])).
fof(f2472,plain,(
  ( ! [X138,X137] : (sP8('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)),X137,X138)) ) | $spl5),
  inference(resolution,[],[f1090,f852])).
fof(f7560,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))))) ) | $spl0),
  inference(subsumption_resolution,[],[f7542,f968])).
fof(f7542,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))))) ) | $spl0),
  inference(resolution,[],[f5254,f911])).
fof(f7557,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))))) ) | $spl0),
  inference(subsumption_resolution,[],[f7541,f968])).
fof(f7541,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))))) ) | $spl0),
  inference(resolution,[],[f5254,f915])).
fof(f19806,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl185),
  inference(cnf_transformation,[],[f19806_D])).
fof(f19806_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl185),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl185])])).
fof(f19831,plain,(
  ( ! [X28] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X28))) ) | $spl76),
  inference(subsumption_resolution,[],[f10314,f953])).
fof(f10314,plain,(
  ( ! [X28] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,X28) | p(s(bool,X28))) ) | $spl76),
  inference(superposition,[],[f8667,f953])).
fof(f19829,plain,(
  ( ! [X30] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X30))) ) | $spl76),
  inference(subsumption_resolution,[],[f10316,f953])).
fof(f10316,plain,(
  ( ! [X30] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,X30) | p(s(bool,X30))) ) | $spl76),
  inference(superposition,[],[f8667,f999])).
fof(f999,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f668])).
fof(f19828,plain,(
  ( ! [X31] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X31))) ) | $spl76),
  inference(subsumption_resolution,[],[f10317,f953])).
fof(f10317,plain,(
  ( ! [X31] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,X31) | p(s(bool,X31))) ) | $spl76),
  inference(superposition,[],[f8667,f1001])).
fof(f1001,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f669])).
fof(f669,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f366])).
fof(f366,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f365])).
fof(f365,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.list.FILTER_NEQ_ID','thm.bool.EQ_CLAUSES|split|3')).
fof(f10347,plain,(
  ( ! [X34] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X34))) ) | $spl76),
  inference(trivial_inequality_removal,[],[f10320])).
fof(f10320,plain,(
  ( ! [X34] : (s(bool,X34) != s(bool,X34) | s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X34))) ) | $spl76),
  inference(superposition,[],[f8667,f1005])).
fof(f10346,plain,(
  ( ! [X35] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X35))) ) | $spl76),
  inference(trivial_inequality_removal,[],[f10321])).
fof(f10321,plain,(
  ( ! [X35] : (s(bool,X35) != s(bool,X35) | s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X35))) ) | $spl76),
  inference(superposition,[],[f8667,f1003])).
fof(f19807,plain,(
  ( ! [X48,X49] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X49)) | ~p(s(bool,X48))) ) | $spl76),
  inference(subsumption_resolution,[],[f19187,f932])).
fof(f19187,plain,(
  ( ! [X48,X49] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X48) != s(bool,X49) | ~p(s(bool,X49)) | ~p(s(bool,X48))) ) | $spl76),
  inference(superposition,[],[f10282,f1003])).
fof(f10282,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X0))) ) | $spl76),
  inference(superposition,[],[f8667,f1005])).
fof(f19804,plain,(
  ( ! [X50,X51] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X51)) | ~p(s(bool,X50))) ) | $spl76),
  inference(subsumption_resolution,[],[f19188,f932])).
fof(f19188,plain,(
  ( ! [X50,X51] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X50) != s(bool,X51) | ~p(s(bool,X51)) | ~p(s(bool,X50))) ) | $spl76),
  inference(superposition,[],[f10282,f1005])).
fof(f19800,plain,(
  ( ! [X0,X1] : (~sP1(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X1) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f17517,f1058])).
fof(f17517,plain,(
  ( ! [X0,X1] : (~sP1(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f11036,f1074])).
fof(f19798,plain,(
  ( ! [X0,X1] : (sP1('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,X1)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f17627,f1058])).
fof(f17627,plain,(
  ( ! [X0,X1] : (sP1('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1) | p(s(bool,X1)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f11045,f1074])).
fof(f19796,plain,(
  ( ! [X0,X1] : (sP1(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X1) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f17672,f1058])).
fof(f17672,plain,(
  ( ! [X0,X1] : (sP1(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f11046,f1074])).
fof(f19794,plain,(
  ( ! [X0,X1] : (sP1('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f17714,f1058])).
fof(f17714,plain,(
  ( ! [X0,X1] : (sP1('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f11047,f1074])).
fof(f19792,plain,(
  ( ! [X0,X1] : (~sP3(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f17761,f1058])).
fof(f17761,plain,(
  ( ! [X0,X1] : (~sP3(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f11051,f1074])).
fof(f19790,plain,(
  ( ! [X0,X1] : (~sP3(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X1) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f17803,f1058])).
fof(f17803,plain,(
  ( ! [X0,X1] : (~sP3(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f11052,f1074])).
fof(f19788,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f17845,f1058])).
fof(f17845,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f11058,f1074])).
fof(f19786,plain,(
  ( ! [X0,X1] : (sP3('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,X1)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f17885,f1058])).
fof(f17885,plain,(
  ( ! [X0,X1] : (sP3('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1) | ~p(s(bool,X1)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f11059,f1074])).
fof(f19784,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X1) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f17934,f1058])).
fof(f17934,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f11060,f1074])).
fof(f19782,plain,(
  ( ! [X0,X1] : (sP3('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f17974,f1058])).
fof(f17974,plain,(
  ( ! [X0,X1] : (sP3('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f11061,f1074])).
fof(f19780,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,X1)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f18021,f1058])).
fof(f18021,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1) | p(s(bool,X1)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f11065,f1074])).
fof(f19778,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f18065,f1058])).
fof(f18065,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f11066,f1074])).
fof(f19776,plain,(
  ( ! [X0,X1] : (~sP10(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f18968,f1058])).
fof(f18968,plain,(
  ( ! [X0,X1] : (~sP10(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f11095,f1074])).
fof(f19774,plain,(
  ( ! [X0,X1] : (~sP10(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X1) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f19122,f1058])).
fof(f19122,plain,(
  ( ! [X0,X1] : (~sP10(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f11096,f1074])).
fof(f19737,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f19736,f1017])).
fof(f19736,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f19564,f7465])).
fof(f7465,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),sK14))) | $spl1),
  inference(resolution,[],[f1516,f953])).
fof(f1516,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) | $spl1),
  inference(forward_demodulation,[],[f1358,f760])).
fof(f1358,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK58(sK15,sK14,sK13))))) | $spl1),
  inference(resolution,[],[f1117,f1054])).
fof(f1054,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK58(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f703])).
fof(f1117,plain,(
  ~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) | $spl1),
  inference(trivial_inequality_removal,[],[f1106])).
fof(f1106,plain,(
  s('type.list.list'(sK13),sK15) != s('type.list.list'(sK13),sK15) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) | $spl1),
  inference(superposition,[],[f1085,f1067])).
fof(f1067,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))) )),
  inference(cnf_transformation,[],[f707])).
fof(f19564,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) | ($spl1 | $spl2)),
  inference(resolution,[],[f7577,f1087])).
fof(f7577,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) | $spl1),
  inference(resolution,[],[f1117,f1053])).
fof(f1053,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK58(X2,X1,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2))))))) )),
  inference(cnf_transformation,[],[f703])).
fof(f19563,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f19457,f19250])).
fof(f19250,plain,(
  ( ! [X19] : (~p(s(bool,X19))) ) | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f19249,f1005])).
fof(f19249,plain,(
  ( ! [X19] : (s(bool,'T_0') != s(bool,X19) | ~p(s(bool,X19))) ) | ($spl6 | $spl46)),
  inference(forward_demodulation,[],[f3464,f11130])).
fof(f19457,plain,(
  ( ! [X549,X548,X550] : (p(s(bool,'const.bool.IN_2'(s(X548,sK58(X549,X550,X548)),s(fun(X548,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X548),X549))))))) ) | ($spl6 | $spl46)),
  inference(resolution,[],[f19250,f1053])).
fof(f19562,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f19405,f19250])).
fof(f19405,plain,(
  ( ! [X377,X376,X375] : (p(s(bool,'const.list.EVERY_2'(s(fun(X375,bool),X376),s('type.list.list'(X375),X377))))) ) | ($spl6 | $spl46)),
  inference(resolution,[],[f19250,f1053])).
fof(f19557,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f19379,f19250])).
fof(f19379,plain,(
  p(s(bool,'F_0')) | ($spl6 | $spl46)),
  inference(resolution,[],[f19250,f1020])).
fof(f1020,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f673])).
fof(f19556,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f19378,f19250])).
fof(f19378,plain,(
  p(s(bool,'F_0')) | ($spl6 | $spl46)),
  inference(resolution,[],[f19250,f1018])).
fof(f1018,plain,(
  p(s(bool,'F_0')) | p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f672])).
fof(f19555,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f19373,f19250])).
fof(f19373,plain,(
  ( ! [X307] : (p(s(bool,X307))) ) | ($spl6 | $spl46)),
  inference(resolution,[],[f19250,f965])).
fof(f965,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f647])).
fof(f647,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f646])).
fof(f646,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f475])).
fof(f475,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f352])).
fof(f352,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X12] : ((p(s(bool,X12)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.list.FILTER_NEQ_ID','thm.bool.IMP_CLAUSES|split|1')).
fof(f19554,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f19371,f19250])).
fof(f19371,plain,(
  p(s(bool,'F_0')) | ($spl6 | $spl46)),
  inference(resolution,[],[f19250,f962])).
fof(f962,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f645])).
fof(f645,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f644])).
fof(f644,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f474])).
fof(f474,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f351])).
fof(f351,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X12] : ((p(s(bool,'F_0')) => p(s(bool,X12))) <=> p(s(bool,'T_0')))),
  file('thm.list.FILTER_NEQ_ID','thm.bool.IMP_CLAUSES|split|2')).
fof(f19553,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f19370,f19250])).
fof(f19370,plain,(
  ( ! [X305] : (p(s(bool,X305))) ) | ($spl6 | $spl46)),
  inference(resolution,[],[f19250,f959])).
fof(f959,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f643])).
fof(f643,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f642])).
fof(f642,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f473])).
fof(f473,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f350])).
fof(f350,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X12] : ((p(s(bool,'T_0')) => p(s(bool,X12))) <=> p(s(bool,X12)))),
  file('thm.list.FILTER_NEQ_ID','thm.bool.IMP_CLAUSES|split|0')).
fof(f19552,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f19368,f19250])).
fof(f19368,plain,(
  ( ! [X303] : (p(s(bool,X303))) ) | ($spl6 | $spl46)),
  inference(resolution,[],[f19250,f956])).
fof(f956,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f641])).
fof(f641,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f640])).
fof(f640,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f472])).
fof(f472,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f349])).
fof(f349,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X12] : ((p(s(bool,X12)) => p(s(bool,X12))) <=> p(s(bool,'T_0')))),
  file('thm.list.FILTER_NEQ_ID','thm.bool.IMP_CLAUSES|split|3')).
fof(f19380,plain,(
  $false | ($spl6 | $spl46)),
  inference(resolution,[],[f19250,f1024])).
fof(f1024,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f12])).
fof(f12,axiom,(
  p(s(bool,'T_0'))),
  file('thm.list.FILTER_NEQ_ID','thm.bool.TRUTH')).
fof(f19381,plain,(
  $false | ($spl6 | $spl46)),
  inference(resolution,[],[f19250,f1025])).
fof(f1025,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.list.FILTER_NEQ_ID',aHL_TRUTH)).
fof(f19382,plain,(
  $false | ($spl6 | $spl46)),
  inference(resolution,[],[f19250,f1080])).
fof(f19410,plain,(
  $false | ($spl4 | $spl6 | $spl46)),
  inference(resolution,[],[f19250,f8706])).
fof(f8706,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.INSERT_2'(s(sK13,X0),s(fun(sK13,bool),sK14))))))) ) | $spl4),
  inference(resolution,[],[f1089,f1040])).
fof(f15818,plain,(
  ( ! [X38,X36] : (p(s(bool,X36)) | s(bool,X36) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.INSERT_2'(s(sK13,X38),s(fun(sK13,bool),sK14)))))) ) | $spl154),
  inference(cnf_transformation,[],[f15818_D])).
fof(f15818_D,plain,(
  ( ! [X38,X36] : (p(s(bool,X36)) | s(bool,X36) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.INSERT_2'(s(sK13,X38),s(fun(sK13,bool),sK14)))))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f15744,plain,(
  ( ! [X306,X305,X307] : (p(s(bool,X305)) | p(s(bool,X306)) | s(bool,X306) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.INSERT_2'(s(sK13,X307),s(fun(sK13,bool),sK14)))))) ) | $spl4),
  inference(resolution,[],[f8706,f1075])).
fof(f15646,plain,(
  ( ! [X64,X62,X65,X63] : (p(s(bool,X62)) | ~p(s(bool,X63)) | s(bool,X62) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.INSERT_2'(s(sK13,X64),s(fun(sK13,bool),sK14))))) | s(bool,X63) != s(bool,X65)) ) | $spl4),
  inference(resolution,[],[f8706,f790])).
fof(f15643,plain,(
  ( ! [X54,X52,X53,X51] : (p(s(bool,X51)) | ~p(s(bool,X52)) | p(s(bool,X53)) | s(bool,X53) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.INSERT_2'(s(sK13,X54),s(fun(sK13,bool),sK14)))))) ) | $spl4),
  inference(resolution,[],[f8706,f789])).
fof(f15642,plain,(
  ( ! [X47,X50,X48,X49] : (p(s(bool,X47)) | p(s(bool,X48)) | s(bool,X47) != s(bool,X49) | s(bool,X48) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.INSERT_2'(s(sK13,X50),s(fun(sK13,bool),sK14)))))) ) | $spl4),
  inference(resolution,[],[f8706,f788])).
fof(f15641,plain,(
  ( ! [X45,X43,X46,X44] : (p(s(bool,X43)) | ~p(s(bool,X44)) | s(bool,X43) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.INSERT_2'(s(sK13,X45),s(fun(sK13,bool),sK14))))) | s(bool,X44) != s(bool,X46)) ) | $spl4),
  inference(resolution,[],[f8706,f786])).
fof(f15639,plain,(
  ( ! [X37,X35,X38,X36] : (p(s(bool,X35)) | p(s(bool,X36)) | s(bool,X35) != s(bool,X37) | s(bool,X36) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.INSERT_2'(s(sK13,X38),s(fun(sK13,bool),sK14)))))) ) | $spl4),
  inference(resolution,[],[f8706,f784])).
fof(f15137,plain,(
  s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),'const.list.NIL_0'))) | s('type.list.list'(sK13),sK60(sK15,sK13)) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK60(sK15,sK13),sK13)),s('type.list.list'(sK13),sK60(sK60(sK15,sK13),sK13)))) | $spl1),
  inference(superposition,[],[f8136,f1074])).
fof(f8136,plain,(
  s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13)))) | $spl1),
  inference(subsumption_resolution,[],[f8115,f1068])).
fof(f1068,plain,(
  ( ! [X0,X1] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f426])).
fof(f426,plain,(
  ! [X0,X1] : s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(flattening,[],[f425])).
fof(f425,plain,(
  ! [X0] : ! [X1] : s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : ! [X7] : s('type.list.list'(X6),'const.list.NIL_0') = s('type.list.list'(X6),'const.list.FILTER_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),'const.list.NIL_0')))),
  file('thm.list.FILTER_NEQ_ID','thm.list.FILTER|split|0')).
fof(f8115,plain,(
  s('type.list.list'(sK13),'const.list.NIL_0') != s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13)))) | $spl1),
  inference(superposition,[],[f1085,f1074])).
fof(f14028,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13)))) | $spl6),
  inference(forward_demodulation,[],[f13580,f1058])).
fof(f13580,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13)))) | $spl6),
  inference(superposition,[],[f11130,f1074])).
fof(f13447,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | $spl151),
  inference(cnf_transformation,[],[f13447_D])).
fof(f13447_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) <=> ~$spl151),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl151])])).
fof(f13445,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13)))) | $spl6),
  inference(forward_demodulation,[],[f13406,f1058])).
fof(f13406,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13)))) | $spl6),
  inference(superposition,[],[f11128,f1074])).
fof(f13396,plain,(
  ( ! [X0] : (sP11(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f13381,f1058])).
fof(f13381,plain,(
  ( ! [X0] : (sP11(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f11105,f1074])).
fof(f13363,plain,(
  ( ! [X0] : (sP11('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f13348,f1058])).
fof(f13348,plain,(
  ( ! [X0] : (sP11('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f11104,f1074])).
fof(f13330,plain,(
  ( ! [X0] : (~sP11(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f13313,f1058])).
fof(f13313,plain,(
  ( ! [X0] : (~sP11(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f11103,f1074])).
fof(f13297,plain,(
  ( ! [X0] : (~sP11('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f13280,f1058])).
fof(f13280,plain,(
  ( ! [X0] : (~sP11('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f11102,f1074])).
fof(f13217,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | ($spl1 | $spl32)),
  inference(subsumption_resolution,[],[f13199,f4399])).
fof(f4399,plain,(
  ( ! [X190] : (sP12('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X190)) ) | $spl32),
  inference(resolution,[],[f1537,f913])).
fof(f1537,plain,(
  p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')))) | $spl32),
  inference(cnf_transformation,[],[f1537_D])).
fof(f1537_D,plain,(
  p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f13199,plain,(
  ( ! [X0] : (~sP12('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl1),
  inference(superposition,[],[f1528,f1074])).
fof(f1528,plain,(
  ( ! [X183] : (~sP12('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),X183) | p(s(bool,X183))) ) | $spl1),
  inference(subsumption_resolution,[],[f1449,f968])).
fof(f1449,plain,(
  ( ! [X183] : (p(s(bool,'F_0')) | p(s(bool,X183)) | ~sP12('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),X183)) ) | $spl1),
  inference(resolution,[],[f1117,f911])).
fof(f13183,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | ($spl1 | $spl32)),
  inference(subsumption_resolution,[],[f13164,f4398])).
fof(f4398,plain,(
  ( ! [X189] : (sP12(X189,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')))) ) | $spl32),
  inference(resolution,[],[f1537,f912])).
fof(f13164,plain,(
  ( ! [X0] : (~sP12(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl1),
  inference(superposition,[],[f1527,f1074])).
fof(f1527,plain,(
  ( ! [X182] : (~sP12(X182,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | p(s(bool,X182))) ) | $spl1),
  inference(subsumption_resolution,[],[f1448,f968])).
fof(f1448,plain,(
  ( ! [X182] : (p(s(bool,'F_0')) | p(s(bool,X182)) | ~sP12(X182,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl1),
  inference(resolution,[],[f1117,f911])).
fof(f13107,plain,(
  s(bool,'T_0') != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | $spl141),
  inference(cnf_transformation,[],[f13107_D])).
fof(f13107_D,plain,(
  s(bool,'T_0') != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) <=> ~$spl141),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl141])])).
fof(f13072,plain,(
  s(bool,'T_0') != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13)))) | $spl1),
  inference(superposition,[],[f1469,f1074])).
fof(f1469,plain,(
  s(bool,'T_0') != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | $spl1),
  inference(resolution,[],[f1117,f1002])).
fof(f1002,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f670])).
fof(f11226,plain,(
  s(bool,'F_0') = s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | $spl110),
  inference(cnf_transformation,[],[f11226_D])).
fof(f11226_D,plain,(
  s(bool,'F_0') = s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f12322,plain,(
  s(bool,'F_0') = s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13)))) | $spl1),
  inference(superposition,[],[f1458,f1074])).
fof(f1458,plain,(
  s(bool,'F_0') = s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | $spl1),
  inference(resolution,[],[f1117,f953])).
fof(f12313,plain,(
  ( ! [X0] : (sP11('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0) | ~p(s(bool,X0))) ) | $spl138),
  inference(cnf_transformation,[],[f12313_D])).
fof(f12313_D,plain,(
  ( ! [X0] : (sP11('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0) | ~p(s(bool,X0))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f12300,plain,(
  ( ! [X0] : (sP11('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl1),
  inference(superposition,[],[f1443,f1074])).
fof(f1443,plain,(
  ( ! [X177] : (sP11('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),X177) | ~p(s(bool,X177))) ) | $spl1),
  inference(resolution,[],[f1117,f900])).
fof(f12286,plain,(
  $false | ($spl4 | $spl42)),
  inference(subsumption_resolution,[],[f12183,f11976])).
fof(f11976,plain,(
  ( ! [X14] : (~p(s(bool,X14))) ) | ($spl4 | $spl42)),
  inference(subsumption_resolution,[],[f11975,f1005])).
fof(f11975,plain,(
  ( ! [X14] : (s(bool,'T_0') != s(bool,X14) | ~p(s(bool,X14))) ) | ($spl4 | $spl42)),
  inference(forward_demodulation,[],[f2572,f8817])).
fof(f8817,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14))) | $spl4),
  inference(resolution,[],[f1089,f1003])).
fof(f12183,plain,(
  ( ! [X549,X548,X550] : (p(s(bool,'const.bool.IN_2'(s(X548,sK58(X549,X550,X548)),s(fun(X548,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X548),X549))))))) ) | ($spl4 | $spl42)),
  inference(resolution,[],[f11976,f1053])).
fof(f12285,plain,(
  $false | ($spl4 | $spl42)),
  inference(subsumption_resolution,[],[f12131,f11976])).
fof(f12131,plain,(
  ( ! [X377,X376,X375] : (p(s(bool,'const.list.EVERY_2'(s(fun(X375,bool),X376),s('type.list.list'(X375),X377))))) ) | ($spl4 | $spl42)),
  inference(resolution,[],[f11976,f1053])).
fof(f12280,plain,(
  $false | ($spl4 | $spl42)),
  inference(subsumption_resolution,[],[f12105,f11976])).
fof(f12105,plain,(
  p(s(bool,'F_0')) | ($spl4 | $spl42)),
  inference(resolution,[],[f11976,f1020])).
fof(f12279,plain,(
  $false | ($spl4 | $spl42)),
  inference(subsumption_resolution,[],[f12104,f11976])).
fof(f12104,plain,(
  p(s(bool,'F_0')) | ($spl4 | $spl42)),
  inference(resolution,[],[f11976,f1018])).
fof(f12278,plain,(
  $false | ($spl4 | $spl42)),
  inference(subsumption_resolution,[],[f12099,f11976])).
fof(f12099,plain,(
  ( ! [X307] : (p(s(bool,X307))) ) | ($spl4 | $spl42)),
  inference(resolution,[],[f11976,f965])).
fof(f12277,plain,(
  $false | ($spl4 | $spl42)),
  inference(subsumption_resolution,[],[f12097,f11976])).
fof(f12097,plain,(
  p(s(bool,'F_0')) | ($spl4 | $spl42)),
  inference(resolution,[],[f11976,f962])).
fof(f12276,plain,(
  $false | ($spl4 | $spl42)),
  inference(subsumption_resolution,[],[f12096,f11976])).
fof(f12096,plain,(
  ( ! [X305] : (p(s(bool,X305))) ) | ($spl4 | $spl42)),
  inference(resolution,[],[f11976,f959])).
fof(f12275,plain,(
  $false | ($spl4 | $spl42)),
  inference(subsumption_resolution,[],[f12094,f11976])).
fof(f12094,plain,(
  ( ! [X303] : (p(s(bool,X303))) ) | ($spl4 | $spl42)),
  inference(resolution,[],[f11976,f956])).
fof(f12106,plain,(
  $false | ($spl4 | $spl42)),
  inference(resolution,[],[f11976,f1024])).
fof(f12107,plain,(
  $false | ($spl4 | $spl42)),
  inference(resolution,[],[f11976,f1025])).
fof(f12108,plain,(
  $false | ($spl4 | $spl42)),
  inference(resolution,[],[f11976,f1080])).
fof(f12136,plain,(
  $false | ($spl4 | $spl6 | $spl42)),
  inference(resolution,[],[f11976,f1092])).
fof(f12186,plain,(
  $false | ($spl4 | $spl32 | $spl42)),
  inference(resolution,[],[f11976,f1537])).
fof(f11974,plain,(
  ( ! [X0] : (sP11(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,X0))) ) | $spl136),
  inference(cnf_transformation,[],[f11974_D])).
fof(f11974_D,plain,(
  ( ! [X0] : (sP11(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,X0))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f11961,plain,(
  ( ! [X0] : (sP11(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl1),
  inference(superposition,[],[f1442,f1074])).
fof(f1442,plain,(
  ( ! [X176] : (sP11(X176,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | ~p(s(bool,X176))) ) | $spl1),
  inference(resolution,[],[f1117,f897])).
fof(f11947,plain,(
  ( ! [X0] : (~sP11('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0) | p(s(bool,X0))) ) | $spl134),
  inference(cnf_transformation,[],[f11947_D])).
fof(f11947_D,plain,(
  ( ! [X0] : (~sP11('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0) | p(s(bool,X0))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f11932,plain,(
  ( ! [X0] : (~sP11('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl1),
  inference(superposition,[],[f1441,f1074])).
fof(f1441,plain,(
  ( ! [X175] : (~sP11('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),X175) | p(s(bool,X175))) ) | $spl1),
  inference(resolution,[],[f1117,f895])).
fof(f895,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | ~sP11(X0,X1)) )),
  inference(cnf_transformation,[],[f616])).
fof(f11920,plain,(
  ( ! [X0] : (~sP11(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X0))) ) | $spl132),
  inference(cnf_transformation,[],[f11920_D])).
fof(f11920_D,plain,(
  ( ! [X0] : (~sP11(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X0))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f11905,plain,(
  ( ! [X0] : (~sP11(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl1),
  inference(superposition,[],[f1440,f1074])).
fof(f1440,plain,(
  ( ! [X174] : (~sP11(X174,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | p(s(bool,X174))) ) | $spl1),
  inference(resolution,[],[f1117,f895])).
fof(f11516,plain,(
  ( ! [X0,X1] : (sP10('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f11503,f1058])).
fof(f11503,plain,(
  ( ! [X0,X1] : (sP10('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f3073,f1074])).
fof(f11489,plain,(
  ( ! [X0,X1] : (sP6('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f11473,f1058])).
fof(f11473,plain,(
  ( ! [X0,X1] : (sP6('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f3066,f1074])).
fof(f11468,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | ($spl4 | $spl6)),
  inference(resolution,[],[f3066,f10496])).
fof(f10496,plain,(
  ( ! [X4,X5] : (~sP6(X5,'F_0',X4) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5)) ) | $spl4),
  inference(subsumption_resolution,[],[f10490,f852])).
fof(f10490,plain,(
  ( ! [X4,X5] : (~sP8(X4,X5,'F_0') | ~sP6(X5,'F_0',X4) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5)) ) | $spl4),
  inference(resolution,[],[f10436,f864])).
fof(f10436,plain,(
  ( ! [X4,X5] : (sP9(X4,'F_0',X5) | ~sP8(X4,X5,'F_0') | ~sP6(X5,'F_0',X4)) ) | $spl4),
  inference(subsumption_resolution,[],[f10434,f1017])).
fof(f10434,plain,(
  ( ! [X4,X5] : (sP9(X4,'F_0',X5) | p(s(bool,'F_0')) | ~sP8(X4,X5,'F_0') | ~sP6(X5,'F_0',X4)) ) | $spl4),
  inference(resolution,[],[f10431,f847])).
fof(f10431,plain,(
  ( ! [X4,X5] : (sP7('F_0',X4,X5)) ) | $spl4),
  inference(subsumption_resolution,[],[f10425,f854])).
fof(f10425,plain,(
  ( ! [X4,X5] : (sP7('F_0',X4,X5) | p(s(bool,X5))) ) | $spl4),
  inference(resolution,[],[f10422,f852])).
fof(f10422,plain,(
  ( ! [X6,X7] : (~sP8(X7,X6,'F_0') | sP7('F_0',X6,X7)) ) | $spl4),
  inference(subsumption_resolution,[],[f10418,f855])).
fof(f10418,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~sP8(X7,X6,'F_0') | sP7('F_0',X6,X7)) ) | $spl4),
  inference(resolution,[],[f10412,f843])).
fof(f10412,plain,(
  ( ! [X2,X3] : (sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0')) ) | $spl4),
  inference(subsumption_resolution,[],[f10411,f1003])).
fof(f10411,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0')) ) | $spl4),
  inference(subsumption_resolution,[],[f10408,f10369])).
fof(f10369,plain,(
  ( ! [X4,X5] : (sP6(X4,'F_0',X5) | s(bool,'T_0') != s(bool,X4)) ) | $spl4),
  inference(subsumption_resolution,[],[f10364,f858])).
fof(f10364,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | sP6(X4,'F_0',X5)) ) | $spl4),
  inference(resolution,[],[f9219,f842])).
fof(f9219,plain,(
  ( ! [X6,X7] : (sP9(X7,'F_0',X6) | s(bool,'T_0') != s(bool,X6) | p(s(bool,X7))) ) | $spl4),
  inference(superposition,[],[f9212,f861])).
fof(f9212,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl4),
  inference(subsumption_resolution,[],[f9196,f1089])).
fof(f9196,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)))) | $spl4),
  inference(superposition,[],[f8815,f1003])).
fof(f8815,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14))) | $spl4),
  inference(resolution,[],[f1089,f998])).
fof(f10408,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0') | ~sP6(X2,'F_0',X3)) ) | $spl4),
  inference(resolution,[],[f10370,f846])).
fof(f10370,plain,(
  ( ! [X6,X7] : (sP7('F_0',X6,X7) | s(bool,'T_0') != s(bool,X6)) ) | $spl4),
  inference(subsumption_resolution,[],[f10365,f854])).
fof(f10365,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | p(s(bool,X7)) | sP7('F_0',X6,X7)) ) | $spl4),
  inference(resolution,[],[f9219,f843])).
fof(f11458,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f11442,f1058])).
fof(f11442,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f3065,f1074])).
fof(f11428,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f11412,f1058])).
fof(f11412,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f3064,f1074])).
fof(f11397,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f11382,f1058])).
fof(f11382,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f3063,f1074])).
fof(f11368,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f11353,f1058])).
fof(f11353,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f3060,f1074])).
fof(f11339,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f11324,f1058])).
fof(f11324,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f3047,f1074])).
fof(f11309,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f11295,f1058])).
fof(f11295,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f3031,f1074])).
fof(f11280,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f11266,f1058])).
fof(f11266,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f3030,f1074])).
fof(f11251,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f11237,f1058])).
fof(f11237,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f3017,f1074])).
fof(f11227,plain,(
  s(bool,'F_0') != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | $spl111),
  inference(cnf_transformation,[],[f11227_D])).
fof(f11227_D,plain,(
  s(bool,'F_0') != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) <=> ~$spl111),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl111])])).
fof(f11221,plain,(
  ( ! [X14] : (p(s(bool,X14)) | s(bool,'F_0') != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')))) ) | ($spl4 | $spl32)),
  inference(resolution,[],[f10496,f4380])).
fof(f4380,plain,(
  ( ! [X161,X162] : (sP6('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X161,X162)) ) | $spl32),
  inference(resolution,[],[f1537,f860])).
fof(f3712,plain,(
  ( ! [X39] : (p(s(bool,X39)) | s(bool,X39) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl48),
  inference(cnf_transformation,[],[f3712_D])).
fof(f3712_D,plain,(
  ( ! [X39] : (p(s(bool,X39)) | s(bool,X39) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f11133,plain,(
  ( ! [X212,X213] : (p(s(bool,X212)) | p(s(bool,X213)) | s(bool,X213) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl6),
  inference(resolution,[],[f1092,f1075])).
fof(f11035,plain,(
  ( ! [X61,X59,X60] : (p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X59) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X60) != s(bool,X61)) ) | $spl6),
  inference(resolution,[],[f1092,f790])).
fof(f11032,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | p(s(bool,X52)) | s(bool,X52) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl6),
  inference(resolution,[],[f1092,f789])).
fof(f11031,plain,(
  ( ! [X47,X48,X49] : (p(s(bool,X47)) | p(s(bool,X48)) | s(bool,X47) != s(bool,X49) | s(bool,X48) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl6),
  inference(resolution,[],[f1092,f788])).
fof(f11030,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,X44)) | ~p(s(bool,X45)) | s(bool,X44) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X45) != s(bool,X46)) ) | $spl6),
  inference(resolution,[],[f1092,f786])).
fof(f11028,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X39) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl6),
  inference(resolution,[],[f1092,f784])).
fof(f11002,plain,(
  ( ! [X0] : (sP12('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f10987,f1058])).
fof(f10987,plain,(
  ( ! [X0] : (sP12('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f3085,f1074])).
fof(f10975,plain,(
  ( ! [X0] : (sP12(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(forward_demodulation,[],[f10960,f1058])).
fof(f10960,plain,(
  ( ! [X0] : (sP12(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl6),
  inference(superposition,[],[f3084,f1074])).
fof(f10952,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f10849,f1322])).
fof(f10849,plain,(
  ( ! [X549,X548,X550] : (p(s(bool,'const.bool.IN_2'(s(X548,sK58(X549,X550,X548)),s(fun(X548,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X548),X549))))))) ) | $spl24),
  inference(resolution,[],[f1322,f1053])).
fof(f10951,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f10798,f1322])).
fof(f10798,plain,(
  ( ! [X377,X376,X375] : (p(s(bool,'const.list.EVERY_2'(s(fun(X375,bool),X376),s('type.list.list'(X375),X377))))) ) | $spl24),
  inference(resolution,[],[f1322,f1053])).
fof(f10946,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f10772,f1322])).
fof(f10772,plain,(
  p(s(bool,'F_0')) | $spl24),
  inference(resolution,[],[f1322,f1020])).
fof(f10945,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f10771,f1322])).
fof(f10771,plain,(
  p(s(bool,'F_0')) | $spl24),
  inference(resolution,[],[f1322,f1018])).
fof(f10944,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f10766,f1322])).
fof(f10766,plain,(
  ( ! [X307] : (p(s(bool,X307))) ) | $spl24),
  inference(resolution,[],[f1322,f965])).
fof(f10943,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f10764,f1322])).
fof(f10764,plain,(
  p(s(bool,'F_0')) | $spl24),
  inference(resolution,[],[f1322,f962])).
fof(f10942,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f10763,f1322])).
fof(f10763,plain,(
  ( ! [X305] : (p(s(bool,X305))) ) | $spl24),
  inference(resolution,[],[f1322,f959])).
fof(f10941,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f10761,f1322])).
fof(f10761,plain,(
  ( ! [X303] : (p(s(bool,X303))) ) | $spl24),
  inference(resolution,[],[f1322,f956])).
fof(f10773,plain,(
  $false | $spl24),
  inference(resolution,[],[f1322,f1024])).
fof(f10774,plain,(
  $false | $spl24),
  inference(resolution,[],[f1322,f1025])).
fof(f10775,plain,(
  $false | $spl24),
  inference(resolution,[],[f1322,f1080])).
fof(f10852,plain,(
  $false | ($spl24 | $spl32)),
  inference(resolution,[],[f1322,f1537])).
fof(f10643,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | $spl104),
  inference(cnf_transformation,[],[f10643_D])).
fof(f10643_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f10634,plain,(
  ( ! [X10] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | ~p(s(bool,X10))) ) | ($spl4 | $spl7)),
  inference(resolution,[],[f10495,f3355])).
fof(f3355,plain,(
  ( ! [X140,X141] : (sP8(X140,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X141)) ) | $spl7),
  inference(resolution,[],[f1091,f851])).
fof(f851,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f594])).
fof(f10495,plain,(
  ( ! [X2,X3] : (~sP8(X2,X3,'F_0') | s(bool,'F_0') = s(bool,X3) | ~p(s(bool,X2))) ) | $spl4),
  inference(subsumption_resolution,[],[f10489,f858])).
fof(f10489,plain,(
  ( ! [X2,X3] : (~sP8(X2,X3,'F_0') | ~sP6(X3,'F_0',X2) | s(bool,'F_0') = s(bool,X3) | ~p(s(bool,X2))) ) | $spl4),
  inference(resolution,[],[f10436,f863])).
fof(f863,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X2)) | ~sP9(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f600])).
fof(f10641,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),sK14))) | $spl102),
  inference(cnf_transformation,[],[f10641_D])).
fof(f10641_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),sK14))) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f10633,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),sK14))) | ~p(s(bool,X9))) ) | ($spl1 | $spl4)),
  inference(resolution,[],[f10495,f7427])).
fof(f7427,plain,(
  ( ! [X136,X135] : (sP8(X135,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)),X136)) ) | $spl1),
  inference(resolution,[],[f1516,f851])).
fof(f10639,plain,(
  s(bool,'F_0') = s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | $spl100),
  inference(cnf_transformation,[],[f10639_D])).
fof(f10639_D,plain,(
  s(bool,'F_0') = s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f10632,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | ~p(s(bool,X8))) ) | ($spl1 | $spl4)),
  inference(resolution,[],[f10495,f1420])).
fof(f1420,plain,(
  ( ! [X134,X135] : (sP8(X134,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),X135)) ) | $spl1),
  inference(resolution,[],[f1117,f851])).
fof(f10623,plain,(
  ( ! [X0,X1] : (sP7(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X1)) ) | $spl98),
  inference(cnf_transformation,[],[f10623_D])).
fof(f10623_D,plain,(
  ( ! [X0,X1] : (sP7(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X1)) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f10622,plain,(
  ( ! [X0,X1] : (sP7(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl7),
  inference(forward_demodulation,[],[f10606,f1058])).
fof(f10606,plain,(
  ( ! [X0,X1] : (sP7(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl7),
  inference(superposition,[],[f3358,f1074])).
fof(f3358,plain,(
  ( ! [X146,X147] : (sP7(X146,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X147)) ) | $spl7),
  inference(resolution,[],[f1091,f855])).
fof(f10585,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | $spl97),
  inference(cnf_transformation,[],[f10585_D])).
fof(f10585_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) <=> ~$spl97),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl97])])).
fof(f10576,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | ~p(s(bool,X10))) ) | ($spl4 | $spl7)),
  inference(resolution,[],[f10569,f3355])).
fof(f10569,plain,(
  ( ! [X14,X15] : (~sP8(X14,X15,'F_0') | s(bool,'T_0') != s(bool,X15) | ~p(s(bool,X14))) ) | $spl4),
  inference(subsumption_resolution,[],[f10568,f858])).
fof(f10568,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | ~sP8(X14,X15,'F_0') | ~sP6(X15,'F_0',X14)) ) | $spl4),
  inference(subsumption_resolution,[],[f10567,f951])).
fof(f10567,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP8(X14,X15,'F_0') | ~sP6(X15,'F_0',X14)) ) | $spl4),
  inference(subsumption_resolution,[],[f10560,f856])).
fof(f10560,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP8(X14,X15,'F_0') | ~sP7('F_0',X15,X14) | ~sP6(X15,'F_0',X14)) ) | $spl4),
  inference(resolution,[],[f9221,f847])).
fof(f9221,plain,(
  ( ! [X10,X11] : (~sP9(X11,'F_0',X10) | ~p(s(bool,X11)) | s(bool,'T_0') != s(bool,X10)) ) | $spl4),
  inference(superposition,[],[f9212,f863])).
fof(f10583,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),sK14))) | $spl95),
  inference(cnf_transformation,[],[f10583_D])).
fof(f10583_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),sK14))) <=> ~$spl95),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl95])])).
fof(f10575,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),sK14))) | ~p(s(bool,X9))) ) | ($spl1 | $spl4)),
  inference(resolution,[],[f10569,f7427])).
fof(f10581,plain,(
  s(bool,'T_0') != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | $spl93),
  inference(cnf_transformation,[],[f10581_D])).
fof(f10581_D,plain,(
  s(bool,'T_0') != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) <=> ~$spl93),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl93])])).
fof(f10574,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | ~p(s(bool,X8))) ) | ($spl1 | $spl4)),
  inference(resolution,[],[f10569,f1420])).
fof(f10533,plain,(
  ( ! [X0,X1] : (sP8('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1)) ) | $spl90),
  inference(cnf_transformation,[],[f10533_D])).
fof(f10533_D,plain,(
  ( ! [X0,X1] : (sP8('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1)) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f10532,plain,(
  ( ! [X0,X1] : (sP8('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl7),
  inference(forward_demodulation,[],[f10516,f1058])).
fof(f10516,plain,(
  ( ! [X0,X1] : (sP8('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl7),
  inference(superposition,[],[f3356,f1074])).
fof(f3356,plain,(
  ( ! [X142,X143] : (sP8('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X142,X143)) ) | $spl7),
  inference(resolution,[],[f1091,f852])).
fof(f10463,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X1)) ) | $spl88),
  inference(cnf_transformation,[],[f10463_D])).
fof(f10463_D,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X1)) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f10462,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl7),
  inference(forward_demodulation,[],[f10447,f1058])).
fof(f10447,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl7),
  inference(superposition,[],[f3355,f1074])).
fof(f10395,plain,(
  ( ! [X0,X1] : (sP4('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1)) ) | $spl86),
  inference(cnf_transformation,[],[f10395_D])).
fof(f10395_D,plain,(
  ( ! [X0,X1] : (sP4('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1)) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f10394,plain,(
  ( ! [X0,X1] : (sP4('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl7),
  inference(forward_demodulation,[],[f10380,f1058])).
fof(f10380,plain,(
  ( ! [X0,X1] : (sP4('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl7),
  inference(superposition,[],[f3343,f1074])).
fof(f3343,plain,(
  ( ! [X116,X117] : (sP4('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X116,X117)) ) | $spl7),
  inference(resolution,[],[f1091,f835])).
fof(f835,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f586])).
fof(f8667,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | s(bool,'F_0') != s(bool,X2)) ) | $spl76),
  inference(cnf_transformation,[],[f8667_D])).
fof(f8667_D,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | s(bool,'F_0') != s(bool,X2)) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f10268,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) ) | $spl4),
  inference(duplicate_literal_removal,[],[f10263])).
fof(f10263,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) ) | $spl4),
  inference(resolution,[],[f9218,f864])).
fof(f9218,plain,(
  ( ! [X4,X5] : (sP9(X5,X4,'F_0') | s(bool,'T_0') != s(bool,X4) | p(s(bool,X5))) ) | $spl4),
  inference(superposition,[],[f9212,f861])).
fof(f10257,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X1)) ) | $spl84),
  inference(cnf_transformation,[],[f10257_D])).
fof(f10257_D,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X1)) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f10256,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl7),
  inference(forward_demodulation,[],[f10242,f1058])).
fof(f10242,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl7),
  inference(superposition,[],[f3342,f1074])).
fof(f3342,plain,(
  ( ! [X114,X115] : (sP4(X114,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X115)) ) | $spl7),
  inference(resolution,[],[f1091,f834])).
fof(f834,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f586])).
fof(f10228,plain,(
  ( ! [X0,X1] : (sP2('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1)) ) | $spl82),
  inference(cnf_transformation,[],[f10228_D])).
fof(f10228_D,plain,(
  ( ! [X0,X1] : (sP2('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1)) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f10227,plain,(
  ( ! [X0,X1] : (sP2('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl7),
  inference(forward_demodulation,[],[f10212,f1058])).
fof(f10212,plain,(
  ( ! [X0,X1] : (sP2('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl7),
  inference(superposition,[],[f3329,f1074])).
fof(f3329,plain,(
  ( ! [X88,X89] : (sP2('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X88,X89)) ) | $spl7),
  inference(resolution,[],[f1091,f818])).
fof(f818,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f579])).
fof(f10177,plain,(
  ( ! [X0,X1] : (sP0('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1)) ) | $spl80),
  inference(cnf_transformation,[],[f10177_D])).
fof(f10177_D,plain,(
  ( ! [X0,X1] : (sP0('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1)) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f10176,plain,(
  ( ! [X0,X1] : (sP0('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl7),
  inference(forward_demodulation,[],[f10162,f1058])).
fof(f10162,plain,(
  ( ! [X0,X1] : (sP0('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl7),
  inference(superposition,[],[f3313,f1074])).
fof(f3313,plain,(
  ( ! [X57,X56] : (sP0('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X56,X57)) ) | $spl7),
  inference(resolution,[],[f1091,f801])).
fof(f801,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f571])).
fof(f10046,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) | $spl78),
  inference(cnf_transformation,[],[f10046_D])).
fof(f10046_D,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f10045,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.pred_set.EMPTY_0'))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl7),
  inference(forward_demodulation,[],[f10031,f1058])).
fof(f10031,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl7),
  inference(superposition,[],[f3312,f1074])).
fof(f3312,plain,(
  ( ! [X54,X55] : (sP0(X54,X55,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl7),
  inference(resolution,[],[f1091,f799])).
fof(f799,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f571])).
fof(f2756,plain,(
  ( ! [X26] : (p(s(bool,X26)) | s(bool,X26) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)))) ) | $spl44),
  inference(cnf_transformation,[],[f2756_D])).
fof(f2756_D,plain,(
  ( ! [X26] : (p(s(bool,X26)) | s(bool,X26) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f8820,plain,(
  ( ! [X200,X199] : (p(s(bool,X199)) | p(s(bool,X200)) | s(bool,X200) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(resolution,[],[f1089,f1075])).
fof(f8722,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,X46) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14))) | s(bool,X47) != s(bool,X48)) ) | $spl4),
  inference(resolution,[],[f1089,f790])).
fof(f8719,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,X37)) | ~p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X39) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(resolution,[],[f1089,f789])).
fof(f8718,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X34) != s(bool,X36) | s(bool,X35) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(resolution,[],[f1089,f788])).
fof(f8717,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14))) | s(bool,X32) != s(bool,X33)) ) | $spl4),
  inference(resolution,[],[f1089,f786])).
fof(f8715,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(resolution,[],[f1089,f784])).
fof(f8666,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) ) | $spl5),
  inference(duplicate_literal_removal,[],[f8661])).
fof(f8661,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) ) | $spl5),
  inference(resolution,[],[f7286,f864])).
fof(f7286,plain,(
  ( ! [X4,X5] : (sP9(X5,X4,'F_0') | s(bool,'T_0') != s(bool,X4) | p(s(bool,X5))) ) | $spl5),
  inference(superposition,[],[f7236,f861])).
fof(f7236,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl5),
  inference(superposition,[],[f2878,f2867])).
fof(f2867,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14))) | $spl5),
  inference(resolution,[],[f1090,f953])).
fof(f2878,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14))) | $spl5),
  inference(resolution,[],[f1090,f1002])).
fof(f8541,plain,(
  ( ! [X196,X195] : (p(s(bool,X195)) | ~p(s(bool,X196)) | s(bool,X196) != s(bool,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) ) | $spl1),
  inference(resolution,[],[f1516,f1075])).
fof(f8437,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X35) != s(bool,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),sK14))) | s(bool,X34) != s(bool,X36)) ) | $spl1),
  inference(resolution,[],[f1516,f790])).
fof(f8436,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) ) | $spl1),
  inference(resolution,[],[f1516,f789])).
fof(f8434,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) ) | $spl1),
  inference(resolution,[],[f1516,f788])).
fof(f8432,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X20) != s(bool,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),sK14))) | s(bool,X19) != s(bool,X21)) ) | $spl1),
  inference(resolution,[],[f1516,f786])).
fof(f8430,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) ) | $spl1),
  inference(resolution,[],[f1516,f784])).
fof(f1320,plain,(
  ( ! [X38,X36] : (p(s(bool,'const.bool.IN_2'(s(sK13,X36),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X38)) | s(bool,X38) != s(bool,'const.bool.IN_2'(s(sK13,X36),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl20),
  inference(cnf_transformation,[],[f1320_D])).
fof(f1320_D,plain,(
  ( ! [X38,X36] : (p(s(bool,'const.bool.IN_2'(s(sK13,X36),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X38)) | s(bool,X38) != s(bool,'const.bool.IN_2'(s(sK13,X36),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f8310,plain,(
  ( ! [X329,X331,X330] : (p(s(bool,'const.bool.IN_2'(s(sK13,X329),s(fun(sK13,bool),sK14)))) | p(s(bool,X330)) | ~p(s(bool,X331)) | s(bool,X331) != s(bool,'const.bool.IN_2'(s(sK13,X329),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl2),
  inference(resolution,[],[f1087,f1075])).
fof(f8206,plain,(
  ( ! [X66,X64,X67,X65] : (p(s(bool,'const.bool.IN_2'(s(sK13,X64),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X65)) | ~p(s(bool,X66)) | s(bool,X66) != s(bool,'const.bool.IN_2'(s(sK13,X64),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X65) != s(bool,X67)) ) | $spl2),
  inference(resolution,[],[f1087,f790])).
fof(f8205,plain,(
  ( ! [X61,X62,X60,X63] : (p(s(bool,'const.bool.IN_2'(s(sK13,X60),s(fun(sK13,bool),sK14)))) | p(s(bool,X61)) | ~p(s(bool,X62)) | ~p(s(bool,X63)) | s(bool,X62) != s(bool,'const.bool.IN_2'(s(sK13,X60),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl2),
  inference(resolution,[],[f1087,f789])).
fof(f8203,plain,(
  ( ! [X54,X52,X55,X53] : (p(s(bool,'const.bool.IN_2'(s(sK13,X52),s(fun(sK13,bool),sK14)))) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X53) != s(bool,X55) | s(bool,X54) != s(bool,'const.bool.IN_2'(s(sK13,X52),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl2),
  inference(resolution,[],[f1087,f788])).
fof(f8201,plain,(
  ( ! [X47,X45,X46,X44] : (p(s(bool,'const.bool.IN_2'(s(sK13,X44),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X46) != s(bool,'const.bool.IN_2'(s(sK13,X44),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X45) != s(bool,X47)) ) | $spl2),
  inference(resolution,[],[f1087,f786])).
fof(f8199,plain,(
  ( ! [X39,X37,X38,X36] : (p(s(bool,'const.bool.IN_2'(s(sK13,X36),s(fun(sK13,bool),sK14)))) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,X39) | s(bool,X38) != s(bool,'const.bool.IN_2'(s(sK13,X36),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl2),
  inference(resolution,[],[f1087,f784])).
fof(f5969,plain,(
  ( ! [X0,X1] : (sP7('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1)) ) | $spl70),
  inference(cnf_transformation,[],[f5969_D])).
fof(f5969_D,plain,(
  ( ! [X0,X1] : (sP7('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1)) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f8171,plain,(
  ( ! [X0,X1] : (sP7('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl1),
  inference(superposition,[],[f1424,f1074])).
fof(f1424,plain,(
  ( ! [X142,X143] : (sP7('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),X142,X143)) ) | $spl1),
  inference(resolution,[],[f1117,f856])).
fof(f5916,plain,(
  ( ! [X0,X1] : (sP7(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X1)) ) | $spl68),
  inference(cnf_transformation,[],[f5916_D])).
fof(f5916_D,plain,(
  ( ! [X0,X1] : (sP7(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X1)) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f8147,plain,(
  ( ! [X0,X1] : (sP7(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl1),
  inference(superposition,[],[f1423,f1074])).
fof(f1423,plain,(
  ( ! [X140,X141] : (sP7(X140,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),X141)) ) | $spl1),
  inference(resolution,[],[f1117,f855])).
fof(f5891,plain,(
  ( ! [X0,X1] : (sP8('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1)) ) | $spl66),
  inference(cnf_transformation,[],[f5891_D])).
fof(f5891_D,plain,(
  ( ! [X0,X1] : (sP8('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f8095,plain,(
  ( ! [X0,X1] : (sP8('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl1),
  inference(superposition,[],[f1421,f1074])).
fof(f1421,plain,(
  ( ! [X136,X137] : (sP8('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),X136,X137)) ) | $spl1),
  inference(resolution,[],[f1117,f852])).
fof(f5866,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X1)) ) | $spl64),
  inference(cnf_transformation,[],[f5866_D])).
fof(f5866_D,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X1)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f8072,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl1),
  inference(superposition,[],[f1420,f1074])).
fof(f5842,plain,(
  ( ! [X0,X1] : (sP4('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1)) ) | $spl62),
  inference(cnf_transformation,[],[f5842_D])).
fof(f5842_D,plain,(
  ( ! [X0,X1] : (sP4('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f8050,plain,(
  ( ! [X0,X1] : (sP4('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl1),
  inference(superposition,[],[f1408,f1074])).
fof(f1408,plain,(
  ( ! [X111,X110] : (sP4('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),X110,X111)) ) | $spl1),
  inference(resolution,[],[f1117,f835])).
fof(f5818,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X1)) ) | $spl60),
  inference(cnf_transformation,[],[f5818_D])).
fof(f5818_D,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X1)) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f8027,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl1),
  inference(superposition,[],[f1407,f1074])).
fof(f1407,plain,(
  ( ! [X109,X108] : (sP4(X108,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),X109)) ) | $spl1),
  inference(resolution,[],[f1117,f834])).
fof(f5606,plain,(
  ( ! [X0,X1] : (sP2('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1)) ) | $spl58),
  inference(cnf_transformation,[],[f5606_D])).
fof(f5606_D,plain,(
  ( ! [X0,X1] : (sP2('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f8003,plain,(
  ( ! [X0,X1] : (sP2('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl1),
  inference(superposition,[],[f1394,f1074])).
fof(f1394,plain,(
  ( ! [X83,X82] : (sP2('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),X82,X83)) ) | $spl1),
  inference(resolution,[],[f1117,f818])).
fof(f3464,plain,(
  ( ! [X19] : (~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl46),
  inference(cnf_transformation,[],[f3464_D])).
fof(f3464_D,plain,(
  ( ! [X19] : (~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f7924,plain,(
  ( ! [X200,X201] : (p(s(bool,X200)) | ~p(s(bool,X201)) | s(bool,X201) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl7),
  inference(resolution,[],[f1091,f1075])).
fof(f7820,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X40) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X39) != s(bool,X41)) ) | $spl7),
  inference(resolution,[],[f1091,f790])).
fof(f7819,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl7),
  inference(resolution,[],[f1091,f789])).
fof(f7817,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | s(bool,X31) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl7),
  inference(resolution,[],[f1091,f788])).
fof(f7815,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X24) != s(bool,X26)) ) | $spl7),
  inference(resolution,[],[f1091,f786])).
fof(f7813,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,X20) | s(bool,X19) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl7),
  inference(resolution,[],[f1091,f784])).
fof(f5581,plain,(
  ( ! [X0,X1] : (sP0('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1)) ) | $spl56),
  inference(cnf_transformation,[],[f5581_D])).
fof(f5581_D,plain,(
  ( ! [X0,X1] : (sP0('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f7792,plain,(
  ( ! [X0,X1] : (sP0('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl1),
  inference(superposition,[],[f1378,f1074])).
fof(f1378,plain,(
  ( ! [X50,X51] : (sP0('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),X50,X51)) ) | $spl1),
  inference(resolution,[],[f1117,f801])).
fof(f4120,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')))) ) | $spl50),
  inference(cnf_transformation,[],[f4120_D])).
fof(f4120_D,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f7769,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl1),
  inference(superposition,[],[f1377,f1074])).
fof(f1377,plain,(
  ( ! [X48,X49] : (sP0(X48,X49,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl1),
  inference(resolution,[],[f1117,f799])).
fof(f1525,plain,(
  ( ! [X13] : (~p(s(bool,X13)) | s(bool,X13) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl28),
  inference(cnf_transformation,[],[f1525_D])).
fof(f1525_D,plain,(
  ( ! [X13] : (~p(s(bool,X13)) | s(bool,X13) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f7694,plain,(
  ( ! [X194,X195] : (p(s(bool,X194)) | ~p(s(bool,X195)) | s(bool,X195) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl1),
  inference(resolution,[],[f1117,f1075])).
fof(f7590,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | s(bool,X33) != s(bool,X35)) ) | $spl1),
  inference(resolution,[],[f1117,f790])).
fof(f7589,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl1),
  inference(resolution,[],[f1117,f789])).
fof(f7587,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl1),
  inference(resolution,[],[f1117,f788])).
fof(f7585,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | s(bool,X18) != s(bool,X20)) ) | $spl1),
  inference(resolution,[],[f1117,f786])).
fof(f7583,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl1),
  inference(resolution,[],[f1117,f784])).
fof(f7481,plain,(
  ( ! [X196,X195] : (p(s(bool,X195)) | ~p(s(bool,X196)) | s(bool,X196) != s(bool,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) ) | $spl1),
  inference(resolution,[],[f1516,f1075])).
fof(f7377,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X35) != s(bool,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),sK14))) | s(bool,X34) != s(bool,X36)) ) | $spl1),
  inference(resolution,[],[f1516,f790])).
fof(f7376,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) ) | $spl1),
  inference(resolution,[],[f1516,f789])).
fof(f7374,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) ) | $spl1),
  inference(resolution,[],[f1516,f788])).
fof(f7372,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X20) != s(bool,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),sK14))) | s(bool,X19) != s(bool,X21)) ) | $spl1),
  inference(resolution,[],[f1516,f786])).
fof(f7370,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(sK13,sK58(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) ) | $spl1),
  inference(resolution,[],[f1516,f784])).
fof(f6095,plain,(
  ( ! [X329,X331,X330] : (p(s(bool,'const.bool.IN_2'(s(sK13,X329),s(fun(sK13,bool),sK14)))) | p(s(bool,X330)) | ~p(s(bool,X331)) | s(bool,X331) != s(bool,'const.bool.IN_2'(s(sK13,X329),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl2),
  inference(resolution,[],[f1087,f1075])).
fof(f5991,plain,(
  ( ! [X66,X64,X67,X65] : (p(s(bool,'const.bool.IN_2'(s(sK13,X64),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X65)) | ~p(s(bool,X66)) | s(bool,X66) != s(bool,'const.bool.IN_2'(s(sK13,X64),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X65) != s(bool,X67)) ) | $spl2),
  inference(resolution,[],[f1087,f790])).
fof(f5990,plain,(
  ( ! [X61,X62,X60,X63] : (p(s(bool,'const.bool.IN_2'(s(sK13,X60),s(fun(sK13,bool),sK14)))) | p(s(bool,X61)) | ~p(s(bool,X62)) | ~p(s(bool,X63)) | s(bool,X62) != s(bool,'const.bool.IN_2'(s(sK13,X60),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl2),
  inference(resolution,[],[f1087,f789])).
fof(f5988,plain,(
  ( ! [X54,X52,X55,X53] : (p(s(bool,'const.bool.IN_2'(s(sK13,X52),s(fun(sK13,bool),sK14)))) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X53) != s(bool,X55) | s(bool,X54) != s(bool,'const.bool.IN_2'(s(sK13,X52),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl2),
  inference(resolution,[],[f1087,f788])).
fof(f5986,plain,(
  ( ! [X47,X45,X46,X44] : (p(s(bool,'const.bool.IN_2'(s(sK13,X44),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X46) != s(bool,'const.bool.IN_2'(s(sK13,X44),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X45) != s(bool,X47)) ) | $spl2),
  inference(resolution,[],[f1087,f786])).
fof(f5984,plain,(
  ( ! [X39,X37,X38,X36] : (p(s(bool,'const.bool.IN_2'(s(sK13,X36),s(fun(sK13,bool),sK14)))) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,X39) | s(bool,X38) != s(bool,'const.bool.IN_2'(s(sK13,X36),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl2),
  inference(resolution,[],[f1087,f784])).
fof(f5955,plain,(
  ( ! [X0,X1] : (sP7('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl1),
  inference(superposition,[],[f1424,f1074])).
fof(f5902,plain,(
  ( ! [X0,X1] : (sP7(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl1),
  inference(superposition,[],[f1423,f1074])).
fof(f5877,plain,(
  ( ! [X0,X1] : (sP8('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl1),
  inference(superposition,[],[f1421,f1074])).
fof(f5853,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl1),
  inference(superposition,[],[f1420,f1074])).
fof(f5830,plain,(
  ( ! [X0,X1] : (sP4('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl1),
  inference(superposition,[],[f1408,f1074])).
fof(f5806,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl1),
  inference(superposition,[],[f1407,f1074])).
fof(f5727,plain,(
  ( ! [X200,X201] : (p(s(bool,X200)) | ~p(s(bool,X201)) | s(bool,X201) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl7),
  inference(resolution,[],[f1091,f1075])).
fof(f5623,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X40) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X39) != s(bool,X41)) ) | $spl7),
  inference(resolution,[],[f1091,f790])).
fof(f5622,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl7),
  inference(resolution,[],[f1091,f789])).
fof(f5620,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | s(bool,X31) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl7),
  inference(resolution,[],[f1091,f788])).
fof(f5618,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X24) != s(bool,X26)) ) | $spl7),
  inference(resolution,[],[f1091,f786])).
fof(f5616,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,X20) | s(bool,X19) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl7),
  inference(resolution,[],[f1091,f784])).
fof(f5593,plain,(
  ( ! [X0,X1] : (sP2('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl1),
  inference(superposition,[],[f1394,f1074])).
fof(f5569,plain,(
  ( ! [X0,X1] : (sP0('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl1),
  inference(superposition,[],[f1378,f1074])).
fof(f5546,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl1),
  inference(superposition,[],[f1377,f1074])).
fof(f5471,plain,(
  ( ! [X194,X195] : (p(s(bool,X194)) | ~p(s(bool,X195)) | s(bool,X195) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl1),
  inference(resolution,[],[f1117,f1075])).
fof(f5367,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | s(bool,X33) != s(bool,X35)) ) | $spl1),
  inference(resolution,[],[f1117,f790])).
fof(f5366,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl1),
  inference(resolution,[],[f1117,f789])).
fof(f5364,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl1),
  inference(resolution,[],[f1117,f788])).
fof(f5362,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | s(bool,X18) != s(bool,X20)) ) | $spl1),
  inference(resolution,[],[f1117,f786])).
fof(f5360,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl1),
  inference(resolution,[],[f1117,f784])).
fof(f5277,plain,(
  ( ! [X200,X199] : (p(s(bool,X199)) | p(s(bool,X200)) | s(bool,X200) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl0),
  inference(resolution,[],[f4976,f1075])).
fof(f5179,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,X46) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | s(bool,X47) != s(bool,X48)) ) | $spl0),
  inference(resolution,[],[f4976,f790])).
fof(f5176,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,X37)) | ~p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X39) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl0),
  inference(resolution,[],[f4976,f789])).
fof(f5175,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X34) != s(bool,X36) | s(bool,X35) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl0),
  inference(resolution,[],[f4976,f788])).
fof(f5174,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | s(bool,X32) != s(bool,X33)) ) | $spl0),
  inference(resolution,[],[f4976,f786])).
fof(f5172,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl0),
  inference(resolution,[],[f4976,f784])).
fof(f4617,plain,(
  ( ! [X212,X213] : (p(s(bool,X212)) | p(s(bool,X213)) | s(bool,X213) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl6),
  inference(resolution,[],[f1092,f1075])).
fof(f4519,plain,(
  ( ! [X61,X59,X60] : (p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X59) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X60) != s(bool,X61)) ) | $spl6),
  inference(resolution,[],[f1092,f790])).
fof(f4516,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | p(s(bool,X52)) | s(bool,X52) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl6),
  inference(resolution,[],[f1092,f789])).
fof(f4515,plain,(
  ( ! [X47,X48,X49] : (p(s(bool,X47)) | p(s(bool,X48)) | s(bool,X47) != s(bool,X49) | s(bool,X48) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl6),
  inference(resolution,[],[f1092,f788])).
fof(f4514,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,X44)) | ~p(s(bool,X45)) | s(bool,X44) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X45) != s(bool,X46)) ) | $spl6),
  inference(resolution,[],[f1092,f786])).
fof(f4512,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X39) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl6),
  inference(resolution,[],[f1092,f784])).
fof(f4483,plain,(
  ( ! [X26] : (p(s(bool,X26)) | s(bool,X26) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')))) ) | $spl52),
  inference(cnf_transformation,[],[f4483_D])).
fof(f4483_D,plain,(
  ( ! [X26] : (p(s(bool,X26)) | s(bool,X26) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f4421,plain,(
  ( ! [X200,X199] : (p(s(bool,X199)) | p(s(bool,X200)) | s(bool,X200) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')))) ) | $spl32),
  inference(resolution,[],[f1537,f1075])).
fof(f4323,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,X46) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | s(bool,X47) != s(bool,X48)) ) | $spl32),
  inference(resolution,[],[f1537,f790])).
fof(f4320,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,X37)) | ~p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X39) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')))) ) | $spl32),
  inference(resolution,[],[f1537,f789])).
fof(f4319,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X34) != s(bool,X36) | s(bool,X35) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')))) ) | $spl32),
  inference(resolution,[],[f1537,f788])).
fof(f4318,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | s(bool,X32) != s(bool,X33)) ) | $spl32),
  inference(resolution,[],[f1537,f786])).
fof(f4316,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')))) ) | $spl32),
  inference(resolution,[],[f1537,f784])).
fof(f4284,plain,(
  $false | $spl33),
  inference(subsumption_resolution,[],[f4283,f1055])).
fof(f4283,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK58('const.list.NIL_0',sK14,sK13)),s(fun(sK13,bool),'const.pred_set.EMPTY_0')))) | $spl33),
  inference(forward_demodulation,[],[f4122,f1058])).
fof(f4122,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK58('const.list.NIL_0',sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))))) | $spl33),
  inference(resolution,[],[f1538,f1053])).
fof(f4282,plain,(
  $false | $spl33),
  inference(subsumption_resolution,[],[f4121,f1068])).
fof(f4121,plain,(
  s('type.list.list'(sK13),'const.list.NIL_0') != s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | $spl33),
  inference(resolution,[],[f1538,f1066])).
fof(f4108,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13))))) ) | $spl1),
  inference(superposition,[],[f1377,f1074])).
fof(f1538,plain,(
  ~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')))) | $spl33),
  inference(cnf_transformation,[],[f1538_D])).
fof(f1538_D,plain,(
  ~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f4042,plain,(
  ~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13)))) | $spl1),
  inference(superposition,[],[f1117,f1074])).
fof(f4034,plain,(
  ( ! [X194,X195] : (p(s(bool,X194)) | ~p(s(bool,X195)) | s(bool,X195) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl1),
  inference(resolution,[],[f1117,f1075])).
fof(f3930,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | s(bool,X33) != s(bool,X35)) ) | $spl1),
  inference(resolution,[],[f1117,f790])).
fof(f3929,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl1),
  inference(resolution,[],[f1117,f789])).
fof(f3927,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl1),
  inference(resolution,[],[f1117,f788])).
fof(f3925,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | s(bool,X18) != s(bool,X20)) ) | $spl1),
  inference(resolution,[],[f1117,f786])).
fof(f3923,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl1),
  inference(resolution,[],[f1117,f784])).
fof(f3848,plain,(
  ( ! [X200,X201] : (p(s(bool,X200)) | ~p(s(bool,X201)) | s(bool,X201) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl7),
  inference(resolution,[],[f1091,f1075])).
fof(f3744,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X40) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X39) != s(bool,X41)) ) | $spl7),
  inference(resolution,[],[f1091,f790])).
fof(f3743,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl7),
  inference(resolution,[],[f1091,f789])).
fof(f3741,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | s(bool,X31) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl7),
  inference(resolution,[],[f1091,f788])).
fof(f3739,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X24) != s(bool,X26)) ) | $spl7),
  inference(resolution,[],[f1091,f786])).
fof(f3737,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,X20) | s(bool,X19) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl7),
  inference(resolution,[],[f1091,f784])).
fof(f3644,plain,(
  ( ! [X212,X213] : (p(s(bool,X212)) | p(s(bool,X213)) | s(bool,X213) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl6),
  inference(resolution,[],[f1092,f1075])).
fof(f3546,plain,(
  ( ! [X61,X59,X60] : (p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X59) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X60) != s(bool,X61)) ) | $spl6),
  inference(resolution,[],[f1092,f790])).
fof(f3543,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | p(s(bool,X52)) | s(bool,X52) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl6),
  inference(resolution,[],[f1092,f789])).
fof(f3542,plain,(
  ( ! [X47,X48,X49] : (p(s(bool,X47)) | p(s(bool,X48)) | s(bool,X47) != s(bool,X49) | s(bool,X48) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl6),
  inference(resolution,[],[f1092,f788])).
fof(f3541,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,X44)) | ~p(s(bool,X45)) | s(bool,X44) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X45) != s(bool,X46)) ) | $spl6),
  inference(resolution,[],[f1092,f786])).
fof(f3539,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X39) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl6),
  inference(resolution,[],[f1092,f784])).
fof(f3409,plain,(
  ( ! [X200,X201] : (p(s(bool,X200)) | ~p(s(bool,X201)) | s(bool,X201) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl7),
  inference(resolution,[],[f1091,f1075])).
fof(f3305,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X40) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X39) != s(bool,X41)) ) | $spl7),
  inference(resolution,[],[f1091,f790])).
fof(f3304,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl7),
  inference(resolution,[],[f1091,f789])).
fof(f3302,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | s(bool,X31) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl7),
  inference(resolution,[],[f1091,f788])).
fof(f3300,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X24) != s(bool,X26)) ) | $spl7),
  inference(resolution,[],[f1091,f786])).
fof(f3298,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,X20) | s(bool,X19) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl7),
  inference(resolution,[],[f1091,f784])).
fof(f3159,plain,(
  $false | ($spl2 | $spl5 | $spl6)),
  inference(subsumption_resolution,[],[f2993,f1090])).
fof(f2993,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)))) | ($spl2 | $spl6)),
  inference(resolution,[],[f1092,f1087])).
fof(f2572,plain,(
  ( ! [X14] : (s(bool,X14) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14))) | ~p(s(bool,X14))) ) | $spl42),
  inference(cnf_transformation,[],[f2572_D])).
fof(f2572_D,plain,(
  ( ! [X14] : (s(bool,X14) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14))) | ~p(s(bool,X14))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f2883,plain,(
  ( ! [X196,X195] : (p(s(bool,X195)) | ~p(s(bool,X196)) | s(bool,X196) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)))) ) | $spl5),
  inference(resolution,[],[f1090,f1075])).
fof(f2779,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X35) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14))) | s(bool,X34) != s(bool,X36)) ) | $spl5),
  inference(resolution,[],[f1090,f790])).
fof(f2778,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)))) ) | $spl5),
  inference(resolution,[],[f1090,f789])).
fof(f2776,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)))) ) | $spl5),
  inference(resolution,[],[f1090,f788])).
fof(f2774,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X20) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14))) | s(bool,X19) != s(bool,X21)) ) | $spl5),
  inference(resolution,[],[f1090,f786])).
fof(f2772,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)))) ) | $spl5),
  inference(resolution,[],[f1090,f784])).
fof(f2700,plain,(
  ( ! [X200,X199] : (p(s(bool,X199)) | p(s(bool,X200)) | s(bool,X200) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(resolution,[],[f1089,f1075])).
fof(f2602,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,X46) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14))) | s(bool,X47) != s(bool,X48)) ) | $spl4),
  inference(resolution,[],[f1089,f790])).
fof(f2599,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,X37)) | ~p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X39) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(resolution,[],[f1089,f789])).
fof(f2598,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X34) != s(bool,X36) | s(bool,X35) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(resolution,[],[f1089,f788])).
fof(f2597,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14))) | s(bool,X32) != s(bool,X33)) ) | $spl4),
  inference(resolution,[],[f1089,f786])).
fof(f2595,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)))) ) | $spl4),
  inference(resolution,[],[f1089,f784])).
fof(f2525,plain,(
  ( ! [X196,X195] : (p(s(bool,X195)) | ~p(s(bool,X196)) | s(bool,X196) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)))) ) | $spl5),
  inference(resolution,[],[f1090,f1075])).
fof(f2421,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X35) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14))) | s(bool,X34) != s(bool,X36)) ) | $spl5),
  inference(resolution,[],[f1090,f790])).
fof(f2420,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)))) ) | $spl5),
  inference(resolution,[],[f1090,f789])).
fof(f2418,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)))) ) | $spl5),
  inference(resolution,[],[f1090,f788])).
fof(f2416,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,X19)) | ~p(s(bool,X20)) | s(bool,X20) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14))) | s(bool,X19) != s(bool,X21)) ) | $spl5),
  inference(resolution,[],[f1090,f786])).
fof(f2414,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)))) ) | $spl5),
  inference(resolution,[],[f1090,f784])).
fof(f2030,plain,(
  s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),'const.list.NIL_0'))) | s('type.list.list'(sK13),sK60(sK15,sK13)) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK60(sK15,sK13),sK13)),s('type.list.list'(sK13),sK60(sK60(sK15,sK13),sK13)))) | $spl1),
  inference(superposition,[],[f1120,f1074])).
fof(f1120,plain,(
  s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13)))) | $spl1),
  inference(subsumption_resolution,[],[f1099,f1068])).
fof(f1099,plain,(
  s('type.list.list'(sK13),'const.list.NIL_0') != s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13)))) | $spl1),
  inference(superposition,[],[f1085,f1074])).
fof(f2009,plain,(
  ( ! [X49] : (s(bool,'T_0') != s(bool,X49)) ) | $spl36),
  inference(cnf_transformation,[],[f2009_D])).
fof(f2009_D,plain,(
  ( ! [X49] : (s(bool,'T_0') != s(bool,X49)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f2003,plain,(
  ( ! [X50,X51] : (s(bool,'T_0') != s(bool,X51) | ~p(s(bool,X50))) ) | $spl22),
  inference(duplicate_literal_removal,[],[f1944])).
fof(f1944,plain,(
  ( ! [X50,X51] : (s(bool,'T_0') != s(bool,X51) | ~p(s(bool,X50)) | ~p(s(bool,X50))) ) | $spl22),
  inference(superposition,[],[f1321,f1005])).
fof(f2004,plain,(
  ( ! [X48,X49] : (s(bool,'T_0') != s(bool,X49) | ~p(s(bool,X48))) ) | $spl22),
  inference(duplicate_literal_removal,[],[f1943])).
fof(f1943,plain,(
  ( ! [X48,X49] : (s(bool,'T_0') != s(bool,X49) | ~p(s(bool,X48)) | ~p(s(bool,X48))) ) | $spl22),
  inference(superposition,[],[f1321,f1003])).
fof(f1908,plain,(
  ( ! [X41] : (s(bool,'F_0') != s(bool,X41)) ) | $spl34),
  inference(cnf_transformation,[],[f1908_D])).
fof(f1908_D,plain,(
  ( ! [X41] : (s(bool,'F_0') != s(bool,X41)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1901,plain,(
  ( ! [X47,X46] : (s(bool,'F_0') != s(bool,X47) | p(s(bool,X46))) ) | $spl18),
  inference(duplicate_literal_removal,[],[f1841])).
fof(f1841,plain,(
  ( ! [X47,X46] : (s(bool,'F_0') != s(bool,X47) | p(s(bool,X46)) | p(s(bool,X46))) ) | $spl18),
  inference(superposition,[],[f1319,f1001])).
fof(f1902,plain,(
  ( ! [X45,X44] : (s(bool,'F_0') != s(bool,X45) | p(s(bool,X44))) ) | $spl18),
  inference(duplicate_literal_removal,[],[f1840])).
fof(f1840,plain,(
  ( ! [X45,X44] : (s(bool,'F_0') != s(bool,X45) | p(s(bool,X44)) | p(s(bool,X44))) ) | $spl18),
  inference(superposition,[],[f1319,f999])).
fof(f1903,plain,(
  ( ! [X41,X40] : (s(bool,'F_0') != s(bool,X41) | p(s(bool,X40))) ) | $spl18),
  inference(duplicate_literal_removal,[],[f1838])).
fof(f1838,plain,(
  ( ! [X41,X40] : (s(bool,'F_0') != s(bool,X41) | p(s(bool,X40)) | p(s(bool,X40))) ) | $spl18),
  inference(superposition,[],[f1319,f953])).
fof(f1822,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f1675,f1323])).
fof(f1675,plain,(
  ~p(s(bool,'T_0')) | $spl26),
  inference(resolution,[],[f1323,f1021])).
fof(f1821,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f1674,f1323])).
fof(f1674,plain,(
  ~p(s(bool,'T_0')) | $spl26),
  inference(resolution,[],[f1323,f1019])).
fof(f1820,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f1671,f1323])).
fof(f1671,plain,(
  ( ! [X326] : (~p(s(bool,X326))) ) | $spl26),
  inference(resolution,[],[f1323,f1013])).
fof(f1819,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f1664,f1323])).
fof(f1664,plain,(
  ( ! [X319] : (~p(s(bool,X319))) ) | $spl26),
  inference(resolution,[],[f1323,f951])).
fof(f1673,plain,(
  $false | $spl26),
  inference(resolution,[],[f1323,f1017])).
fof(f1695,plain,(
  $false | $spl26),
  inference(resolution,[],[f1323,f1055])).
fof(f1747,plain,(
  $false | ($spl1 | $spl26)),
  inference(resolution,[],[f1323,f1117])).
fof(f1482,plain,(
  ~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59(sK15,sK13)),s('type.list.list'(sK13),sK60(sK15,sK13)))) | $spl1),
  inference(superposition,[],[f1117,f1074])).
fof(f1474,plain,(
  ( ! [X194,X195] : (p(s(bool,X194)) | ~p(s(bool,X195)) | s(bool,X195) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl1),
  inference(resolution,[],[f1117,f1075])).
fof(f1370,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | s(bool,X33) != s(bool,X35)) ) | $spl1),
  inference(resolution,[],[f1117,f790])).
fof(f1369,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl1),
  inference(resolution,[],[f1117,f789])).
fof(f1367,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl1),
  inference(resolution,[],[f1117,f788])).
fof(f1365,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | s(bool,X18) != s(bool,X20)) ) | $spl1),
  inference(resolution,[],[f1117,f786])).
fof(f1363,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl1),
  inference(resolution,[],[f1117,f784])).
fof(f1257,plain,(
  ( ! [X329,X331,X330] : (p(s(bool,'const.bool.IN_2'(s(sK13,X329),s(fun(sK13,bool),sK14)))) | p(s(bool,X330)) | ~p(s(bool,X331)) | s(bool,X331) != s(bool,'const.bool.IN_2'(s(sK13,X329),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl2),
  inference(resolution,[],[f1087,f1075])).
fof(f1153,plain,(
  ( ! [X66,X64,X67,X65] : (p(s(bool,'const.bool.IN_2'(s(sK13,X64),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X65)) | ~p(s(bool,X66)) | s(bool,X66) != s(bool,'const.bool.IN_2'(s(sK13,X64),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X65) != s(bool,X67)) ) | $spl2),
  inference(resolution,[],[f1087,f790])).
fof(f1152,plain,(
  ( ! [X61,X62,X60,X63] : (p(s(bool,'const.bool.IN_2'(s(sK13,X60),s(fun(sK13,bool),sK14)))) | p(s(bool,X61)) | ~p(s(bool,X62)) | ~p(s(bool,X63)) | s(bool,X62) != s(bool,'const.bool.IN_2'(s(sK13,X60),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl2),
  inference(resolution,[],[f1087,f789])).
fof(f1150,plain,(
  ( ! [X54,X52,X55,X53] : (p(s(bool,'const.bool.IN_2'(s(sK13,X52),s(fun(sK13,bool),sK14)))) | p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X53) != s(bool,X55) | s(bool,X54) != s(bool,'const.bool.IN_2'(s(sK13,X52),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl2),
  inference(resolution,[],[f1087,f788])).
fof(f1148,plain,(
  ( ! [X47,X45,X46,X44] : (p(s(bool,'const.bool.IN_2'(s(sK13,X44),s(fun(sK13,bool),sK14)))) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X46) != s(bool,'const.bool.IN_2'(s(sK13,X44),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X45) != s(bool,X47)) ) | $spl2),
  inference(resolution,[],[f1087,f786])).
fof(f1146,plain,(
  ( ! [X39,X37,X38,X36] : (p(s(bool,'const.bool.IN_2'(s(sK13,X36),s(fun(sK13,bool),sK14)))) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,X39) | s(bool,X38) != s(bool,'const.bool.IN_2'(s(sK13,X36),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl2),
  inference(resolution,[],[f1087,f784])).
fof(f1131,plain,(
  s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)),s('type.list.list'(sK13),sK60('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)))) | $spl16),
  inference(cnf_transformation,[],[f1131_D])).
fof(f1131_D,plain,(
  s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)),s('type.list.list'(sK13),sK60('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1110,plain,(
  s('type.list.list'(sK13),'const.list.NIL_0') != s('type.list.list'(sK13),sK15) | s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK59('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)),s('type.list.list'(sK13),sK60('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)))) | $spl1),
  inference(superposition,[],[f1085,f1074])).
fof(f1127,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),sK14)))) | s('type.list.list'(sK13),'const.list.NIL_0') != s('type.list.list'(sK13),sK15) | $spl1),
  inference(forward_demodulation,[],[f1109,f760])).
fof(f1109,plain,(
  s('type.list.list'(sK13),'const.list.NIL_0') != s('type.list.list'(sK13),sK15) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK53(sK15,sK14,sK13))))) | $spl1),
  inference(superposition,[],[f1085,f1042])).
fof(f1108,plain,(
  s('type.list.list'(sK13),'const.list.NIL_0') != s('type.list.list'(sK13),sK15) | p(s(bool,'const.bool.IN_2'(s(sK13,sK53(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) | $spl1),
  inference(superposition,[],[f1085,f1041])).
fof(f1107,plain,(
  s('type.list.list'(sK13),'const.list.NIL_0') != s('type.list.list'(sK13),sK15) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(sK13,bool),sK14))),s('type.list.list'(sK13),sK15)))) | $spl1),
  inference(superposition,[],[f1085,f1065])).
fof(f1085,plain,(
  s('type.list.list'(sK13),sK15) != s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | $spl1),
  inference(cnf_transformation,[],[f1085_D])).
fof(f1085_D,plain,(
  s('type.list.list'(sK13),sK15) != s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f1092,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) | $spl6),
  inference(cnf_transformation,[],[f1092_D])).
fof(f1092_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f711,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) | s('type.list.list'(sK13),sK15) != s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))),
  inference(cnf_transformation,[],[f509])).
fof(f509,plain,(
  (s('type.list.list'(sK13),sK15) != s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | (p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) & ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16)))))) & (s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(sK13,X3),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13,sK14,sK15,sK16])],[f508])).
fof(f508,plain,(
  ? [X0,X1,X2] : ((s('type.list.list'(X0),X2) != s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) | ? [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))))) & (s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) | ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(rectify,[],[f507])).
fof(f507,plain,(
  ? [X0,X1,X2] : ((s('type.list.list'(X0),X2) != s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) | ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))) & (s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) | ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f435])).
fof(f435,plain,(
  ? [X0,X1,X2] : (s('type.list.list'(X0),X2) != s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) <~> ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))))),
  inference(ennf_transformation,[],[f264])).
fof(f264,plain,(
  ~! [X0,X1,X2] : (s('type.list.list'(X0),X2) != s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) <=> ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))))),
  inference(flattening,[],[f263])).
fof(f263,plain,(
  ~! [X0] : ! [X1,X2] : (~s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) <=> ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : (~s('type.list.list'(X6),X8) = s('type.list.list'(X6),'const.list.FILTER_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X8))) <=> ? [X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X8)))))) & ~p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X9))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : (~s('type.list.list'(X6),X8) = s('type.list.list'(X6),'const.list.FILTER_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X8))) <=> ? [X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X8)))))) & ~p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X9))))))),
  file('thm.list.FILTER_NEQ_ID','thm.list.FILTER_NEQ_ID')).
fof(f1088,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),sK14)))) | s('type.list.list'(sK13),sK15) != s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))),
  inference(forward_demodulation,[],[f712,f760])).
fof(f712,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16)))) | s('type.list.list'(sK13),sK15) != s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))),
  inference(cnf_transformation,[],[f509])).
fof(f1086,plain,(
  s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | $spl0),
  inference(cnf_transformation,[],[f1086_D])).
fof(f1086_D,plain,(
  s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f1087,plain,(
  ( ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(sK13,X3),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) | p(s(bool,'const.bool.IN_2'(s(sK13,X3),s(fun(sK13,bool),sK14))))) ) | $spl2),
  inference(cnf_transformation,[],[f1087_D])).
fof(f1087_D,plain,(
  ( ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(sK13,X3),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) | p(s(bool,'const.bool.IN_2'(s(sK13,X3),s(fun(sK13,bool),sK14))))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f1084,plain,(
  ( ! [X3] : (p(s(bool,'const.bool.IN_2'(s(sK13,X3),s(fun(sK13,bool),sK14)))) | ~p(s(bool,'const.bool.IN_2'(s(sK13,X3),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) )),
  inference(forward_demodulation,[],[f713,f760])).
fof(f713,plain,(
  ( ! [X3] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X3)))) | ~p(s(bool,'const.bool.IN_2'(s(sK13,X3),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) )),
  inference(cnf_transformation,[],[f509])).
% SZS output end Proof for thm.list
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 735
% Passive clauses: 12753
% Generated clauses: 34078
% Final active clauses: 463
% Final passive clauses: 5782
% Input formulas: 261
% Initial clauses: 364
% 
% Selected by SInE selection: 105
% SInE iterations: 1
% 
% Duplicate literals: 731
% Trivial inequalities: 107
% Fw subsumption resolutions: 2990
% Fw demodulations: 6505
% Bw demodulations: 61
% 
% Simple tautologies: 1372
% Equational tautologies: 5
% Forward subsumptions: 9353
% 
% Binary resolution: 8945
% Forward superposition: 7947
% Backward superposition: 5676
% Equality resolution: 6
% 
% Split clauses: 456
% Split components: 972
% SAT solver clauses: 523
% SAT solver unit clauses: 19
% SAT solver binary clauses: 39
% 
% Sat splits: 456
% Sat splitting refutations: 57
% 
% Memory used [KB]: 12153
% Time elapsed: 0.641 s
% ------------------------------
% ------------------------------
% Success in time 12.353 s
