Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.ringNorm
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ringNorm
% SZS output start Proof for thm.ringNorm
fof(f303,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f302])).
fof(f302,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_spolynom_simplify_2'(s('type.ring.ring'(sK0),sK1),s('type.canonical.spolynom'(sK0),sK2))) != s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_spolynom_simplify_2'(s('type.ring.ring'(sK0),sK1),s('type.canonical.spolynom'(sK0),sK2)))),
  inference(superposition,[],[f50,f296])).
fof(f296,plain,(
  ( ! [X4,X5,X3] : (s('type.canonical.canonical_sum'(X3),'const.ringNorm.r_spolynom_simplify_2'(s('type.ring.ring'(X3),X4),s('type.canonical.spolynom'(X3),X5))) = s('type.canonical.canonical_sum'(X3),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(X3),X4),s('type.canonical.canonical_sum'(X3),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X3),X4),s('type.canonical.spolynom'(X3),X5)))))) )),
  inference(forward_demodulation,[],[f295,f241])).
fof(f241,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_simplify_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_simplify_2'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))) )),
  inference(forward_demodulation,[],[f238,f59])).
fof(f59,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_simplify_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_spolynom_simplify_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))) )),
  inference(cnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_simplify_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_spolynom_simplify_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))),
  inference(flattening,[],[f31])).
fof(f31,plain,(
  ! [X0] : ! [X1,X2] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_simplify_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_spolynom_simplify_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : ! [X11,X12] : s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_spolynom_simplify_2'(s('type.ring.ring'(X6),X11),s('type.canonical.spolynom'(X6),X12))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.spolynom'(X6),'type.canonical.canonical_sum'(X6)),'const.ringNorm.r_spolynom_simplify_1'(s('type.ring.ring'(X6),X11))),s('type.canonical.spolynom'(X6),X12)))),
  file('thm.ringNorm.spolynom_simplify_def','ARITY_1_2_const.ringNorm.r_spolynom_simplify')).
fof(f238,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_simplify_2'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_spolynom_simplify_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))) )),
  inference(superposition,[],[f60,f56])).
fof(f56,plain,(
  ( ! [X0,X1] : (s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_spolynom_simplify_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.spolynom_simplify_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f26])).
fof(f26,plain,(
  ! [X0,X1] : s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_spolynom_simplify_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.spolynom_simplify_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f25])).
fof(f25,plain,(
  ! [X0] : ! [X1] : s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_spolynom_simplify_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.spolynom_simplify_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : ! [X7] : s(fun('type.canonical.spolynom'(X6),'type.canonical.canonical_sum'(X6)),'const.ringNorm.r_spolynom_simplify_1'(s('type.ring.ring'(X6),X7))) = s(fun('type.canonical.spolynom'(X6),'type.canonical.canonical_sum'(X6)),'const.canonical.spolynom_simplify_1'(s('type.semi_ring.semi_ring'(X6),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X6),X7)))))),
  file('thm.ringNorm.spolynom_simplify_def','thm.ringNorm.r_spolynom_simplify_def')).
fof(f60,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.spolynom_simplify_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))) )),
  inference(cnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.spolynom_simplify_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ! [X0] : ! [X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.spolynom_simplify_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6] : ! [X11,X12] : s('type.canonical.canonical_sum'(X6),'const.canonical.spolynom_simplify_2'(s('type.semi_ring.semi_ring'(X6),X11),s('type.canonical.spolynom'(X6),X12))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.spolynom'(X6),'type.canonical.canonical_sum'(X6)),'const.canonical.spolynom_simplify_1'(s('type.semi_ring.semi_ring'(X6),X11))),s('type.canonical.spolynom'(X6),X12)))),
  file('thm.ringNorm.spolynom_simplify_def','ARITY_1_2_const.canonical.spolynom_simplify')).
fof(f295,plain,(
  ( ! [X4,X5,X3] : (s('type.canonical.canonical_sum'(X3),'const.canonical.spolynom_simplify_2'(s('type.semi_ring.semi_ring'(X3),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X3),X4))),s('type.canonical.spolynom'(X3),X5))) = s('type.canonical.canonical_sum'(X3),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(X3),X4),s('type.canonical.canonical_sum'(X3),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X3),X4),s('type.canonical.spolynom'(X3),X5)))))) )),
  inference(forward_demodulation,[],[f291,f249])).
fof(f249,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))) )),
  inference(forward_demodulation,[],[f246,f61])).
fof(f61,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_spolynom_normalize_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))) )),
  inference(cnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_spolynom_normalize_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))),
  inference(flattening,[],[f35])).
fof(f35,plain,(
  ! [X0] : ! [X1,X2] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_spolynom_normalize_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6] : ! [X11,X12] : s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X6),X11),s('type.canonical.spolynom'(X6),X12))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.spolynom'(X6),'type.canonical.canonical_sum'(X6)),'const.ringNorm.r_spolynom_normalize_1'(s('type.ring.ring'(X6),X11))),s('type.canonical.spolynom'(X6),X12)))),
  file('thm.ringNorm.spolynom_simplify_def','ARITY_1_2_const.ringNorm.r_spolynom_normalize')).
fof(f246,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_spolynom_normalize_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))) )),
  inference(superposition,[],[f62,f55])).
fof(f55,plain,(
  ( ! [X0,X1] : (s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_spolynom_normalize_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f24])).
fof(f24,plain,(
  ! [X0,X1] : s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_spolynom_normalize_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f23])).
fof(f23,plain,(
  ! [X0] : ! [X1] : s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_spolynom_normalize_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : ! [X7] : s(fun('type.canonical.spolynom'(X6),'type.canonical.canonical_sum'(X6)),'const.ringNorm.r_spolynom_normalize_1'(s('type.ring.ring'(X6),X7))) = s(fun('type.canonical.spolynom'(X6),'type.canonical.canonical_sum'(X6)),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'(X6),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X6),X7)))))),
  file('thm.ringNorm.spolynom_simplify_def','thm.ringNorm.r_spolynom_normalize_def')).
fof(f62,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))) )),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ! [X0] : ! [X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.spolynom'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.spolynom'(X0),X2)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6] : ! [X11,X12] : s('type.canonical.canonical_sum'(X6),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X6),X11),s('type.canonical.spolynom'(X6),X12))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.spolynom'(X6),'type.canonical.canonical_sum'(X6)),'const.canonical.spolynom_normalize_1'(s('type.semi_ring.semi_ring'(X6),X11))),s('type.canonical.spolynom'(X6),X12)))),
  file('thm.ringNorm.spolynom_simplify_def','ARITY_1_2_const.canonical.spolynom_normalize')).
fof(f291,plain,(
  ( ! [X4,X5,X3] : (s('type.canonical.canonical_sum'(X3),'const.canonical.spolynom_simplify_2'(s('type.semi_ring.semi_ring'(X3),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X3),X4))),s('type.canonical.spolynom'(X3),X5))) = s('type.canonical.canonical_sum'(X3),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(X3),X4),s('type.canonical.canonical_sum'(X3),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X3),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X3),X4))),s('type.canonical.spolynom'(X3),X5)))))) )),
  inference(superposition,[],[f280,f63])).
fof(f63,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X2)))))) )),
  inference(cnf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X2)))))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ! [X0] : ! [X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.spolynom'(X0),X2)))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X10,X8] : s('type.canonical.canonical_sum'(X6),'const.canonical.spolynom_simplify_2'(s('type.semi_ring.semi_ring'(X6),X10),s('type.canonical.spolynom'(X6),X8))) = s('type.canonical.canonical_sum'(X6),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X6),X10),s('type.canonical.canonical_sum'(X6),'const.canonical.spolynom_normalize_2'(s('type.semi_ring.semi_ring'(X6),X10),s('type.canonical.spolynom'(X6),X8)))))),
  file('thm.ringNorm.spolynom_simplify_def','thm.canonical.spolynom_simplify_def')).
fof(f280,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2))) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2)))) )),
  inference(forward_demodulation,[],[f231,f58])).
fof(f58,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_canonical_sum_simplify_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2)))) )),
  inference(cnf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_canonical_sum_simplify_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2)))),
  inference(flattening,[],[f29])).
fof(f29,plain,(
  ! [X0] : ! [X1,X2] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_canonical_sum_simplify_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6] : ! [X11,X12] : s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(X6),X11),s('type.canonical.canonical_sum'(X6),X12))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)),'const.ringNorm.r_canonical_sum_simplify_1'(s('type.ring.ring'(X6),X11))),s('type.canonical.canonical_sum'(X6),X12)))),
  file('thm.ringNorm.spolynom_simplify_def','ARITY_1_2_const.ringNorm.r_canonical_sum_simplify')).
fof(f231,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_canonical_sum_simplify_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2)))) )),
  inference(superposition,[],[f57,f54])).
fof(f54,plain,(
  ( ! [X0,X1] : (s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_canonical_sum_simplify_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.canonical_sum_simplify_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f22])).
fof(f22,plain,(
  ! [X0,X1] : s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_canonical_sum_simplify_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.canonical_sum_simplify_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f21])).
fof(f21,plain,(
  ! [X0] : ! [X1] : s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_canonical_sum_simplify_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.canonical_sum_simplify_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : ! [X7] : s(fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)),'const.ringNorm.r_canonical_sum_simplify_1'(s('type.ring.ring'(X6),X7))) = s(fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)),'const.canonical.canonical_sum_simplify_1'(s('type.semi_ring.semi_ring'(X6),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X6),X7)))))),
  file('thm.ringNorm.spolynom_simplify_def','thm.ringNorm.r_canonical_sum_simplify_def')).
fof(f57,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.canonical_sum_simplify_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2)))) )),
  inference(cnf_transformation,[],[f28])).
fof(f28,plain,(
  ! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.canonical_sum_simplify_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2)))),
  inference(flattening,[],[f27])).
fof(f27,plain,(
  ! [X0] : ! [X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.canonical_sum_simplify_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X6] : ! [X11,X12] : s('type.canonical.canonical_sum'(X6),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X6),X11),s('type.canonical.canonical_sum'(X6),X12))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)),'const.canonical.canonical_sum_simplify_1'(s('type.semi_ring.semi_ring'(X6),X11))),s('type.canonical.canonical_sum'(X6),X12)))),
  file('thm.ringNorm.spolynom_simplify_def','ARITY_1_2_const.canonical.canonical_sum_simplify')).
fof(f50,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_spolynom_simplify_2'(s('type.ring.ring'(sK0),sK1),s('type.canonical.spolynom'(sK0),sK2))) != s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(sK0),sK1),s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(sK0),sK1),s('type.canonical.spolynom'(sK0),sK2)))))),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_spolynom_simplify_2'(s('type.ring.ring'(sK0),sK1),s('type.canonical.spolynom'(sK0),sK2))) != s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(sK0),sK1),s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(sK0),sK1),s('type.canonical.spolynom'(sK0),sK2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f45])).
fof(f45,plain,(
  ? [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_simplify_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) != s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),X2)))))),
  inference(ennf_transformation,[],[f19])).
fof(f19,plain,(
  ~! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_simplify_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),X2)))))),
  inference(flattening,[],[f18])).
fof(f18,plain,(
  ~! [X0] : ! [X1,X2] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_simplify_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),X2))) = s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X0),X1),s('type.canonical.spolynom'(X0),X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_spolynom_simplify_2'(s('type.ring.ring'(X6),X7),s('type.canonical.spolynom'(X6),X8))) = s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(X6),X7),s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X6),X7),s('type.canonical.spolynom'(X6),X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_spolynom_simplify_2'(s('type.ring.ring'(X6),X7),s('type.canonical.spolynom'(X6),X8))) = s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(X6),X7),s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_spolynom_normalize_2'(s('type.ring.ring'(X6),X7),s('type.canonical.spolynom'(X6),X8)))))),
  file('thm.ringNorm.spolynom_simplify_def','thm.ringNorm.spolynom_simplify_def')).
% SZS output end Proof for thm.ringNorm
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 26
% Passive clauses: 47
% Generated clauses: 238
% Final active clauses: 24
% Final passive clauses: 16
% Input formulas: 16
% Initial clauses: 17
% 
% Duplicate literals: 42
% Trivial inequalities: 3
% Fw demodulations: 8
% 
% Simple tautologies: 18
% Equational tautologies: 14
% Forward subsumptions: 105
% 
% Binary resolution: 4
% Forward superposition: 54
% Backward superposition: 58
% Self superposition: 27
% Equality factoring: 26
% 
% Split clauses: 1
% Split components: 2
% SAT solver clauses: 2
% SAT solver unit clauses: 1
% SAT solver binary clauses: 1
% 
% Sat splits: 1
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 2
% TWLsolver calls for satisfiability: 2
% 
% Memory used [KB]: 895
% Time elapsed: 0.018 s
% ------------------------------
% ------------------------------
% Success in time 0.02 s
