Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.extreal
% (72174)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 717
% Passive clauses: 11598
% Generated clauses: 20141
% Final active clauses: 710
% Final passive clauses: 10487
% Input formulas: 45
% Initial clauses: 192
% 
% Duplicate literals: 303
% Trivial inequalities: 16
% Fw subsumption resolutions: 1701
% Fw demodulations: 245
% Bw demodulations: 6
% 
% Simple tautologies: 239
% Equational tautologies: 54
% Forward subsumptions: 5711
% 
% Binary resolution: 12982
% Factoring: 1035
% Forward superposition: 1729
% Backward superposition: 1723
% Self superposition: 22
% Equality factoring: 26
% Equality resolution: 4
% 
% Split clauses: 333
% Split components: 667
% SAT solver clauses: 671
% SAT solver unit clauses: 17
% SAT solver binary clauses: 182
% 
% Sat splits: 333
% 
% TWLsolver clauses: 671
% TWLsolver calls for satisfiability: 414
% 
% Memory used [KB]: 9466
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.extreal
% (72191)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 509
% Passive clauses: 3560
% Generated clauses: 13308
% Final active clauses: 471
% Final passive clauses: 3016
% Input formulas: 45
% Initial clauses: 192
% 
% Duplicate literals: 205
% Trivial inequalities: 26
% Fw subsumption resolutions: 943
% Fw demodulations: 374
% Bw demodulations: 44
% Condensations: 256
% Global subsumptions: 13
% 
% Simple tautologies: 266
% Equational tautologies: 17
% Forward subsumptions: 3708
% 
% Binary resolution: 4608
% Factoring: 105
% Forward superposition: 1338
% Backward superposition: 1459
% Self superposition: 22
% Equality factoring: 26
% Equality resolution: 4
% 
% Split clauses: 58
% Split components: 116
% SAT solver clauses: 4115
% SAT solver unit clauses: 189
% SAT solver binary clauses: 94
% SAT solver learnt clauses: 27
% SAT solver learnt literals: 11
% 
% Sat splits: 58
% Sat splitting refutations: 9
% 
% TWLsolver clauses: 4062
% TWLsolver calls for satisfiability: 1876
% 
% Memory used [KB]: 14839
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.extreal
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.extreal
% SZS output start Proof for thm.extreal
fof(f58816,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f262,f463,f461,f261,f465,f260,f462,f467,f2553,f2557,f2561,f2564,f2567,f2570,f2573,f2576,f2579,f2582,f2585,f2588,f2591,f2594,f2597,f2600,f2603,f2643,f2554,f2645,f2558,f2648,f2562,f2650,f2565,f2652,f2568,f2654,f2571,f2656,f2574,f2658,f2577,f2660,f2580,f2662,f2583,f2664,f2586,f2666,f2589,f2668,f2592,f2670,f2595,f2672,f2598,f2674,f2601,f2676,f2604,f2769,f2770,f2785,f2787,f2789,f2791,f2793,f2795,f2797,f2799,f2801,f2803,f2805,f2807,f2809,f2811,f2813,f2815,f2817,f2819,f2821,f2823,f3882,f3885,f3883,f5004,f5006,f5005,f5008,f5010,f5012,f5018,f5020,f5022,f5024,f5028,f5030,f5032,f5034,f5036,f5038,f5052,f5054,f5068,f5070,f5081,f5092,f5096,f5098,f5103,f5108,f5112,f5114,f5118,f5120,f5124,f5126,f5146,f5176,f5227,f5229,f5473,f5477,f5485,f5502,f5507,f5519,f5529,f5541,f5551,f5553,f5565,f5575,f5587,f5597,f5599,f5601,f5603,f5605,f6613,f6614,f5629,f5631,f5633,f5635,f5637,f5639,f5641,f5643,f5645,f5647,f5677,f5679,f5681,f5683,f5698,f5713,f5731,f5735,f5745,f5747,f5754,f5756,f5774,f5779,f5789,f5791,f5798,f5800,f5812,f5822,f5834,f5844,f5846,f5863,f5865,f5876,f5878,f5896,f5898,f5909,f5911,f5931,f5933,f5944,f5946,f5960,f5975,f5989,f6004,f6023,f6042,f6064,f6066,f6092,f6096,f6098,f6113,f6115,f6117,f6125,f6127,f6135,f6137,f6145,f6147,f6155,f6157,f6178,f6195,f6212,f6214,f6226,f6228,f6230,f6261,f6263,f6265,f6267,f6269,f6271,f6273,f6292,f6295,f6314,f6317,f6338,f6340,f6360,f6362,f6364,f6373,f6381,f6389,f6391,f6399,f6401,f6417,f6434,f6453,f6470,f6478,f6480,f6482,f6484,f6486,f6488,f6490,f6499,f6501,f6503,f6505,f6507,f6509,f6511,f6513,f6515,f6518,f6531,f6533,f6540,f6542,f6549,f6551,f6556,f6558,f6560,f6574,f6591,f6615,f6616,f6617,f6618,f6619,f6620,f6621,f6622,f6623,f6624,f6625,f6626,f6627,f6628,f6629,f6630,f6631,f6632,f6633,f6634,f6635,f6636,f6637,f6638,f6639,f6640,f6641,f6642,f6643,f6644,f6645,f6646,f6647,f6648,f6649,f6650,f6651,f6652,f6653,f6654,f6655,f6656,f6657,f6658,f6659,f6660,f6661,f6662,f6663,f6664,f6665,f6666,f6667,f6668,f6669,f6670,f6671,f6672,f6673,f6674,f6675,f6676,f6677,f6678,f6679,f6680,f6681,f6682,f6683,f6684,f6685,f6686,f6687,f6688,f6689,f6690,f6691,f6692,f6693,f6694,f6695,f6696,f6697,f6698,f6699,f6700,f6701,f6702,f6703,f6704,f6705,f6706,f6707,f6708,f6709,f6710,f6711,f6712,f6713,f6714,f6715,f6716,f6717,f6718,f6719,f6720,f6721,f6722,f6723,f6724,f6725,f6726,f6727,f6728,f6729,f6730,f6731,f6732,f6733,f6734,f6735,f6736,f6737,f6738,f6739,f6740,f6741,f6742,f6743,f6744,f6745,f6746,f6747,f6748,f6749,f6750,f6751,f6752,f6753,f6754,f6755,f6756,f6757,f6758,f6759,f6760,f6761,f6762,f6763,f6764,f6765,f6766,f6767,f6768,f6769,f6770,f6771,f6772,f6773,f6774,f6776,f6778,f6779,f6781,f6783,f6785,f6787,f6788,f6789,f6790,f6792,f6793,f6794,f6795,f6796,f6797,f6798,f6799,f6800,f6801,f6805,f6806,f6807,f6808,f6809,f6810,f6811,f6812,f6813,f6814,f6815,f6816,f6817,f6818,f6819,f6820,f6821,f6822,f6823,f6824,f6825,f6826,f6827,f6828,f6829,f6830,f6831,f6832,f6833,f6834,f6835,f6836,f6837,f6838,f6839,f6840,f6841,f6842,f6843,f6844,f6845,f6846,f6847,f6848,f6849,f6850,f6851,f6852,f6853,f6854,f6855,f6856,f6857,f6858,f6859,f6860,f6861,f6862,f6863,f6864,f6865,f6866,f6867,f6868,f6869,f6870,f6871,f6872,f6873,f6874,f6875,f6876,f6877,f6878,f6879,f6880,f6881,f6882,f6883,f6884,f6885,f6886,f6887,f6888,f7090,f7091,f7098,f7095,f7092,f7093,f7101,f7094,f7103,f7096,f7100,f7208,f7209,f9417,f9418,f9419,f9420,f9484,f9485,f9486,f9487,f10139,f11737,f21558,f21561,f31760,f31761,f31762,f31764,f31766,f31768,f31820,f31824,f31826,f31828,f31830,f31832,f31834,f31838,f31842,f31844,f31845,f31987,f31763,f31988,f31765,f31989,f31767,f31990,f31769,f32041,f31821,f32044,f31825,f32045,f31827,f32046,f31829,f32047,f31831,f32048,f31833,f32049,f31835,f32052,f31839,f32055,f31843,f32735,f32740,f32736,f32868,f32873,f32869,f33954,f33959,f33955,f33992,f33997,f33993,f34048,f34053,f34049,f36198,f36199,f36200,f36201,f36202,f36203,f36204,f36205,f37031,f37010,f37051,f37052,f37053,f37050,f37057,f37058,f36989,f37060,f37062,f37063,f37064,f36620,f37065,f36621,f37066,f37067,f36623,f37068,f36624,f37069,f36625,f37070,f36626,f37071,f36627,f37072,f36628,f37073,f36629,f37074,f36630,f37075,f36631,f37076,f36632,f37077,f36633,f37078,f36634,f37079,f37080,f37081,f36966,f36660,f37082,f36661,f37083,f36662,f37084,f36663,f37085,f36664,f37086,f37088,f36680,f37091,f36684,f37092,f36685,f37093,f36686,f37094,f37096,f37097,f37098,f37099,f37100,f37101,f37102,f37103,f37104,f37105,f37106,f37107,f37108,f36714,f37110,f36715,f37111,f36716,f37112,f36717,f37113,f36718,f37114,f36719,f37115,f36720,f37116,f36721,f37117,f36722,f37118,f36723,f37119,f37120,f37122,f37121,f36731,f37125,f37126,f37127,f36738,f37129,f36745,f37059,f37130,f37140,f37141,f37142,f37143,f37144,f37145,f37146,f37147,f37148,f37149,f37615,f37620,f37616,f39690,f39697,f39701,f39705,f39709,f39711,f39713,f39670,f39669,f39668,f39710,f39707,f39715,f39666,f39716,f39665,f39664,f39720,f39663,f39695,f39693,f39717,f39660,f39719,f39691,f39688,f39692,f30081,f29441,f39721,f28801,f28605,f39722,f7219,f21522,f39724,f21559,f21523,f7221,f39725,f44645,f44647,f44646,f45274,f45275,f45310,f45311,f45312,f45313,f45314,f45315,f45316,f45317,f45318,f45319,f46140,f46141,f46171,f46172,f46179,f46180,f46181,f46182,f46183,f46184,f46187,f46188,f46189,f46190,f47088,f47089,f47122,f47123,f47124,f47125,f47126,f47127,f47128,f47129,f47130,f3884,f47131,f58798,f58804,f2555,f58806])).
fof(f58806,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.min.@_1'(s(fun('type.extreal.extreal',bool),sK13))),s('type.extreal.extreal',sK14)))) | $spl296),
  inference(cnf_transformation,[],[f58806_D])).
fof(f58806_D,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.min.@_1'(s(fun('type.extreal.extreal',bool),sK13))),s('type.extreal.extreal',sK14)))) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f2555,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl10),
  inference(cnf_transformation,[],[f2555_D])).
fof(f2555_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f58804,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.min.@_1'(s(fun('type.extreal.extreal',bool),sK13))),s('type.extreal.extreal',sK14))))) ) | ($spl3 | $spl6)),
  inference(duplicate_literal_removal,[],[f58803])).
fof(f58803,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.min.@_1'(s(fun('type.extreal.extreal',bool),sK13))),s('type.extreal.extreal',sK14)))) | s(bool,'T_0') = s(bool,X2)) ) | ($spl3 | $spl6)),
  inference(forward_demodulation,[],[f58802,f1457])).
fof(f1457,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13))) | $spl6),
  inference(unit_resulting_resolution,[],[f263,f1452,f321])).
fof(f321,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f128])).
fof(f128,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X10,X11] : ((p(s(bool,X10)) => p(s(bool,X11))) => ((p(s(bool,X11)) => p(s(bool,X10))) => s(bool,X10) = s(bool,X11)))),
  file('thm.extreal.sup_lt','thm.bool.IMP_ANTISYM_AX')).
fof(f1452,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)))) | $spl6),
  inference(subsumption_resolution,[],[f1450,f263])).
fof(f1450,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)))) | $spl6),
  inference(superposition,[],[f469,f530])).
fof(f530,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK15))) | $spl6),
  inference(unit_resulting_resolution,[],[f467,f277])).
fof(f277,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.extreal.sup_lt','thm.bool.EQ_CLAUSES|split|1')).
fof(f469,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(forward_demodulation,[],[f306,f303])).
fof(f303,plain,(
  ( ! [X0,X1] : (s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) )),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1] : s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ! [X0] : ! [X1] : s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X9] : ! [X5] : s(bool,'|dollar|exists_1'(s(fun(X9,bool),X5))) = s(bool,happ(s(fun(X9,bool),X5),s(X9,'const.min.@_1'(s(fun(X9,bool),X5)))))),
  file('thm.extreal.sup_lt','thm.bool.EXISTS_DEF')).
fof(f306,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(ennf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(flattening,[],[f75])).
fof(f75,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X9] : ! [X6,X8] : (p(s(bool,happ(s(fun(X9,bool),X6),s(X9,X8)))) => p(s(bool,happ(s(fun(X9,bool),X6),s(X9,'const.min.@_1'(s(fun(X9,bool),X6)))))))),
  file('thm.extreal.sup_lt','thm.bool.SELECT_AX')).
fof(f263,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.extreal.sup_lt',aHL_TRUTH)).
fof(f58802,plain,(
  ( ! [X2] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.min.@_1'(s(fun('type.extreal.extreal',bool),sK13))),s('type.extreal.extreal',sK14)))) | s(bool,X2) = s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13))) | s(bool,'T_0') = s(bool,X2)) ) | $spl3),
  inference(subsumption_resolution,[],[f58771,f263])).
fof(f58771,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.min.@_1'(s(fun('type.extreal.extreal',bool),sK13))),s('type.extreal.extreal',sK14)))) | s(bool,X2) = s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13))) | s(bool,'T_0') = s(bool,X2)) ) | $spl3),
  inference(superposition,[],[f16608,f2039])).
fof(f2039,plain,(
  ( ! [X23,X21,X22] : (s(bool,X23) = s(bool,'|dollar|exists_1'(s(fun(X21,bool),X22))) | s(bool,'T_0') = s(bool,happ(s(fun(X21,bool),X22),s(X21,'const.min.@_1'(s(fun(X21,bool),X22))))) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f783,f303])).
fof(f783,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f270,f270])).
fof(f270,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.extreal.sup_lt',aHL_BOOL_CASES)).
fof(f16608,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',X0)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',sK14))))) ) | $spl3),
  inference(subsumption_resolution,[],[f16555,f263])).
fof(f16555,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',X0)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',sK14))))) ) | $spl3),
  inference(superposition,[],[f326,f12256])).
fof(f12256,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13))),s('type.extreal.extreal',sK14))) | $spl3),
  inference(unit_resulting_resolution,[],[f463,f1296])).
fof(f1296,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) )),
  inference(subsumption_resolution,[],[f1295,f468])).
fof(f468,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f266,f263])).
fof(f266,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f48])).
fof(f48,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f18])).
fof(f18,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.extreal.sup_lt','thm.bool.NOT_CLAUSES|split|2')).
fof(f1295,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) )),
  inference(superposition,[],[f325,f270])).
fof(f325,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))) )),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))))) & (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(nnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) <=> ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))))),
  inference(flattening,[],[f88])).
fof(f88,plain,(
  ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) <=> ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X8,X7] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X8),s('type.extreal.extreal',X7)))) <=> ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X8)))))),
  file('thm.extreal.sup_lt','thm.extreal.extreal_lt_def')).
fof(f326,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X0))),s('type.extreal.extreal',X1)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X3)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X1))))) )),
  inference(cnf_transformation,[],[f202])).
fof(f202,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X0))),s('type.extreal.extreal',X1)))) | ! [X3] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X3)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X1)))))) & ((p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',sK22(X1,X0))))) & ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK22(X1,X0)),s('type.extreal.extreal',X1))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X0))),s('type.extreal.extreal',X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK22])],[f201])).
fof(f201,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X0))),s('type.extreal.extreal',X1)))) | ! [X3] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X3)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X1)))))) & (? [X2] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X2)))) & ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X1))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X0))),s('type.extreal.extreal',X1))))))),
  inference(rectify,[],[f200])).
fof(f200,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X0))),s('type.extreal.extreal',X1)))) | ! [X2] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X2)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X1)))))) & (? [X2] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X2)))) & ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X1))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X0))),s('type.extreal.extreal',X1))))))),
  inference(nnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X0))),s('type.extreal.extreal',X1)))) <=> ! [X2] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X2)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X1))))))),
  inference(ennf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0,X1] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X0))),s('type.extreal.extreal',X1)))) <=> ! [X2] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X2)))) => p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X1))))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X17,X8] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X17))),s('type.extreal.extreal',X8)))) <=> ! [X7] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),X17),s('type.extreal.extreal',X7)))) => p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X8))))))),
  file('thm.extreal.sup_lt','thm.extreal.sup_le')).
fof(f58798,plain,(
  $false | ($spl3 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f58797,f468])).
fof(f58797,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f58796,f13097])).
fof(f13097,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal',sK14))) | $spl4),
  inference(unit_resulting_resolution,[],[f525,f11742,f846])).
fof(f846,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f836,f468])).
fof(f836,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f318,f270])).
fof(f318,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f129])).
fof(f11742,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal',sK14)))) | $spl4),
  inference(unit_resulting_resolution,[],[f465,f324])).
fof(f324,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))) )),
  inference(cnf_transformation,[],[f199])).
fof(f525,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f468,f276])).
fof(f276,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f171])).
fof(f58796,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal',sK14)))) | ($spl3 | $spl6)),
  inference(subsumption_resolution,[],[f58766,f263])).
fof(f58766,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal',sK14)))) | ($spl3 | $spl6)),
  inference(superposition,[],[f16608,f530])).
fof(f47131,plain,(
  ( ! [X212,X215,X214] : (sP1('const.min.@_1'(s(fun(bool,bool),X212)),X214) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X212))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X212),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X215) | s(bool,X214) = s(bool,X215)) ) | $spl294),
  inference(cnf_transformation,[],[f47131_D])).
fof(f47131_D,plain,(
  ( ! [X212,X215,X214] : (sP1('const.min.@_1'(s(fun(bool,bool),X212)),X214) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X212))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X212),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X215) | s(bool,X214) = s(bool,X215)) ) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f3884,plain,(
  ( ! [X8] : (p(s(bool,X8))) ) | $spl46),
  inference(cnf_transformation,[],[f3884_D])).
fof(f3884_D,plain,(
  ( ! [X8] : (p(s(bool,X8))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f47130,plain,(
  ( ! [X212,X215,X213,X214] : (p(s(bool,X213)) | sP1('const.min.@_1'(s(fun(bool,bool),X212)),X214) | s(bool,X214) = s(bool,X215) | s(bool,'T_0') = s(bool,X215) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X212),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X212)))) )),
  inference(subsumption_resolution,[],[f46772,f276])).
fof(f46772,plain,(
  ( ! [X212,X215,X213,X214] : (p(s(bool,X213)) | sP1('const.min.@_1'(s(fun(bool,bool),X212)),X214) | s(bool,X214) = s(bool,X215) | s(bool,'T_0') = s(bool,X215) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X212),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X212))) | s(bool,'T_0') = s(bool,X213)) )),
  inference(superposition,[],[f2610,f4907])).
fof(f4907,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0'))) | s(bool,X1) = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f1769,f783])).
fof(f1769,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,happ(s(fun(X2,bool),X3),s(X2,'const.min.@_1'(s(fun(X2,bool),X3))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X2,bool),X3)))) )),
  inference(superposition,[],[f303,f270])).
fof(f2610,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f2247,f263])).
fof(f2247,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f353,f783])).
fof(f353,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f213])).
fof(f213,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f212])).
fof(f212,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f211])).
fof(f211,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f47129,plain,(
  ( ! [X210,X208,X211] : (sP1(X210,'const.min.@_1'(s(fun(bool,bool),X208))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X208))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X208),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X211) | s(bool,X210) = s(bool,X211)) ) | $spl292),
  inference(cnf_transformation,[],[f47129_D])).
fof(f47129_D,plain,(
  ( ! [X210,X208,X211] : (sP1(X210,'const.min.@_1'(s(fun(bool,bool),X208))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X208))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X208),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X211) | s(bool,X210) = s(bool,X211)) ) <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f47128,plain,(
  ( ! [X210,X208,X211,X209] : (p(s(bool,X209)) | sP1(X210,'const.min.@_1'(s(fun(bool,bool),X208))) | s(bool,X210) = s(bool,X211) | s(bool,'T_0') = s(bool,X211) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X208),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X208)))) )),
  inference(subsumption_resolution,[],[f46771,f276])).
fof(f46771,plain,(
  ( ! [X210,X208,X211,X209] : (p(s(bool,X209)) | sP1(X210,'const.min.@_1'(s(fun(bool,bool),X208))) | s(bool,X210) = s(bool,X211) | s(bool,'T_0') = s(bool,X211) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X208),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X208))) | s(bool,'T_0') = s(bool,X209)) )),
  inference(superposition,[],[f2609,f4907])).
fof(f2609,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(subsumption_resolution,[],[f2246,f263])).
fof(f2246,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f350,f783])).
fof(f350,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f213])).
fof(f47127,plain,(
  ( ! [X199] : (sP1('const.min.@_1'(s(fun(bool,bool),X199)),'T_0') | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X199))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X199),s(bool,'T_0')))) ) | $spl290),
  inference(cnf_transformation,[],[f47127_D])).
fof(f47127_D,plain,(
  ( ! [X199] : (sP1('const.min.@_1'(s(fun(bool,bool),X199)),'T_0') | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X199))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X199),s(bool,'T_0')))) ) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f47126,plain,(
  ( ! [X200,X199] : (p(s(bool,X200)) | sP1('const.min.@_1'(s(fun(bool,bool),X199)),'T_0') | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X199),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X199)))) ) | $spl12),
  inference(subsumption_resolution,[],[f46768,f276])).
fof(f46768,plain,(
  ( ! [X200,X199] : (p(s(bool,X200)) | sP1('const.min.@_1'(s(fun(bool,bool),X199)),'T_0') | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X199),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X199))) | s(bool,'T_0') = s(bool,X200)) ) | $spl12),
  inference(superposition,[],[f2558,f4907])).
fof(f47125,plain,(
  ( ! [X197] : (sP1('T_0','const.min.@_1'(s(fun(bool,bool),X197))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X197))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X197),s(bool,'T_0')))) ) | $spl288),
  inference(cnf_transformation,[],[f47125_D])).
fof(f47125_D,plain,(
  ( ! [X197] : (sP1('T_0','const.min.@_1'(s(fun(bool,bool),X197))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X197))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X197),s(bool,'T_0')))) ) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f47124,plain,(
  ( ! [X198,X197] : (p(s(bool,X198)) | sP1('T_0','const.min.@_1'(s(fun(bool,bool),X197))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X197),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X197)))) ) | $spl8),
  inference(subsumption_resolution,[],[f46767,f276])).
fof(f46767,plain,(
  ( ! [X198,X197] : (p(s(bool,X198)) | sP1('T_0','const.min.@_1'(s(fun(bool,bool),X197))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X197),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X197))) | s(bool,'T_0') = s(bool,X198)) ) | $spl8),
  inference(superposition,[],[f2554,f4907])).
fof(f47123,plain,(
  ( ! [X192,X190,X193] : (s(bool,X192) != s(bool,'const.min.@_1'(s(fun(bool,bool),X190))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X190))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X190),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X193) | s(bool,X192) = s(bool,X193)) ) | $spl286),
  inference(cnf_transformation,[],[f47123_D])).
fof(f47123_D,plain,(
  ( ! [X192,X190,X193] : (s(bool,X192) != s(bool,'const.min.@_1'(s(fun(bool,bool),X190))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X190))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X190),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X193) | s(bool,X192) = s(bool,X193)) ) <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f47122,plain,(
  ( ! [X191,X192,X190,X193] : (p(s(bool,X191)) | s(bool,X192) != s(bool,'const.min.@_1'(s(fun(bool,bool),X190))) | s(bool,X192) = s(bool,X193) | s(bool,'T_0') = s(bool,X193) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X190),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X190)))) )),
  inference(subsumption_resolution,[],[f46765,f276])).
fof(f46765,plain,(
  ( ! [X191,X192,X190,X193] : (p(s(bool,X191)) | s(bool,X192) != s(bool,'const.min.@_1'(s(fun(bool,bool),X190))) | s(bool,X192) = s(bool,X193) | s(bool,'T_0') = s(bool,X193) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X190),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X190))) | s(bool,'T_0') = s(bool,X191)) )),
  inference(superposition,[],[f2204,f4907])).
fof(f2204,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f276,f783])).
fof(f47089,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'const.min.@_1'(s(fun(bool,bool),X3))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X3),s(bool,'T_0')))) ) | $spl284),
  inference(cnf_transformation,[],[f47089_D])).
fof(f47089_D,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'const.min.@_1'(s(fun(bool,bool),X3))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X3),s(bool,'T_0')))) ) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f47088,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'const.min.@_1'(s(fun(bool,bool),X3))) | p(s(bool,X2)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X3),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3)))) )),
  inference(subsumption_resolution,[],[f46451,f276])).
fof(f46451,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'const.min.@_1'(s(fun(bool,bool),X3))) | p(s(bool,X2)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X3),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f276,f4907])).
fof(f46190,plain,(
  ( ! [X325,X324,X322] : (sP1('|dollar|exists_1'(s(fun(bool,bool),X322)),X324) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X322))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X322),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X325) | s(bool,X324) = s(bool,X325)) ) | $spl282),
  inference(cnf_transformation,[],[f46190_D])).
fof(f46190_D,plain,(
  ( ! [X325,X324,X322] : (sP1('|dollar|exists_1'(s(fun(bool,bool),X322)),X324) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X322))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X322),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X325) | s(bool,X324) = s(bool,X325)) ) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f46189,plain,(
  ( ! [X325,X323,X324,X322] : (p(s(bool,X323)) | sP1('|dollar|exists_1'(s(fun(bool,bool),X322)),X324) | s(bool,X324) = s(bool,X325) | s(bool,'T_0') = s(bool,X325) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X322),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X322)))) )),
  inference(subsumption_resolution,[],[f45797,f276])).
fof(f45797,plain,(
  ( ! [X325,X323,X324,X322] : (p(s(bool,X323)) | sP1('|dollar|exists_1'(s(fun(bool,bool),X322)),X324) | s(bool,X324) = s(bool,X325) | s(bool,'T_0') = s(bool,X325) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X322),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X322))) | s(bool,'T_0') = s(bool,X323)) )),
  inference(superposition,[],[f2610,f4862])).
fof(f4862,plain,(
  ( ! [X6,X5] : (s(bool,X6) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X5))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X5),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X5))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1768,f783])).
fof(f1768,plain,(
  ( ! [X0] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) = s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f303,f270])).
fof(f46188,plain,(
  ( ! [X321,X318,X320] : (sP1(X320,'|dollar|exists_1'(s(fun(bool,bool),X318))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X318))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X318),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X321) | s(bool,X320) = s(bool,X321)) ) | $spl280),
  inference(cnf_transformation,[],[f46188_D])).
fof(f46188_D,plain,(
  ( ! [X321,X318,X320] : (sP1(X320,'|dollar|exists_1'(s(fun(bool,bool),X318))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X318))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X318),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X321) | s(bool,X320) = s(bool,X321)) ) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f46187,plain,(
  ( ! [X321,X318,X320,X319] : (p(s(bool,X319)) | sP1(X320,'|dollar|exists_1'(s(fun(bool,bool),X318))) | s(bool,X320) = s(bool,X321) | s(bool,'T_0') = s(bool,X321) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X318),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X318)))) )),
  inference(subsumption_resolution,[],[f45796,f276])).
fof(f45796,plain,(
  ( ! [X321,X318,X320,X319] : (p(s(bool,X319)) | sP1(X320,'|dollar|exists_1'(s(fun(bool,bool),X318))) | s(bool,X320) = s(bool,X321) | s(bool,'T_0') = s(bool,X321) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X318),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X318))) | s(bool,'T_0') = s(bool,X319)) )),
  inference(superposition,[],[f2609,f4862])).
fof(f46184,plain,(
  ( ! [X309] : (sP1('|dollar|exists_1'(s(fun(bool,bool),X309)),'T_0') | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X309))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X309),s(bool,'F_0')))) ) | $spl278),
  inference(cnf_transformation,[],[f46184_D])).
fof(f46184_D,plain,(
  ( ! [X309] : (sP1('|dollar|exists_1'(s(fun(bool,bool),X309)),'T_0') | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X309))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X309),s(bool,'F_0')))) ) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f46183,plain,(
  ( ! [X310,X309] : (p(s(bool,X310)) | sP1('|dollar|exists_1'(s(fun(bool,bool),X309)),'T_0') | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X309),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X309)))) ) | $spl12),
  inference(subsumption_resolution,[],[f45793,f276])).
fof(f45793,plain,(
  ( ! [X310,X309] : (p(s(bool,X310)) | sP1('|dollar|exists_1'(s(fun(bool,bool),X309)),'T_0') | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X309),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X309))) | s(bool,'T_0') = s(bool,X310)) ) | $spl12),
  inference(superposition,[],[f2558,f4862])).
fof(f46182,plain,(
  ( ! [X307] : (sP1('T_0','|dollar|exists_1'(s(fun(bool,bool),X307))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X307))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X307),s(bool,'F_0')))) ) | $spl276),
  inference(cnf_transformation,[],[f46182_D])).
fof(f46182_D,plain,(
  ( ! [X307] : (sP1('T_0','|dollar|exists_1'(s(fun(bool,bool),X307))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X307))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X307),s(bool,'F_0')))) ) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f46181,plain,(
  ( ! [X308,X307] : (p(s(bool,X308)) | sP1('T_0','|dollar|exists_1'(s(fun(bool,bool),X307))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X307),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X307)))) ) | $spl8),
  inference(subsumption_resolution,[],[f45792,f276])).
fof(f45792,plain,(
  ( ! [X308,X307] : (p(s(bool,X308)) | sP1('T_0','|dollar|exists_1'(s(fun(bool,bool),X307))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X307),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X307))) | s(bool,'T_0') = s(bool,X308)) ) | $spl8),
  inference(superposition,[],[f2554,f4862])).
fof(f46180,plain,(
  ( ! [X300,X302,X303] : (s(bool,X302) != s(bool,'|dollar|exists_1'(s(fun(bool,bool),X300))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X300))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X300),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X303) | s(bool,X302) = s(bool,X303)) ) | $spl274),
  inference(cnf_transformation,[],[f46180_D])).
fof(f46180_D,plain,(
  ( ! [X300,X302,X303] : (s(bool,X302) != s(bool,'|dollar|exists_1'(s(fun(bool,bool),X300))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X300))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X300),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X303) | s(bool,X302) = s(bool,X303)) ) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f46179,plain,(
  ( ! [X300,X302,X301,X303] : (p(s(bool,X301)) | s(bool,X302) != s(bool,'|dollar|exists_1'(s(fun(bool,bool),X300))) | s(bool,X302) = s(bool,X303) | s(bool,'T_0') = s(bool,X303) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X300),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X300)))) )),
  inference(subsumption_resolution,[],[f45790,f276])).
fof(f45790,plain,(
  ( ! [X300,X302,X301,X303] : (p(s(bool,X301)) | s(bool,X302) != s(bool,'|dollar|exists_1'(s(fun(bool,bool),X300))) | s(bool,X302) = s(bool,X303) | s(bool,'T_0') = s(bool,X303) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X300),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X300))) | s(bool,'T_0') = s(bool,X301)) )),
  inference(superposition,[],[f2204,f4862])).
fof(f46172,plain,(
  ( ! [X111,X109,X112] : (s(bool,X111) = s(bool,happ(s(fun(bool,bool),X109),s(bool,X112))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X109))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X109),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X111)) ) | $spl272),
  inference(cnf_transformation,[],[f46172_D])).
fof(f46172_D,plain,(
  ( ! [X111,X109,X112] : (s(bool,X111) = s(bool,happ(s(fun(bool,bool),X109),s(bool,X112))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X109))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X109),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X111)) ) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f46171,plain,(
  ( ! [X111,X109,X112,X110] : (p(s(bool,X110)) | s(bool,X111) = s(bool,happ(s(fun(bool,bool),X109),s(bool,X112))) | s(bool,'T_0') = s(bool,X111) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X109),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X109)))) )),
  inference(subsumption_resolution,[],[f45740,f276])).
fof(f45740,plain,(
  ( ! [X111,X109,X112,X110] : (p(s(bool,X110)) | s(bool,X111) = s(bool,happ(s(fun(bool,bool),X109),s(bool,X112))) | s(bool,'T_0') = s(bool,X111) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X109),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X109))) | s(bool,'T_0') = s(bool,X110)) )),
  inference(superposition,[],[f2629,f4862])).
fof(f2629,plain,(
  ( ! [X118,X116,X117,X115] : (p(s(bool,'|dollar|exists_1'(s(fun(X115,bool),X116)))) | s(bool,X118) = s(bool,happ(s(fun(X115,bool),X116),s(X115,X117))) | s(bool,'T_0') = s(bool,X118)) )),
  inference(subsumption_resolution,[],[f2273,f263])).
fof(f2273,plain,(
  ( ! [X118,X116,X117,X115] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X115,bool),X116)))) | s(bool,X118) = s(bool,happ(s(fun(X115,bool),X116),s(X115,X117))) | s(bool,'T_0') = s(bool,X118)) )),
  inference(superposition,[],[f469,f783])).
fof(f46141,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X3),s(bool,'F_0')))) ) | $spl270),
  inference(cnf_transformation,[],[f46141_D])).
fof(f46141_D,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X3),s(bool,'F_0')))) ) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f46140,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) | p(s(bool,X2)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X3),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X3)))) )),
  inference(subsumption_resolution,[],[f45448,f276])).
fof(f45448,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) | p(s(bool,X2)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X3),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f276,f4862])).
fof(f45319,plain,(
  ( ! [X255,X253,X256] : (sP1(happ(s(fun(bool,bool),X253),s(bool,'F_0')),X255) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X253))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X253))) | s(bool,'T_0') = s(bool,X256) | s(bool,X255) = s(bool,X256)) ) | $spl268),
  inference(cnf_transformation,[],[f45319_D])).
fof(f45319_D,plain,(
  ( ! [X255,X253,X256] : (sP1(happ(s(fun(bool,bool),X253),s(bool,'F_0')),X255) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X253))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X253))) | s(bool,'T_0') = s(bool,X256) | s(bool,X255) = s(bool,X256)) ) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f45318,plain,(
  ( ! [X255,X253,X254,X256] : (p(s(bool,X254)) | sP1(happ(s(fun(bool,bool),X253),s(bool,'F_0')),X255) | s(bool,X255) = s(bool,X256) | s(bool,'T_0') = s(bool,X256) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X253))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X253)))) )),
  inference(subsumption_resolution,[],[f45098,f276])).
fof(f45098,plain,(
  ( ! [X255,X253,X254,X256] : (p(s(bool,X254)) | sP1(happ(s(fun(bool,bool),X253),s(bool,'F_0')),X255) | s(bool,X255) = s(bool,X256) | s(bool,'T_0') = s(bool,X256) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X253))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X253))) | s(bool,'T_0') = s(bool,X254)) )),
  inference(superposition,[],[f2610,f4860])).
fof(f4860,plain,(
  ( ! [X2,X1] : (s(bool,X2) = s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X1))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1768,f783])).
fof(f45317,plain,(
  ( ! [X251,X249,X252] : (sP1(X251,happ(s(fun(bool,bool),X249),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X249))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X249))) | s(bool,'T_0') = s(bool,X252) | s(bool,X251) = s(bool,X252)) ) | $spl266),
  inference(cnf_transformation,[],[f45317_D])).
fof(f45317_D,plain,(
  ( ! [X251,X249,X252] : (sP1(X251,happ(s(fun(bool,bool),X249),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X249))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X249))) | s(bool,'T_0') = s(bool,X252) | s(bool,X251) = s(bool,X252)) ) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f45316,plain,(
  ( ! [X251,X249,X252,X250] : (p(s(bool,X250)) | sP1(X251,happ(s(fun(bool,bool),X249),s(bool,'F_0'))) | s(bool,X251) = s(bool,X252) | s(bool,'T_0') = s(bool,X252) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X249))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X249)))) )),
  inference(subsumption_resolution,[],[f45097,f276])).
fof(f45097,plain,(
  ( ! [X251,X249,X252,X250] : (p(s(bool,X250)) | sP1(X251,happ(s(fun(bool,bool),X249),s(bool,'F_0'))) | s(bool,X251) = s(bool,X252) | s(bool,'T_0') = s(bool,X252) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X249))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X249))) | s(bool,'T_0') = s(bool,X250)) )),
  inference(superposition,[],[f2609,f4860])).
fof(f45315,plain,(
  ( ! [X240] : (sP1(happ(s(fun(bool,bool),X240),s(bool,'F_0')),'T_0') | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X240))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X240)))) ) | $spl264),
  inference(cnf_transformation,[],[f45315_D])).
fof(f45315_D,plain,(
  ( ! [X240] : (sP1(happ(s(fun(bool,bool),X240),s(bool,'F_0')),'T_0') | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X240))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X240)))) ) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f45314,plain,(
  ( ! [X241,X240] : (p(s(bool,X241)) | sP1(happ(s(fun(bool,bool),X240),s(bool,'F_0')),'T_0') | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X240))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X240)))) ) | $spl12),
  inference(subsumption_resolution,[],[f45094,f276])).
fof(f45094,plain,(
  ( ! [X241,X240] : (p(s(bool,X241)) | sP1(happ(s(fun(bool,bool),X240),s(bool,'F_0')),'T_0') | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X240))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X240))) | s(bool,'T_0') = s(bool,X241)) ) | $spl12),
  inference(superposition,[],[f2558,f4860])).
fof(f45313,plain,(
  ( ! [X238] : (sP1('T_0',happ(s(fun(bool,bool),X238),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X238))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X238)))) ) | $spl262),
  inference(cnf_transformation,[],[f45313_D])).
fof(f45313_D,plain,(
  ( ! [X238] : (sP1('T_0',happ(s(fun(bool,bool),X238),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X238))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X238)))) ) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f45312,plain,(
  ( ! [X239,X238] : (p(s(bool,X239)) | sP1('T_0',happ(s(fun(bool,bool),X238),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X238))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X238)))) ) | $spl8),
  inference(subsumption_resolution,[],[f45093,f276])).
fof(f45093,plain,(
  ( ! [X239,X238] : (p(s(bool,X239)) | sP1('T_0',happ(s(fun(bool,bool),X238),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X238))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X238))) | s(bool,'T_0') = s(bool,X239)) ) | $spl8),
  inference(superposition,[],[f2554,f4860])).
fof(f45311,plain,(
  ( ! [X231,X233,X234] : (s(bool,X233) != s(bool,happ(s(fun(bool,bool),X231),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X231))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X231))) | s(bool,'T_0') = s(bool,X234) | s(bool,X233) = s(bool,X234)) ) | $spl260),
  inference(cnf_transformation,[],[f45311_D])).
fof(f45311_D,plain,(
  ( ! [X231,X233,X234] : (s(bool,X233) != s(bool,happ(s(fun(bool,bool),X231),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X231))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X231))) | s(bool,'T_0') = s(bool,X234) | s(bool,X233) = s(bool,X234)) ) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f45310,plain,(
  ( ! [X231,X233,X234,X232] : (p(s(bool,X232)) | s(bool,X233) != s(bool,happ(s(fun(bool,bool),X231),s(bool,'F_0'))) | s(bool,X233) = s(bool,X234) | s(bool,'T_0') = s(bool,X234) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X231))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X231)))) )),
  inference(subsumption_resolution,[],[f45091,f276])).
fof(f45091,plain,(
  ( ! [X231,X233,X234,X232] : (p(s(bool,X232)) | s(bool,X233) != s(bool,happ(s(fun(bool,bool),X231),s(bool,'F_0'))) | s(bool,X233) = s(bool,X234) | s(bool,'T_0') = s(bool,X234) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X231))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X231))) | s(bool,'T_0') = s(bool,X232)) )),
  inference(superposition,[],[f2204,f4860])).
fof(f45275,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X3),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X3)))) ) | $spl258),
  inference(cnf_transformation,[],[f45275_D])).
fof(f45275_D,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X3),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X3)))) ) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f45274,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X3),s(bool,'F_0'))) | p(s(bool,X2)) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3)))) )),
  inference(subsumption_resolution,[],[f44773,f276])).
fof(f44773,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X3),s(bool,'F_0'))) | p(s(bool,X2)) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f276,f4860])).
fof(f44646,plain,(
  ( ! [X88,X87,X91,X86] : (s(fun(X86,bool),X87) = s(fun(X86,bool),X88) | s(bool,'T_0') = s(bool,X91) | s(bool,X91) = s(bool,g(s(X86,sK23(X87,X88,bool,X86))))) ) | $spl254),
  inference(cnf_transformation,[],[f44646_D])).
fof(f44646_D,plain,(
  ( ! [X88,X87,X91,X86] : (s(fun(X86,bool),X87) = s(fun(X86,bool),X88) | s(bool,'T_0') = s(bool,X91) | s(bool,X91) = s(bool,g(s(X86,sK23(X87,X88,bool,X86))))) ) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f44647,plain,(
  ( ! [X90,X89] : (s(bool,'T_0') != s(bool,happ(s(fun(X89,bool),X90),s(X89,'const.min.@_1'(s(fun(X89,bool),X90))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X89,bool),X90)))) ) | $spl256),
  inference(cnf_transformation,[],[f44647_D])).
fof(f44647_D,plain,(
  ( ! [X90,X89] : (s(bool,'T_0') != s(bool,happ(s(fun(X89,bool),X90),s(X89,'const.min.@_1'(s(fun(X89,bool),X90))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X89,bool),X90)))) ) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f44645,plain,(
  ( ! [X90,X88,X87,X91,X89,X86] : (s(bool,'T_0') != s(bool,happ(s(fun(X89,bool),X90),s(X89,'const.min.@_1'(s(fun(X89,bool),X90))))) | s(fun(X86,bool),X87) = s(fun(X86,bool),X88) | s(bool,X91) = s(bool,g(s(X86,sK23(X87,X88,bool,X86)))) | s(bool,'T_0') = s(bool,X91) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X89,bool),X90)))) )),
  inference(subsumption_resolution,[],[f44626,f2290])).
fof(f2290,plain,(
  ( ! [X177,X178,X176,X179] : (s(bool,'T_0') != s(bool,f(s(X176,sK23(X177,X178,bool,X176)))) | s(fun(X176,bool),X177) = s(fun(X176,bool),X178) | s(bool,X179) = s(bool,g(s(X176,sK23(X177,X178,bool,X176)))) | s(bool,'T_0') = s(bool,X179)) )),
  inference(superposition,[],[f366,f783])).
fof(f366,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK23(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK23(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK23(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK23(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK23])],[f142])).
fof(f142,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f103])).
fof(f103,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.extreal.sup_lt',aHL_EXT)).
fof(f44626,plain,(
  ( ! [X90,X88,X87,X91,X89,X86] : (s(bool,'T_0') != s(bool,happ(s(fun(X89,bool),X90),s(X89,'const.min.@_1'(s(fun(X89,bool),X90))))) | s(fun(X86,bool),X87) = s(fun(X86,bool),X88) | s(bool,X91) = s(bool,g(s(X86,sK23(X87,X88,bool,X86)))) | s(bool,'T_0') = s(bool,X91) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X89,bool),X90))) | s(bool,'T_0') = s(bool,f(s(X86,sK23(X87,X88,bool,X86))))) )),
  inference(superposition,[],[f2290,f2064])).
fof(f2064,plain,(
  ( ! [X14,X15,X16] : (s(bool,X16) = s(bool,happ(s(fun(X14,bool),X15),s(X14,'const.min.@_1'(s(fun(X14,bool),X15))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X14,bool),X15))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f783,f303])).
fof(f39725,plain,(
  ( ! [X144,X145] : (s(fun(bool,bool),X144) = s(fun(bool,bool),X145) | s(bool,'T_0') = s(bool,g(s(bool,sK23(X144,X145,bool,bool))))) ) | $spl252),
  inference(cnf_transformation,[],[f39725_D])).
fof(f39725_D,plain,(
  ( ! [X144,X145] : (s(fun(bool,bool),X144) = s(fun(bool,bool),X145) | s(bool,'T_0') = s(bool,g(s(bool,sK23(X144,X145,bool,bool))))) ) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f7221,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl65),
  inference(cnf_transformation,[],[f7221_D])).
fof(f7221_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl65),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl65])])).
fof(f21523,plain,(
  ( ! [X144,X145] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X144) = s(fun(bool,bool),X145) | s(bool,'T_0') = s(bool,g(s(bool,sK23(X144,X145,bool,bool))))) ) | $spl52),
  inference(duplicate_literal_removal,[],[f21517])).
fof(f21517,plain,(
  ( ! [X144,X145] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X144) = s(fun(bool,bool),X145) | s(bool,'T_0') = s(bool,g(s(bool,sK23(X144,X145,bool,bool)))) | s(fun(bool,bool),X144) = s(fun(bool,bool),X145)) ) | $spl52),
  inference(superposition,[],[f1976,f7096])).
fof(f1976,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK23(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK23(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f366,f270])).
fof(f21559,plain,(
  ( ! [X146,X147] : (s(fun(bool,bool),X146) = s(fun(bool,bool),X147)) ) | $spl74),
  inference(cnf_transformation,[],[f21559_D])).
fof(f21559_D,plain,(
  ( ! [X146,X147] : (s(fun(bool,bool),X146) = s(fun(bool,bool),X147)) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f39724,plain,(
  s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | $spl251),
  inference(cnf_transformation,[],[f39724_D])).
fof(f39724_D,plain,(
  s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) <=> ~$spl251),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl251])])).
fof(f21522,plain,(
  ( ! [X146,X147] : (s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X146) = s(fun(bool,bool),X147)) ) | $spl52),
  inference(duplicate_literal_removal,[],[f21518])).
fof(f21518,plain,(
  ( ! [X146,X147] : (s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X146) = s(fun(bool,bool),X147) | s(fun(bool,bool),X146) = s(fun(bool,bool),X147)) ) | $spl52),
  inference(superposition,[],[f366,f7096])).
fof(f7219,plain,(
  ( ! [X2,X0,X1] : (s(bool,X2) = s(bool,sK23(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK23(X0,X1,bool,bool)))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2)) ) | $spl62),
  inference(cnf_transformation,[],[f7219_D])).
fof(f7219_D,plain,(
  ( ! [X2,X0,X1] : (s(bool,X2) = s(bool,sK23(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK23(X0,X1,bool,bool)))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f39722,plain,(
  ( ! [X1000,X998] : (s(bool,'F_0') != s(bool,f(s(bool,X998))) | s(bool,'T_0') = s(bool,X1000) | s(bool,X998) = s(bool,X1000)) ) | $spl248),
  inference(cnf_transformation,[],[f39722_D])).
fof(f39722_D,plain,(
  ( ! [X1000,X998] : (s(bool,'F_0') != s(bool,f(s(bool,X998))) | s(bool,'T_0') = s(bool,X1000) | s(bool,X998) = s(bool,X1000)) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f28605,plain,(
  ( ! [X1000,X998,X996,X999,X997] : (s(bool,'F_0') != s(bool,f(s(bool,X998))) | s(fun(bool,bool),X996) = s(fun(bool,bool),X997) | s(bool,'T_0') = s(bool,g(s(bool,sK23(X996,X997,bool,bool)))) | s(bool,X998) = s(bool,X999) | s(bool,'T_0') = s(bool,X999) | s(bool,X1000) = s(bool,sK23(X996,X997,bool,bool)) | s(bool,'T_0') = s(bool,X1000)) )),
  inference(superposition,[],[f1976,f2054])).
fof(f2054,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f783,f783])).
fof(f28801,plain,(
  ( ! [X1000,X998,X996,X999,X997] : (s(bool,'F_0') != s(bool,f(s(bool,X998))) | s(fun(bool,bool),X996) = s(fun(bool,bool),X997) | s(bool,'T_0') = s(bool,g(s(bool,sK23(X996,X997,bool,bool)))) | s(bool,X999) = s(bool,sK23(X996,X997,bool,bool)) | s(bool,'T_0') = s(bool,X999) | s(bool,X998) = s(bool,X1000) | s(bool,'T_0') = s(bool,X1000)) )),
  inference(superposition,[],[f1976,f2054])).
fof(f39721,plain,(
  ( ! [X957,X954,X958,X956,X955] : (s(fun(bool,bool),X954) = s(fun(bool,bool),X955) | s(bool,X956) = s(bool,sK23(X954,X955,bool,bool)) | s(bool,X956) = s(bool,X957) | s(bool,'T_0') = s(bool,X958) | s(bool,X957) = s(bool,X958) | s(bool,'T_0') = s(bool,g(s(bool,sK23(X954,X955,bool,bool))))) ) | $spl246),
  inference(cnf_transformation,[],[f39721_D])).
fof(f39721_D,plain,(
  ( ! [X957,X954,X958,X956,X955] : (s(fun(bool,bool),X954) = s(fun(bool,bool),X955) | s(bool,X956) = s(bool,sK23(X954,X955,bool,bool)) | s(bool,X956) = s(bool,X957) | s(bool,'T_0') = s(bool,X958) | s(bool,X957) = s(bool,X958) | s(bool,'T_0') = s(bool,g(s(bool,sK23(X954,X955,bool,bool))))) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f29441,plain,(
  ( ! [X957,X954,X958,X956,X955] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X954) = s(fun(bool,bool),X955) | s(bool,'T_0') = s(bool,g(s(bool,sK23(X954,X955,bool,bool)))) | s(bool,X956) = s(bool,X957) | s(bool,X957) = s(bool,sK23(X954,X955,bool,bool)) | s(bool,X956) = s(bool,X958) | s(bool,'T_0') = s(bool,X958)) )),
  inference(superposition,[],[f1976,f2054])).
fof(f30081,plain,(
  ( ! [X957,X954,X958,X956,X955] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X954) = s(fun(bool,bool),X955) | s(bool,'T_0') = s(bool,g(s(bool,sK23(X954,X955,bool,bool)))) | s(bool,X956) = s(bool,X957) | s(bool,X957) = s(bool,X958) | s(bool,'T_0') = s(bool,X958) | s(bool,X956) = s(bool,sK23(X954,X955,bool,bool))) )),
  inference(superposition,[],[f1976,f2054])).
fof(f39692,plain,(
  ( ! [X4,X5,X3] : (s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,X5) | s(bool,X5) = s(bool,sK23(X3,X4,bool,bool))) ) | $spl232),
  inference(cnf_transformation,[],[f39692_D])).
fof(f39692_D,plain,(
  ( ! [X4,X5,X3] : (s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,X5) | s(bool,X5) = s(bool,sK23(X3,X4,bool,bool))) ) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f39688,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,X2) = s(bool,sK23(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,X2)) ) | $spl77),
  inference(subsumption_resolution,[],[f39659,f21561])).
fof(f39659,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,X2) = s(bool,sK23(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,g(s(bool,'T_0')))) )),
  inference(superposition,[],[f2296,f270])).
fof(f2296,plain,(
  ( ! [X198,X200,X199,X201] : (s(X200,f(s(bool,'T_0'))) != s(X200,g(s(bool,'T_0'))) | s(fun(bool,X200),X198) = s(fun(bool,X200),X199) | s(bool,X201) = s(bool,sK23(X198,X199,X200,bool)) | s(bool,'T_0') = s(bool,X201)) )),
  inference(superposition,[],[f366,f783])).
fof(f39691,plain,(
  ( ! [X6] : (s(bool,X6) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl230),
  inference(cnf_transformation,[],[f39691_D])).
fof(f39691_D,plain,(
  ( ! [X6] : (s(bool,X6) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f39719,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | $spl243),
  inference(cnf_transformation,[],[f39719_D])).
fof(f39719_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl243),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl243])])).
fof(f39660,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,X5) = s(bool,sK23(X3,X4,bool,bool)) | s(bool,'T_0') = s(bool,X5) | s(bool,X6) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f2296,f783])).
fof(f39717,plain,(
  ( ! [X57] : (s(bool,X57) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X57)) ) | $spl240),
  inference(cnf_transformation,[],[f39717_D])).
fof(f39717_D,plain,(
  ( ! [X57] : (s(bool,X57) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X57)) ) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f39693,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,X7) != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X8) = s(fun(bool,bool),X9) | s(bool,X10) = s(bool,sK23(X8,X9,bool,bool)) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X7)) ) | $spl77),
  inference(subsumption_resolution,[],[f39661,f21561])).
fof(f39661,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,X7) != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X8) = s(fun(bool,bool),X9) | s(bool,X10) = s(bool,sK23(X8,X9,bool,bool)) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,g(s(bool,'T_0')))) )),
  inference(superposition,[],[f2296,f783])).
fof(f39695,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,X11) != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,X14) = s(bool,sK23(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,X11)) ) | $spl77),
  inference(subsumption_resolution,[],[f39662,f21561])).
fof(f39662,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,X11) != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,X14) = s(bool,sK23(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f2296,f783])).
fof(f39663,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X15) = s(fun(bool,bool),X16) | s(bool,X17) = s(bool,sK23(X15,X16,bool,bool)) | s(bool,'T_0') = s(bool,X17) | s(bool,X18) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f2296,f783])).
fof(f39720,plain,(
  ( ! [X27,X32] : (s(bool,X27) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X32) | s(bool,X27) = s(bool,X32)) ) | $spl244),
  inference(cnf_transformation,[],[f39720_D])).
fof(f39720_D,plain,(
  ( ! [X27,X32] : (s(bool,X27) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X32) | s(bool,X27) = s(bool,X32)) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f39664,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (s(bool,X21) != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X22) = s(fun(bool,bool),X23) | s(bool,X24) = s(bool,sK23(X22,X23,bool,bool)) | s(bool,'T_0') = s(bool,X24) | s(bool,X21) = s(bool,X25) | s(bool,'T_0') = s(bool,X25) | s(bool,X26) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f2296,f2054])).
fof(f39665,plain,(
  ( ! [X30,X28,X31,X29,X27,X32] : (s(bool,X27) != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X28) = s(fun(bool,bool),X29) | s(bool,X30) = s(bool,sK23(X28,X29,bool,bool)) | s(bool,'T_0') = s(bool,X30) | s(bool,X31) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X31) | s(bool,X27) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f2296,f2054])).
fof(f39716,plain,(
  ( ! [X61,X62] : (s(bool,X61) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X62) | s(bool,X61) = s(bool,X62)) ) | $spl238),
  inference(cnf_transformation,[],[f39716_D])).
fof(f39716_D,plain,(
  ( ! [X61,X62] : (s(bool,X61) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X62) | s(bool,X61) = s(bool,X62)) ) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f39666,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (s(bool,X33) != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X34) = s(fun(bool,bool),X35) | s(bool,X36) = s(bool,sK23(X34,X35,bool,bool)) | s(bool,'T_0') = s(bool,X36) | s(bool,X37) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X33) | s(bool,X37) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f2296,f2054])).
fof(f39715,plain,(
  ( ! [X68,X67,X63] : (s(bool,X63) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X68) | s(bool,X67) = s(bool,X68) | s(bool,X63) = s(bool,X67)) ) | $spl236),
  inference(cnf_transformation,[],[f39715_D])).
fof(f39715_D,plain,(
  ( ! [X68,X67,X63] : (s(bool,X63) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X68) | s(bool,X67) = s(bool,X68) | s(bool,X63) = s(bool,X67)) ) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f39707,plain,(
  ( ! [X39,X43,X41,X44,X42,X40] : (s(bool,X39) != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X40) = s(fun(bool,bool),X41) | s(bool,X42) = s(bool,sK23(X40,X41,bool,bool)) | s(bool,'T_0') = s(bool,X42) | s(bool,X39) = s(bool,X43) | s(bool,X43) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) ) | $spl77),
  inference(subsumption_resolution,[],[f39667,f21561])).
fof(f39667,plain,(
  ( ! [X39,X43,X41,X44,X42,X40] : (s(bool,X39) != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X40) = s(fun(bool,bool),X41) | s(bool,X42) = s(bool,sK23(X40,X41,bool,bool)) | s(bool,'T_0') = s(bool,X42) | s(bool,X39) = s(bool,X43) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,X43) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f2296,f2054])).
fof(f39710,plain,(
  ( ! [X50,X48,X49] : (s(bool,X48) = s(bool,X49) | s(bool,'T_0') = s(bool,X50) | s(bool,X48) = s(bool,X50) | s(bool,X49) = s(bool,g(s(bool,'T_0')))) ) | $spl234),
  inference(cnf_transformation,[],[f39710_D])).
fof(f39710_D,plain,(
  ( ! [X50,X48,X49] : (s(bool,X48) = s(bool,X49) | s(bool,'T_0') = s(bool,X50) | s(bool,X48) = s(bool,X50) | s(bool,X49) = s(bool,g(s(bool,'T_0')))) ) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f39668,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X45) = s(fun(bool,bool),X46) | s(bool,X47) = s(bool,sK23(X45,X46,bool,bool)) | s(bool,'T_0') = s(bool,X47) | s(bool,X48) = s(bool,X49) | s(bool,X49) = s(bool,g(s(bool,'T_0'))) | s(bool,X48) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f2296,f2054])).
fof(f39669,plain,(
  ( ! [X54,X52,X56,X55,X53,X51] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X51) = s(fun(bool,bool),X52) | s(bool,X53) = s(bool,sK23(X51,X52,bool,bool)) | s(bool,'T_0') = s(bool,X53) | s(bool,X54) = s(bool,X55) | s(bool,X55) = s(bool,X56) | s(bool,'T_0') = s(bool,X56) | s(bool,X54) = s(bool,g(s(bool,'T_0')))) )),
  inference(superposition,[],[f2296,f2054])).
fof(f39670,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (s(bool,X57) != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X58) = s(fun(bool,bool),X59) | s(bool,X60) = s(bool,sK23(X58,X59,bool,bool)) | s(bool,'T_0') = s(bool,X60) | s(bool,X61) = s(bool,g(s(bool,'T_0'))) | s(bool,X61) = s(bool,X62) | s(bool,'T_0') = s(bool,X62) | s(bool,'T_0') = s(bool,X57)) )),
  inference(superposition,[],[f2296,f2054])).
fof(f39713,plain,(
  ( ! [X68,X66,X64,X67,X65,X63] : (s(bool,X63) != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X64) = s(fun(bool,bool),X65) | s(bool,X66) = s(bool,sK23(X64,X65,bool,bool)) | s(bool,'T_0') = s(bool,X66) | s(bool,X63) = s(bool,X67) | s(bool,X67) = s(bool,X68) | s(bool,'T_0') = s(bool,X68)) ) | $spl77),
  inference(subsumption_resolution,[],[f39671,f21561])).
fof(f39671,plain,(
  ( ! [X68,X66,X64,X67,X65,X63] : (s(bool,X63) != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X64) = s(fun(bool,bool),X65) | s(bool,X66) = s(bool,sK23(X64,X65,bool,bool)) | s(bool,'T_0') = s(bool,X66) | s(bool,X63) = s(bool,X67) | s(bool,X67) = s(bool,X68) | s(bool,'T_0') = s(bool,X68) | s(bool,'T_0') = s(bool,g(s(bool,'T_0')))) )),
  inference(superposition,[],[f2296,f2054])).
fof(f39711,plain,(
  ( ! [X54,X52,X56,X55,X53,X51] : (s(fun(bool,bool),X51) = s(fun(bool,bool),X52) | s(bool,X53) = s(bool,sK23(X51,X52,bool,bool)) | s(bool,'T_0') = s(bool,X53) | s(bool,X54) = s(bool,X55) | s(bool,X55) = s(bool,X56) | s(bool,'T_0') = s(bool,X56) | s(bool,X54) = s(bool,g(s(bool,'T_0')))) ) | $spl65),
  inference(subsumption_resolution,[],[f39669,f8380])).
fof(f8380,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | $spl65),
  inference(unit_resulting_resolution,[],[f468,f7221,f846])).
fof(f39709,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (s(fun(bool,bool),X45) = s(fun(bool,bool),X46) | s(bool,X47) = s(bool,sK23(X45,X46,bool,bool)) | s(bool,'T_0') = s(bool,X47) | s(bool,X48) = s(bool,X49) | s(bool,X49) = s(bool,g(s(bool,'T_0'))) | s(bool,X48) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) ) | $spl65),
  inference(subsumption_resolution,[],[f39668,f8380])).
fof(f39705,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X4) = s(bool,sK23(X5,X6,bool,bool)) | s(fun(bool,bool),X5) = s(fun(bool,bool),X6) | s(bool,'T_0') = s(bool,X4) | s(bool,X7) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X7)) ) | $spl65),
  inference(duplicate_literal_removal,[],[f39704])).
fof(f39704,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X4) = s(bool,sK23(X5,X6,bool,bool)) | s(fun(bool,bool),X5) = s(fun(bool,bool),X6) | s(bool,'T_0') = s(bool,X4) | s(bool,X7) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,X4)) ) | $spl65),
  inference(condensation,[],[f39703])).
fof(f39703,plain,(
  ( ! [X30,X28,X31,X29,X27,X32] : (s(fun(bool,bool),X28) = s(fun(bool,bool),X29) | s(bool,X30) = s(bool,sK23(X28,X29,bool,bool)) | s(bool,'T_0') = s(bool,X30) | s(bool,X31) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X31) | s(bool,X27) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) ) | $spl65),
  inference(subsumption_resolution,[],[f39702,f783])).
fof(f39702,plain,(
  ( ! [X30,X28,X31,X29,X27,X32] : (s(bool,'T_0') != s(bool,X27) | s(fun(bool,bool),X28) = s(fun(bool,bool),X29) | s(bool,X30) = s(bool,sK23(X28,X29,bool,bool)) | s(bool,'T_0') = s(bool,X30) | s(bool,X31) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X31) | s(bool,X27) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) ) | $spl65),
  inference(forward_demodulation,[],[f39665,f8380])).
fof(f39701,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,X3) = s(bool,sK23(X4,X5,bool,bool)) | s(fun(bool,bool),X4) = s(fun(bool,bool),X5) | s(bool,'T_0') = s(bool,X3) | s(bool,X6) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl65),
  inference(duplicate_literal_removal,[],[f39700])).
fof(f39700,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,X3) = s(bool,sK23(X4,X5,bool,bool)) | s(fun(bool,bool),X4) = s(fun(bool,bool),X5) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,X6) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl65),
  inference(condensation,[],[f39699])).
fof(f39699,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (s(fun(bool,bool),X22) = s(fun(bool,bool),X23) | s(bool,X24) = s(bool,sK23(X22,X23,bool,bool)) | s(bool,'T_0') = s(bool,X24) | s(bool,X21) = s(bool,X25) | s(bool,'T_0') = s(bool,X25) | s(bool,X26) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X26)) ) | $spl65),
  inference(subsumption_resolution,[],[f39698,f783])).
fof(f39698,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (s(bool,'T_0') != s(bool,X21) | s(fun(bool,bool),X22) = s(fun(bool,bool),X23) | s(bool,X24) = s(bool,sK23(X22,X23,bool,bool)) | s(bool,'T_0') = s(bool,X24) | s(bool,X21) = s(bool,X25) | s(bool,'T_0') = s(bool,X25) | s(bool,X26) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X26)) ) | $spl65),
  inference(forward_demodulation,[],[f39664,f8380])).
fof(f39697,plain,(
  ( ! [X17,X15,X18,X16] : (s(fun(bool,bool),X15) = s(fun(bool,bool),X16) | s(bool,X17) = s(bool,sK23(X15,X16,bool,bool)) | s(bool,'T_0') = s(bool,X17) | s(bool,X18) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X18)) ) | $spl65),
  inference(subsumption_resolution,[],[f39663,f8380])).
fof(f39690,plain,(
  ( ! [X6,X4,X5,X3] : (s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,X5) = s(bool,sK23(X3,X4,bool,bool)) | s(bool,'T_0') = s(bool,X5) | s(bool,X6) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl65),
  inference(subsumption_resolution,[],[f39660,f8380])).
fof(f37616,plain,(
  ( ! [X30,X33,X34,X32] : (s(bool,X32) = s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),X30))) | s(bool,'T_0') = s(bool,X34) | s(bool,X33) = s(bool,X34) | s(bool,X33) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.min.@_1'(s(fun('type.extreal.extreal',bool),X30))),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X30))))) | s(bool,'T_0') = s(bool,X32)) ) | $spl228),
  inference(cnf_transformation,[],[f37616_D])).
fof(f37616_D,plain,(
  ( ! [X30,X33,X34,X32] : (s(bool,X32) = s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),X30))) | s(bool,'T_0') = s(bool,X34) | s(bool,X33) = s(bool,X34) | s(bool,X33) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.min.@_1'(s(fun('type.extreal.extreal',bool),X30))),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X30))))) | s(bool,'T_0') = s(bool,X32)) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f37620,plain,(
  ( ! [X54,X52,X50,X53,X51] : (p(s(bool,X51)) | s(bool,X52) = s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),X50))) | s(bool,'T_0') = s(bool,X52) | s(bool,X53) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.min.@_1'(s(fun('type.extreal.extreal',bool),X50))),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X50))))) | s(bool,X53) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f37606,f276])).
fof(f37606,plain,(
  ( ! [X54,X52,X50,X53,X51] : (p(s(bool,X51)) | s(bool,X52) = s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),X50))) | s(bool,'T_0') = s(bool,X52) | s(bool,X53) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.min.@_1'(s(fun('type.extreal.extreal',bool),X50))),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X50))))) | s(bool,X53) = s(bool,X54) | s(bool,'T_0') = s(bool,X54) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f4747,f2054])).
fof(f4747,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.min.@_1'(s(fun('type.extreal.extreal',bool),X0))),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X0)))))) | s(bool,X1) = s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),X0))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(subsumption_resolution,[],[f4741,f263])).
fof(f4741,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.min.@_1'(s(fun('type.extreal.extreal',bool),X0))),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X0)))))) | s(bool,X1) = s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),X0))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f1840,f783])).
fof(f1840,plain,(
  ( ! [X0] : (~p(s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),X0)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.min.@_1'(s(fun('type.extreal.extreal',bool),X0))),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X0))))))) )),
  inference(superposition,[],[f317,f303])).
fof(f317,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X1)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X0))))))) )),
  inference(cnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0,X1] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X1)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X0)))))))),
  inference(ennf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X1)))) => p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X0)))))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X17,X8] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),X17),s('type.extreal.extreal',X8)))) => p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X8),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X17)))))))),
  file('thm.extreal.sup_lt','thm.extreal.le_sup_imp')).
fof(f37615,plain,(
  ( ! [X30,X33,X31,X34,X32] : (p(s(bool,X31)) | s(bool,X32) = s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),X30))) | s(bool,'T_0') = s(bool,X32) | s(bool,X33) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.min.@_1'(s(fun('type.extreal.extreal',bool),X30))),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X30))))) | s(bool,X33) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f37602,f276])).
fof(f37602,plain,(
  ( ! [X30,X33,X31,X34,X32] : (p(s(bool,X31)) | s(bool,X32) = s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),X30))) | s(bool,'T_0') = s(bool,X32) | s(bool,X33) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.min.@_1'(s(fun('type.extreal.extreal',bool),X30))),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X30))))) | s(bool,'T_0') = s(bool,X31) | s(bool,X33) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f4747,f2054])).
fof(f37149,plain,(
  ( ! [X178,X176,X179] : (sP1('const.min.@_1'(s(fun(bool,bool),X176)),X178) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X176))) = s(bool,happ(s(fun(bool,bool),X176),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X179) | s(bool,X178) = s(bool,X179)) ) | $spl226),
  inference(cnf_transformation,[],[f37149_D])).
fof(f37149_D,plain,(
  ( ! [X178,X176,X179] : (sP1('const.min.@_1'(s(fun(bool,bool),X176)),X178) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X176))) = s(bool,happ(s(fun(bool,bool),X176),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X179) | s(bool,X178) = s(bool,X179)) ) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f37148,plain,(
  ( ! [X177,X178,X176,X179] : (p(s(bool,X177)) | sP1('const.min.@_1'(s(fun(bool,bool),X176)),X178) | s(bool,X178) = s(bool,X179) | s(bool,'T_0') = s(bool,X179) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X176))) = s(bool,happ(s(fun(bool,bool),X176),s(bool,'T_0')))) )),
  inference(subsumption_resolution,[],[f36859,f276])).
fof(f36859,plain,(
  ( ! [X177,X178,X176,X179] : (p(s(bool,X177)) | sP1('const.min.@_1'(s(fun(bool,bool),X176)),X178) | s(bool,X178) = s(bool,X179) | s(bool,'T_0') = s(bool,X179) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X176))) = s(bool,happ(s(fun(bool,bool),X176),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X177)) )),
  inference(superposition,[],[f2610,f2291])).
fof(f2291,plain,(
  ( ! [X180,X181] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X180))) = s(bool,happ(s(fun(bool,bool),X180),s(bool,'T_0'))) | s(bool,X181) = s(bool,'const.min.@_1'(s(fun(bool,bool),X180))) | s(bool,'T_0') = s(bool,X181)) )),
  inference(superposition,[],[f303,f783])).
fof(f37147,plain,(
  ( ! [X175,X174,X172] : (sP1(X174,'const.min.@_1'(s(fun(bool,bool),X172))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X172))) = s(bool,happ(s(fun(bool,bool),X172),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X175) | s(bool,X174) = s(bool,X175)) ) | $spl224),
  inference(cnf_transformation,[],[f37147_D])).
fof(f37147_D,plain,(
  ( ! [X175,X174,X172] : (sP1(X174,'const.min.@_1'(s(fun(bool,bool),X172))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X172))) = s(bool,happ(s(fun(bool,bool),X172),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X175) | s(bool,X174) = s(bool,X175)) ) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f37146,plain,(
  ( ! [X175,X173,X174,X172] : (p(s(bool,X173)) | sP1(X174,'const.min.@_1'(s(fun(bool,bool),X172))) | s(bool,X174) = s(bool,X175) | s(bool,'T_0') = s(bool,X175) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X172))) = s(bool,happ(s(fun(bool,bool),X172),s(bool,'T_0')))) )),
  inference(subsumption_resolution,[],[f36858,f276])).
fof(f36858,plain,(
  ( ! [X175,X173,X174,X172] : (p(s(bool,X173)) | sP1(X174,'const.min.@_1'(s(fun(bool,bool),X172))) | s(bool,X174) = s(bool,X175) | s(bool,'T_0') = s(bool,X175) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X172))) = s(bool,happ(s(fun(bool,bool),X172),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X173)) )),
  inference(superposition,[],[f2609,f2291])).
fof(f37145,plain,(
  ( ! [X163] : (sP1('const.min.@_1'(s(fun(bool,bool),X163)),'T_0') | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X163))) = s(bool,happ(s(fun(bool,bool),X163),s(bool,'T_0')))) ) | $spl222),
  inference(cnf_transformation,[],[f37145_D])).
fof(f37145_D,plain,(
  ( ! [X163] : (sP1('const.min.@_1'(s(fun(bool,bool),X163)),'T_0') | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X163))) = s(bool,happ(s(fun(bool,bool),X163),s(bool,'T_0')))) ) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f37144,plain,(
  ( ! [X163,X164] : (p(s(bool,X164)) | sP1('const.min.@_1'(s(fun(bool,bool),X163)),'T_0') | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X163))) = s(bool,happ(s(fun(bool,bool),X163),s(bool,'T_0')))) ) | $spl12),
  inference(subsumption_resolution,[],[f36855,f276])).
fof(f36855,plain,(
  ( ! [X163,X164] : (p(s(bool,X164)) | sP1('const.min.@_1'(s(fun(bool,bool),X163)),'T_0') | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X163))) = s(bool,happ(s(fun(bool,bool),X163),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X164)) ) | $spl12),
  inference(superposition,[],[f2558,f2291])).
fof(f37143,plain,(
  ( ! [X161] : (sP1('T_0','const.min.@_1'(s(fun(bool,bool),X161))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X161))) = s(bool,happ(s(fun(bool,bool),X161),s(bool,'T_0')))) ) | $spl220),
  inference(cnf_transformation,[],[f37143_D])).
fof(f37143_D,plain,(
  ( ! [X161] : (sP1('T_0','const.min.@_1'(s(fun(bool,bool),X161))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X161))) = s(bool,happ(s(fun(bool,bool),X161),s(bool,'T_0')))) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f37142,plain,(
  ( ! [X161,X162] : (p(s(bool,X162)) | sP1('T_0','const.min.@_1'(s(fun(bool,bool),X161))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X161))) = s(bool,happ(s(fun(bool,bool),X161),s(bool,'T_0')))) ) | $spl8),
  inference(subsumption_resolution,[],[f36854,f276])).
fof(f36854,plain,(
  ( ! [X161,X162] : (p(s(bool,X162)) | sP1('T_0','const.min.@_1'(s(fun(bool,bool),X161))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X161))) = s(bool,happ(s(fun(bool,bool),X161),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X162)) ) | $spl8),
  inference(superposition,[],[f2554,f2291])).
fof(f37141,plain,(
  ( ! [X159,X157,X160] : (s(bool,X159) != s(bool,'const.min.@_1'(s(fun(bool,bool),X157))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X157))) = s(bool,happ(s(fun(bool,bool),X157),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X160) | s(bool,X159) = s(bool,X160)) ) | $spl218),
  inference(cnf_transformation,[],[f37141_D])).
fof(f37141_D,plain,(
  ( ! [X159,X157,X160] : (s(bool,X159) != s(bool,'const.min.@_1'(s(fun(bool,bool),X157))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X157))) = s(bool,happ(s(fun(bool,bool),X157),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X160) | s(bool,X159) = s(bool,X160)) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f37140,plain,(
  ( ! [X158,X159,X157,X160] : (p(s(bool,X158)) | s(bool,X159) != s(bool,'const.min.@_1'(s(fun(bool,bool),X157))) | s(bool,X159) = s(bool,X160) | s(bool,'T_0') = s(bool,X160) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X157))) = s(bool,happ(s(fun(bool,bool),X157),s(bool,'T_0')))) )),
  inference(subsumption_resolution,[],[f36853,f276])).
fof(f36853,plain,(
  ( ! [X158,X159,X157,X160] : (p(s(bool,X158)) | s(bool,X159) != s(bool,'const.min.@_1'(s(fun(bool,bool),X157))) | s(bool,X159) = s(bool,X160) | s(bool,'T_0') = s(bool,X160) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X157))) = s(bool,happ(s(fun(bool,bool),X157),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X158)) )),
  inference(superposition,[],[f2204,f2291])).
fof(f37130,plain,(
  ( ! [X518,X517] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X518))),s('type.extreal.extreal',X517)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK22(X517,X518)),s('type.extreal.extreal',X517)))) ) | $spl216),
  inference(cnf_transformation,[],[f37130_D])).
fof(f37130_D,plain,(
  ( ! [X518,X517] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X518))),s('type.extreal.extreal',X517)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK22(X517,X518)),s('type.extreal.extreal',X517)))) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f37059,plain,(
  ( ! [X5] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X5)))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X5))) = s(bool,happ(s(fun(bool,bool),X5),s(bool,'T_0')))) ) | $spl132),
  inference(cnf_transformation,[],[f37059_D])).
fof(f37059_D,plain,(
  ( ! [X5] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X5)))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X5))) = s(bool,happ(s(fun(bool,bool),X5),s(bool,'T_0')))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f36745,plain,(
  ( ! [X518,X517,X519] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X519)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X518))),s('type.extreal.extreal',X517)))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X519))) = s(bool,happ(s(fun(bool,bool),X519),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK22(X517,X518)),s('type.extreal.extreal',X517)))) )),
  inference(superposition,[],[f328,f2291])).
fof(f328,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK22(X1,X0)),s('type.extreal.extreal',X1)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X0))),s('type.extreal.extreal',X1))))) )),
  inference(cnf_transformation,[],[f202])).
fof(f37129,plain,(
  ( ! [X504,X506,X503] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X503),s('type.extreal.extreal',X506)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X503))),s('type.extreal.extreal',X504))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X506),s('type.extreal.extreal',X504))))) ) | $spl214),
  inference(cnf_transformation,[],[f37129_D])).
fof(f37129_D,plain,(
  ( ! [X504,X506,X503] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X503),s('type.extreal.extreal',X506)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X503))),s('type.extreal.extreal',X504))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X506),s('type.extreal.extreal',X504))))) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f36738,plain,(
  ( ! [X505,X504,X506,X503] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X505)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X503),s('type.extreal.extreal',X506)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X506),s('type.extreal.extreal',X504)))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X505))) = s(bool,happ(s(fun(bool,bool),X505),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X503))),s('type.extreal.extreal',X504)))) )),
  inference(superposition,[],[f326,f2291])).
fof(f37127,plain,(
  ( ! [X486,X487] : (s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X486),s('type.extreal.extreal',X487)))) ) | $spl212),
  inference(cnf_transformation,[],[f37127_D])).
fof(f37127_D,plain,(
  ( ! [X486,X487] : (s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X486),s('type.extreal.extreal',X487)))) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f37126,plain,(
  ( ! [X486,X488,X487] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X488)))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X488))) = s(bool,happ(s(fun(bool,bool),X488),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X486),s('type.extreal.extreal',X487)))) )),
  inference(subsumption_resolution,[],[f36733,f1296])).
fof(f36733,plain,(
  ( ! [X486,X488,X487] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X488)))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X487),s('type.extreal.extreal',X486)))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X488))) = s(bool,happ(s(fun(bool,bool),X488),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X486),s('type.extreal.extreal',X487)))) )),
  inference(superposition,[],[f324,f2291])).
fof(f37125,plain,(
  ( ! [X480,X482,X479] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X482),s('type.extreal.extreal',X480)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X479),s('type.extreal.extreal',X480))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X482),s('type.extreal.extreal',X479))))) ) | $spl210),
  inference(cnf_transformation,[],[f37125_D])).
fof(f37125_D,plain,(
  ( ! [X480,X482,X479] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X482),s('type.extreal.extreal',X480)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X479),s('type.extreal.extreal',X480))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X482),s('type.extreal.extreal',X479))))) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f36731,plain,(
  ( ! [X480,X482,X479,X481] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X481)))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X482),s('type.extreal.extreal',X480)))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X482),s('type.extreal.extreal',X479)))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X481))) = s(bool,happ(s(fun(bool,bool),X481),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X479),s('type.extreal.extreal',X480)))) )),
  inference(superposition,[],[f367,f2291])).
fof(f367,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2)))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2)))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))) )),
  inference(cnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2)))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2)))))),
  inference(flattening,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2)))))),
  inference(ennf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) & p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))))) => p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2)))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X8,X7,X16] : ((p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X8),s('type.extreal.extreal',X7)))) & p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X16))))) => p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X8),s('type.extreal.extreal',X16)))))),
  file('thm.extreal.sup_lt','thm.extreal.lte_trans')).
fof(f37121,plain,(
  ( ! [X470] : (s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X470))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X470)))) ) | $spl208),
  inference(cnf_transformation,[],[f37121_D])).
fof(f37121_D,plain,(
  ( ! [X470] : (s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X470))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X470)))) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f37122,plain,(
  ( ! [X472,X473] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X473))) = s(bool,happ(s(fun(bool,bool),X473),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X472))) | ~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X473)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X472)))) )),
  inference(global_subsumption,[],[f262,f261,f260,f263,f270,f277,f276,f321,f318,f335,f334,f470,f353,f350,f349,f348,f366,f373,f372,f371,f455,f369,f368,f384,f383,f382,f381,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f396,f395,f394,f393,f400,f399,f398,f397,f454,f404,f403,f402,f401,f411,f410,f409,f408,f417,f414,f413,f412,f424,f423,f422,f453,f420,f419,f418,f428,f427,f426,f425,f432,f431,f430,f429,f441,f440,f439,f452,f437,f436,f435,f445,f444,f443,f442,f449,f448,f447,f446,f471,f474,f477,f480,f482,f484,f487,f489,f491,f494,f497,f500,f503,f505,f507,f510,f513,f515,f517,f520,f522,f525,f565,f575,f578,f586,f602,f613,f622,f639,f644,f645,f647,f654,f658,f659,f662,f674,f688,f690,f698,f710,f783,f846,f878,f876,f904,f901,f947,f972,f987,f995,f1008,f1012,f1026,f1041,f1111,f1126,f1142,f1201,f1218,f1233,f1313,f1361,f1767,f1766,f1765,f1764,f2054,f2140,f2204,f2553,f2557,f2561,f2564,f2567,f2570,f2573,f2576,f2579,f2582,f2585,f2588,f2591,f2594,f2597,f2600,f2603,f2605,f2606,f2607,f2608,f2609,f2610,f2611,f2612,f2613,f2614,f2615,f2616,f2617,f2618,f2619,f2620,f2621,f2622,f2623,f2624,f2625,f2296,f2643,f2645,f2648,f2650,f2652,f2654,f2656,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2636,f328,f1840,f2064,f2039,f1797,f1796,f1795,f1794,f1793,f1792,f1791,f1790,f1789,f1788,f1787,f1786,f1785,f1784,f1783,f1782,f1781,f1780,f1779,f2635,f2634,f1296,f367,f325,f324,f2291,f2295,f2294,f2293,f2292,f2139,f2138,f2137,f2136,f2135,f327,f1866,f1867,f317,f326,f2626,f2637,f1831,f1817,f303,f2633,f2632,f2631,f2630,f2629,f2628,f2627,f469,f316,f314,f312,f310,f305,f304,f2134,f2290,f1976,f1975,f1959,f1947,f1923,f1898,f1769,f1768,f468,f281,f280,f3740,f3742,f3744,f3757,f3760,f3762,f3764,f3849,f3851,f3853,f3862,f3865,f3867,f3869,f3882,f3885,f3886,f3897,f4303,f4308,f4310,f4312,f4317,f4319,f4321,f4722,f4724,f4726,f4731,f4733,f4735,f4747,f4752,f4754,f4756,f4807,f4809,f4811,f4821,f4823,f4825,f4838,f4840,f4842,f4860,f4862,f4864,f4869,f4870,f4871,f4872,f4873,f4875,f4880,f4881,f4882,f4883,f4884,f4885,f4888,f4889,f4892,f4893,f4894,f4907,f4909,f4965,f4966,f4967,f4969,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f5004,f5006,f5007,f6917,f6931,f6948,f6967,f6996,f7007,f7009,f7045,f7069,f7071,f7090,f7091,f7092,f7093,f7094,f7124,f7130,f7147,f7153,f7173,f7179,f7193,f7199,f7208,f7209,f7223,f7240,f7242,f7244,f7261,f7263,f7265,f7288,f9309,f9324,f9325,f9417,f9398,f9399,f9484,f9465,f9466,f9489,f9490,f18495,f18500,f18505,f18466,f18467,f19646,f19651,f19656,f19617,f19618,f19906,f19911,f19916,f19878,f19879,f21428,f21429,f21430,f21457,f21458,f21459,f21931,f22080,f22041,f23768,f23773,f23780,f23727,f23728,f24311,f24316,f24323,f24270,f24271,f24460,f24461,f24462,f24450,f24463,f24464,f24465,f24466,f24467,f24468,f24469,f24470,f24471,f24472,f24473,f24474,f24475,f24476,f24477,f24478,f24479,f24480,f24481,f24482,f24483,f24484,f24485,f24486,f24487,f24488,f24489,f24506,f24507,f24508,f24523,f24538,f24539,f24540,f24555,f24596,f24613,f24585,f24621,f24648,f24649,f24650,f24672,f24673,f24674,f25900,f25921,f25973,f25992,f25993,f25994,f26033,f26034,f26035,f26075,f26076,f26077,f26088,f26089,f26090,f26105,f27929,f27930,f27931,f27946,f27968,f27969,f27970,f27991,f27992,f27993,f28005,f28006,f28007,f28031,f28032,f28033,f31391,f28156,f28157,f28226,f28227,f31209,f28261,f28262,f28272,f28300,f28301,f28557,f28559,f28560,f28561,f28562,f28563,f28564,f28565,f28566,f28567,f28568,f28569,f28570,f28571,f28572,f28573,f28574,f28575,f28576,f28577,f28578,f28579,f28580,f28581,f28582,f28583,f28584,f28585,f28586,f28587,f28596,f28597,f28598,f28599,f28600,f28601,f28602,f28603,f28604,f31756,f31760,f31766,f31768,f31819,f31820,f31823,f31824,f31826,f31828,f31830,f31832,f31834,f31837,f31838,f31841,f31842,f28949,f28951,f28952,f28953,f28954,f28955,f28956,f28957,f28958,f28959,f28960,f28961,f28962,f28963,f28964,f28965,f28966,f28967,f28968,f28969,f28970,f28971,f28972,f28973,f28974,f28975,f28976,f28977,f28978,f28979,f28988,f28989,f28990,f28991,f28992,f28993,f28994,f28995,f28996,f31844,f31859,f31860,f31861,f31862,f31863,f31864,f31865,f31866,f31867,f31868,f31869,f31870,f31871,f31872,f29145,f29152,f29155,f29156,f29157,f29158,f29159,f29160,f29161,f29162,f29163,f29164,f29165,f29166,f29167,f29168,f29169,f29170,f29171,f29172,f29173,f29174,f29175,f29184,f29185,f29186,f29187,f29195,f31922,f31923,f31924,f31925,f31926,f31927,f31928,f31929,f31930,f31931,f31932,f31933,f31934,f31935,f31936,f31937,f31938,f31939,f31940,f31941,f31942,f31943,f31944,f31945,f31946,f31947,f31948,f31952,f31953,f31954,f31955,f31956,f31957,f31958,f31959,f31960,f31961,f31962,f31963,f31964,f31965,f31966,f31967,f31968,f31969,f31970,f31971,f31972,f31973,f31974,f31975,f31976,f31977,f31978,f29393,f29395,f29396,f29397,f29398,f29399,f29400,f29401,f29402,f29403,f29404,f29405,f29406,f29407,f29408,f29409,f29410,f29411,f29412,f29413,f29414,f29415,f29416,f29417,f29418,f29419,f29420,f29421,f29422,f29423,f29432,f29433,f29434,f29435,f29436,f29437,f29438,f29439,f29440,f31989,f31990,f32041,f32044,f32045,f32046,f32047,f32048,f32049,f32052,f32055,f32125,f30120,f32469,f32470,f32471,f32472,f32473,f32474,f32475,f32476,f32477,f32478,f32479,f32480,f32481,f32482,f32483,f32484,f32485,f32486,f32487,f32488,f32489,f32490,f32491,f32492,f32493,f32494,f32495,f32496,f32497,f32498,f32499,f32500,f32501,f32502,f32503,f32504,f32505,f32506,f32507,f32508,f32509,f32510,f32511,f32512,f32513,f32514,f32515,f32516,f32517,f32518,f32519,f32520,f32521,f32522,f32523,f32524,f32525,f32526,f32527,f32528,f32529,f32691,f32692,f32693,f32696,f32698,f32699,f32700,f32732,f32735,f32740,f32824,f32825,f32826,f32829,f32831,f32832,f32833,f32865,f32868,f32873,f32932,f33004,f33066,f33132,f33183,f33335,f33469,f33623,f33852,f33850,f33916,f33923,f33951,f33954,f33959,f33989,f33992,f33997,f34045,f34048,f34053,f34210,f34211,f34212,f34215,f34217,f34218,f34219,f34272,f34273,f34274,f34277,f34279,f34280,f34281,f34411,f34412,f34413,f34416,f34418,f34419,f34420,f34493,f34494,f34495,f34498,f34500,f34501,f34502,f34520,f34521,f34522,f34525,f34527,f34528,f34529,f34549,f34550,f34551,f34554,f34556,f34557,f34558,f34578,f34579,f34580,f34583,f34585,f34586,f34587,f34664,f34665,f34666,f34669,f34671,f34672,f34673,f35269,f34731,f34762,f34763,f34766,f34767,f34769,f34770,f34771,f34924,f34925,f34930,f34945,f34980,f35440,f35443,f35444,f35445,f35446,f35447,f35448,f35449,f35450,f35451,f35452,f35453,f35454,f35455,f35456,f35457,f35458,f35459,f35460,f36038,f36032,f36018,f35621,f35789,f36202,f36204,f35923,f36266,f36267,f36287,f36290,f36291,f36292,f36296,f36321,f36322,f36323,f36324,f36325,f36327,f36332,f36333,f36334,f36335,f36336,f36337,f36341,f36345,f36346,f36362,f36365,f36366,f36367,f36368,f36369,f36370,f36371,f36372,f36373,f36374,f36375,f36376,f36377,f36378,f36445,f36466,f36469,f36470,f36471,f36475,f36476,f37031,f37026,f37012,f37010,f36505,f36506,f36542,f36543,f36544,f36545,f36546,f36548,f36553,f36554,f36555,f36556,f36557,f36558,f36562,f36566,f36567,f36583,f36584,f36587,f36588,f36589,f36590,f36591,f36592,f36593,f36594,f36595,f36596,f36597,f36598,f36599,f36600,f37057,f36989,f37060,f37062,f36618,f37063,f36619,f37064,f36620,f36621,f36622,f37067,f36623,f36624,f36625,f36626,f36627,f36628,f36629,f36630,f36631,f36632,f36633,f36634,f36635,f37080,f36636,f37081,f36968,f36966,f36660,f36661,f36662,f36663,f36664,f36963,f36961,f36959,f36957,f37087,f37088,f36680,f36684,f36685,f36686,f36698,f37096,f36699,f37098,f36700,f37100,f36701,f37101,f36702,f37102,f36703,f37103,f36704,f37104,f36705,f37106,f36708,f37108,f36714,f36715,f36716,f36717,f36718,f36719,f36720,f36721,f36722,f36723,f37120,f36728])).
fof(f36728,plain,(
  ( ! [X472,X473] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X473)))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X472),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X472))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X473))) = s(bool,happ(s(fun(bool,bool),X473),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X472)))) )),
  inference(superposition,[],[f4870,f2291])).
fof(f36708,plain,(
  ( ! [X405,X403,X404,X406] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X405)))) | s(bool,X406) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X404),s('type.extreal.extreal',X403))) | s(bool,'T_0') = s(bool,X406) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X405))) = s(bool,happ(s(fun(bool,bool),X405),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X403),s('type.extreal.extreal',X404)))) )),
  inference(superposition,[],[f2635,f2291])).
fof(f36705,plain,(
  ( ! [X397,X395,X396,X394] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X395)))) | p(s(bool,happ(s(fun(bool,bool),X394),s(bool,X396)))) | ~p(s(bool,X397)) | s(bool,'T_0') = s(bool,sK18(X394,X397,bool)) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X395))) = s(bool,happ(s(fun(bool,bool),X395),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X394),s(bool,'F_0')))) )),
  inference(superposition,[],[f1898,f2291])).
fof(f36704,plain,(
  ( ! [X391,X393,X390,X392] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X391)))) | p(s(bool,happ(s(fun(bool,bool),X390),s(bool,X392)))) | ~p(s(bool,X393)) | s(bool,'T_0') = s(bool,sK19(X390,X393,bool)) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X391))) = s(bool,happ(s(fun(bool,bool),X391),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X390),s(bool,'F_0')))) )),
  inference(superposition,[],[f1923,f2291])).
fof(f36703,plain,(
  ( ! [X389,X387,X388,X386] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X387)))) | p(s(bool,X388)) | p(s(bool,happ(s(fun(bool,bool),X386),s(bool,X389)))) | s(bool,'T_0') = s(bool,sK20(X386,X388,bool)) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X387))) = s(bool,happ(s(fun(bool,bool),X387),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X386),s(bool,'F_0')))) )),
  inference(superposition,[],[f1947,f2291])).
fof(f36702,plain,(
  ( ! [X383,X385,X382,X384] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X383)))) | p(s(bool,X384)) | p(s(bool,happ(s(fun(bool,bool),X382),s(bool,X385)))) | s(bool,'T_0') = s(bool,sK21(X382,X384,bool)) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X383))) = s(bool,happ(s(fun(bool,bool),X383),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X382),s(bool,'F_0')))) )),
  inference(superposition,[],[f1959,f2291])).
fof(f36701,plain,(
  ( ! [X377,X379,X380,X378,X381] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X380)))) | p(s(bool,X379)) | p(s(bool,happ(s(fun(X377,bool),X378),s(X377,X381)))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X380))) = s(bool,happ(s(fun(bool,bool),X380),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(X377,bool),X378),s(X377,sK21(X378,X379,X377))))) )),
  inference(superposition,[],[f316,f2291])).
fof(f36700,plain,(
  ( ! [X372,X374,X376,X373,X375] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X375)))) | p(s(bool,X374)) | p(s(bool,happ(s(fun(X372,bool),X373),s(X372,X376)))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X375))) = s(bool,happ(s(fun(bool,bool),X375),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(X372,bool),X373),s(X372,sK20(X373,X374,X372))))) )),
  inference(superposition,[],[f314,f2291])).
fof(f36699,plain,(
  ( ! [X368,X370,X367,X369,X371] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X370)))) | p(s(bool,happ(s(fun(X367,bool),X368),s(X367,X371)))) | ~p(s(bool,X369)) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X370))) = s(bool,happ(s(fun(bool,bool),X370),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(X367,bool),X368),s(X367,sK19(X368,X369,X367))))) )),
  inference(superposition,[],[f312,f2291])).
fof(f36698,plain,(
  ( ! [X364,X366,X362,X365,X363] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X365)))) | p(s(bool,happ(s(fun(X362,bool),X363),s(X362,X366)))) | ~p(s(bool,X364)) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X365))) = s(bool,happ(s(fun(bool,bool),X365),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(X362,bool),X363),s(X362,sK18(X363,X364,X362))))) )),
  inference(superposition,[],[f310,f2291])).
fof(f37087,plain,(
  ( ! [X218,X220,X219,X217] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X218)))) | s(bool,X217) = s(bool,'|dollar|exists_1'(s(fun(X219,bool),X220))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X218))) = s(bool,happ(s(fun(bool,bool),X218),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X217)) )),
  inference(subsumption_resolution,[],[f36669,f2605])).
fof(f36669,plain,(
  ( ! [X218,X220,X219,X217] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X218)))) | s(bool,X217) = s(bool,'|dollar|exists_1'(s(fun(X219,bool),X220))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X219,bool),X220))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X218))) = s(bool,happ(s(fun(bool,bool),X218),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X217)) )),
  inference(superposition,[],[f4971,f2291])).
fof(f36957,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | sP6('T_0','|dollar|exists_1'(s(fun(bool,bool),X0)),happ(s(fun(bool,bool),X0),s(bool,'const.min.@_1'(s(fun(bool,bool),X1))))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1))) = s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) )),
  inference(condensation,[],[f36956])).
fof(f36956,plain,(
  ( ! [X216,X215,X214] : (sP6('T_0','|dollar|exists_1'(s(fun(bool,bool),X216)),happ(s(fun(bool,bool),X216),s(bool,'const.min.@_1'(s(fun(bool,bool),X215))))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X216))) | s(bool,'T_0') = s(bool,X214) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X215))) = s(bool,happ(s(fun(bool,bool),X215),s(bool,'T_0')))) )),
  inference(duplicate_literal_removal,[],[f36668])).
fof(f36668,plain,(
  ( ! [X216,X215,X214] : (sP6('T_0','|dollar|exists_1'(s(fun(bool,bool),X216)),happ(s(fun(bool,bool),X216),s(bool,'const.min.@_1'(s(fun(bool,bool),X215))))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X216))) | s(bool,'T_0') = s(bool,X214) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X215))) = s(bool,happ(s(fun(bool,bool),X215),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X214)) )),
  inference(superposition,[],[f3867,f2291])).
fof(f36959,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | ~sP6('F_0','|dollar|exists_1'(s(fun(bool,bool),X0)),happ(s(fun(bool,bool),X0),s(bool,'const.min.@_1'(s(fun(bool,bool),X1))))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1))) = s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) )),
  inference(condensation,[],[f36958])).
fof(f36958,plain,(
  ( ! [X212,X213,X211] : (~sP6('F_0','|dollar|exists_1'(s(fun(bool,bool),X213)),happ(s(fun(bool,bool),X213),s(bool,'const.min.@_1'(s(fun(bool,bool),X212))))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X213))) | s(bool,'T_0') = s(bool,X211) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X212))) = s(bool,happ(s(fun(bool,bool),X212),s(bool,'T_0')))) )),
  inference(duplicate_literal_removal,[],[f36667])).
fof(f36667,plain,(
  ( ! [X212,X213,X211] : (~sP6('F_0','|dollar|exists_1'(s(fun(bool,bool),X213)),happ(s(fun(bool,bool),X213),s(bool,'const.min.@_1'(s(fun(bool,bool),X212))))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X213))) | s(bool,'T_0') = s(bool,X211) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X212))) = s(bool,happ(s(fun(bool,bool),X212),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X211)) )),
  inference(superposition,[],[f3851,f2291])).
fof(f36961,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | sP6('T_0',happ(s(fun(bool,bool),X0),s(bool,'const.min.@_1'(s(fun(bool,bool),X1)))),'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1))) = s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) )),
  inference(condensation,[],[f36960])).
fof(f36960,plain,(
  ( ! [X210,X208,X209] : (sP6('T_0',happ(s(fun(bool,bool),X210),s(bool,'const.min.@_1'(s(fun(bool,bool),X209)))),'|dollar|exists_1'(s(fun(bool,bool),X210))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X210))) | s(bool,'T_0') = s(bool,X208) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X209))) = s(bool,happ(s(fun(bool,bool),X209),s(bool,'T_0')))) )),
  inference(duplicate_literal_removal,[],[f36666])).
fof(f36666,plain,(
  ( ! [X210,X208,X209] : (sP6('T_0',happ(s(fun(bool,bool),X210),s(bool,'const.min.@_1'(s(fun(bool,bool),X209)))),'|dollar|exists_1'(s(fun(bool,bool),X210))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X210))) | s(bool,'T_0') = s(bool,X208) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X209))) = s(bool,happ(s(fun(bool,bool),X209),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X208)) )),
  inference(superposition,[],[f3762,f2291])).
fof(f36963,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | ~sP6('F_0',happ(s(fun(bool,bool),X0),s(bool,'const.min.@_1'(s(fun(bool,bool),X1)))),'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1))) = s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) )),
  inference(condensation,[],[f36962])).
fof(f36962,plain,(
  ( ! [X206,X207,X205] : (~sP6('F_0',happ(s(fun(bool,bool),X207),s(bool,'const.min.@_1'(s(fun(bool,bool),X206)))),'|dollar|exists_1'(s(fun(bool,bool),X207))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X207))) | s(bool,'T_0') = s(bool,X205) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X206))) = s(bool,happ(s(fun(bool,bool),X206),s(bool,'T_0')))) )),
  inference(duplicate_literal_removal,[],[f36665])).
fof(f36665,plain,(
  ( ! [X206,X207,X205] : (~sP6('F_0',happ(s(fun(bool,bool),X207),s(bool,'const.min.@_1'(s(fun(bool,bool),X206)))),'|dollar|exists_1'(s(fun(bool,bool),X207))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X207))) | s(bool,'T_0') = s(bool,X205) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X206))) = s(bool,happ(s(fun(bool,bool),X206),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X205)) )),
  inference(superposition,[],[f3742,f2291])).
fof(f36968,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) = s(bool,happ(s(fun(bool,bool),X0),s(bool,'const.min.@_1'(s(fun(bool,bool),X1))))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1))) = s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) )),
  inference(condensation,[],[f36967])).
fof(f36967,plain,(
  ( ! [X156,X154,X155] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X156))) = s(bool,happ(s(fun(bool,bool),X156),s(bool,'const.min.@_1'(s(fun(bool,bool),X155))))) | s(bool,'T_0') = s(bool,X154) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X156))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X155))) = s(bool,happ(s(fun(bool,bool),X155),s(bool,'T_0')))) )),
  inference(duplicate_literal_removal,[],[f36652])).
fof(f36652,plain,(
  ( ! [X156,X154,X155] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X156))) = s(bool,happ(s(fun(bool,bool),X156),s(bool,'const.min.@_1'(s(fun(bool,bool),X155))))) | s(bool,'T_0') = s(bool,X154) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X156))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X155))) = s(bool,happ(s(fun(bool,bool),X155),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X154)) )),
  inference(superposition,[],[f2135,f2291])).
fof(f36636,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X88)))) | ~sP9(X87,X89,X90) | sP10(X87,X89,X90) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X88))) = s(bool,happ(s(fun(bool,bool),X88),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X87)) )),
  inference(superposition,[],[f453,f2291])).
fof(f36635,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X84)))) | ~sP11(X83,X85,X86) | sP12(X83,X85,X86) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X84))) = s(bool,happ(s(fun(bool,bool),X84),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X83)) )),
  inference(superposition,[],[f452,f2291])).
fof(f36622,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X32)))) | sP6(X33,X34,X31) | ~sP5(X33,X31,X34) | ~sP4(X34,X31,X33) | ~sP3(X31,X34,X33) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X32))) = s(bool,happ(s(fun(bool,bool),X32),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f382,f2291])).
fof(f36619,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X23)))) | sP0(X22,X24) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X23))) = s(bool,happ(s(fun(bool,bool),X23),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f335,f2291])).
fof(f36618,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X20)))) | sP0(X21,X19) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X20))) = s(bool,happ(s(fun(bool,bool),X20),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f334,f2291])).
fof(f36600,plain,(
  ( ! [X227,X225,X226,X224] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X224)))) | s(bool,'|dollar|exists_1'(s(fun(X225,bool),X226))) = s(bool,happ(s(fun(bool,bool),X224),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X225,bool),X226))) | s(bool,X227) = s(bool,'const.min.@_1'(s(fun(bool,bool),X224))) | s(bool,'T_0') = s(bool,X227)) )),
  inference(superposition,[],[f4971,f2291])).
fof(f36599,plain,(
  ( ! [X222,X223,X221] : (sP6('T_0','|dollar|exists_1'(s(fun(bool,bool),X222)),happ(s(fun(bool,bool),X222),s(bool,'|dollar|exists_1'(s(fun(bool,bool),X221))))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X222))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X221),s(bool,'T_0'))) | s(bool,X223) = s(bool,'const.min.@_1'(s(fun(bool,bool),X221))) | s(bool,'T_0') = s(bool,X223)) )),
  inference(superposition,[],[f3867,f2291])).
fof(f36598,plain,(
  ( ! [X218,X220,X219] : (~sP6('F_0','|dollar|exists_1'(s(fun(bool,bool),X219)),happ(s(fun(bool,bool),X219),s(bool,'|dollar|exists_1'(s(fun(bool,bool),X218))))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X219))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X218),s(bool,'T_0'))) | s(bool,X220) = s(bool,'const.min.@_1'(s(fun(bool,bool),X218))) | s(bool,'T_0') = s(bool,X220)) )),
  inference(superposition,[],[f3851,f2291])).
fof(f36597,plain,(
  ( ! [X216,X215,X217] : (sP6('T_0',happ(s(fun(bool,bool),X216),s(bool,'|dollar|exists_1'(s(fun(bool,bool),X215)))),'|dollar|exists_1'(s(fun(bool,bool),X216))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X216))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X215),s(bool,'T_0'))) | s(bool,X217) = s(bool,'const.min.@_1'(s(fun(bool,bool),X215))) | s(bool,'T_0') = s(bool,X217)) )),
  inference(superposition,[],[f3762,f2291])).
fof(f36596,plain,(
  ( ! [X212,X213,X214] : (~sP6('F_0',happ(s(fun(bool,bool),X213),s(bool,'|dollar|exists_1'(s(fun(bool,bool),X212)))),'|dollar|exists_1'(s(fun(bool,bool),X213))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X213))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X212),s(bool,'T_0'))) | s(bool,X214) = s(bool,'const.min.@_1'(s(fun(bool,bool),X212))) | s(bool,'T_0') = s(bool,X214)) )),
  inference(superposition,[],[f3742,f2291])).
fof(f36595,plain,(
  ( ! [X210,X208,X211,X209,X207] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X207)))) | p(s(bool,X208)) | ~sP12(X209,X208,happ(s(fun(bool,bool),X207),s(bool,'T_0'))) | s(bool,X209) = s(bool,X210) | s(bool,'T_0') = s(bool,X210) | s(bool,X211) = s(bool,'const.min.@_1'(s(fun(bool,bool),X207))) | s(bool,'T_0') = s(bool,X211)) )),
  inference(superposition,[],[f2623,f2291])).
fof(f36594,plain,(
  ( ! [X206,X204,X202,X205,X203] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X202)))) | sP12(X203,happ(s(fun(bool,bool),X202),s(bool,'T_0')),X204) | s(bool,X203) = s(bool,X205) | s(bool,'T_0') = s(bool,X205) | s(bool,X206) = s(bool,'const.min.@_1'(s(fun(bool,bool),X202))) | s(bool,'T_0') = s(bool,X206)) )),
  inference(superposition,[],[f2622,f2291])).
fof(f36593,plain,(
  ( ! [X198,X200,X199,X197,X201] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X197)))) | sP10(X198,happ(s(fun(bool,bool),X197),s(bool,'T_0')),X199) | s(bool,X198) = s(bool,X200) | s(bool,'T_0') = s(bool,X200) | s(bool,X201) = s(bool,'const.min.@_1'(s(fun(bool,bool),X197))) | s(bool,'T_0') = s(bool,X201)) )),
  inference(superposition,[],[f2618,f2291])).
fof(f36592,plain,(
  ( ! [X196,X194,X192,X195,X193] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X192)))) | sP10(X193,X194,happ(s(fun(bool,bool),X192),s(bool,'T_0'))) | s(bool,X193) = s(bool,X195) | s(bool,'T_0') = s(bool,X195) | s(bool,X196) = s(bool,'const.min.@_1'(s(fun(bool,bool),X192))) | s(bool,'T_0') = s(bool,X196)) )),
  inference(superposition,[],[f2617,f2291])).
fof(f36591,plain,(
  ( ! [X191,X189,X187,X190,X188] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X187)))) | p(s(bool,X188)) | ~sP8(X188,happ(s(fun(bool,bool),X187),s(bool,'T_0')),X189) | s(bool,X189) = s(bool,X190) | s(bool,'T_0') = s(bool,X190) | s(bool,X191) = s(bool,'const.min.@_1'(s(fun(bool,bool),X187))) | s(bool,'T_0') = s(bool,X191)) )),
  inference(superposition,[],[f2615,f2291])).
fof(f36590,plain,(
  ( ! [X185,X186,X184,X183] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X183)))) | sP1(happ(s(fun(bool,bool),X183),s(bool,'T_0')),X184) | s(bool,X184) = s(bool,X185) | s(bool,'T_0') = s(bool,X185) | s(bool,X186) = s(bool,'const.min.@_1'(s(fun(bool,bool),X183))) | s(bool,'T_0') = s(bool,X186)) )),
  inference(superposition,[],[f2610,f2291])).
fof(f36589,plain,(
  ( ! [X182,X180,X181,X179] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X179)))) | sP1(X180,happ(s(fun(bool,bool),X179),s(bool,'T_0'))) | s(bool,X180) = s(bool,X181) | s(bool,'T_0') = s(bool,X181) | s(bool,X182) = s(bool,'const.min.@_1'(s(fun(bool,bool),X179))) | s(bool,'T_0') = s(bool,X182)) )),
  inference(superposition,[],[f2609,f2291])).
fof(f36588,plain,(
  ( ! [X177,X175,X178,X176] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X175)))) | s(bool,X176) = s(bool,happ(s(fun(bool,bool),X175),s(bool,'T_0'))) | s(bool,X176) = s(bool,X177) | s(bool,'T_0') = s(bool,X177) | s(bool,X178) = s(bool,'const.min.@_1'(s(fun(bool,bool),X175))) | s(bool,'T_0') = s(bool,X178)) )),
  inference(superposition,[],[f2606,f2291])).
fof(f36587,plain,(
  ( ! [X173,X174,X172] : (s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),X172))) | s(bool,X173) = s(bool,happ(s(fun(bool,bool),X172),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X173) | s(bool,X174) = s(bool,'const.min.@_1'(s(fun(bool,bool),X172))) | s(bool,'T_0') = s(bool,X174)) )),
  inference(superposition,[],[f2605,f2291])).
fof(f36584,plain,(
  ( ! [X167,X165,X166,X164] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X164)))) | s(bool,X165) != s(bool,happ(s(fun(bool,bool),X164),s(bool,'T_0'))) | s(bool,X165) = s(bool,X166) | s(bool,'T_0') = s(bool,X166) | s(bool,X167) = s(bool,'const.min.@_1'(s(fun(bool,bool),X164))) | s(bool,'T_0') = s(bool,X167)) )),
  inference(superposition,[],[f2204,f2291])).
fof(f36583,plain,(
  ( ! [X163,X161,X162] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X162))) = s(bool,happ(s(fun(bool,bool),X162),s(bool,'|dollar|exists_1'(s(fun(bool,bool),X161))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X161),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X162))) | s(bool,X163) = s(bool,'const.min.@_1'(s(fun(bool,bool),X161))) | s(bool,'T_0') = s(bool,X163)) )),
  inference(superposition,[],[f2135,f2291])).
fof(f36567,plain,(
  ( ! [X94,X97,X95,X96] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X94)))) | ~sP9(happ(s(fun(bool,bool),X94),s(bool,'T_0')),X95,X96) | sP10(happ(s(fun(bool,bool),X94),s(bool,'T_0')),X95,X96) | s(bool,X97) = s(bool,'const.min.@_1'(s(fun(bool,bool),X94))) | s(bool,'T_0') = s(bool,X97)) )),
  inference(superposition,[],[f453,f2291])).
fof(f36566,plain,(
  ( ! [X92,X90,X93,X91] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X90)))) | ~sP11(happ(s(fun(bool,bool),X90),s(bool,'T_0')),X91,X92) | sP12(happ(s(fun(bool,bool),X90),s(bool,'T_0')),X91,X92) | s(bool,X93) = s(bool,'const.min.@_1'(s(fun(bool,bool),X90))) | s(bool,'T_0') = s(bool,X93)) )),
  inference(superposition,[],[f452,f2291])).
fof(f36562,plain,(
  ( ! [X76,X74,X77,X75] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X74)))) | p(s(bool,X75)) | ~p(s(bool,X76)) | ~sP11(X76,X75,happ(s(fun(bool,bool),X74),s(bool,'T_0'))) | s(bool,X77) = s(bool,'const.min.@_1'(s(fun(bool,bool),X74))) | s(bool,'T_0') = s(bool,X77)) )),
  inference(superposition,[],[f442,f2291])).
fof(f36558,plain,(
  ( ! [X61,X59,X60,X58] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X58)))) | p(s(bool,X59)) | p(s(bool,X60)) | ~sP9(happ(s(fun(bool,bool),X58),s(bool,'T_0')),X59,X60) | s(bool,X61) = s(bool,'const.min.@_1'(s(fun(bool,bool),X58))) | s(bool,'T_0') = s(bool,X61)) )),
  inference(superposition,[],[f425,f2291])).
fof(f36557,plain,(
  ( ! [X57,X54,X56,X55] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X54)))) | ~p(s(bool,X55)) | p(s(bool,X56)) | ~sP8(X56,X55,happ(s(fun(bool,bool),X54),s(bool,'T_0'))) | s(bool,X57) = s(bool,'const.min.@_1'(s(fun(bool,bool),X54))) | s(bool,'T_0') = s(bool,X57)) )),
  inference(superposition,[],[f417,f2291])).
fof(f36556,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X50)))) | ~p(s(bool,X51)) | p(s(bool,X52)) | ~sP7(X51,happ(s(fun(bool,bool),X50),s(bool,'T_0')),X52) | s(bool,X53) = s(bool,'const.min.@_1'(s(fun(bool,bool),X50))) | s(bool,'T_0') = s(bool,X53)) )),
  inference(superposition,[],[f408,f2291])).
fof(f36555,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X46)))) | ~p(s(bool,X47)) | p(s(bool,X48)) | ~sP4(X47,happ(s(fun(bool,bool),X46),s(bool,'T_0')),X48) | s(bool,X49) = s(bool,'const.min.@_1'(s(fun(bool,bool),X46))) | s(bool,'T_0') = s(bool,X49)) )),
  inference(superposition,[],[f389,f2291])).
fof(f36554,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X42)))) | ~p(s(bool,X43)) | p(s(bool,X44)) | ~sP5(X43,happ(s(fun(bool,bool),X42),s(bool,'T_0')),X44) | s(bool,X45) = s(bool,'const.min.@_1'(s(fun(bool,bool),X42))) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f385,f2291])).
fof(f36553,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X38)))) | sP6(X39,X40,happ(s(fun(bool,bool),X38),s(bool,'T_0'))) | ~sP5(X39,happ(s(fun(bool,bool),X38),s(bool,'T_0')),X40) | ~sP4(X40,happ(s(fun(bool,bool),X38),s(bool,'T_0')),X39) | ~sP3(happ(s(fun(bool,bool),X38),s(bool,'T_0')),X40,X39) | s(bool,X41) = s(bool,'const.min.@_1'(s(fun(bool,bool),X38))) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f382,f2291])).
fof(f36548,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X23)))) | ~p(s(bool,X24)) | s(bool,X24) = s(bool,happ(s(fun(bool,bool),X23),s(bool,'T_0'))) | s(bool,X25) = s(bool,'const.min.@_1'(s(fun(bool,bool),X23))) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f321,f2291])).
fof(f36546,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X17)))) | p(s(bool,X18)) | s(bool,X18) = s(bool,happ(s(fun(bool,bool),X17),s(bool,'T_0'))) | s(bool,X19) = s(bool,'const.min.@_1'(s(fun(bool,bool),X17))) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f318,f2291])).
fof(f36545,plain,(
  ( ! [X15,X16] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X15)))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X15),s(bool,'T_0'))) | s(bool,X16) = s(bool,'const.min.@_1'(s(fun(bool,bool),X15))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f281,f2291])).
fof(f36544,plain,(
  ( ! [X14,X13] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X13)))) | s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X13),s(bool,'T_0'))) | s(bool,X14) = s(bool,'const.min.@_1'(s(fun(bool,bool),X13))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f280,f2291])).
fof(f36543,plain,(
  ( ! [X12,X11] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X11)))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X11),s(bool,'T_0'))) | s(bool,X12) = s(bool,'const.min.@_1'(s(fun(bool,bool),X11))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f277,f2291])).
fof(f36542,plain,(
  ( ! [X10,X9] : (s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),X9))) | p(s(bool,happ(s(fun(bool,bool),X9),s(bool,'T_0')))) | s(bool,X10) = s(bool,'const.min.@_1'(s(fun(bool,bool),X9))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f276,f2291])).
fof(f36506,plain,(
  ( ! [X103,X101,X102] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X101))) = s(bool,'const.min.@_1'(s(fun(bool,bool),X103))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X103))) = s(bool,happ(s(fun(bool,bool),X103),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X101),s(bool,X102))) | s(bool,'T_0') = s(bool,X102) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X101)))) )),
  inference(superposition,[],[f2291,f2135])).
fof(f36505,plain,(
  ( ! [X99,X100] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X99))) = s(bool,'const.min.@_1'(s(fun(bool,bool),X100))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X100))) = s(bool,happ(s(fun(bool,bool),X100),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X99),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X99)))) )),
  inference(superposition,[],[f2291,f1768])).
fof(f37012,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(bool,'const.min.@_1'(s(fun(bool,bool),X2))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1))))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X2))) = s(bool,happ(s(fun(bool,bool),X2),s(bool,'T_0')))) )),
  inference(condensation,[],[f37011])).
fof(f37011,plain,(
  ( ! [X61,X59,X60,X58] : (s(bool,'const.min.@_1'(s(fun(bool,bool),X61))) = s(bool,happ(s(fun(X59,bool),X60),s(X59,'const.min.@_1'(s(fun(X59,bool),X60))))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X61))) = s(bool,happ(s(fun(bool,bool),X61),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X58) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X59,bool),X60)))) )),
  inference(duplicate_literal_removal,[],[f36493])).
fof(f36493,plain,(
  ( ! [X61,X59,X60,X58] : (s(bool,'const.min.@_1'(s(fun(bool,bool),X61))) = s(bool,happ(s(fun(X59,bool),X60),s(X59,'const.min.@_1'(s(fun(X59,bool),X60))))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X61))) = s(bool,happ(s(fun(bool,bool),X61),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X58) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X59,bool),X60))) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f2291,f2064])).
fof(f37026,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1))))) | s(bool,'const.min.@_1'(s(fun(bool,bool),X2))) = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X2))) = s(bool,happ(s(fun(bool,bool),X2),s(bool,'T_0')))) )),
  inference(condensation,[],[f37025])).
fof(f37025,plain,(
  ( ! [X14,X17,X15,X16] : (s(bool,'const.min.@_1'(s(fun(bool,bool),X17))) = s(bool,'|dollar|exists_1'(s(fun(X15,bool),X16))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X17))) = s(bool,happ(s(fun(bool,bool),X17),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,happ(s(fun(X15,bool),X16),s(X15,'const.min.@_1'(s(fun(X15,bool),X16)))))) )),
  inference(duplicate_literal_removal,[],[f36484])).
fof(f36484,plain,(
  ( ! [X14,X17,X15,X16] : (s(bool,'const.min.@_1'(s(fun(bool,bool),X17))) = s(bool,'|dollar|exists_1'(s(fun(X15,bool),X16))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X17))) = s(bool,happ(s(fun(bool,bool),X17),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,happ(s(fun(X15,bool),X16),s(X15,'const.min.@_1'(s(fun(X15,bool),X16))))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f2291,f2039])).
fof(f36476,plain,(
  ( ! [X66,X64,X65] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X64))) = s(bool,'const.min.@_1'(s(fun(bool,bool),X65))) | s(bool,X66) = s(bool,'const.min.@_1'(s(fun(bool,bool),X64))) | s(bool,'T_0') = s(bool,X66) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X65))) = s(bool,happ(s(fun(bool,bool),X65),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X64),s(bool,'T_0')))) )),
  inference(superposition,[],[f2291,f2291])).
fof(f36475,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X60))) = s(bool,happ(s(fun(X61,bool),X62),s(X61,'const.min.@_1'(s(fun(X61,bool),X62))))) | s(bool,X63) = s(bool,'const.min.@_1'(s(fun(bool,bool),X60))) | s(bool,'T_0') = s(bool,X63) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X61,bool),X62))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X60),s(bool,'T_0')))) )),
  inference(superposition,[],[f2291,f2064])).
fof(f36471,plain,(
  ( ! [X43,X41,X44,X42,X40] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X40))) | s(bool,X41) = s(bool,'const.min.@_1'(s(fun(bool,bool),X40))) | s(bool,'T_0') = s(bool,X41) | s(bool,X42) = s(bool,X43) | s(bool,X43) = s(bool,happ(s(fun(bool,bool),X40),s(bool,'T_0'))) | s(bool,X42) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f2291,f2054])).
fof(f36470,plain,(
  ( ! [X39,X37,X35,X38,X36] : (s(bool,X36) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X35))) | s(bool,X37) = s(bool,'const.min.@_1'(s(fun(bool,bool),X35))) | s(bool,'T_0') = s(bool,X37) | s(bool,X36) = s(bool,X38) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X35),s(bool,'T_0'))) | s(bool,X38) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f2291,f2054])).
fof(f36469,plain,(
  ( ! [X30,X33,X31,X34,X32] : (s(bool,X31) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X30))) | s(bool,X32) = s(bool,'const.min.@_1'(s(fun(bool,bool),X30))) | s(bool,'T_0') = s(bool,X32) | s(bool,X33) = s(bool,happ(s(fun(bool,bool),X30),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X31) | s(bool,X33) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f2291,f2054])).
fof(f36466,plain,(
  ( ! [X19,X17,X18,X16] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X16))) = s(bool,'|dollar|exists_1'(s(fun(X17,bool),X18))) | s(bool,X19) = s(bool,'const.min.@_1'(s(fun(bool,bool),X16))) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,happ(s(fun(X17,bool),X18),s(X17,'const.min.@_1'(s(fun(X17,bool),X18))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X16),s(bool,'T_0')))) )),
  inference(superposition,[],[f2291,f2039])).
fof(f36445,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1))) = s(bool,happ(s(fun(bool,bool),X1),s(bool,X0))) | s(bool,X2) = s(bool,'const.min.@_1'(s(fun(bool,bool),X1))) | s(bool,'T_0') = s(bool,X2) | s(bool,X3) = s(bool,X4) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X5) | s(bool,X0) = s(bool,X3)) )),
  inference(superposition,[],[f2291,f2054])).
fof(f36378,plain,(
  ( ! [X265,X267,X264,X266] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X264)))) | s(bool,'|dollar|exists_1'(s(fun(X266,bool),X267))) = s(bool,happ(s(fun(bool,bool),X264),s(bool,X265))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X266,bool),X267))) | s(bool,'T_0') = s(bool,X265) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X264)))) )),
  inference(superposition,[],[f4971,f2135])).
fof(f36377,plain,(
  ( ! [X261,X263,X262] : (sP6('T_0','|dollar|exists_1'(s(fun(bool,bool),X263)),happ(s(fun(bool,bool),X263),s(bool,'|dollar|exists_1'(s(fun(bool,bool),X261))))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X263))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X261),s(bool,X262))) | s(bool,'T_0') = s(bool,X262) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X261)))) )),
  inference(superposition,[],[f3867,f2135])).
fof(f36376,plain,(
  ( ! [X259,X260,X258] : (~sP6('F_0','|dollar|exists_1'(s(fun(bool,bool),X260)),happ(s(fun(bool,bool),X260),s(bool,'|dollar|exists_1'(s(fun(bool,bool),X258))))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X260))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X258),s(bool,X259))) | s(bool,'T_0') = s(bool,X259) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X258)))) )),
  inference(superposition,[],[f3851,f2135])).
fof(f36375,plain,(
  ( ! [X255,X257,X256] : (sP6('T_0',happ(s(fun(bool,bool),X257),s(bool,'|dollar|exists_1'(s(fun(bool,bool),X255)))),'|dollar|exists_1'(s(fun(bool,bool),X257))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X257))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X255),s(bool,X256))) | s(bool,'T_0') = s(bool,X256) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X255)))) )),
  inference(superposition,[],[f3762,f2135])).
fof(f36374,plain,(
  ( ! [X253,X254,X252] : (~sP6('F_0',happ(s(fun(bool,bool),X254),s(bool,'|dollar|exists_1'(s(fun(bool,bool),X252)))),'|dollar|exists_1'(s(fun(bool,bool),X254))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X254))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X252),s(bool,X253))) | s(bool,'T_0') = s(bool,X253) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X252)))) )),
  inference(superposition,[],[f3742,f2135])).
fof(f36373,plain,(
  ( ! [X251,X249,X250,X248,X247] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X247)))) | p(s(bool,X249)) | ~sP12(X250,X249,happ(s(fun(bool,bool),X247),s(bool,X248))) | s(bool,X250) = s(bool,X251) | s(bool,'T_0') = s(bool,X251) | s(bool,'T_0') = s(bool,X248) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X247)))) )),
  inference(superposition,[],[f2623,f2135])).
fof(f36372,plain,(
  ( ! [X243,X246,X244,X242,X245] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X242)))) | sP12(X244,happ(s(fun(bool,bool),X242),s(bool,X243)),X245) | s(bool,X244) = s(bool,X246) | s(bool,'T_0') = s(bool,X246) | s(bool,'T_0') = s(bool,X243) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X242)))) )),
  inference(superposition,[],[f2622,f2135])).
fof(f36371,plain,(
  ( ! [X241,X239,X237,X240,X238] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X237)))) | sP10(X239,happ(s(fun(bool,bool),X237),s(bool,X238)),X240) | s(bool,X239) = s(bool,X241) | s(bool,'T_0') = s(bool,X241) | s(bool,'T_0') = s(bool,X238) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X237)))) )),
  inference(superposition,[],[f2618,f2135])).
fof(f36370,plain,(
  ( ! [X235,X233,X236,X234,X232] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X232)))) | sP10(X234,X235,happ(s(fun(bool,bool),X232),s(bool,X233))) | s(bool,X234) = s(bool,X236) | s(bool,'T_0') = s(bool,X236) | s(bool,'T_0') = s(bool,X233) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X232)))) )),
  inference(superposition,[],[f2617,f2135])).
fof(f36369,plain,(
  ( ! [X231,X229,X227,X230,X228] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X227)))) | p(s(bool,X229)) | ~sP8(X229,happ(s(fun(bool,bool),X227),s(bool,X228)),X230) | s(bool,X230) = s(bool,X231) | s(bool,'T_0') = s(bool,X231) | s(bool,'T_0') = s(bool,X228) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X227)))) )),
  inference(superposition,[],[f2615,f2135])).
fof(f36368,plain,(
  ( ! [X225,X223,X226,X224] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X223)))) | sP1(happ(s(fun(bool,bool),X223),s(bool,X224)),X225) | s(bool,X225) = s(bool,X226) | s(bool,'T_0') = s(bool,X226) | s(bool,'T_0') = s(bool,X224) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X223)))) )),
  inference(superposition,[],[f2610,f2135])).
fof(f36367,plain,(
  ( ! [X222,X220,X219,X221] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X219)))) | sP1(X221,happ(s(fun(bool,bool),X219),s(bool,X220))) | s(bool,X221) = s(bool,X222) | s(bool,'T_0') = s(bool,X222) | s(bool,'T_0') = s(bool,X220) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X219)))) )),
  inference(superposition,[],[f2609,f2135])).
fof(f36366,plain,(
  ( ! [X218,X216,X215,X217] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X215)))) | s(bool,X217) = s(bool,happ(s(fun(bool,bool),X215),s(bool,X216))) | s(bool,X217) = s(bool,X218) | s(bool,'T_0') = s(bool,X218) | s(bool,'T_0') = s(bool,X216) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X215)))) )),
  inference(superposition,[],[f2606,f2135])).
fof(f36365,plain,(
  ( ! [X212,X213,X214] : (s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),X212))) | s(bool,X214) = s(bool,happ(s(fun(bool,bool),X212),s(bool,X213))) | s(bool,'T_0') = s(bool,X214) | s(bool,'T_0') = s(bool,X213) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X212)))) )),
  inference(superposition,[],[f2605,f2135])).
fof(f36362,plain,(
  ( ! [X206,X204,X207,X205] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X204)))) | s(bool,X206) != s(bool,happ(s(fun(bool,bool),X204),s(bool,X205))) | s(bool,X206) = s(bool,X207) | s(bool,'T_0') = s(bool,X207) | s(bool,'T_0') = s(bool,X205) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X204)))) )),
  inference(superposition,[],[f2204,f2135])).
fof(f36346,plain,(
  ( ! [X140,X138,X139,X137] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X137)))) | ~sP9(happ(s(fun(bool,bool),X137),s(bool,X138)),X139,X140) | sP10(happ(s(fun(bool,bool),X137),s(bool,X138)),X139,X140) | s(bool,'T_0') = s(bool,X138) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X137)))) )),
  inference(superposition,[],[f453,f2135])).
fof(f36345,plain,(
  ( ! [X134,X136,X135,X133] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X133)))) | ~sP11(happ(s(fun(bool,bool),X133),s(bool,X134)),X135,X136) | sP12(happ(s(fun(bool,bool),X133),s(bool,X134)),X135,X136) | s(bool,'T_0') = s(bool,X134) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X133)))) )),
  inference(superposition,[],[f452,f2135])).
fof(f36341,plain,(
  ( ! [X118,X120,X119,X117] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X117)))) | p(s(bool,X119)) | ~p(s(bool,X120)) | ~sP11(X120,X119,happ(s(fun(bool,bool),X117),s(bool,X118))) | s(bool,'T_0') = s(bool,X118) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X117)))) )),
  inference(superposition,[],[f442,f2135])).
fof(f36337,plain,(
  ( ! [X103,X101,X102,X104] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X101)))) | p(s(bool,X103)) | p(s(bool,X104)) | ~sP9(happ(s(fun(bool,bool),X101),s(bool,X102)),X103,X104) | s(bool,'T_0') = s(bool,X102) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X101)))) )),
  inference(superposition,[],[f425,f2135])).
fof(f36336,plain,(
  ( ! [X99,X97,X100,X98] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X97)))) | ~p(s(bool,X99)) | p(s(bool,X100)) | ~sP8(X100,X99,happ(s(fun(bool,bool),X97),s(bool,X98))) | s(bool,'T_0') = s(bool,X98) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X97)))) )),
  inference(superposition,[],[f417,f2135])).
fof(f36335,plain,(
  ( ! [X94,X95,X93,X96] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X93)))) | ~p(s(bool,X95)) | p(s(bool,X96)) | ~sP7(X95,happ(s(fun(bool,bool),X93),s(bool,X94)),X96) | s(bool,'T_0') = s(bool,X94) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X93)))) )),
  inference(superposition,[],[f408,f2135])).
fof(f36334,plain,(
  ( ! [X92,X90,X91,X89] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X89)))) | ~p(s(bool,X91)) | p(s(bool,X92)) | ~sP4(X91,happ(s(fun(bool,bool),X89),s(bool,X90)),X92) | s(bool,'T_0') = s(bool,X90) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X89)))) )),
  inference(superposition,[],[f389,f2135])).
fof(f36333,plain,(
  ( ! [X88,X87,X85,X86] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X85)))) | ~p(s(bool,X87)) | p(s(bool,X88)) | ~sP5(X87,happ(s(fun(bool,bool),X85),s(bool,X86)),X88) | s(bool,'T_0') = s(bool,X86) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X85)))) )),
  inference(superposition,[],[f385,f2135])).
fof(f36332,plain,(
  ( ! [X83,X81,X84,X82] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X81)))) | sP6(X83,X84,happ(s(fun(bool,bool),X81),s(bool,X82))) | ~sP5(X83,happ(s(fun(bool,bool),X81),s(bool,X82)),X84) | ~sP4(X84,happ(s(fun(bool,bool),X81),s(bool,X82)),X83) | ~sP3(happ(s(fun(bool,bool),X81),s(bool,X82)),X84,X83) | s(bool,'T_0') = s(bool,X82) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X81)))) )),
  inference(superposition,[],[f382,f2135])).
fof(f36327,plain,(
  ( ! [X68,X66,X67] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X66)))) | ~p(s(bool,X68)) | s(bool,X68) = s(bool,happ(s(fun(bool,bool),X66),s(bool,X67))) | s(bool,'T_0') = s(bool,X67) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X66)))) )),
  inference(superposition,[],[f321,f2135])).
fof(f36325,plain,(
  ( ! [X61,X62,X60] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X60)))) | p(s(bool,X62)) | s(bool,X62) = s(bool,happ(s(fun(bool,bool),X60),s(bool,X61))) | s(bool,'T_0') = s(bool,X61) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X60)))) )),
  inference(superposition,[],[f318,f2135])).
fof(f36324,plain,(
  ( ! [X59,X58] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X58)))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X58),s(bool,X59))) | s(bool,'T_0') = s(bool,X59) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X58)))) )),
  inference(superposition,[],[f281,f2135])).
fof(f36323,plain,(
  ( ! [X57,X56] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X56)))) | s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X56),s(bool,X57))) | s(bool,'T_0') = s(bool,X57) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X56)))) )),
  inference(superposition,[],[f280,f2135])).
fof(f36322,plain,(
  ( ! [X54,X55] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X54)))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X54),s(bool,X55))) | s(bool,'T_0') = s(bool,X55) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X54)))) )),
  inference(superposition,[],[f277,f2135])).
fof(f36321,plain,(
  ( ! [X52,X53] : (s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),X52))) | p(s(bool,happ(s(fun(bool,bool),X52),s(bool,X53)))) | s(bool,'T_0') = s(bool,X53) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X52)))) )),
  inference(superposition,[],[f276,f2135])).
fof(f36296,plain,(
  ( ! [X70,X68,X69,X67] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X67))) = s(bool,happ(s(fun(X69,bool),X70),s(X69,'const.min.@_1'(s(fun(X69,bool),X70))))) | s(bool,'T_0') = s(bool,X68) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X67))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X69,bool),X70))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X67),s(bool,X68)))) )),
  inference(superposition,[],[f2135,f2064])).
fof(f36292,plain,(
  ( ! [X47,X50,X48,X51,X49] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X47))) | s(bool,'T_0') = s(bool,X48) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X47))) | s(bool,X49) = s(bool,X50) | s(bool,X50) = s(bool,happ(s(fun(bool,bool),X47),s(bool,X48))) | s(bool,X49) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f2135,f2054])).
fof(f36291,plain,(
  ( ! [X45,X43,X46,X44,X42] : (s(bool,X44) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X42))) | s(bool,'T_0') = s(bool,X43) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X42))) | s(bool,X44) = s(bool,X45) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X42),s(bool,X43))) | s(bool,X45) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f2135,f2054])).
fof(f36290,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s(bool,X39) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X37))) | s(bool,'T_0') = s(bool,X38) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X37))) | s(bool,X40) = s(bool,happ(s(fun(bool,bool),X37),s(bool,X38))) | s(bool,'T_0') = s(bool,X39) | s(bool,X40) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f2135,f2054])).
fof(f36287,plain,(
  ( ! [X26,X24,X23,X25] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X23))) = s(bool,'|dollar|exists_1'(s(fun(X25,bool),X26))) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X23))) | s(bool,'T_0') = s(bool,happ(s(fun(X25,bool),X26),s(X25,'const.min.@_1'(s(fun(X25,bool),X26))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X23),s(bool,X24)))) )),
  inference(superposition,[],[f2135,f2039])).
fof(f36267,plain,(
  ( ! [X99,X100,X98] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X100))) = s(bool,happ(s(fun(bool,bool),X100),s(bool,'|dollar|exists_1'(s(fun(bool,bool),X98))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X98),s(bool,X99))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X100))) | s(bool,'T_0') = s(bool,X99) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X98)))) )),
  inference(superposition,[],[f2135,f2135])).
fof(f36266,plain,(
  ( ! [X97,X96] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X97))) = s(bool,happ(s(fun(bool,bool),X97),s(bool,'|dollar|exists_1'(s(fun(bool,bool),X96))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X96),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X97))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X96)))) )),
  inference(superposition,[],[f2135,f1768])).
fof(f35923,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,happ(s(fun(X1,bool),X2),s(X1,'const.min.@_1'(s(fun(X1,bool),X2))))) | s(bool,X0) = s(bool,happ(s(fun(X1,bool),X2),s(X1,'const.min.@_1'(s(fun(X1,bool),X2))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X1,bool),X2)))) )),
  inference(equality_factoring,[],[f2064])).
fof(f35789,plain,(
  ( ! [X716,X714,X717,X713,X715] : (s(bool,f(s(X713,sK23(X714,X715,bool,X713)))) != s(bool,happ(s(fun(X716,bool),X717),s(X716,'const.min.@_1'(s(fun(X716,bool),X717))))) | s(fun(X713,bool),X714) = s(fun(X713,bool),X715) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X716,bool),X717))) | s(bool,'T_0') = s(bool,g(s(X713,sK23(X714,X715,bool,X713))))) )),
  inference(superposition,[],[f366,f2064])).
fof(f35621,plain,(
  ( ! [X127,X128,X126] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X126))) = s(bool,happ(s(fun(X127,bool),X128),s(X127,'const.min.@_1'(s(fun(X127,bool),X128))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X127,bool),X128))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X126),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X126)))) )),
  inference(superposition,[],[f2064,f1768])).
fof(f36018,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1))))) = s(bool,happ(s(fun(X2,bool),X3),s(X2,'const.min.@_1'(s(fun(X2,bool),X3))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X2,bool),X3)))) )),
  inference(condensation,[],[f36017])).
fof(f36017,plain,(
  ( ! [X76,X74,X72,X75,X73] : (s(bool,happ(s(fun(X73,bool),X74),s(X73,'const.min.@_1'(s(fun(X73,bool),X74))))) = s(bool,happ(s(fun(X75,bool),X76),s(X75,'const.min.@_1'(s(fun(X75,bool),X76))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X75,bool),X76))) | s(bool,'T_0') = s(bool,X72) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X73,bool),X74)))) )),
  inference(duplicate_literal_removal,[],[f35613])).
fof(f35613,plain,(
  ( ! [X76,X74,X72,X75,X73] : (s(bool,happ(s(fun(X73,bool),X74),s(X73,'const.min.@_1'(s(fun(X73,bool),X74))))) = s(bool,happ(s(fun(X75,bool),X76),s(X75,'const.min.@_1'(s(fun(X75,bool),X76))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X75,bool),X76))) | s(bool,'T_0') = s(bool,X72) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X73,bool),X74))) | s(bool,'T_0') = s(bool,X72)) )),
  inference(superposition,[],[f2064,f2064])).
fof(f36032,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(bool,'|dollar|exists_1'(s(fun(X2,bool),X3))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1))))) | s(bool,'T_0') = s(bool,happ(s(fun(X2,bool),X3),s(X2,'const.min.@_1'(s(fun(X2,bool),X3)))))) )),
  inference(condensation,[],[f36031])).
fof(f36031,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s(bool,'|dollar|exists_1'(s(fun(X20,bool),X21))) = s(bool,happ(s(fun(X22,bool),X23),s(X22,'const.min.@_1'(s(fun(X22,bool),X23))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X22,bool),X23))) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,happ(s(fun(X20,bool),X21),s(X20,'const.min.@_1'(s(fun(X20,bool),X21)))))) )),
  inference(duplicate_literal_removal,[],[f35604])).
fof(f35604,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s(bool,'|dollar|exists_1'(s(fun(X20,bool),X21))) = s(bool,happ(s(fun(X22,bool),X23),s(X22,'const.min.@_1'(s(fun(X22,bool),X23))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X22,bool),X23))) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,happ(s(fun(X20,bool),X21),s(X20,'const.min.@_1'(s(fun(X20,bool),X21))))) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f2064,f2039])).
fof(f36038,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(bool,'F_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) )),
  inference(condensation,[],[f36037])).
fof(f36037,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun(X1,bool),X2),s(X1,'const.min.@_1'(s(fun(X1,bool),X2))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X1,bool),X2))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f35599])).
fof(f35599,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun(X1,bool),X2),s(X1,'const.min.@_1'(s(fun(X1,bool),X2))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X1,bool),X2))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2064,f270])).
fof(f35460,plain,(
  ( ! [X269,X271,X273,X270,X272] : (s(bool,'|dollar|exists_1'(s(fun(X271,bool),X272))) = s(bool,happ(s(fun(X269,bool),X270),s(X269,'const.min.@_1'(s(fun(X269,bool),X270))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X271,bool),X272))) | s(bool,X273) = s(bool,'|dollar|exists_1'(s(fun(X269,bool),X270))) | s(bool,'T_0') = s(bool,X273)) )),
  inference(subsumption_resolution,[],[f35122,f263])).
fof(f35122,plain,(
  ( ! [X269,X271,X273,X270,X272] : (~p(s(bool,'T_0')) | s(bool,'|dollar|exists_1'(s(fun(X271,bool),X272))) = s(bool,happ(s(fun(X269,bool),X270),s(X269,'const.min.@_1'(s(fun(X269,bool),X270))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X271,bool),X272))) | s(bool,X273) = s(bool,'|dollar|exists_1'(s(fun(X269,bool),X270))) | s(bool,'T_0') = s(bool,X273)) )),
  inference(superposition,[],[f4971,f2039])).
fof(f35459,plain,(
  ( ! [X251,X249,X252,X250,X248,X247] : (p(s(bool,X249)) | ~sP12(X250,X249,happ(s(fun(X247,bool),X248),s(X247,'const.min.@_1'(s(fun(X247,bool),X248))))) | s(bool,X250) = s(bool,X251) | s(bool,'T_0') = s(bool,X251) | s(bool,X252) = s(bool,'|dollar|exists_1'(s(fun(X247,bool),X248))) | s(bool,'T_0') = s(bool,X252)) )),
  inference(subsumption_resolution,[],[f35117,f263])).
fof(f35117,plain,(
  ( ! [X251,X249,X252,X250,X248,X247] : (~p(s(bool,'T_0')) | p(s(bool,X249)) | ~sP12(X250,X249,happ(s(fun(X247,bool),X248),s(X247,'const.min.@_1'(s(fun(X247,bool),X248))))) | s(bool,X250) = s(bool,X251) | s(bool,'T_0') = s(bool,X251) | s(bool,X252) = s(bool,'|dollar|exists_1'(s(fun(X247,bool),X248))) | s(bool,'T_0') = s(bool,X252)) )),
  inference(superposition,[],[f2623,f2039])).
fof(f35458,plain,(
  ( ! [X243,X241,X246,X244,X242,X245] : (sP12(X243,happ(s(fun(X241,bool),X242),s(X241,'const.min.@_1'(s(fun(X241,bool),X242)))),X244) | s(bool,X243) = s(bool,X245) | s(bool,'T_0') = s(bool,X245) | s(bool,X246) = s(bool,'|dollar|exists_1'(s(fun(X241,bool),X242))) | s(bool,'T_0') = s(bool,X246)) )),
  inference(subsumption_resolution,[],[f35116,f263])).
fof(f35116,plain,(
  ( ! [X243,X241,X246,X244,X242,X245] : (~p(s(bool,'T_0')) | sP12(X243,happ(s(fun(X241,bool),X242),s(X241,'const.min.@_1'(s(fun(X241,bool),X242)))),X244) | s(bool,X243) = s(bool,X245) | s(bool,'T_0') = s(bool,X245) | s(bool,X246) = s(bool,'|dollar|exists_1'(s(fun(X241,bool),X242))) | s(bool,'T_0') = s(bool,X246)) )),
  inference(superposition,[],[f2622,f2039])).
fof(f35457,plain,(
  ( ! [X239,X237,X235,X240,X238,X236] : (sP10(X237,happ(s(fun(X235,bool),X236),s(X235,'const.min.@_1'(s(fun(X235,bool),X236)))),X238) | s(bool,X237) = s(bool,X239) | s(bool,'T_0') = s(bool,X239) | s(bool,X240) = s(bool,'|dollar|exists_1'(s(fun(X235,bool),X236))) | s(bool,'T_0') = s(bool,X240)) )),
  inference(subsumption_resolution,[],[f35115,f263])).
fof(f35115,plain,(
  ( ! [X239,X237,X235,X240,X238,X236] : (~p(s(bool,'T_0')) | sP10(X237,happ(s(fun(X235,bool),X236),s(X235,'const.min.@_1'(s(fun(X235,bool),X236)))),X238) | s(bool,X237) = s(bool,X239) | s(bool,'T_0') = s(bool,X239) | s(bool,X240) = s(bool,'|dollar|exists_1'(s(fun(X235,bool),X236))) | s(bool,'T_0') = s(bool,X240)) )),
  inference(superposition,[],[f2618,f2039])).
fof(f35456,plain,(
  ( ! [X231,X229,X233,X230,X234,X232] : (sP10(X231,X232,happ(s(fun(X229,bool),X230),s(X229,'const.min.@_1'(s(fun(X229,bool),X230))))) | s(bool,X231) = s(bool,X233) | s(bool,'T_0') = s(bool,X233) | s(bool,X234) = s(bool,'|dollar|exists_1'(s(fun(X229,bool),X230))) | s(bool,'T_0') = s(bool,X234)) )),
  inference(subsumption_resolution,[],[f35114,f263])).
fof(f35114,plain,(
  ( ! [X231,X229,X233,X230,X234,X232] : (~p(s(bool,'T_0')) | sP10(X231,X232,happ(s(fun(X229,bool),X230),s(X229,'const.min.@_1'(s(fun(X229,bool),X230))))) | s(bool,X231) = s(bool,X233) | s(bool,'T_0') = s(bool,X233) | s(bool,X234) = s(bool,'|dollar|exists_1'(s(fun(X229,bool),X230))) | s(bool,'T_0') = s(bool,X234)) )),
  inference(superposition,[],[f2617,f2039])).
fof(f35455,plain,(
  ( ! [X227,X225,X223,X228,X226,X224] : (p(s(bool,X225)) | ~sP8(X225,happ(s(fun(X223,bool),X224),s(X223,'const.min.@_1'(s(fun(X223,bool),X224)))),X226) | s(bool,X226) = s(bool,X227) | s(bool,'T_0') = s(bool,X227) | s(bool,X228) = s(bool,'|dollar|exists_1'(s(fun(X223,bool),X224))) | s(bool,'T_0') = s(bool,X228)) )),
  inference(subsumption_resolution,[],[f35113,f263])).
fof(f35113,plain,(
  ( ! [X227,X225,X223,X228,X226,X224] : (~p(s(bool,'T_0')) | p(s(bool,X225)) | ~sP8(X225,happ(s(fun(X223,bool),X224),s(X223,'const.min.@_1'(s(fun(X223,bool),X224)))),X226) | s(bool,X226) = s(bool,X227) | s(bool,'T_0') = s(bool,X227) | s(bool,X228) = s(bool,'|dollar|exists_1'(s(fun(X223,bool),X224))) | s(bool,'T_0') = s(bool,X228)) )),
  inference(superposition,[],[f2615,f2039])).
fof(f35454,plain,(
  ( ! [X134,X132,X136,X135,X133] : (~sP9(happ(s(fun(X132,bool),X133),s(X132,'const.min.@_1'(s(fun(X132,bool),X133)))),X134,X135) | sP10(happ(s(fun(X132,bool),X133),s(X132,'const.min.@_1'(s(fun(X132,bool),X133)))),X134,X135) | s(bool,X136) = s(bool,'|dollar|exists_1'(s(fun(X132,bool),X133))) | s(bool,'T_0') = s(bool,X136)) )),
  inference(subsumption_resolution,[],[f35096,f263])).
fof(f35096,plain,(
  ( ! [X134,X132,X136,X135,X133] : (~p(s(bool,'T_0')) | ~sP9(happ(s(fun(X132,bool),X133),s(X132,'const.min.@_1'(s(fun(X132,bool),X133)))),X134,X135) | sP10(happ(s(fun(X132,bool),X133),s(X132,'const.min.@_1'(s(fun(X132,bool),X133)))),X134,X135) | s(bool,X136) = s(bool,'|dollar|exists_1'(s(fun(X132,bool),X133))) | s(bool,'T_0') = s(bool,X136)) )),
  inference(superposition,[],[f453,f2039])).
fof(f35453,plain,(
  ( ! [X127,X130,X128,X131,X129] : (~sP11(happ(s(fun(X127,bool),X128),s(X127,'const.min.@_1'(s(fun(X127,bool),X128)))),X129,X130) | sP12(happ(s(fun(X127,bool),X128),s(X127,'const.min.@_1'(s(fun(X127,bool),X128)))),X129,X130) | s(bool,X131) = s(bool,'|dollar|exists_1'(s(fun(X127,bool),X128))) | s(bool,'T_0') = s(bool,X131)) )),
  inference(subsumption_resolution,[],[f35095,f263])).
fof(f35095,plain,(
  ( ! [X127,X130,X128,X131,X129] : (~p(s(bool,'T_0')) | ~sP11(happ(s(fun(X127,bool),X128),s(X127,'const.min.@_1'(s(fun(X127,bool),X128)))),X129,X130) | sP12(happ(s(fun(X127,bool),X128),s(X127,'const.min.@_1'(s(fun(X127,bool),X128)))),X129,X130) | s(bool,X131) = s(bool,'|dollar|exists_1'(s(fun(X127,bool),X128))) | s(bool,'T_0') = s(bool,X131)) )),
  inference(superposition,[],[f452,f2039])).
fof(f35452,plain,(
  ( ! [X125,X123,X126,X124,X122] : (~p(s(bool,X124)) | p(s(bool,X125)) | ~sP12(happ(s(fun(X122,bool),X123),s(X122,'const.min.@_1'(s(fun(X122,bool),X123)))),X125,X124) | s(bool,X126) = s(bool,'|dollar|exists_1'(s(fun(X122,bool),X123))) | s(bool,'T_0') = s(bool,X126)) )),
  inference(subsumption_resolution,[],[f35094,f263])).
fof(f35094,plain,(
  ( ! [X125,X123,X126,X124,X122] : (~p(s(bool,'T_0')) | ~p(s(bool,X124)) | p(s(bool,X125)) | ~sP12(happ(s(fun(X122,bool),X123),s(X122,'const.min.@_1'(s(fun(X122,bool),X123)))),X125,X124) | s(bool,X126) = s(bool,'|dollar|exists_1'(s(fun(X122,bool),X123))) | s(bool,'T_0') = s(bool,X126)) )),
  inference(superposition,[],[f449,f2039])).
fof(f35451,plain,(
  ( ! [X121,X118,X120,X119,X117] : (~p(s(bool,X119)) | sP12(happ(s(fun(X117,bool),X118),s(X117,'const.min.@_1'(s(fun(X117,bool),X118)))),X119,X120) | s(bool,X121) = s(bool,'|dollar|exists_1'(s(fun(X117,bool),X118))) | s(bool,'T_0') = s(bool,X121)) )),
  inference(subsumption_resolution,[],[f35093,f263])).
fof(f35093,plain,(
  ( ! [X121,X118,X120,X119,X117] : (~p(s(bool,'T_0')) | ~p(s(bool,X119)) | sP12(happ(s(fun(X117,bool),X118),s(X117,'const.min.@_1'(s(fun(X117,bool),X118)))),X119,X120) | s(bool,X121) = s(bool,'|dollar|exists_1'(s(fun(X117,bool),X118))) | s(bool,'T_0') = s(bool,X121)) )),
  inference(superposition,[],[f448,f2039])).
fof(f35450,plain,(
  ( ! [X116,X114,X112,X115,X113] : (p(s(bool,X114)) | sP12(happ(s(fun(X112,bool),X113),s(X112,'const.min.@_1'(s(fun(X112,bool),X113)))),X115,X114) | s(bool,X116) = s(bool,'|dollar|exists_1'(s(fun(X112,bool),X113))) | s(bool,'T_0') = s(bool,X116)) )),
  inference(subsumption_resolution,[],[f35092,f263])).
fof(f35092,plain,(
  ( ! [X116,X114,X112,X115,X113] : (~p(s(bool,'T_0')) | p(s(bool,X114)) | sP12(happ(s(fun(X112,bool),X113),s(X112,'const.min.@_1'(s(fun(X112,bool),X113)))),X115,X114) | s(bool,X116) = s(bool,'|dollar|exists_1'(s(fun(X112,bool),X113))) | s(bool,'T_0') = s(bool,X116)) )),
  inference(superposition,[],[f447,f2039])).
fof(f35449,plain,(
  ( ! [X103,X105,X102,X106,X104] : (p(s(bool,X104)) | p(s(bool,X105)) | ~sP10(happ(s(fun(X102,bool),X103),s(X102,'const.min.@_1'(s(fun(X102,bool),X103)))),X105,X104) | s(bool,X106) = s(bool,'|dollar|exists_1'(s(fun(X102,bool),X103))) | s(bool,'T_0') = s(bool,X106)) )),
  inference(subsumption_resolution,[],[f35090,f263])).
fof(f35090,plain,(
  ( ! [X103,X105,X102,X106,X104] : (~p(s(bool,'T_0')) | p(s(bool,X104)) | p(s(bool,X105)) | ~sP10(happ(s(fun(X102,bool),X103),s(X102,'const.min.@_1'(s(fun(X102,bool),X103)))),X105,X104) | s(bool,X106) = s(bool,'|dollar|exists_1'(s(fun(X102,bool),X103))) | s(bool,'T_0') = s(bool,X106)) )),
  inference(superposition,[],[f432,f2039])).
fof(f35448,plain,(
  ( ! [X101,X99,X97,X100,X98] : (~p(s(bool,X99)) | sP10(happ(s(fun(X97,bool),X98),s(X97,'const.min.@_1'(s(fun(X97,bool),X98)))),X99,X100) | s(bool,X101) = s(bool,'|dollar|exists_1'(s(fun(X97,bool),X98))) | s(bool,'T_0') = s(bool,X101)) )),
  inference(subsumption_resolution,[],[f35089,f263])).
fof(f35089,plain,(
  ( ! [X101,X99,X97,X100,X98] : (~p(s(bool,'T_0')) | ~p(s(bool,X99)) | sP10(happ(s(fun(X97,bool),X98),s(X97,'const.min.@_1'(s(fun(X97,bool),X98)))),X99,X100) | s(bool,X101) = s(bool,'|dollar|exists_1'(s(fun(X97,bool),X98))) | s(bool,'T_0') = s(bool,X101)) )),
  inference(superposition,[],[f431,f2039])).
fof(f35447,plain,(
  ( ! [X94,X92,X95,X93,X96] : (~p(s(bool,X94)) | sP10(happ(s(fun(X92,bool),X93),s(X92,'const.min.@_1'(s(fun(X92,bool),X93)))),X95,X94) | s(bool,X96) = s(bool,'|dollar|exists_1'(s(fun(X92,bool),X93))) | s(bool,'T_0') = s(bool,X96)) )),
  inference(subsumption_resolution,[],[f35088,f263])).
fof(f35088,plain,(
  ( ! [X94,X92,X95,X93,X96] : (~p(s(bool,'T_0')) | ~p(s(bool,X94)) | sP10(happ(s(fun(X92,bool),X93),s(X92,'const.min.@_1'(s(fun(X92,bool),X93)))),X95,X94) | s(bool,X96) = s(bool,'|dollar|exists_1'(s(fun(X92,bool),X93))) | s(bool,'T_0') = s(bool,X96)) )),
  inference(superposition,[],[f430,f2039])).
fof(f35446,plain,(
  ( ! [X61,X59,X60,X58] : (p(s(bool,X60)) | sP1(X60,happ(s(fun(X58,bool),X59),s(X58,'const.min.@_1'(s(fun(X58,bool),X59))))) | s(bool,X61) = s(bool,'|dollar|exists_1'(s(fun(X58,bool),X59))) | s(bool,'T_0') = s(bool,X61)) )),
  inference(subsumption_resolution,[],[f35081,f263])).
fof(f35081,plain,(
  ( ! [X61,X59,X60,X58] : (~p(s(bool,'T_0')) | p(s(bool,X60)) | sP1(X60,happ(s(fun(X58,bool),X59),s(X58,'const.min.@_1'(s(fun(X58,bool),X59))))) | s(bool,X61) = s(bool,'|dollar|exists_1'(s(fun(X58,bool),X59))) | s(bool,'T_0') = s(bool,X61)) )),
  inference(superposition,[],[f353,f2039])).
fof(f35445,plain,(
  ( ! [X57,X54,X56,X55] : (p(s(bool,X56)) | sP1(happ(s(fun(X54,bool),X55),s(X54,'const.min.@_1'(s(fun(X54,bool),X55)))),X56) | s(bool,X57) = s(bool,'|dollar|exists_1'(s(fun(X54,bool),X55))) | s(bool,'T_0') = s(bool,X57)) )),
  inference(subsumption_resolution,[],[f35080,f263])).
fof(f35080,plain,(
  ( ! [X57,X54,X56,X55] : (~p(s(bool,'T_0')) | p(s(bool,X56)) | sP1(happ(s(fun(X54,bool),X55),s(X54,'const.min.@_1'(s(fun(X54,bool),X55)))),X56) | s(bool,X57) = s(bool,'|dollar|exists_1'(s(fun(X54,bool),X55))) | s(bool,'T_0') = s(bool,X57)) )),
  inference(superposition,[],[f350,f2039])).
fof(f35444,plain,(
  ( ! [X52,X50,X53,X51] : (sP0(happ(s(fun(X50,bool),X51),s(X50,'const.min.@_1'(s(fun(X50,bool),X51)))),X52) | s(bool,X53) = s(bool,'|dollar|exists_1'(s(fun(X50,bool),X51))) | s(bool,'T_0') = s(bool,X53)) )),
  inference(subsumption_resolution,[],[f35079,f263])).
fof(f35079,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,'T_0')) | sP0(happ(s(fun(X50,bool),X51),s(X50,'const.min.@_1'(s(fun(X50,bool),X51)))),X52) | s(bool,X53) = s(bool,'|dollar|exists_1'(s(fun(X50,bool),X51))) | s(bool,'T_0') = s(bool,X53)) )),
  inference(superposition,[],[f335,f2039])).
fof(f35443,plain,(
  ( ! [X47,X48,X46,X49] : (sP0(X48,happ(s(fun(X46,bool),X47),s(X46,'const.min.@_1'(s(fun(X46,bool),X47))))) | s(bool,X49) = s(bool,'|dollar|exists_1'(s(fun(X46,bool),X47))) | s(bool,'T_0') = s(bool,X49)) )),
  inference(subsumption_resolution,[],[f35078,f263])).
fof(f35078,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,'T_0')) | sP0(X48,happ(s(fun(X46,bool),X47),s(X46,'const.min.@_1'(s(fun(X46,bool),X47))))) | s(bool,X49) = s(bool,'|dollar|exists_1'(s(fun(X46,bool),X47))) | s(bool,'T_0') = s(bool,X49)) )),
  inference(superposition,[],[f334,f2039])).
fof(f35440,plain,(
  ( ! [X30,X28,X29] : (s(bool,'F_0') != s(bool,happ(s(fun(X28,bool),X29),s(X28,'const.min.@_1'(s(fun(X28,bool),X29))))) | s(bool,X30) = s(bool,'|dollar|exists_1'(s(fun(X28,bool),X29))) | s(bool,'T_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f35073,f263])).
fof(f35073,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,happ(s(fun(X28,bool),X29),s(X28,'const.min.@_1'(s(fun(X28,bool),X29))))) | s(bool,X30) = s(bool,'|dollar|exists_1'(s(fun(X28,bool),X29))) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f280,f2039])).
fof(f34980,plain,(
  ( ! [X938,X936,X934,X937,X935] : (s(X936,f(s(bool,'|dollar|exists_1'(s(fun(X937,bool),X938))))) != s(X936,g(s(bool,'|dollar|exists_1'(s(fun(X937,bool),X938))))) | s(fun(bool,X936),X934) = s(fun(bool,X936),X935) | s(bool,'T_0') = s(bool,happ(s(fun(X937,bool),X938),s(X937,'const.min.@_1'(s(fun(X937,bool),X938))))) | s(bool,'T_0') = s(bool,sK23(X934,X935,X936,bool))) )),
  inference(superposition,[],[f366,f2039])).
fof(f34945,plain,(
  ( ! [X788,X787,X789] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X787),s(bool,'|dollar|exists_1'(s(fun(X788,bool),X789))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X787))) | s(bool,'T_0') = s(bool,happ(s(fun(X788,bool),X789),s(X788,'const.min.@_1'(s(fun(X788,bool),X789))))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X787)))) )),
  inference(superposition,[],[f4969,f2039])).
fof(f34930,plain,(
  ( ! [X724,X722,X723] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X722),s(bool,'|dollar|exists_1'(s(fun(X723,bool),X724))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X722))) | s(bool,'T_0') = s(bool,happ(s(fun(X723,bool),X724),s(X723,'const.min.@_1'(s(fun(X723,bool),X724))))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X722)))) )),
  inference(superposition,[],[f1769,f2039])).
fof(f34925,plain,(
  ( ! [X707,X708,X709] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X707))) = s(bool,happ(s(fun(bool,bool),X707),s(bool,'|dollar|exists_1'(s(fun(X708,bool),X709))))) | s(bool,'T_0') = s(bool,happ(s(fun(X708,bool),X709),s(X708,'const.min.@_1'(s(fun(X708,bool),X709))))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X707)))) )),
  inference(superposition,[],[f303,f2039])).
fof(f34924,plain,(
  ( ! [X700,X704,X702,X701,X703] : (s(bool,'|dollar|exists_1'(s(fun(X703,bool),X704))) != s(bool,f(s(X700,sK23(X701,X702,bool,X700)))) | s(fun(X700,bool),X701) = s(fun(X700,bool),X702) | s(bool,'T_0') = s(bool,happ(s(fun(X703,bool),X704),s(X703,'const.min.@_1'(s(fun(X703,bool),X704))))) | s(bool,'T_0') = s(bool,g(s(X700,sK23(X701,X702,bool,X700))))) )),
  inference(superposition,[],[f366,f2039])).
fof(f34771,plain,(
  ( ! [X39,X41,X38,X42,X40] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X38),s(bool,'T_0'))) | s(bool,X39) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X38))) | s(bool,'T_0') = s(bool,X39) | s(bool,X40) = s(bool,X41) | s(bool,X41) = s(bool,'const.min.@_1'(s(fun(bool,bool),X38))) | s(bool,X40) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f2039,f2054])).
fof(f34770,plain,(
  ( ! [X37,X35,X33,X36,X34] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X33),s(bool,X34))) | s(bool,X35) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X33))) | s(bool,'T_0') = s(bool,X35) | s(bool,X34) = s(bool,X36) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X33))) | s(bool,X36) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f2039,f2054])).
fof(f34769,plain,(
  ( ! [X30,X28,X31,X29,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X28),s(bool,X29))) | s(bool,X30) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X28))) | s(bool,'T_0') = s(bool,X30) | s(bool,X31) = s(bool,'const.min.@_1'(s(fun(bool,bool),X28))) | s(bool,'T_0') = s(bool,X29) | s(bool,X31) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f2039,f2054])).
fof(f34767,plain,(
  ( ! [X21,X19,X22,X20,X18] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X18),s(bool,X19))) | s(bool,X20) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X18))) | s(bool,'T_0') = s(bool,X20) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21) | s(bool,X22) = s(bool,'const.min.@_1'(s(fun(bool,bool),X18))) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f2039,f2054])).
fof(f34766,plain,(
  ( ! [X14,X17,X15,X16] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X14),s(bool,'|dollar|exists_1'(s(fun(X15,bool),X16))))) | s(bool,X17) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X14))) | s(bool,'T_0') = s(bool,X17) | s(bool,'T_0') = s(bool,happ(s(fun(X15,bool),X16),s(X15,'const.min.@_1'(s(fun(X15,bool),X16))))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X14)))) )),
  inference(superposition,[],[f2039,f2039])).
fof(f34763,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X5),s(bool,X6))) | s(bool,X7) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X5))) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X5)))) )),
  inference(superposition,[],[f2039,f783])).
fof(f34762,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X2),s(bool,'T_0'))) | s(bool,X3) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X2))) | s(bool,'T_0') = s(bool,X3) | s(bool,X4) = s(bool,'const.min.@_1'(s(fun(bool,bool),X2))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f2039,f783])).
fof(f34731,plain,(
  ( ! [X118,X116,X117] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X116))) = s(bool,'|dollar|exists_1'(s(fun(X117,bool),X118))) | s(bool,'T_0') = s(bool,happ(s(fun(X117,bool),X118),s(X117,'const.min.@_1'(s(fun(X117,bool),X118))))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X116),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X116)))) )),
  inference(superposition,[],[f2039,f1768])).
fof(f35269,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1))))) | s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,'|dollar|exists_1'(s(fun(X2,bool),X3))) | s(bool,'T_0') = s(bool,happ(s(fun(X2,bool),X3),s(X2,'const.min.@_1'(s(fun(X2,bool),X3)))))) )),
  inference(condensation,[],[f35268])).
fof(f35268,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s(bool,'|dollar|exists_1'(s(fun(X20,bool),X21))) = s(bool,'|dollar|exists_1'(s(fun(X22,bool),X23))) | s(bool,'T_0') = s(bool,happ(s(fun(X22,bool),X23),s(X22,'const.min.@_1'(s(fun(X22,bool),X23))))) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,happ(s(fun(X20,bool),X21),s(X20,'const.min.@_1'(s(fun(X20,bool),X21)))))) )),
  inference(duplicate_literal_removal,[],[f34713])).
fof(f34713,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s(bool,'|dollar|exists_1'(s(fun(X20,bool),X21))) = s(bool,'|dollar|exists_1'(s(fun(X22,bool),X23))) | s(bool,'T_0') = s(bool,happ(s(fun(X22,bool),X23),s(X22,'const.min.@_1'(s(fun(X22,bool),X23))))) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,happ(s(fun(X20,bool),X21),s(X20,'const.min.@_1'(s(fun(X20,bool),X21))))) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f2039,f2039])).
fof(f34673,plain,(
  ( ! [X47,X45,X43,X48,X46,X44] : (~sP12(happ(s(fun(bool,bool),X43),s(bool,'T_0')),X44,X45) | p(s(bool,X44)) | ~p(s(bool,X45)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X43))) | s(bool,X46) = s(bool,X47) | s(bool,X47) = s(bool,'const.min.@_1'(s(fun(bool,bool),X43))) | s(bool,X46) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f4988,f2054])).
fof(f34672,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (~sP12(happ(s(fun(bool,bool),X37),s(bool,X38)),X39,X40) | p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X37))) | s(bool,X38) = s(bool,X41) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X37))) | s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f4988,f2054])).
fof(f34671,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (~sP12(happ(s(fun(bool,bool),X31),s(bool,X32)),X33,X34) | p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X31))) | s(bool,X35) = s(bool,'const.min.@_1'(s(fun(bool,bool),X31))) | s(bool,'T_0') = s(bool,X32) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f4988,f2054])).
fof(f34669,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (~sP12(happ(s(fun(bool,bool),X19),s(bool,X20)),X21,X22) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X19))) | s(bool,X20) = s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,X24) = s(bool,'const.min.@_1'(s(fun(bool,bool),X19))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f4988,f2054])).
fof(f34666,plain,(
  ( ! [X10,X8,X7,X9] : (~sP12(happ(s(fun(bool,bool),X7),s(bool,X8)),X9,X10) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X7))) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X7)))) )),
  inference(superposition,[],[f4988,f783])).
fof(f34665,plain,(
  ( ! [X6,X4,X5,X3] : (~sP12(happ(s(fun(bool,bool),X3),s(bool,'T_0')),X4,X5) | p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) | s(bool,X6) = s(bool,'const.min.@_1'(s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f4988,f783])).
fof(f34664,plain,(
  ( ! [X2,X0,X1] : (~sP12(happ(s(fun(bool,bool),X0),s(bool,'F_0')),X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f4988,f270])).
fof(f34587,plain,(
  ( ! [X47,X45,X43,X48,X46,X44] : (~sP11(X44,X45,happ(s(fun(bool,bool),X43),s(bool,'T_0'))) | ~p(s(bool,X44)) | p(s(bool,X45)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X43))) | s(bool,X46) = s(bool,X47) | s(bool,X47) = s(bool,'const.min.@_1'(s(fun(bool,bool),X43))) | s(bool,X46) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f4985,f2054])).
fof(f34586,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (~sP11(X39,X40,happ(s(fun(bool,bool),X37),s(bool,X38))) | ~p(s(bool,X39)) | p(s(bool,X40)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X37))) | s(bool,X38) = s(bool,X41) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X37))) | s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f4985,f2054])).
fof(f34585,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (~sP11(X33,X34,happ(s(fun(bool,bool),X31),s(bool,X32))) | ~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X31))) | s(bool,X35) = s(bool,'const.min.@_1'(s(fun(bool,bool),X31))) | s(bool,'T_0') = s(bool,X32) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f4985,f2054])).
fof(f34583,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (~sP11(X21,X22,happ(s(fun(bool,bool),X19),s(bool,X20))) | ~p(s(bool,X21)) | p(s(bool,X22)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X19))) | s(bool,X20) = s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,X24) = s(bool,'const.min.@_1'(s(fun(bool,bool),X19))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f4985,f2054])).
fof(f34580,plain,(
  ( ! [X10,X8,X7,X9] : (~sP11(X9,X10,happ(s(fun(bool,bool),X7),s(bool,X8))) | ~p(s(bool,X9)) | p(s(bool,X10)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X7))) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X7)))) )),
  inference(superposition,[],[f4985,f783])).
fof(f34579,plain,(
  ( ! [X6,X4,X5,X3] : (~sP11(X4,X5,happ(s(fun(bool,bool),X3),s(bool,'T_0'))) | ~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) | s(bool,X6) = s(bool,'const.min.@_1'(s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f4985,f783])).
fof(f34578,plain,(
  ( ! [X2,X0,X1] : (~sP11(X1,X2,happ(s(fun(bool,bool),X0),s(bool,'F_0'))) | ~p(s(bool,X1)) | p(s(bool,X2)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f4985,f270])).
fof(f34558,plain,(
  ( ! [X47,X45,X43,X48,X46,X44] : (~sP10(happ(s(fun(bool,bool),X43),s(bool,'T_0')),X44,X45) | p(s(bool,X44)) | p(s(bool,X45)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X43))) | s(bool,X46) = s(bool,X47) | s(bool,X47) = s(bool,'const.min.@_1'(s(fun(bool,bool),X43))) | s(bool,X46) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f4984,f2054])).
fof(f34557,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (~sP10(happ(s(fun(bool,bool),X37),s(bool,X38)),X39,X40) | p(s(bool,X39)) | p(s(bool,X40)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X37))) | s(bool,X38) = s(bool,X41) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X37))) | s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f4984,f2054])).
fof(f34556,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (~sP10(happ(s(fun(bool,bool),X31),s(bool,X32)),X33,X34) | p(s(bool,X33)) | p(s(bool,X34)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X31))) | s(bool,X35) = s(bool,'const.min.@_1'(s(fun(bool,bool),X31))) | s(bool,'T_0') = s(bool,X32) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f4984,f2054])).
fof(f34554,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (~sP10(happ(s(fun(bool,bool),X19),s(bool,X20)),X21,X22) | p(s(bool,X21)) | p(s(bool,X22)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X19))) | s(bool,X20) = s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,X24) = s(bool,'const.min.@_1'(s(fun(bool,bool),X19))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f4984,f2054])).
fof(f34551,plain,(
  ( ! [X10,X8,X7,X9] : (~sP10(happ(s(fun(bool,bool),X7),s(bool,X8)),X9,X10) | p(s(bool,X9)) | p(s(bool,X10)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X7))) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X7)))) )),
  inference(superposition,[],[f4984,f783])).
fof(f34550,plain,(
  ( ! [X6,X4,X5,X3] : (~sP10(happ(s(fun(bool,bool),X3),s(bool,'T_0')),X4,X5) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) | s(bool,X6) = s(bool,'const.min.@_1'(s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f4984,f783])).
fof(f34549,plain,(
  ( ! [X2,X0,X1] : (~sP10(happ(s(fun(bool,bool),X0),s(bool,'F_0')),X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f4984,f270])).
fof(f34529,plain,(
  ( ! [X47,X45,X43,X48,X46,X44] : (~sP9(happ(s(fun(bool,bool),X43),s(bool,'T_0')),X44,X45) | p(s(bool,X45)) | p(s(bool,X44)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X43))) | s(bool,X46) = s(bool,X47) | s(bool,X47) = s(bool,'const.min.@_1'(s(fun(bool,bool),X43))) | s(bool,X46) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f4981,f2054])).
fof(f34528,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (~sP9(happ(s(fun(bool,bool),X37),s(bool,X38)),X39,X40) | p(s(bool,X40)) | p(s(bool,X39)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X37))) | s(bool,X38) = s(bool,X41) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X37))) | s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f4981,f2054])).
fof(f34527,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (~sP9(happ(s(fun(bool,bool),X31),s(bool,X32)),X33,X34) | p(s(bool,X34)) | p(s(bool,X33)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X31))) | s(bool,X35) = s(bool,'const.min.@_1'(s(fun(bool,bool),X31))) | s(bool,'T_0') = s(bool,X32) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f4981,f2054])).
fof(f34525,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (~sP9(happ(s(fun(bool,bool),X19),s(bool,X20)),X21,X22) | p(s(bool,X22)) | p(s(bool,X21)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X19))) | s(bool,X20) = s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,X24) = s(bool,'const.min.@_1'(s(fun(bool,bool),X19))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f4981,f2054])).
fof(f34522,plain,(
  ( ! [X10,X8,X7,X9] : (~sP9(happ(s(fun(bool,bool),X7),s(bool,X8)),X9,X10) | p(s(bool,X10)) | p(s(bool,X9)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X7))) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X7)))) )),
  inference(superposition,[],[f4981,f783])).
fof(f34521,plain,(
  ( ! [X6,X4,X5,X3] : (~sP9(happ(s(fun(bool,bool),X3),s(bool,'T_0')),X4,X5) | p(s(bool,X5)) | p(s(bool,X4)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) | s(bool,X6) = s(bool,'const.min.@_1'(s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f4981,f783])).
fof(f34520,plain,(
  ( ! [X2,X0,X1] : (~sP9(happ(s(fun(bool,bool),X0),s(bool,'F_0')),X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f4981,f270])).
fof(f34502,plain,(
  ( ! [X47,X45,X43,X48,X46,X44] : (~sP8(X44,X45,happ(s(fun(bool,bool),X43),s(bool,'T_0'))) | p(s(bool,X44)) | ~p(s(bool,X45)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X43))) | s(bool,X46) = s(bool,X47) | s(bool,X47) = s(bool,'const.min.@_1'(s(fun(bool,bool),X43))) | s(bool,X46) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f4980,f2054])).
fof(f34501,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (~sP8(X39,X40,happ(s(fun(bool,bool),X37),s(bool,X38))) | p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X37))) | s(bool,X38) = s(bool,X41) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X37))) | s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f4980,f2054])).
fof(f34500,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (~sP8(X33,X34,happ(s(fun(bool,bool),X31),s(bool,X32))) | p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X31))) | s(bool,X35) = s(bool,'const.min.@_1'(s(fun(bool,bool),X31))) | s(bool,'T_0') = s(bool,X32) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f4980,f2054])).
fof(f34498,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (~sP8(X21,X22,happ(s(fun(bool,bool),X19),s(bool,X20))) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X19))) | s(bool,X20) = s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,X24) = s(bool,'const.min.@_1'(s(fun(bool,bool),X19))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f4980,f2054])).
fof(f34495,plain,(
  ( ! [X10,X8,X7,X9] : (~sP8(X9,X10,happ(s(fun(bool,bool),X7),s(bool,X8))) | p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X7))) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X7)))) )),
  inference(superposition,[],[f4980,f783])).
fof(f34494,plain,(
  ( ! [X6,X4,X5,X3] : (~sP8(X4,X5,happ(s(fun(bool,bool),X3),s(bool,'T_0'))) | p(s(bool,X4)) | ~p(s(bool,X5)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) | s(bool,X6) = s(bool,'const.min.@_1'(s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f4980,f783])).
fof(f34493,plain,(
  ( ! [X2,X0,X1] : (~sP8(X1,X2,happ(s(fun(bool,bool),X0),s(bool,'F_0'))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f4980,f270])).
fof(f34420,plain,(
  ( ! [X47,X45,X43,X48,X46,X44] : (~sP7(X44,happ(s(fun(bool,bool),X43),s(bool,'T_0')),X45) | p(s(bool,X45)) | ~p(s(bool,X44)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X43))) | s(bool,X46) = s(bool,X47) | s(bool,X47) = s(bool,'const.min.@_1'(s(fun(bool,bool),X43))) | s(bool,X46) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f4979,f2054])).
fof(f34419,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (~sP7(X39,happ(s(fun(bool,bool),X37),s(bool,X38)),X40) | p(s(bool,X40)) | ~p(s(bool,X39)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X37))) | s(bool,X38) = s(bool,X41) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X37))) | s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f4979,f2054])).
fof(f34418,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (~sP7(X33,happ(s(fun(bool,bool),X31),s(bool,X32)),X34) | p(s(bool,X34)) | ~p(s(bool,X33)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X31))) | s(bool,X35) = s(bool,'const.min.@_1'(s(fun(bool,bool),X31))) | s(bool,'T_0') = s(bool,X32) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f4979,f2054])).
fof(f34416,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (~sP7(X21,happ(s(fun(bool,bool),X19),s(bool,X20)),X22) | p(s(bool,X22)) | ~p(s(bool,X21)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X19))) | s(bool,X20) = s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,X24) = s(bool,'const.min.@_1'(s(fun(bool,bool),X19))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f4979,f2054])).
fof(f34413,plain,(
  ( ! [X10,X8,X7,X9] : (~sP7(X9,happ(s(fun(bool,bool),X7),s(bool,X8)),X10) | p(s(bool,X10)) | ~p(s(bool,X9)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X7))) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X7)))) )),
  inference(superposition,[],[f4979,f783])).
fof(f34412,plain,(
  ( ! [X6,X4,X5,X3] : (~sP7(X4,happ(s(fun(bool,bool),X3),s(bool,'T_0')),X5) | p(s(bool,X5)) | ~p(s(bool,X4)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) | s(bool,X6) = s(bool,'const.min.@_1'(s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f4979,f783])).
fof(f34411,plain,(
  ( ! [X2,X0,X1] : (~sP7(X1,happ(s(fun(bool,bool),X0),s(bool,'F_0')),X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f4979,f270])).
fof(f34281,plain,(
  ( ! [X47,X45,X43,X48,X46,X44] : (~sP4(X44,happ(s(fun(bool,bool),X43),s(bool,'T_0')),X45) | p(s(bool,X45)) | ~p(s(bool,X44)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X43))) | s(bool,X46) = s(bool,X47) | s(bool,X47) = s(bool,'const.min.@_1'(s(fun(bool,bool),X43))) | s(bool,X46) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f4978,f2054])).
fof(f34280,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (~sP4(X39,happ(s(fun(bool,bool),X37),s(bool,X38)),X40) | p(s(bool,X40)) | ~p(s(bool,X39)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X37))) | s(bool,X38) = s(bool,X41) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X37))) | s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f4978,f2054])).
fof(f34279,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (~sP4(X33,happ(s(fun(bool,bool),X31),s(bool,X32)),X34) | p(s(bool,X34)) | ~p(s(bool,X33)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X31))) | s(bool,X35) = s(bool,'const.min.@_1'(s(fun(bool,bool),X31))) | s(bool,'T_0') = s(bool,X32) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f4978,f2054])).
fof(f34277,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (~sP4(X21,happ(s(fun(bool,bool),X19),s(bool,X20)),X22) | p(s(bool,X22)) | ~p(s(bool,X21)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X19))) | s(bool,X20) = s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,X24) = s(bool,'const.min.@_1'(s(fun(bool,bool),X19))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f4978,f2054])).
fof(f34274,plain,(
  ( ! [X10,X8,X7,X9] : (~sP4(X9,happ(s(fun(bool,bool),X7),s(bool,X8)),X10) | p(s(bool,X10)) | ~p(s(bool,X9)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X7))) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X7)))) )),
  inference(superposition,[],[f4978,f783])).
fof(f34273,plain,(
  ( ! [X6,X4,X5,X3] : (~sP4(X4,happ(s(fun(bool,bool),X3),s(bool,'T_0')),X5) | p(s(bool,X5)) | ~p(s(bool,X4)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) | s(bool,X6) = s(bool,'const.min.@_1'(s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f4978,f783])).
fof(f34272,plain,(
  ( ! [X2,X0,X1] : (~sP4(X1,happ(s(fun(bool,bool),X0),s(bool,'F_0')),X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f4978,f270])).
fof(f34219,plain,(
  ( ! [X47,X45,X43,X48,X46,X44] : (~sP5(X44,happ(s(fun(bool,bool),X43),s(bool,'T_0')),X45) | p(s(bool,X45)) | ~p(s(bool,X44)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X43))) | s(bool,X46) = s(bool,X47) | s(bool,X47) = s(bool,'const.min.@_1'(s(fun(bool,bool),X43))) | s(bool,X46) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f4977,f2054])).
fof(f34218,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (~sP5(X39,happ(s(fun(bool,bool),X37),s(bool,X38)),X40) | p(s(bool,X40)) | ~p(s(bool,X39)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X37))) | s(bool,X38) = s(bool,X41) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X37))) | s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f4977,f2054])).
fof(f34217,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (~sP5(X33,happ(s(fun(bool,bool),X31),s(bool,X32)),X34) | p(s(bool,X34)) | ~p(s(bool,X33)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X31))) | s(bool,X35) = s(bool,'const.min.@_1'(s(fun(bool,bool),X31))) | s(bool,'T_0') = s(bool,X32) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f4977,f2054])).
fof(f34215,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (~sP5(X21,happ(s(fun(bool,bool),X19),s(bool,X20)),X22) | p(s(bool,X22)) | ~p(s(bool,X21)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X19))) | s(bool,X20) = s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,X24) = s(bool,'const.min.@_1'(s(fun(bool,bool),X19))) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f4977,f2054])).
fof(f34212,plain,(
  ( ! [X10,X8,X7,X9] : (~sP5(X9,happ(s(fun(bool,bool),X7),s(bool,X8)),X10) | p(s(bool,X10)) | ~p(s(bool,X9)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X7))) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X7)))) )),
  inference(superposition,[],[f4977,f783])).
fof(f34211,plain,(
  ( ! [X6,X4,X5,X3] : (~sP5(X4,happ(s(fun(bool,bool),X3),s(bool,'T_0')),X5) | p(s(bool,X5)) | ~p(s(bool,X4)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) | s(bool,X6) = s(bool,'const.min.@_1'(s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f4977,f783])).
fof(f34210,plain,(
  ( ! [X2,X0,X1] : (~sP5(X1,happ(s(fun(bool,bool),X0),s(bool,'F_0')),X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f4977,f270])).
fof(f34045,plain,(
  ( ! [X8,X7,X9] : (s(bool,X9) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK22(X8,X7)),s('type.extreal.extreal',X8))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X7))),s('type.extreal.extreal',X8)))) )),
  inference(subsumption_resolution,[],[f34030,f468])).
fof(f34030,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,'F_0')) | s(bool,X9) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK22(X8,X7)),s('type.extreal.extreal',X8))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X7))),s('type.extreal.extreal',X8)))) )),
  inference(superposition,[],[f2636,f270])).
fof(f33989,plain,(
  ( ! [X6,X8,X7] : (s(bool,X8) = s(bool,happ(s(fun('type.extreal.extreal',bool),X7),s('type.extreal.extreal',X6))) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X7)))))) )),
  inference(subsumption_resolution,[],[f33974,f468])).
fof(f33974,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,'F_0')) | s(bool,X8) = s(bool,happ(s(fun('type.extreal.extreal',bool),X7),s('type.extreal.extreal',X6))) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X7)))))) )),
  inference(superposition,[],[f2626,f270])).
fof(f33951,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0)))) )),
  inference(subsumption_resolution,[],[f33936,f468])).
fof(f33936,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f4872,f270])).
fof(f33923,plain,(
  ( ! [X28,X26,X27,X25] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X25),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X25))) | s(bool,X26) = s(bool,X27) | s(bool,X27) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X25))) | s(bool,X26) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f33910,f263])).
fof(f33910,plain,(
  ( ! [X28,X26,X27,X25] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X25),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X25))) | s(bool,X26) = s(bool,X27) | s(bool,X27) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X25))) | s(bool,X26) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f4871,f2054])).
fof(f33916,plain,(
  ( ! [X2,X1] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X1))) | s(bool,X2) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f33902,f263])).
fof(f33902,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X1))) | s(bool,X2) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f4871,f783])).
fof(f33850,plain,(
  ( ! [X28,X26,X27,X25] : (p(s(bool,happ(s(fun(bool,bool),X25),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X25))) | s(bool,X26) = s(bool,X27) | s(bool,X27) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X25))) | s(bool,X26) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(trivial_inequality_removal,[],[f33843])).
fof(f33843,plain,(
  ( ! [X28,X26,X27,X25] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,happ(s(fun(bool,bool),X25),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X25))) | s(bool,X26) = s(bool,X27) | s(bool,X27) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X25))) | s(bool,X26) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f4869,f2054])).
fof(f33852,plain,(
  ( ! [X2,X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X1))) | s(bool,X2) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(trivial_inequality_removal,[],[f33835])).
fof(f33835,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X1))) | s(bool,X2) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f4869,f783])).
fof(f33623,plain,(
  ( ! [X80,X81] : (sP6('T_0','|dollar|exists_1'(s(fun(bool,bool),X81)),happ(s(fun(bool,bool),X81),s(bool,'|dollar|exists_1'(s(fun(bool,bool),X80))))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X81))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X80),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X80)))) )),
  inference(superposition,[],[f3867,f1768])).
fof(f33469,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP6('T_0','|dollar|exists_1'(s(fun(bool,bool),X1)),happ(s(fun(bool,bool),X1),s(bool,X0))) | s(bool,X2) = s(bool,'const.min.@_1'(s(fun(bool,bool),X1))) | s(bool,'T_0') = s(bool,X2) | s(bool,X3) = s(bool,X4) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X5) | s(bool,X0) = s(bool,X3)) )),
  inference(superposition,[],[f3865,f2054])).
fof(f33335,plain,(
  ( ! [X80,X81] : (~sP6('F_0','|dollar|exists_1'(s(fun(bool,bool),X81)),happ(s(fun(bool,bool),X81),s(bool,'|dollar|exists_1'(s(fun(bool,bool),X80))))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X81))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X80),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X80)))) )),
  inference(superposition,[],[f3851,f1768])).
fof(f33183,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP6('F_0','|dollar|exists_1'(s(fun(bool,bool),X1)),happ(s(fun(bool,bool),X1),s(bool,X0))) | s(bool,X2) = s(bool,'const.min.@_1'(s(fun(bool,bool),X1))) | s(bool,'T_0') = s(bool,X2) | s(bool,X3) = s(bool,X4) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X5) | s(bool,X0) = s(bool,X3)) )),
  inference(superposition,[],[f3849,f2054])).
fof(f33132,plain,(
  ( ! [X80,X81] : (sP6('T_0',happ(s(fun(bool,bool),X81),s(bool,'|dollar|exists_1'(s(fun(bool,bool),X80)))),'|dollar|exists_1'(s(fun(bool,bool),X81))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X81))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X80),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X80)))) )),
  inference(superposition,[],[f3762,f1768])).
fof(f33066,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP6('T_0',happ(s(fun(bool,bool),X1),s(bool,X0)),'|dollar|exists_1'(s(fun(bool,bool),X1))) | s(bool,X2) = s(bool,'const.min.@_1'(s(fun(bool,bool),X1))) | s(bool,'T_0') = s(bool,X2) | s(bool,X3) = s(bool,X4) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X5) | s(bool,X0) = s(bool,X3)) )),
  inference(superposition,[],[f3760,f2054])).
fof(f33004,plain,(
  ( ! [X80,X81] : (~sP6('F_0',happ(s(fun(bool,bool),X81),s(bool,'|dollar|exists_1'(s(fun(bool,bool),X80)))),'|dollar|exists_1'(s(fun(bool,bool),X81))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X81))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X80),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X80)))) )),
  inference(superposition,[],[f3742,f1768])).
fof(f32932,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP6('F_0',happ(s(fun(bool,bool),X1),s(bool,X0)),'|dollar|exists_1'(s(fun(bool,bool),X1))) | s(bool,X2) = s(bool,'const.min.@_1'(s(fun(bool,bool),X1))) | s(bool,'T_0') = s(bool,X2) | s(bool,X3) = s(bool,X4) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X5) | s(bool,X0) = s(bool,X3)) )),
  inference(superposition,[],[f3740,f2054])).
fof(f32865,plain,(
  ( ! [X4,X2,X5,X3] : (s(bool,X4) = s(bool,happ(s(fun(X2,bool),X3),s(X2,X5))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,happ(s(fun(X2,bool),X3),s(X2,sK16(X3,X2))))) )),
  inference(subsumption_resolution,[],[f32840,f468])).
fof(f32840,plain,(
  ( ! [X4,X2,X5,X3] : (p(s(bool,'F_0')) | s(bool,X4) = s(bool,happ(s(fun(X2,bool),X3),s(X2,X5))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,happ(s(fun(X2,bool),X3),s(X2,sK16(X3,X2))))) )),
  inference(superposition,[],[f2628,f270])).
fof(f32833,plain,(
  ( ! [X47,X45,X43,X48,X46,X44] : (p(s(bool,happ(s(fun(bool,bool),X43),s(bool,'T_0')))) | s(bool,X44) = s(bool,happ(s(fun(bool,bool),X43),s(bool,X45))) | s(bool,'T_0') = s(bool,X44) | s(bool,X46) = s(bool,X47) | s(bool,X47) = s(bool,sK16(X43,bool)) | s(bool,X46) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f2628,f2054])).
fof(f32832,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (p(s(bool,happ(s(fun(bool,bool),X37),s(bool,X38)))) | s(bool,X39) = s(bool,happ(s(fun(bool,bool),X37),s(bool,X40))) | s(bool,'T_0') = s(bool,X39) | s(bool,X38) = s(bool,X41) | s(bool,'T_0') = s(bool,sK16(X37,bool)) | s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f2628,f2054])).
fof(f32831,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (p(s(bool,happ(s(fun(bool,bool),X31),s(bool,X32)))) | s(bool,X33) = s(bool,happ(s(fun(bool,bool),X31),s(bool,X34))) | s(bool,'T_0') = s(bool,X33) | s(bool,X35) = s(bool,sK16(X31,bool)) | s(bool,'T_0') = s(bool,X32) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f2628,f2054])).
fof(f32829,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (p(s(bool,happ(s(fun(bool,bool),X19),s(bool,X20)))) | s(bool,X21) = s(bool,happ(s(fun(bool,bool),X19),s(bool,X22))) | s(bool,'T_0') = s(bool,X21) | s(bool,X20) = s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,X24) = s(bool,sK16(X19,bool)) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f2628,f2054])).
fof(f32826,plain,(
  ( ! [X10,X8,X7,X9] : (p(s(bool,happ(s(fun(bool,bool),X7),s(bool,X8)))) | s(bool,X9) = s(bool,happ(s(fun(bool,bool),X7),s(bool,X10))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,sK16(X7,bool))) )),
  inference(superposition,[],[f2628,f783])).
fof(f32825,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,happ(s(fun(bool,bool),X3),s(bool,'T_0')))) | s(bool,X4) = s(bool,happ(s(fun(bool,bool),X3),s(bool,X5))) | s(bool,'T_0') = s(bool,X4) | s(bool,X6) = s(bool,sK16(X3,bool)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f2628,f783])).
fof(f32824,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | s(bool,X1) = s(bool,happ(s(fun(bool,bool),X0),s(bool,X2))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,sK16(X0,bool))) )),
  inference(superposition,[],[f2628,f270])).
fof(f32732,plain,(
  ( ! [X4,X2,X5,X3] : (s(bool,X4) = s(bool,happ(s(fun(X2,bool),X3),s(X2,X5))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,happ(s(fun(X2,bool),X3),s(X2,sK17(X3,X2))))) )),
  inference(subsumption_resolution,[],[f32707,f468])).
fof(f32707,plain,(
  ( ! [X4,X2,X5,X3] : (p(s(bool,'F_0')) | s(bool,X4) = s(bool,happ(s(fun(X2,bool),X3),s(X2,X5))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,happ(s(fun(X2,bool),X3),s(X2,sK17(X3,X2))))) )),
  inference(superposition,[],[f2627,f270])).
fof(f32700,plain,(
  ( ! [X47,X45,X43,X48,X46,X44] : (p(s(bool,happ(s(fun(bool,bool),X43),s(bool,'T_0')))) | s(bool,X44) = s(bool,happ(s(fun(bool,bool),X43),s(bool,X45))) | s(bool,'T_0') = s(bool,X44) | s(bool,X46) = s(bool,X47) | s(bool,X47) = s(bool,sK17(X43,bool)) | s(bool,X46) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f2627,f2054])).
fof(f32699,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (p(s(bool,happ(s(fun(bool,bool),X37),s(bool,X38)))) | s(bool,X39) = s(bool,happ(s(fun(bool,bool),X37),s(bool,X40))) | s(bool,'T_0') = s(bool,X39) | s(bool,X38) = s(bool,X41) | s(bool,'T_0') = s(bool,sK17(X37,bool)) | s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f2627,f2054])).
fof(f32698,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (p(s(bool,happ(s(fun(bool,bool),X31),s(bool,X32)))) | s(bool,X33) = s(bool,happ(s(fun(bool,bool),X31),s(bool,X34))) | s(bool,'T_0') = s(bool,X33) | s(bool,X35) = s(bool,sK17(X31,bool)) | s(bool,'T_0') = s(bool,X32) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f2627,f2054])).
fof(f32696,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (p(s(bool,happ(s(fun(bool,bool),X19),s(bool,X20)))) | s(bool,X21) = s(bool,happ(s(fun(bool,bool),X19),s(bool,X22))) | s(bool,'T_0') = s(bool,X21) | s(bool,X20) = s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,X24) = s(bool,sK17(X19,bool)) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f2627,f2054])).
fof(f32693,plain,(
  ( ! [X10,X8,X7,X9] : (p(s(bool,happ(s(fun(bool,bool),X7),s(bool,X8)))) | s(bool,X9) = s(bool,happ(s(fun(bool,bool),X7),s(bool,X10))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,sK17(X7,bool))) )),
  inference(superposition,[],[f2627,f783])).
fof(f32692,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,happ(s(fun(bool,bool),X3),s(bool,'T_0')))) | s(bool,X4) = s(bool,happ(s(fun(bool,bool),X3),s(bool,X5))) | s(bool,'T_0') = s(bool,X4) | s(bool,X6) = s(bool,sK17(X3,bool)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f2627,f783])).
fof(f32691,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | s(bool,X1) = s(bool,happ(s(fun(bool,bool),X0),s(bool,X2))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,sK17(X0,bool))) )),
  inference(superposition,[],[f2627,f270])).
fof(f32529,plain,(
  ( ! [X198,X200,X199,X201] : (s(bool,'|dollar|exists_1'(s(fun(X200,bool),X201))) = s(bool,happ(s(fun('type.extreal.extreal',bool),X198),s('type.extreal.extreal',sK22(X199,X198)))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X200,bool),X201))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X198))),s('type.extreal.extreal',X199)))) )),
  inference(subsumption_resolution,[],[f32391,f263])).
fof(f32391,plain,(
  ( ! [X198,X200,X199,X201] : (~p(s(bool,'T_0')) | s(bool,'|dollar|exists_1'(s(fun(X200,bool),X201))) = s(bool,happ(s(fun('type.extreal.extreal',bool),X198),s('type.extreal.extreal',sK22(X199,X198)))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X200,bool),X201))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X198))),s('type.extreal.extreal',X199)))) )),
  inference(superposition,[],[f4971,f5007])).
fof(f32528,plain,(
  ( ! [X196,X194,X197,X195,X193] : (p(s(bool,X195)) | ~sP12(X196,X195,happ(s(fun('type.extreal.extreal',bool),X193),s('type.extreal.extreal',sK22(X194,X193)))) | s(bool,X196) = s(bool,X197) | s(bool,'T_0') = s(bool,X197) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X193))),s('type.extreal.extreal',X194)))) )),
  inference(subsumption_resolution,[],[f32390,f263])).
fof(f32390,plain,(
  ( ! [X196,X194,X197,X195,X193] : (~p(s(bool,'T_0')) | p(s(bool,X195)) | ~sP12(X196,X195,happ(s(fun('type.extreal.extreal',bool),X193),s('type.extreal.extreal',sK22(X194,X193)))) | s(bool,X196) = s(bool,X197) | s(bool,'T_0') = s(bool,X197) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X193))),s('type.extreal.extreal',X194)))) )),
  inference(superposition,[],[f2623,f5007])).
fof(f32527,plain,(
  ( ! [X191,X189,X192,X190,X188] : (sP12(X190,happ(s(fun('type.extreal.extreal',bool),X188),s('type.extreal.extreal',sK22(X189,X188))),X191) | s(bool,X190) = s(bool,X192) | s(bool,'T_0') = s(bool,X192) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X188))),s('type.extreal.extreal',X189)))) )),
  inference(subsumption_resolution,[],[f32389,f263])).
fof(f32389,plain,(
  ( ! [X191,X189,X192,X190,X188] : (~p(s(bool,'T_0')) | sP12(X190,happ(s(fun('type.extreal.extreal',bool),X188),s('type.extreal.extreal',sK22(X189,X188))),X191) | s(bool,X190) = s(bool,X192) | s(bool,'T_0') = s(bool,X192) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X188))),s('type.extreal.extreal',X189)))) )),
  inference(superposition,[],[f2622,f5007])).
fof(f32526,plain,(
  ( ! [X187,X185,X186,X184,X183] : (sP10(X185,happ(s(fun('type.extreal.extreal',bool),X183),s('type.extreal.extreal',sK22(X184,X183))),X186) | s(bool,X185) = s(bool,X187) | s(bool,'T_0') = s(bool,X187) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X183))),s('type.extreal.extreal',X184)))) )),
  inference(subsumption_resolution,[],[f32388,f263])).
fof(f32388,plain,(
  ( ! [X187,X185,X186,X184,X183] : (~p(s(bool,'T_0')) | sP10(X185,happ(s(fun('type.extreal.extreal',bool),X183),s('type.extreal.extreal',sK22(X184,X183))),X186) | s(bool,X185) = s(bool,X187) | s(bool,'T_0') = s(bool,X187) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X183))),s('type.extreal.extreal',X184)))) )),
  inference(superposition,[],[f2618,f5007])).
fof(f32525,plain,(
  ( ! [X182,X180,X178,X181,X179] : (sP10(X180,X181,happ(s(fun('type.extreal.extreal',bool),X178),s('type.extreal.extreal',sK22(X179,X178)))) | s(bool,X180) = s(bool,X182) | s(bool,'T_0') = s(bool,X182) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X178))),s('type.extreal.extreal',X179)))) )),
  inference(subsumption_resolution,[],[f32387,f263])).
fof(f32387,plain,(
  ( ! [X182,X180,X178,X181,X179] : (~p(s(bool,'T_0')) | sP10(X180,X181,happ(s(fun('type.extreal.extreal',bool),X178),s('type.extreal.extreal',sK22(X179,X178)))) | s(bool,X180) = s(bool,X182) | s(bool,'T_0') = s(bool,X182) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X178))),s('type.extreal.extreal',X179)))) )),
  inference(superposition,[],[f2617,f5007])).
fof(f32524,plain,(
  ( ! [X177,X175,X173,X176,X174] : (p(s(bool,X175)) | ~sP8(X175,happ(s(fun('type.extreal.extreal',bool),X173),s('type.extreal.extreal',sK22(X174,X173))),X176) | s(bool,X176) = s(bool,X177) | s(bool,'T_0') = s(bool,X177) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X173))),s('type.extreal.extreal',X174)))) )),
  inference(subsumption_resolution,[],[f32386,f263])).
fof(f32386,plain,(
  ( ! [X177,X175,X173,X176,X174] : (~p(s(bool,'T_0')) | p(s(bool,X175)) | ~sP8(X175,happ(s(fun('type.extreal.extreal',bool),X173),s('type.extreal.extreal',sK22(X174,X173))),X176) | s(bool,X176) = s(bool,X177) | s(bool,'T_0') = s(bool,X177) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X173))),s('type.extreal.extreal',X174)))) )),
  inference(superposition,[],[f2615,f5007])).
fof(f32523,plain,(
  ( ! [X101,X99,X102,X100] : (~sP9(happ(s(fun('type.extreal.extreal',bool),X99),s('type.extreal.extreal',sK22(X100,X99))),X101,X102) | sP10(happ(s(fun('type.extreal.extreal',bool),X99),s('type.extreal.extreal',sK22(X100,X99))),X101,X102) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X99))),s('type.extreal.extreal',X100)))) )),
  inference(subsumption_resolution,[],[f32369,f263])).
fof(f32369,plain,(
  ( ! [X101,X99,X102,X100] : (~p(s(bool,'T_0')) | ~sP9(happ(s(fun('type.extreal.extreal',bool),X99),s('type.extreal.extreal',sK22(X100,X99))),X101,X102) | sP10(happ(s(fun('type.extreal.extreal',bool),X99),s('type.extreal.extreal',sK22(X100,X99))),X101,X102) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X99))),s('type.extreal.extreal',X100)))) )),
  inference(superposition,[],[f453,f5007])).
fof(f32522,plain,(
  ( ! [X97,X95,X98,X96] : (~sP11(happ(s(fun('type.extreal.extreal',bool),X95),s('type.extreal.extreal',sK22(X96,X95))),X97,X98) | sP12(happ(s(fun('type.extreal.extreal',bool),X95),s('type.extreal.extreal',sK22(X96,X95))),X97,X98) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X95))),s('type.extreal.extreal',X96)))) )),
  inference(subsumption_resolution,[],[f32368,f263])).
fof(f32368,plain,(
  ( ! [X97,X95,X98,X96] : (~p(s(bool,'T_0')) | ~sP11(happ(s(fun('type.extreal.extreal',bool),X95),s('type.extreal.extreal',sK22(X96,X95))),X97,X98) | sP12(happ(s(fun('type.extreal.extreal',bool),X95),s('type.extreal.extreal',sK22(X96,X95))),X97,X98) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X95))),s('type.extreal.extreal',X96)))) )),
  inference(superposition,[],[f452,f5007])).
fof(f32521,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,X93)) | p(s(bool,X94)) | ~sP12(happ(s(fun('type.extreal.extreal',bool),X91),s('type.extreal.extreal',sK22(X92,X91))),X94,X93) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X91))),s('type.extreal.extreal',X92)))) )),
  inference(subsumption_resolution,[],[f32367,f263])).
fof(f32367,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,'T_0')) | ~p(s(bool,X93)) | p(s(bool,X94)) | ~sP12(happ(s(fun('type.extreal.extreal',bool),X91),s('type.extreal.extreal',sK22(X92,X91))),X94,X93) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X91))),s('type.extreal.extreal',X92)))) )),
  inference(superposition,[],[f449,f5007])).
fof(f32520,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,X89)) | sP12(happ(s(fun('type.extreal.extreal',bool),X87),s('type.extreal.extreal',sK22(X88,X87))),X89,X90) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X87))),s('type.extreal.extreal',X88)))) )),
  inference(subsumption_resolution,[],[f32366,f263])).
fof(f32366,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,'T_0')) | ~p(s(bool,X89)) | sP12(happ(s(fun('type.extreal.extreal',bool),X87),s('type.extreal.extreal',sK22(X88,X87))),X89,X90) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X87))),s('type.extreal.extreal',X88)))) )),
  inference(superposition,[],[f448,f5007])).
fof(f32519,plain,(
  ( ! [X85,X83,X86,X84] : (p(s(bool,X85)) | sP12(happ(s(fun('type.extreal.extreal',bool),X83),s('type.extreal.extreal',sK22(X84,X83))),X86,X85) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X83))),s('type.extreal.extreal',X84)))) )),
  inference(subsumption_resolution,[],[f32365,f263])).
fof(f32365,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,'T_0')) | p(s(bool,X85)) | sP12(happ(s(fun('type.extreal.extreal',bool),X83),s('type.extreal.extreal',sK22(X84,X83))),X86,X85) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X83))),s('type.extreal.extreal',X84)))) )),
  inference(superposition,[],[f447,f5007])).
fof(f32518,plain,(
  ( ! [X80,X81,X79,X82] : (p(s(bool,X81)) | ~p(s(bool,X82)) | ~sP11(X82,X81,happ(s(fun('type.extreal.extreal',bool),X79),s('type.extreal.extreal',sK22(X80,X79)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X79))),s('type.extreal.extreal',X80)))) )),
  inference(subsumption_resolution,[],[f32364,f263])).
fof(f32364,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,'T_0')) | p(s(bool,X81)) | ~p(s(bool,X82)) | ~sP11(X82,X81,happ(s(fun('type.extreal.extreal',bool),X79),s('type.extreal.extreal',sK22(X80,X79)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X79))),s('type.extreal.extreal',X80)))) )),
  inference(superposition,[],[f442,f5007])).
fof(f32517,plain,(
  ( ! [X78,X76,X77,X75] : (p(s(bool,X77)) | p(s(bool,X78)) | ~sP10(happ(s(fun('type.extreal.extreal',bool),X75),s('type.extreal.extreal',sK22(X76,X75))),X78,X77) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X75))),s('type.extreal.extreal',X76)))) )),
  inference(subsumption_resolution,[],[f32363,f263])).
fof(f32363,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,'T_0')) | p(s(bool,X77)) | p(s(bool,X78)) | ~sP10(happ(s(fun('type.extreal.extreal',bool),X75),s('type.extreal.extreal',sK22(X76,X75))),X78,X77) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X75))),s('type.extreal.extreal',X76)))) )),
  inference(superposition,[],[f432,f5007])).
fof(f32516,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X73)) | sP10(happ(s(fun('type.extreal.extreal',bool),X71),s('type.extreal.extreal',sK22(X72,X71))),X73,X74) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X71))),s('type.extreal.extreal',X72)))) )),
  inference(subsumption_resolution,[],[f32362,f263])).
fof(f32362,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,'T_0')) | ~p(s(bool,X73)) | sP10(happ(s(fun('type.extreal.extreal',bool),X71),s('type.extreal.extreal',sK22(X72,X71))),X73,X74) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X71))),s('type.extreal.extreal',X72)))) )),
  inference(superposition,[],[f431,f5007])).
fof(f32515,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X69)) | sP10(happ(s(fun('type.extreal.extreal',bool),X67),s('type.extreal.extreal',sK22(X68,X67))),X70,X69) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X67))),s('type.extreal.extreal',X68)))) )),
  inference(subsumption_resolution,[],[f32361,f263])).
fof(f32361,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | ~p(s(bool,X69)) | sP10(happ(s(fun('type.extreal.extreal',bool),X67),s('type.extreal.extreal',sK22(X68,X67))),X70,X69) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X67))),s('type.extreal.extreal',X68)))) )),
  inference(superposition,[],[f430,f5007])).
fof(f32514,plain,(
  ( ! [X66,X64,X65,X63] : (p(s(bool,X65)) | p(s(bool,X66)) | ~sP9(happ(s(fun('type.extreal.extreal',bool),X63),s('type.extreal.extreal',sK22(X64,X63))),X65,X66) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X63))),s('type.extreal.extreal',X64)))) )),
  inference(subsumption_resolution,[],[f32360,f263])).
fof(f32360,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'T_0')) | p(s(bool,X65)) | p(s(bool,X66)) | ~sP9(happ(s(fun('type.extreal.extreal',bool),X63),s('type.extreal.extreal',sK22(X64,X63))),X65,X66) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X63))),s('type.extreal.extreal',X64)))) )),
  inference(superposition,[],[f425,f5007])).
fof(f32513,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X61)) | p(s(bool,X62)) | ~sP8(X62,X61,happ(s(fun('type.extreal.extreal',bool),X59),s('type.extreal.extreal',sK22(X60,X59)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X59))),s('type.extreal.extreal',X60)))) )),
  inference(subsumption_resolution,[],[f32359,f263])).
fof(f32359,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | ~p(s(bool,X61)) | p(s(bool,X62)) | ~sP8(X62,X61,happ(s(fun('type.extreal.extreal',bool),X59),s('type.extreal.extreal',sK22(X60,X59)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X59))),s('type.extreal.extreal',X60)))) )),
  inference(superposition,[],[f417,f5007])).
fof(f32512,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X57)) | p(s(bool,X58)) | ~sP7(X57,happ(s(fun('type.extreal.extreal',bool),X55),s('type.extreal.extreal',sK22(X56,X55))),X58) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X55))),s('type.extreal.extreal',X56)))) )),
  inference(subsumption_resolution,[],[f32358,f263])).
fof(f32358,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X57)) | p(s(bool,X58)) | ~sP7(X57,happ(s(fun('type.extreal.extreal',bool),X55),s('type.extreal.extreal',sK22(X56,X55))),X58) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X55))),s('type.extreal.extreal',X56)))) )),
  inference(superposition,[],[f408,f5007])).
fof(f32511,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X53)) | p(s(bool,X54)) | ~sP4(X53,happ(s(fun('type.extreal.extreal',bool),X51),s('type.extreal.extreal',sK22(X52,X51))),X54) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X51))),s('type.extreal.extreal',X52)))) )),
  inference(subsumption_resolution,[],[f32357,f263])).
fof(f32357,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X53)) | p(s(bool,X54)) | ~sP4(X53,happ(s(fun('type.extreal.extreal',bool),X51),s('type.extreal.extreal',sK22(X52,X51))),X54) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X51))),s('type.extreal.extreal',X52)))) )),
  inference(superposition,[],[f389,f5007])).
fof(f32510,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X49)) | p(s(bool,X50)) | ~sP5(X49,happ(s(fun('type.extreal.extreal',bool),X47),s('type.extreal.extreal',sK22(X48,X47))),X50) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X47))),s('type.extreal.extreal',X48)))) )),
  inference(subsumption_resolution,[],[f32356,f263])).
fof(f32356,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X49)) | p(s(bool,X50)) | ~sP5(X49,happ(s(fun('type.extreal.extreal',bool),X47),s('type.extreal.extreal',sK22(X48,X47))),X50) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X47))),s('type.extreal.extreal',X48)))) )),
  inference(superposition,[],[f385,f5007])).
fof(f32509,plain,(
  ( ! [X45,X43,X46,X44] : (sP6(X45,X46,happ(s(fun('type.extreal.extreal',bool),X43),s('type.extreal.extreal',sK22(X44,X43)))) | ~sP5(X45,happ(s(fun('type.extreal.extreal',bool),X43),s('type.extreal.extreal',sK22(X44,X43))),X46) | ~sP4(X46,happ(s(fun('type.extreal.extreal',bool),X43),s('type.extreal.extreal',sK22(X44,X43))),X45) | ~sP3(happ(s(fun('type.extreal.extreal',bool),X43),s('type.extreal.extreal',sK22(X44,X43))),X46,X45) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X43))),s('type.extreal.extreal',X44)))) )),
  inference(subsumption_resolution,[],[f32355,f263])).
fof(f32355,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | sP6(X45,X46,happ(s(fun('type.extreal.extreal',bool),X43),s('type.extreal.extreal',sK22(X44,X43)))) | ~sP5(X45,happ(s(fun('type.extreal.extreal',bool),X43),s('type.extreal.extreal',sK22(X44,X43))),X46) | ~sP4(X46,happ(s(fun('type.extreal.extreal',bool),X43),s('type.extreal.extreal',sK22(X44,X43))),X45) | ~sP3(happ(s(fun('type.extreal.extreal',bool),X43),s('type.extreal.extreal',sK22(X44,X43))),X46,X45) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X43))),s('type.extreal.extreal',X44)))) )),
  inference(superposition,[],[f382,f5007])).
fof(f32508,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,X42)) | sP1(X42,happ(s(fun('type.extreal.extreal',bool),X40),s('type.extreal.extreal',sK22(X41,X40)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X40))),s('type.extreal.extreal',X41)))) )),
  inference(subsumption_resolution,[],[f32354,f263])).
fof(f32354,plain,(
  ( ! [X41,X42,X40] : (~p(s(bool,'T_0')) | p(s(bool,X42)) | sP1(X42,happ(s(fun('type.extreal.extreal',bool),X40),s('type.extreal.extreal',sK22(X41,X40)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X40))),s('type.extreal.extreal',X41)))) )),
  inference(superposition,[],[f353,f5007])).
fof(f32507,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,X39)) | sP1(happ(s(fun('type.extreal.extreal',bool),X37),s('type.extreal.extreal',sK22(X38,X37))),X39) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X37))),s('type.extreal.extreal',X38)))) )),
  inference(subsumption_resolution,[],[f32353,f263])).
fof(f32353,plain,(
  ( ! [X39,X37,X38] : (~p(s(bool,'T_0')) | p(s(bool,X39)) | sP1(happ(s(fun('type.extreal.extreal',bool),X37),s('type.extreal.extreal',sK22(X38,X37))),X39) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X37))),s('type.extreal.extreal',X38)))) )),
  inference(superposition,[],[f350,f5007])).
fof(f32506,plain,(
  ( ! [X35,X36,X34] : (sP0(happ(s(fun('type.extreal.extreal',bool),X34),s('type.extreal.extreal',sK22(X35,X34))),X36) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X34))),s('type.extreal.extreal',X35)))) )),
  inference(subsumption_resolution,[],[f32352,f263])).
fof(f32352,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,'T_0')) | sP0(happ(s(fun('type.extreal.extreal',bool),X34),s('type.extreal.extreal',sK22(X35,X34))),X36) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X34))),s('type.extreal.extreal',X35)))) )),
  inference(superposition,[],[f335,f5007])).
fof(f32505,plain,(
  ( ! [X33,X31,X32] : (sP0(X33,happ(s(fun('type.extreal.extreal',bool),X31),s('type.extreal.extreal',sK22(X32,X31)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X31))),s('type.extreal.extreal',X32)))) )),
  inference(subsumption_resolution,[],[f32351,f263])).
fof(f32351,plain,(
  ( ! [X33,X31,X32] : (~p(s(bool,'T_0')) | sP0(X33,happ(s(fun('type.extreal.extreal',bool),X31),s('type.extreal.extreal',sK22(X32,X31)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X31))),s('type.extreal.extreal',X32)))) )),
  inference(superposition,[],[f334,f5007])).
fof(f32504,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,X30)) | s(bool,X30) = s(bool,happ(s(fun('type.extreal.extreal',bool),X28),s('type.extreal.extreal',sK22(X29,X28)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X28))),s('type.extreal.extreal',X29)))) )),
  inference(subsumption_resolution,[],[f32350,f263])).
fof(f32350,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | ~p(s(bool,X30)) | s(bool,X30) = s(bool,happ(s(fun('type.extreal.extreal',bool),X28),s('type.extreal.extreal',sK22(X29,X28)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X28))),s('type.extreal.extreal',X29)))) )),
  inference(superposition,[],[f321,f5007])).
fof(f32503,plain,(
  ( ! [X19,X18] : (s(bool,'F_0') != s(bool,happ(s(fun('type.extreal.extreal',bool),X18),s('type.extreal.extreal',sK22(X19,X18)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X18))),s('type.extreal.extreal',X19)))) )),
  inference(subsumption_resolution,[],[f32346,f263])).
fof(f32346,plain,(
  ( ! [X19,X18] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,happ(s(fun('type.extreal.extreal',bool),X18),s('type.extreal.extreal',sK22(X19,X18)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X18))),s('type.extreal.extreal',X19)))) )),
  inference(superposition,[],[f280,f5007])).
fof(f32502,plain,(
  ( ! [X10,X11] : (p(s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),X10)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X10))),s('type.extreal.extreal',X11)))) )),
  inference(subsumption_resolution,[],[f32342,f263])).
fof(f32342,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),X10)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X10))),s('type.extreal.extreal',X11)))) )),
  inference(superposition,[],[f469,f5007])).
fof(f32501,plain,(
  ( ! [X8,X9] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),X8),s('type.extreal.extreal',sK16(X8,'type.extreal.extreal'))))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X8))),s('type.extreal.extreal',X9)))) )),
  inference(subsumption_resolution,[],[f32341,f263])).
fof(f32341,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.extreal.extreal',bool),X8),s('type.extreal.extreal',sK16(X8,'type.extreal.extreal'))))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X8))),s('type.extreal.extreal',X9)))) )),
  inference(superposition,[],[f304,f5007])).
fof(f32500,plain,(
  ( ! [X6,X7] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),X6),s('type.extreal.extreal',sK17(X6,'type.extreal.extreal'))))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X6))),s('type.extreal.extreal',X7)))) )),
  inference(subsumption_resolution,[],[f32340,f263])).
fof(f32340,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.extreal.extreal',bool),X6),s('type.extreal.extreal',sK17(X6,'type.extreal.extreal'))))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X6))),s('type.extreal.extreal',X7)))) )),
  inference(superposition,[],[f305,f5007])).
fof(f32499,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK22(X5,X4)),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X4)))))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X4))),s('type.extreal.extreal',X5)))) )),
  inference(subsumption_resolution,[],[f32339,f263])).
fof(f32339,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK22(X5,X4)),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X4)))))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X4))),s('type.extreal.extreal',X5)))) )),
  inference(superposition,[],[f317,f5007])).
fof(f32498,plain,(
  ( ! [X204,X202,X205,X203] : (s(bool,'|dollar|exists_1'(s(fun(X204,bool),X205))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X202))),s('type.extreal.extreal',X203))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X204,bool),X205))) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X202),s('type.extreal.extreal',sK22(X203,X202))))) )),
  inference(subsumption_resolution,[],[f32327,f263])).
fof(f32327,plain,(
  ( ! [X204,X202,X205,X203] : (~p(s(bool,'T_0')) | s(bool,'|dollar|exists_1'(s(fun(X204,bool),X205))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X202))),s('type.extreal.extreal',X203))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X204,bool),X205))) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X202),s('type.extreal.extreal',sK22(X203,X202))))) )),
  inference(superposition,[],[f4971,f5007])).
fof(f32497,plain,(
  ( ! [X198,X200,X199,X197,X201] : (p(s(bool,X199)) | ~sP12(X200,X199,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X197))),s('type.extreal.extreal',X198))) | s(bool,X200) = s(bool,X201) | s(bool,'T_0') = s(bool,X201) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X197),s('type.extreal.extreal',sK22(X198,X197))))) )),
  inference(subsumption_resolution,[],[f32326,f263])).
fof(f32326,plain,(
  ( ! [X198,X200,X199,X197,X201] : (~p(s(bool,'T_0')) | p(s(bool,X199)) | ~sP12(X200,X199,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X197))),s('type.extreal.extreal',X198))) | s(bool,X200) = s(bool,X201) | s(bool,'T_0') = s(bool,X201) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X197),s('type.extreal.extreal',sK22(X198,X197))))) )),
  inference(superposition,[],[f2623,f5007])).
fof(f32496,plain,(
  ( ! [X196,X194,X192,X195,X193] : (sP12(X194,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X192))),s('type.extreal.extreal',X193)),X195) | s(bool,X194) = s(bool,X196) | s(bool,'T_0') = s(bool,X196) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X192),s('type.extreal.extreal',sK22(X193,X192))))) )),
  inference(subsumption_resolution,[],[f32325,f263])).
fof(f32325,plain,(
  ( ! [X196,X194,X192,X195,X193] : (~p(s(bool,'T_0')) | sP12(X194,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X192))),s('type.extreal.extreal',X193)),X195) | s(bool,X194) = s(bool,X196) | s(bool,'T_0') = s(bool,X196) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X192),s('type.extreal.extreal',sK22(X193,X192))))) )),
  inference(superposition,[],[f2622,f5007])).
fof(f32495,plain,(
  ( ! [X191,X189,X187,X190,X188] : (sP10(X189,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X187))),s('type.extreal.extreal',X188)),X190) | s(bool,X189) = s(bool,X191) | s(bool,'T_0') = s(bool,X191) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X187),s('type.extreal.extreal',sK22(X188,X187))))) )),
  inference(subsumption_resolution,[],[f32324,f263])).
fof(f32324,plain,(
  ( ! [X191,X189,X187,X190,X188] : (~p(s(bool,'T_0')) | sP10(X189,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X187))),s('type.extreal.extreal',X188)),X190) | s(bool,X189) = s(bool,X191) | s(bool,'T_0') = s(bool,X191) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X187),s('type.extreal.extreal',sK22(X188,X187))))) )),
  inference(superposition,[],[f2618,f5007])).
fof(f32494,plain,(
  ( ! [X185,X182,X186,X184,X183] : (sP10(X184,X185,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X182))),s('type.extreal.extreal',X183))) | s(bool,X184) = s(bool,X186) | s(bool,'T_0') = s(bool,X186) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X182),s('type.extreal.extreal',sK22(X183,X182))))) )),
  inference(subsumption_resolution,[],[f32323,f263])).
fof(f32323,plain,(
  ( ! [X185,X182,X186,X184,X183] : (~p(s(bool,'T_0')) | sP10(X184,X185,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X182))),s('type.extreal.extreal',X183))) | s(bool,X184) = s(bool,X186) | s(bool,'T_0') = s(bool,X186) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X182),s('type.extreal.extreal',sK22(X183,X182))))) )),
  inference(superposition,[],[f2617,f5007])).
fof(f32493,plain,(
  ( ! [X177,X180,X178,X181,X179] : (p(s(bool,X179)) | ~sP8(X179,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X177))),s('type.extreal.extreal',X178)),X180) | s(bool,X180) = s(bool,X181) | s(bool,'T_0') = s(bool,X181) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X177),s('type.extreal.extreal',sK22(X178,X177))))) )),
  inference(subsumption_resolution,[],[f32322,f263])).
fof(f32322,plain,(
  ( ! [X177,X180,X178,X181,X179] : (~p(s(bool,'T_0')) | p(s(bool,X179)) | ~sP8(X179,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X177))),s('type.extreal.extreal',X178)),X180) | s(bool,X180) = s(bool,X181) | s(bool,'T_0') = s(bool,X181) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X177),s('type.extreal.extreal',sK22(X178,X177))))) )),
  inference(superposition,[],[f2615,f5007])).
fof(f32492,plain,(
  ( ! [X103,X105,X106,X104] : (~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X103))),s('type.extreal.extreal',X104)),X105,X106) | sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X103))),s('type.extreal.extreal',X104)),X105,X106) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X103),s('type.extreal.extreal',sK22(X104,X103))))) )),
  inference(subsumption_resolution,[],[f32305,f263])).
fof(f32305,plain,(
  ( ! [X103,X105,X106,X104] : (~p(s(bool,'T_0')) | ~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X103))),s('type.extreal.extreal',X104)),X105,X106) | sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X103))),s('type.extreal.extreal',X104)),X105,X106) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X103),s('type.extreal.extreal',sK22(X104,X103))))) )),
  inference(superposition,[],[f453,f5007])).
fof(f32491,plain,(
  ( ! [X101,X99,X102,X100] : (~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X99))),s('type.extreal.extreal',X100)),X101,X102) | sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X99))),s('type.extreal.extreal',X100)),X101,X102) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X99),s('type.extreal.extreal',sK22(X100,X99))))) )),
  inference(subsumption_resolution,[],[f32304,f263])).
fof(f32304,plain,(
  ( ! [X101,X99,X102,X100] : (~p(s(bool,'T_0')) | ~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X99))),s('type.extreal.extreal',X100)),X101,X102) | sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X99))),s('type.extreal.extreal',X100)),X101,X102) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X99),s('type.extreal.extreal',sK22(X100,X99))))) )),
  inference(superposition,[],[f452,f5007])).
fof(f32490,plain,(
  ( ! [X97,X95,X98,X96] : (~p(s(bool,X97)) | p(s(bool,X98)) | ~sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X95))),s('type.extreal.extreal',X96)),X98,X97) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X95),s('type.extreal.extreal',sK22(X96,X95))))) )),
  inference(subsumption_resolution,[],[f32303,f263])).
fof(f32303,plain,(
  ( ! [X97,X95,X98,X96] : (~p(s(bool,'T_0')) | ~p(s(bool,X97)) | p(s(bool,X98)) | ~sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X95))),s('type.extreal.extreal',X96)),X98,X97) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X95),s('type.extreal.extreal',sK22(X96,X95))))) )),
  inference(superposition,[],[f449,f5007])).
fof(f32489,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,X93)) | sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X91))),s('type.extreal.extreal',X92)),X93,X94) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X91),s('type.extreal.extreal',sK22(X92,X91))))) )),
  inference(subsumption_resolution,[],[f32302,f263])).
fof(f32302,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,'T_0')) | ~p(s(bool,X93)) | sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X91))),s('type.extreal.extreal',X92)),X93,X94) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X91),s('type.extreal.extreal',sK22(X92,X91))))) )),
  inference(superposition,[],[f448,f5007])).
fof(f32488,plain,(
  ( ! [X90,X88,X87,X89] : (p(s(bool,X89)) | sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X87))),s('type.extreal.extreal',X88)),X90,X89) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X87),s('type.extreal.extreal',sK22(X88,X87))))) )),
  inference(subsumption_resolution,[],[f32301,f263])).
fof(f32301,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,'T_0')) | p(s(bool,X89)) | sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X87))),s('type.extreal.extreal',X88)),X90,X89) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X87),s('type.extreal.extreal',sK22(X88,X87))))) )),
  inference(superposition,[],[f447,f5007])).
fof(f32487,plain,(
  ( ! [X85,X83,X86,X84] : (p(s(bool,X85)) | ~p(s(bool,X86)) | ~sP11(X86,X85,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X83))),s('type.extreal.extreal',X84))) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X83),s('type.extreal.extreal',sK22(X84,X83))))) )),
  inference(subsumption_resolution,[],[f32300,f263])).
fof(f32300,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,'T_0')) | p(s(bool,X85)) | ~p(s(bool,X86)) | ~sP11(X86,X85,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X83))),s('type.extreal.extreal',X84))) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X83),s('type.extreal.extreal',sK22(X84,X83))))) )),
  inference(superposition,[],[f442,f5007])).
fof(f32486,plain,(
  ( ! [X80,X81,X79,X82] : (p(s(bool,X81)) | p(s(bool,X82)) | ~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X79))),s('type.extreal.extreal',X80)),X82,X81) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X79),s('type.extreal.extreal',sK22(X80,X79))))) )),
  inference(subsumption_resolution,[],[f32299,f263])).
fof(f32299,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,'T_0')) | p(s(bool,X81)) | p(s(bool,X82)) | ~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X79))),s('type.extreal.extreal',X80)),X82,X81) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X79),s('type.extreal.extreal',sK22(X80,X79))))) )),
  inference(superposition,[],[f432,f5007])).
fof(f32485,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X77)) | sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X75))),s('type.extreal.extreal',X76)),X77,X78) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X75),s('type.extreal.extreal',sK22(X76,X75))))) )),
  inference(subsumption_resolution,[],[f32298,f263])).
fof(f32298,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,'T_0')) | ~p(s(bool,X77)) | sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X75))),s('type.extreal.extreal',X76)),X77,X78) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X75),s('type.extreal.extreal',sK22(X76,X75))))) )),
  inference(superposition,[],[f431,f5007])).
fof(f32484,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X73)) | sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X71))),s('type.extreal.extreal',X72)),X74,X73) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X71),s('type.extreal.extreal',sK22(X72,X71))))) )),
  inference(subsumption_resolution,[],[f32297,f263])).
fof(f32297,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,'T_0')) | ~p(s(bool,X73)) | sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X71))),s('type.extreal.extreal',X72)),X74,X73) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X71),s('type.extreal.extreal',sK22(X72,X71))))) )),
  inference(superposition,[],[f430,f5007])).
fof(f32483,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,X69)) | p(s(bool,X70)) | ~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X67))),s('type.extreal.extreal',X68)),X69,X70) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X67),s('type.extreal.extreal',sK22(X68,X67))))) )),
  inference(subsumption_resolution,[],[f32296,f263])).
fof(f32296,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | p(s(bool,X69)) | p(s(bool,X70)) | ~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X67))),s('type.extreal.extreal',X68)),X69,X70) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X67),s('type.extreal.extreal',sK22(X68,X67))))) )),
  inference(superposition,[],[f425,f5007])).
fof(f32482,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X65)) | p(s(bool,X66)) | ~sP8(X66,X65,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X63))),s('type.extreal.extreal',X64))) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X63),s('type.extreal.extreal',sK22(X64,X63))))) )),
  inference(subsumption_resolution,[],[f32295,f263])).
fof(f32295,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'T_0')) | ~p(s(bool,X65)) | p(s(bool,X66)) | ~sP8(X66,X65,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X63))),s('type.extreal.extreal',X64))) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X63),s('type.extreal.extreal',sK22(X64,X63))))) )),
  inference(superposition,[],[f417,f5007])).
fof(f32481,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X61)) | p(s(bool,X62)) | ~sP7(X61,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X59))),s('type.extreal.extreal',X60)),X62) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X59),s('type.extreal.extreal',sK22(X60,X59))))) )),
  inference(subsumption_resolution,[],[f32294,f263])).
fof(f32294,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | ~p(s(bool,X61)) | p(s(bool,X62)) | ~sP7(X61,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X59))),s('type.extreal.extreal',X60)),X62) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X59),s('type.extreal.extreal',sK22(X60,X59))))) )),
  inference(superposition,[],[f408,f5007])).
fof(f32480,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X57)) | p(s(bool,X58)) | ~sP4(X57,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X55))),s('type.extreal.extreal',X56)),X58) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X55),s('type.extreal.extreal',sK22(X56,X55))))) )),
  inference(subsumption_resolution,[],[f32293,f263])).
fof(f32293,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X57)) | p(s(bool,X58)) | ~sP4(X57,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X55))),s('type.extreal.extreal',X56)),X58) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X55),s('type.extreal.extreal',sK22(X56,X55))))) )),
  inference(superposition,[],[f389,f5007])).
fof(f32479,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X53)) | p(s(bool,X54)) | ~sP5(X53,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X51))),s('type.extreal.extreal',X52)),X54) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X51),s('type.extreal.extreal',sK22(X52,X51))))) )),
  inference(subsumption_resolution,[],[f32292,f263])).
fof(f32292,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X53)) | p(s(bool,X54)) | ~sP5(X53,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X51))),s('type.extreal.extreal',X52)),X54) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X51),s('type.extreal.extreal',sK22(X52,X51))))) )),
  inference(superposition,[],[f385,f5007])).
fof(f32478,plain,(
  ( ! [X47,X50,X48,X49] : (sP6(X49,X50,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X47))),s('type.extreal.extreal',X48))) | ~sP5(X49,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X47))),s('type.extreal.extreal',X48)),X50) | ~sP4(X50,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X47))),s('type.extreal.extreal',X48)),X49) | ~sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X47))),s('type.extreal.extreal',X48)),X50,X49) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X47),s('type.extreal.extreal',sK22(X48,X47))))) )),
  inference(subsumption_resolution,[],[f32291,f263])).
fof(f32291,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | sP6(X49,X50,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X47))),s('type.extreal.extreal',X48))) | ~sP5(X49,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X47))),s('type.extreal.extreal',X48)),X50) | ~sP4(X50,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X47))),s('type.extreal.extreal',X48)),X49) | ~sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X47))),s('type.extreal.extreal',X48)),X50,X49) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X47),s('type.extreal.extreal',sK22(X48,X47))))) )),
  inference(superposition,[],[f382,f5007])).
fof(f32477,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,X46)) | sP1(X46,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X44))),s('type.extreal.extreal',X45))) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X44),s('type.extreal.extreal',sK22(X45,X44))))) )),
  inference(subsumption_resolution,[],[f32290,f263])).
fof(f32290,plain,(
  ( ! [X45,X46,X44] : (~p(s(bool,'T_0')) | p(s(bool,X46)) | sP1(X46,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X44))),s('type.extreal.extreal',X45))) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X44),s('type.extreal.extreal',sK22(X45,X44))))) )),
  inference(superposition,[],[f353,f5007])).
fof(f32476,plain,(
  ( ! [X43,X41,X42] : (p(s(bool,X43)) | sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X41))),s('type.extreal.extreal',X42)),X43) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X41),s('type.extreal.extreal',sK22(X42,X41))))) )),
  inference(subsumption_resolution,[],[f32289,f263])).
fof(f32289,plain,(
  ( ! [X43,X41,X42] : (~p(s(bool,'T_0')) | p(s(bool,X43)) | sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X41))),s('type.extreal.extreal',X42)),X43) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X41),s('type.extreal.extreal',sK22(X42,X41))))) )),
  inference(superposition,[],[f350,f5007])).
fof(f32475,plain,(
  ( ! [X39,X38,X40] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X38))),s('type.extreal.extreal',X39)),X40) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X38),s('type.extreal.extreal',sK22(X39,X38))))) )),
  inference(subsumption_resolution,[],[f32288,f263])).
fof(f32288,plain,(
  ( ! [X39,X38,X40] : (~p(s(bool,'T_0')) | sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X38))),s('type.extreal.extreal',X39)),X40) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X38),s('type.extreal.extreal',sK22(X39,X38))))) )),
  inference(superposition,[],[f335,f5007])).
fof(f32474,plain,(
  ( ! [X37,X35,X36] : (sP0(X37,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X35))),s('type.extreal.extreal',X36))) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X35),s('type.extreal.extreal',sK22(X36,X35))))) )),
  inference(subsumption_resolution,[],[f32287,f263])).
fof(f32287,plain,(
  ( ! [X37,X35,X36] : (~p(s(bool,'T_0')) | sP0(X37,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X35))),s('type.extreal.extreal',X36))) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X35),s('type.extreal.extreal',sK22(X36,X35))))) )),
  inference(superposition,[],[f334,f5007])).
fof(f32473,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,X34)) | s(bool,X34) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X32))),s('type.extreal.extreal',X33))) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X32),s('type.extreal.extreal',sK22(X33,X32))))) )),
  inference(subsumption_resolution,[],[f32286,f263])).
fof(f32286,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X34)) | s(bool,X34) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X32))),s('type.extreal.extreal',X33))) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X32),s('type.extreal.extreal',sK22(X33,X32))))) )),
  inference(superposition,[],[f321,f5007])).
fof(f32472,plain,(
  ( ! [X23,X22] : (s(bool,'F_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X22))),s('type.extreal.extreal',X23))) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X22),s('type.extreal.extreal',sK22(X23,X22))))) )),
  inference(subsumption_resolution,[],[f32282,f263])).
fof(f32282,plain,(
  ( ! [X23,X22] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X22))),s('type.extreal.extreal',X23))) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X22),s('type.extreal.extreal',sK22(X23,X22))))) )),
  inference(superposition,[],[f280,f5007])).
fof(f32471,plain,(
  ( ! [X14,X15] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X15),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X14)))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X14),s('type.extreal.extreal',sK22(X15,X14))))) )),
  inference(subsumption_resolution,[],[f32278,f263])).
fof(f32278,plain,(
  ( ! [X14,X15] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X15),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X14)))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X14),s('type.extreal.extreal',sK22(X15,X14))))) )),
  inference(superposition,[],[f324,f5007])).
fof(f32470,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X11),s('type.extreal.extreal',X10)))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X11),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X9)))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X9),s('type.extreal.extreal',sK22(X10,X9))))) )),
  inference(subsumption_resolution,[],[f32276,f263])).
fof(f32276,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'T_0')) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X11),s('type.extreal.extreal',X10)))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X11),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X9)))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X9),s('type.extreal.extreal',sK22(X10,X9))))) )),
  inference(superposition,[],[f367,f5007])).
fof(f32469,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X2),s('type.extreal.extreal',X4)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X4),s('type.extreal.extreal',X3)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X2),s('type.extreal.extreal',sK22(X3,X2))))) )),
  inference(subsumption_resolution,[],[f32274,f263])).
fof(f32274,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X2),s('type.extreal.extreal',X4)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X4),s('type.extreal.extreal',X3)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X2),s('type.extreal.extreal',sK22(X3,X2))))) )),
  inference(superposition,[],[f326,f5007])).
fof(f30120,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X2) = s(bool,X3) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2)) )),
  inference(equality_factoring,[],[f2054])).
fof(f32125,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f30113,f783])).
fof(f30113,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,X2)) )),
  inference(equality_factoring,[],[f2054])).
fof(f29440,plain,(
  ( ! [X951,X949,X952,X950,X948,X953] : (s(X950,f(s(bool,'T_0'))) != s(X950,g(s(bool,'T_0'))) | s(fun(bool,X950),X948) = s(fun(bool,X950),X949) | s(bool,X951) = s(bool,X952) | s(bool,X952) = s(bool,sK23(X948,X949,X950,bool)) | s(bool,X951) = s(bool,X953) | s(bool,'T_0') = s(bool,X953)) )),
  inference(superposition,[],[f366,f2054])).
fof(f29439,plain,(
  ( ! [X947,X945,X942,X946,X944,X943] : (~p(s(bool,happ(s(fun(bool,bool),X942),s(bool,'T_0')))) | p(s(bool,X943)) | p(s(bool,happ(s(fun(bool,bool),X942),s(bool,X944)))) | s(bool,X945) = s(bool,X946) | s(bool,X946) = s(bool,sK21(X942,X943,bool)) | s(bool,X945) = s(bool,X947) | s(bool,'T_0') = s(bool,X947)) )),
  inference(superposition,[],[f316,f2054])).
fof(f29438,plain,(
  ( ! [X938,X936,X940,X939,X937,X941] : (~p(s(bool,happ(s(fun(bool,bool),X936),s(bool,'T_0')))) | p(s(bool,X937)) | p(s(bool,happ(s(fun(bool,bool),X936),s(bool,X938)))) | s(bool,X939) = s(bool,X940) | s(bool,X940) = s(bool,sK20(X936,X937,bool)) | s(bool,X939) = s(bool,X941) | s(bool,'T_0') = s(bool,X941)) )),
  inference(superposition,[],[f314,f2054])).
fof(f29437,plain,(
  ( ! [X930,X934,X932,X931,X935,X933] : (~p(s(bool,happ(s(fun(bool,bool),X930),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X930),s(bool,X932)))) | ~p(s(bool,X931)) | s(bool,X933) = s(bool,X934) | s(bool,X934) = s(bool,sK19(X930,X931,bool)) | s(bool,X933) = s(bool,X935) | s(bool,'T_0') = s(bool,X935)) )),
  inference(superposition,[],[f312,f2054])).
fof(f29436,plain,(
  ( ! [X926,X924,X928,X927,X925,X929] : (~p(s(bool,happ(s(fun(bool,bool),X924),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X924),s(bool,X926)))) | ~p(s(bool,X925)) | s(bool,X927) = s(bool,X928) | s(bool,X928) = s(bool,sK18(X924,X925,bool)) | s(bool,X927) = s(bool,X929) | s(bool,'T_0') = s(bool,X929)) )),
  inference(superposition,[],[f310,f2054])).
fof(f29435,plain,(
  ( ! [X922,X920,X923,X921] : (p(s(bool,happ(s(fun(bool,bool),X920),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X920))) | s(bool,X921) = s(bool,X922) | s(bool,X922) = s(bool,sK17(X920,bool)) | s(bool,X921) = s(bool,X923) | s(bool,'T_0') = s(bool,X923)) )),
  inference(superposition,[],[f4966,f2054])).
fof(f29434,plain,(
  ( ! [X918,X916,X915,X919,X917] : (p(s(bool,happ(s(fun(bool,bool),X915),s(bool,'T_0')))) | s(bool,X916) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X915))) | s(bool,'T_0') = s(bool,X916) | s(bool,X917) = s(bool,X918) | s(bool,X918) = s(bool,sK17(X915,bool)) | s(bool,X917) = s(bool,X919) | s(bool,'T_0') = s(bool,X919)) )),
  inference(superposition,[],[f4303,f2054])).
fof(f29433,plain,(
  ( ! [X912,X914,X911,X913] : (p(s(bool,happ(s(fun(bool,bool),X911),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X911))) | s(bool,X912) = s(bool,X913) | s(bool,X913) = s(bool,sK16(X911,bool)) | s(bool,X912) = s(bool,X914) | s(bool,'T_0') = s(bool,X914)) )),
  inference(superposition,[],[f4967,f2054])).
fof(f29432,plain,(
  ( ! [X907,X909,X906,X910,X908] : (p(s(bool,happ(s(fun(bool,bool),X906),s(bool,'T_0')))) | s(bool,X907) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X906))) | s(bool,'T_0') = s(bool,X907) | s(bool,X908) = s(bool,X909) | s(bool,X909) = s(bool,sK16(X906,bool)) | s(bool,X908) = s(bool,X910) | s(bool,'T_0') = s(bool,X910)) )),
  inference(superposition,[],[f3897,f2054])).
fof(f29423,plain,(
  ( ! [X872,X870,X868,X873,X871,X869] : (sP12(happ(s(fun(bool,bool),X868),s(bool,'T_0')),X869,X870) | ~p(s(bool,X869)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X868))) | s(bool,X871) = s(bool,X872) | s(bool,X872) = s(bool,'const.min.@_1'(s(fun(bool,bool),X868))) | s(bool,X871) = s(bool,X873) | s(bool,'T_0') = s(bool,X873)) )),
  inference(superposition,[],[f4987,f2054])).
fof(f29422,plain,(
  ( ! [X866,X864,X863,X867,X865,X862] : (sP12(happ(s(fun(bool,bool),X862),s(bool,'T_0')),X863,X864) | p(s(bool,X864)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X862))) | s(bool,X865) = s(bool,X866) | s(bool,X866) = s(bool,'const.min.@_1'(s(fun(bool,bool),X862))) | s(bool,X865) = s(bool,X867) | s(bool,'T_0') = s(bool,X867)) )),
  inference(superposition,[],[f4986,f2054])).
fof(f29421,plain,(
  ( ! [X860,X859,X857,X861,X858,X856] : (sP10(happ(s(fun(bool,bool),X856),s(bool,'T_0')),X857,X858) | ~p(s(bool,X857)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X856))) | s(bool,X859) = s(bool,X860) | s(bool,X860) = s(bool,'const.min.@_1'(s(fun(bool,bool),X856))) | s(bool,X859) = s(bool,X861) | s(bool,'T_0') = s(bool,X861)) )),
  inference(superposition,[],[f4983,f2054])).
fof(f29420,plain,(
  ( ! [X850,X854,X852,X851,X855,X853] : (sP10(happ(s(fun(bool,bool),X850),s(bool,'T_0')),X851,X852) | ~p(s(bool,X852)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X850))) | s(bool,X853) = s(bool,X854) | s(bool,X854) = s(bool,'const.min.@_1'(s(fun(bool,bool),X850))) | s(bool,X853) = s(bool,X855) | s(bool,'T_0') = s(bool,X855)) )),
  inference(superposition,[],[f4982,f2054])).
fof(f29419,plain,(
  ( ! [X848,X847,X845,X849,X846] : (sP1(X846,happ(s(fun(bool,bool),X845),s(bool,'T_0'))) | p(s(bool,X846)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X845))) | s(bool,X847) = s(bool,X848) | s(bool,X848) = s(bool,'const.min.@_1'(s(fun(bool,bool),X845))) | s(bool,X847) = s(bool,X849) | s(bool,'T_0') = s(bool,X849)) )),
  inference(superposition,[],[f4975,f2054])).
fof(f29418,plain,(
  ( ! [X843,X841,X842,X840,X844] : (sP1(happ(s(fun(bool,bool),X840),s(bool,'T_0')),X841) | p(s(bool,X841)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X840))) | s(bool,X842) = s(bool,X843) | s(bool,X843) = s(bool,'const.min.@_1'(s(fun(bool,bool),X840))) | s(bool,X842) = s(bool,X844) | s(bool,'T_0') = s(bool,X844)) )),
  inference(superposition,[],[f4974,f2054])).
fof(f29417,plain,(
  ( ! [X835,X839,X837,X838,X836] : (sP0(happ(s(fun(bool,bool),X835),s(bool,'T_0')),X836) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X835))) | s(bool,X837) = s(bool,X838) | s(bool,X838) = s(bool,'const.min.@_1'(s(fun(bool,bool),X835))) | s(bool,X837) = s(bool,X839) | s(bool,'T_0') = s(bool,X839)) )),
  inference(superposition,[],[f4973,f2054])).
fof(f29416,plain,(
  ( ! [X831,X833,X830,X834,X832] : (sP0(X831,happ(s(fun(bool,bool),X830),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X830))) | s(bool,X832) = s(bool,X833) | s(bool,X833) = s(bool,'const.min.@_1'(s(fun(bool,bool),X830))) | s(bool,X832) = s(bool,X834) | s(bool,'T_0') = s(bool,X834)) )),
  inference(superposition,[],[f4972,f2054])).
fof(f29415,plain,(
  ( ! [X829,X826,X828,X827] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X826),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X826))) | s(bool,X827) = s(bool,X828) | s(bool,X828) = s(bool,'const.min.@_1'(s(fun(bool,bool),X826))) | s(bool,X827) = s(bool,X829) | s(bool,'T_0') = s(bool,X829)) )),
  inference(superposition,[],[f4969,f2054])).
fof(f29414,plain,(
  ( ! [X825,X823,X824,X822] : (sP5('T_0',happ(s(fun(bool,bool),X822),s(bool,'T_0')),'|dollar|exists_1'(s(fun(bool,bool),X822))) | s(bool,X823) = s(bool,X824) | s(bool,X824) = s(bool,'const.min.@_1'(s(fun(bool,bool),X822))) | s(bool,X823) = s(bool,X825) | s(bool,'T_0') = s(bool,X825)) )),
  inference(superposition,[],[f3862,f2054])).
fof(f29413,plain,(
  ( ! [X819,X821,X818,X820] : (sP5('T_0','|dollar|exists_1'(s(fun(bool,bool),X818)),happ(s(fun(bool,bool),X818),s(bool,'T_0'))) | s(bool,X819) = s(bool,X820) | s(bool,X820) = s(bool,'const.min.@_1'(s(fun(bool,bool),X818))) | s(bool,X819) = s(bool,X821) | s(bool,'T_0') = s(bool,X821)) )),
  inference(superposition,[],[f3757,f2054])).
fof(f29412,plain,(
  ( ! [X817,X814,X812,X816,X815,X813] : (sP10(happ(s(fun(bool,bool),X812),s(bool,'T_0')),X813,X814) | ~sP9(happ(s(fun(bool,bool),X812),s(bool,'const.min.@_1'(s(fun(bool,bool),X812)))),X813,X814) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X812)))) | s(bool,X815) = s(bool,X816) | s(bool,X816) = s(bool,'const.min.@_1'(s(fun(bool,bool),X812))) | s(bool,X815) = s(bool,X817) | s(bool,'T_0') = s(bool,X817)) )),
  inference(superposition,[],[f1797,f2054])).
fof(f29411,plain,(
  ( ! [X810,X808,X806,X811,X809,X807] : (sP12(happ(s(fun(bool,bool),X806),s(bool,'T_0')),X807,X808) | ~sP11(happ(s(fun(bool,bool),X806),s(bool,'const.min.@_1'(s(fun(bool,bool),X806)))),X807,X808) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X806)))) | s(bool,X809) = s(bool,X810) | s(bool,X810) = s(bool,'const.min.@_1'(s(fun(bool,bool),X806))) | s(bool,X809) = s(bool,X811) | s(bool,'T_0') = s(bool,X811)) )),
  inference(superposition,[],[f1796,f2054])).
fof(f29410,plain,(
  ( ! [X802,X800,X804,X803,X801,X805] : (~sP12(happ(s(fun(bool,bool),X800),s(bool,'T_0')),X801,X802) | ~p(s(bool,X802)) | p(s(bool,X801)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X800)))) | s(bool,X803) = s(bool,X804) | s(bool,X804) = s(bool,'const.min.@_1'(s(fun(bool,bool),X800))) | s(bool,X803) = s(bool,X805) | s(bool,'T_0') = s(bool,X805)) )),
  inference(superposition,[],[f1795,f2054])).
fof(f29409,plain,(
  ( ! [X798,X796,X795,X799,X797,X794] : (sP12(happ(s(fun(bool,bool),X794),s(bool,'T_0')),X795,X796) | ~p(s(bool,X795)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X794)))) | s(bool,X797) = s(bool,X798) | s(bool,X798) = s(bool,'const.min.@_1'(s(fun(bool,bool),X794))) | s(bool,X797) = s(bool,X799) | s(bool,'T_0') = s(bool,X799)) )),
  inference(superposition,[],[f1794,f2054])).
fof(f29408,plain,(
  ( ! [X790,X788,X793,X791,X789,X792] : (sP12(happ(s(fun(bool,bool),X788),s(bool,'T_0')),X789,X790) | p(s(bool,X790)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X788)))) | s(bool,X791) = s(bool,X792) | s(bool,X792) = s(bool,'const.min.@_1'(s(fun(bool,bool),X788))) | s(bool,X791) = s(bool,X793) | s(bool,'T_0') = s(bool,X793)) )),
  inference(superposition,[],[f1793,f2054])).
fof(f29407,plain,(
  ( ! [X786,X784,X783,X787,X785,X782] : (~sP10(happ(s(fun(bool,bool),X782),s(bool,'T_0')),X783,X784) | p(s(bool,X784)) | p(s(bool,X783)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X782)))) | s(bool,X785) = s(bool,X786) | s(bool,X786) = s(bool,'const.min.@_1'(s(fun(bool,bool),X782))) | s(bool,X785) = s(bool,X787) | s(bool,'T_0') = s(bool,X787)) )),
  inference(superposition,[],[f1791,f2054])).
fof(f29406,plain,(
  ( ! [X779,X777,X781,X778,X776,X780] : (sP10(happ(s(fun(bool,bool),X776),s(bool,'T_0')),X777,X778) | ~p(s(bool,X777)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X776)))) | s(bool,X779) = s(bool,X780) | s(bool,X780) = s(bool,'const.min.@_1'(s(fun(bool,bool),X776))) | s(bool,X779) = s(bool,X781) | s(bool,'T_0') = s(bool,X781)) )),
  inference(superposition,[],[f1790,f2054])).
fof(f29405,plain,(
  ( ! [X771,X775,X773,X770,X774,X772] : (sP10(happ(s(fun(bool,bool),X770),s(bool,'T_0')),X771,X772) | ~p(s(bool,X772)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X770)))) | s(bool,X773) = s(bool,X774) | s(bool,X774) = s(bool,'const.min.@_1'(s(fun(bool,bool),X770))) | s(bool,X773) = s(bool,X775) | s(bool,'T_0') = s(bool,X775)) )),
  inference(superposition,[],[f1789,f2054])).
fof(f29404,plain,(
  ( ! [X765,X767,X769,X766,X768] : (sP1(X766,happ(s(fun(bool,bool),X765),s(bool,'T_0'))) | p(s(bool,X766)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X765)))) | s(bool,X767) = s(bool,X768) | s(bool,X768) = s(bool,'const.min.@_1'(s(fun(bool,bool),X765))) | s(bool,X767) = s(bool,X769) | s(bool,'T_0') = s(bool,X769)) )),
  inference(superposition,[],[f1782,f2054])).
fof(f29403,plain,(
  ( ! [X761,X763,X764,X760,X762] : (sP1(happ(s(fun(bool,bool),X760),s(bool,'T_0')),X761) | p(s(bool,X761)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X760)))) | s(bool,X762) = s(bool,X763) | s(bool,X763) = s(bool,'const.min.@_1'(s(fun(bool,bool),X760))) | s(bool,X762) = s(bool,X764) | s(bool,'T_0') = s(bool,X764)) )),
  inference(superposition,[],[f1781,f2054])).
fof(f29402,plain,(
  ( ! [X757,X759,X755,X756,X758] : (sP0(happ(s(fun(bool,bool),X755),s(bool,'T_0')),X756) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X755)))) | s(bool,X757) = s(bool,X758) | s(bool,X758) = s(bool,'const.min.@_1'(s(fun(bool,bool),X755))) | s(bool,X757) = s(bool,X759) | s(bool,'T_0') = s(bool,X759)) )),
  inference(superposition,[],[f1780,f2054])).
fof(f29401,plain,(
  ( ! [X753,X751,X752,X754,X750] : (sP0(X751,happ(s(fun(bool,bool),X750),s(bool,'T_0'))) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X750)))) | s(bool,X752) = s(bool,X753) | s(bool,X753) = s(bool,'const.min.@_1'(s(fun(bool,bool),X750))) | s(bool,X752) = s(bool,X754) | s(bool,'T_0') = s(bool,X754)) )),
  inference(superposition,[],[f1779,f2054])).
fof(f29400,plain,(
  ( ! [X749,X747,X748,X746] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X746),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X746))) | s(bool,X747) = s(bool,X748) | s(bool,X748) = s(bool,'const.min.@_1'(s(fun(bool,bool),X746))) | s(bool,X747) = s(bool,X749) | s(bool,'T_0') = s(bool,X749)) )),
  inference(superposition,[],[f1769,f2054])).
fof(f29399,plain,(
  ( ! [X745,X743,X744,X742] : (sP6('T_0','|dollar|exists_1'(s(fun(bool,bool),X742)),happ(s(fun(bool,bool),X742),s(bool,'T_0'))) | s(bool,X743) = s(bool,X744) | s(bool,X744) = s(bool,'const.min.@_1'(s(fun(bool,bool),X742))) | s(bool,X743) = s(bool,X745) | s(bool,'T_0') = s(bool,X745)) )),
  inference(superposition,[],[f1767,f2054])).
fof(f29398,plain,(
  ( ! [X738,X741,X739,X740] : (~sP6('F_0','|dollar|exists_1'(s(fun(bool,bool),X738)),happ(s(fun(bool,bool),X738),s(bool,'T_0'))) | s(bool,X739) = s(bool,X740) | s(bool,X740) = s(bool,'const.min.@_1'(s(fun(bool,bool),X738))) | s(bool,X739) = s(bool,X741) | s(bool,'T_0') = s(bool,X741)) )),
  inference(superposition,[],[f1766,f2054])).
fof(f29397,plain,(
  ( ! [X736,X734,X737,X735] : (sP6('T_0',happ(s(fun(bool,bool),X734),s(bool,'T_0')),'|dollar|exists_1'(s(fun(bool,bool),X734))) | s(bool,X735) = s(bool,X736) | s(bool,X736) = s(bool,'const.min.@_1'(s(fun(bool,bool),X734))) | s(bool,X735) = s(bool,X737) | s(bool,'T_0') = s(bool,X737)) )),
  inference(superposition,[],[f1765,f2054])).
fof(f29396,plain,(
  ( ! [X732,X730,X733,X731] : (~sP6('F_0',happ(s(fun(bool,bool),X730),s(bool,'T_0')),'|dollar|exists_1'(s(fun(bool,bool),X730))) | s(bool,X731) = s(bool,X732) | s(bool,X732) = s(bool,'const.min.@_1'(s(fun(bool,bool),X730))) | s(bool,X731) = s(bool,X733) | s(bool,'T_0') = s(bool,X733)) )),
  inference(superposition,[],[f1764,f2054])).
fof(f29395,plain,(
  ( ! [X728,X726,X729,X727] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X726))) = s(bool,happ(s(fun(bool,bool),X726),s(bool,'T_0'))) | s(bool,X727) = s(bool,X728) | s(bool,X728) = s(bool,'const.min.@_1'(s(fun(bool,bool),X726))) | s(bool,X727) = s(bool,X729) | s(bool,'T_0') = s(bool,X729)) )),
  inference(superposition,[],[f303,f2054])).
fof(f29393,plain,(
  ( ! [X720,X716,X718,X717,X719,X715] : (s(bool,'T_0') != s(bool,f(s(X715,sK23(X716,X717,bool,X715)))) | s(fun(X715,bool),X716) = s(fun(X715,bool),X717) | s(bool,X718) = s(bool,X719) | s(bool,X719) = s(bool,g(s(X715,sK23(X716,X717,bool,X715)))) | s(bool,X718) = s(bool,X720) | s(bool,'T_0') = s(bool,X720)) )),
  inference(superposition,[],[f366,f2054])).
fof(f31978,plain,(
  ( ! [X693,X691,X692,X694,X690] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X691))),s('type.extreal.extreal',X690)))) | s(bool,X692) = s(bool,X693) | s(bool,X693) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK22(X690,X691)),s('type.extreal.extreal',X690))) | s(bool,X692) = s(bool,X694) | s(bool,'T_0') = s(bool,X694)) )),
  inference(subsumption_resolution,[],[f29389,f263])).
fof(f29389,plain,(
  ( ! [X693,X691,X692,X694,X690] : (~p(s(bool,'T_0')) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X691))),s('type.extreal.extreal',X690)))) | s(bool,X692) = s(bool,X693) | s(bool,X693) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK22(X690,X691)),s('type.extreal.extreal',X690))) | s(bool,X692) = s(bool,X694) | s(bool,'T_0') = s(bool,X694)) )),
  inference(superposition,[],[f328,f2054])).
fof(f31977,plain,(
  ( ! [X666,X668,X670,X667,X669,X665] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X665),s('type.extreal.extreal',X667)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X667),s('type.extreal.extreal',X666)))) | s(bool,X668) = s(bool,X669) | s(bool,X669) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X665))),s('type.extreal.extreal',X666))) | s(bool,X668) = s(bool,X670) | s(bool,'T_0') = s(bool,X670)) )),
  inference(subsumption_resolution,[],[f29383,f263])).
fof(f29383,plain,(
  ( ! [X666,X668,X670,X667,X669,X665] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X665),s('type.extreal.extreal',X667)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X667),s('type.extreal.extreal',X666)))) | s(bool,X668) = s(bool,X669) | s(bool,X669) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X665))),s('type.extreal.extreal',X666))) | s(bool,X668) = s(bool,X670) | s(bool,'T_0') = s(bool,X670)) )),
  inference(superposition,[],[f326,f2054])).
fof(f31976,plain,(
  ( ! [X656,X658,X657,X659,X655] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X656),s('type.extreal.extreal',X655)))) | s(bool,X657) = s(bool,X658) | s(bool,X658) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X655),s('type.extreal.extreal',X656))) | s(bool,X657) = s(bool,X659) | s(bool,'T_0') = s(bool,X659)) )),
  inference(subsumption_resolution,[],[f29381,f263])).
fof(f29381,plain,(
  ( ! [X656,X658,X657,X659,X655] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X656),s('type.extreal.extreal',X655)))) | s(bool,X657) = s(bool,X658) | s(bool,X658) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X655),s('type.extreal.extreal',X656))) | s(bool,X657) = s(bool,X659) | s(bool,'T_0') = s(bool,X659)) )),
  inference(superposition,[],[f324,f2054])).
fof(f31975,plain,(
  ( ! [X648,X644,X646,X649,X645,X647] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X646),s('type.extreal.extreal',X645)))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X646),s('type.extreal.extreal',X644)))) | s(bool,X647) = s(bool,X648) | s(bool,X648) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X644),s('type.extreal.extreal',X645))) | s(bool,X647) = s(bool,X649) | s(bool,'T_0') = s(bool,X649)) )),
  inference(subsumption_resolution,[],[f29379,f263])).
fof(f29379,plain,(
  ( ! [X648,X644,X646,X649,X645,X647] : (~p(s(bool,'T_0')) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X646),s('type.extreal.extreal',X645)))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X646),s('type.extreal.extreal',X644)))) | s(bool,X647) = s(bool,X648) | s(bool,X648) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X644),s('type.extreal.extreal',X645))) | s(bool,X647) = s(bool,X649) | s(bool,'T_0') = s(bool,X649)) )),
  inference(superposition,[],[f367,f2054])).
fof(f31974,plain,(
  ( ! [X633,X631,X632,X634,X630] : (p(s(bool,happ(s(fun(X630,bool),X631),s(X630,sK17(X631,X630))))) | s(bool,X632) = s(bool,X633) | s(bool,X633) = s(bool,'|dollar|exists_1'(s(fun(X630,bool),X631))) | s(bool,X632) = s(bool,X634) | s(bool,'T_0') = s(bool,X634)) )),
  inference(subsumption_resolution,[],[f29377,f263])).
fof(f29377,plain,(
  ( ! [X633,X631,X632,X634,X630] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X630,bool),X631),s(X630,sK17(X631,X630))))) | s(bool,X632) = s(bool,X633) | s(bool,X633) = s(bool,'|dollar|exists_1'(s(fun(X630,bool),X631))) | s(bool,X632) = s(bool,X634) | s(bool,'T_0') = s(bool,X634)) )),
  inference(superposition,[],[f1831,f2054])).
fof(f31973,plain,(
  ( ! [X629,X625,X627,X628,X626] : (p(s(bool,happ(s(fun(X625,bool),X626),s(X625,sK16(X626,X625))))) | s(bool,X627) = s(bool,X628) | s(bool,X628) = s(bool,'|dollar|exists_1'(s(fun(X625,bool),X626))) | s(bool,X627) = s(bool,X629) | s(bool,'T_0') = s(bool,X629)) )),
  inference(subsumption_resolution,[],[f29376,f263])).
fof(f29376,plain,(
  ( ! [X629,X625,X627,X628,X626] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X625,bool),X626),s(X625,sK16(X626,X625))))) | s(bool,X627) = s(bool,X628) | s(bool,X628) = s(bool,'|dollar|exists_1'(s(fun(X625,bool),X626))) | s(bool,X627) = s(bool,X629) | s(bool,'T_0') = s(bool,X629)) )),
  inference(superposition,[],[f1817,f2054])).
fof(f31972,plain,(
  ( ! [X621,X623,X619,X624,X620,X622,X618] : (p(s(bool,X620)) | ~p(s(bool,X621)) | ~sP11(X621,X620,happ(s(fun(X618,bool),X619),s(X618,'const.min.@_1'(s(fun(X618,bool),X619))))) | s(bool,X622) = s(bool,X623) | s(bool,X623) = s(bool,'|dollar|exists_1'(s(fun(X618,bool),X619))) | s(bool,X622) = s(bool,X624) | s(bool,'T_0') = s(bool,X624)) )),
  inference(subsumption_resolution,[],[f29375,f263])).
fof(f29375,plain,(
  ( ! [X621,X623,X619,X624,X620,X622,X618] : (~p(s(bool,'T_0')) | p(s(bool,X620)) | ~p(s(bool,X621)) | ~sP11(X621,X620,happ(s(fun(X618,bool),X619),s(X618,'const.min.@_1'(s(fun(X618,bool),X619))))) | s(bool,X622) = s(bool,X623) | s(bool,X623) = s(bool,'|dollar|exists_1'(s(fun(X618,bool),X619))) | s(bool,X622) = s(bool,X624) | s(bool,'T_0') = s(bool,X624)) )),
  inference(superposition,[],[f1792,f2054])).
fof(f31971,plain,(
  ( ! [X617,X613,X615,X611,X616,X612,X614] : (p(s(bool,X613)) | p(s(bool,X614)) | ~sP9(happ(s(fun(X611,bool),X612),s(X611,'const.min.@_1'(s(fun(X611,bool),X612)))),X613,X614) | s(bool,X615) = s(bool,X616) | s(bool,X616) = s(bool,'|dollar|exists_1'(s(fun(X611,bool),X612))) | s(bool,X615) = s(bool,X617) | s(bool,'T_0') = s(bool,X617)) )),
  inference(subsumption_resolution,[],[f29374,f263])).
fof(f29374,plain,(
  ( ! [X617,X613,X615,X611,X616,X612,X614] : (~p(s(bool,'T_0')) | p(s(bool,X613)) | p(s(bool,X614)) | ~sP9(happ(s(fun(X611,bool),X612),s(X611,'const.min.@_1'(s(fun(X611,bool),X612)))),X613,X614) | s(bool,X615) = s(bool,X616) | s(bool,X616) = s(bool,'|dollar|exists_1'(s(fun(X611,bool),X612))) | s(bool,X615) = s(bool,X617) | s(bool,'T_0') = s(bool,X617)) )),
  inference(superposition,[],[f1788,f2054])).
fof(f31970,plain,(
  ( ! [X606,X610,X609,X605,X607,X608,X604] : (~p(s(bool,X606)) | p(s(bool,X607)) | ~sP8(X607,X606,happ(s(fun(X604,bool),X605),s(X604,'const.min.@_1'(s(fun(X604,bool),X605))))) | s(bool,X608) = s(bool,X609) | s(bool,X609) = s(bool,'|dollar|exists_1'(s(fun(X604,bool),X605))) | s(bool,X608) = s(bool,X610) | s(bool,'T_0') = s(bool,X610)) )),
  inference(subsumption_resolution,[],[f29373,f263])).
fof(f29373,plain,(
  ( ! [X606,X610,X609,X605,X607,X608,X604] : (~p(s(bool,'T_0')) | ~p(s(bool,X606)) | p(s(bool,X607)) | ~sP8(X607,X606,happ(s(fun(X604,bool),X605),s(X604,'const.min.@_1'(s(fun(X604,bool),X605))))) | s(bool,X608) = s(bool,X609) | s(bool,X609) = s(bool,'|dollar|exists_1'(s(fun(X604,bool),X605))) | s(bool,X608) = s(bool,X610) | s(bool,'T_0') = s(bool,X610)) )),
  inference(superposition,[],[f1787,f2054])).
fof(f31969,plain,(
  ( ! [X600,X602,X598,X601,X603,X597,X599] : (~p(s(bool,X599)) | p(s(bool,X600)) | ~sP7(X599,happ(s(fun(X597,bool),X598),s(X597,'const.min.@_1'(s(fun(X597,bool),X598)))),X600) | s(bool,X601) = s(bool,X602) | s(bool,X602) = s(bool,'|dollar|exists_1'(s(fun(X597,bool),X598))) | s(bool,X601) = s(bool,X603) | s(bool,'T_0') = s(bool,X603)) )),
  inference(subsumption_resolution,[],[f29372,f263])).
fof(f29372,plain,(
  ( ! [X600,X602,X598,X601,X603,X597,X599] : (~p(s(bool,'T_0')) | ~p(s(bool,X599)) | p(s(bool,X600)) | ~sP7(X599,happ(s(fun(X597,bool),X598),s(X597,'const.min.@_1'(s(fun(X597,bool),X598)))),X600) | s(bool,X601) = s(bool,X602) | s(bool,X602) = s(bool,'|dollar|exists_1'(s(fun(X597,bool),X598))) | s(bool,X601) = s(bool,X603) | s(bool,'T_0') = s(bool,X603)) )),
  inference(superposition,[],[f1786,f2054])).
fof(f31968,plain,(
  ( ! [X596,X592,X594,X590,X593,X595,X591] : (~p(s(bool,X592)) | p(s(bool,X593)) | ~sP4(X592,happ(s(fun(X590,bool),X591),s(X590,'const.min.@_1'(s(fun(X590,bool),X591)))),X593) | s(bool,X594) = s(bool,X595) | s(bool,X595) = s(bool,'|dollar|exists_1'(s(fun(X590,bool),X591))) | s(bool,X594) = s(bool,X596) | s(bool,'T_0') = s(bool,X596)) )),
  inference(subsumption_resolution,[],[f29371,f263])).
fof(f29371,plain,(
  ( ! [X596,X592,X594,X590,X593,X595,X591] : (~p(s(bool,'T_0')) | ~p(s(bool,X592)) | p(s(bool,X593)) | ~sP4(X592,happ(s(fun(X590,bool),X591),s(X590,'const.min.@_1'(s(fun(X590,bool),X591)))),X593) | s(bool,X594) = s(bool,X595) | s(bool,X595) = s(bool,'|dollar|exists_1'(s(fun(X590,bool),X591))) | s(bool,X594) = s(bool,X596) | s(bool,'T_0') = s(bool,X596)) )),
  inference(superposition,[],[f1785,f2054])).
fof(f31967,plain,(
  ( ! [X588,X584,X586,X589,X585,X587,X583] : (~p(s(bool,X585)) | p(s(bool,X586)) | ~sP5(X585,happ(s(fun(X583,bool),X584),s(X583,'const.min.@_1'(s(fun(X583,bool),X584)))),X586) | s(bool,X587) = s(bool,X588) | s(bool,X588) = s(bool,'|dollar|exists_1'(s(fun(X583,bool),X584))) | s(bool,X587) = s(bool,X589) | s(bool,'T_0') = s(bool,X589)) )),
  inference(subsumption_resolution,[],[f29370,f263])).
fof(f29370,plain,(
  ( ! [X588,X584,X586,X589,X585,X587,X583] : (~p(s(bool,'T_0')) | ~p(s(bool,X585)) | p(s(bool,X586)) | ~sP5(X585,happ(s(fun(X583,bool),X584),s(X583,'const.min.@_1'(s(fun(X583,bool),X584)))),X586) | s(bool,X587) = s(bool,X588) | s(bool,X588) = s(bool,'|dollar|exists_1'(s(fun(X583,bool),X584))) | s(bool,X587) = s(bool,X589) | s(bool,'T_0') = s(bool,X589)) )),
  inference(superposition,[],[f1784,f2054])).
fof(f31966,plain,(
  ( ! [X580,X582,X576,X578,X581,X577,X579] : (sP6(X578,X579,happ(s(fun(X576,bool),X577),s(X576,'const.min.@_1'(s(fun(X576,bool),X577))))) | ~sP5(X578,happ(s(fun(X576,bool),X577),s(X576,'const.min.@_1'(s(fun(X576,bool),X577)))),X579) | ~sP4(X579,happ(s(fun(X576,bool),X577),s(X576,'const.min.@_1'(s(fun(X576,bool),X577)))),X578) | ~sP3(happ(s(fun(X576,bool),X577),s(X576,'const.min.@_1'(s(fun(X576,bool),X577)))),X579,X578) | s(bool,X580) = s(bool,X581) | s(bool,X581) = s(bool,'|dollar|exists_1'(s(fun(X576,bool),X577))) | s(bool,X580) = s(bool,X582) | s(bool,'T_0') = s(bool,X582)) )),
  inference(subsumption_resolution,[],[f29369,f263])).
fof(f29369,plain,(
  ( ! [X580,X582,X576,X578,X581,X577,X579] : (~p(s(bool,'T_0')) | sP6(X578,X579,happ(s(fun(X576,bool),X577),s(X576,'const.min.@_1'(s(fun(X576,bool),X577))))) | ~sP5(X578,happ(s(fun(X576,bool),X577),s(X576,'const.min.@_1'(s(fun(X576,bool),X577)))),X579) | ~sP4(X579,happ(s(fun(X576,bool),X577),s(X576,'const.min.@_1'(s(fun(X576,bool),X577)))),X578) | ~sP3(happ(s(fun(X576,bool),X577),s(X576,'const.min.@_1'(s(fun(X576,bool),X577)))),X579,X578) | s(bool,X580) = s(bool,X581) | s(bool,X581) = s(bool,'|dollar|exists_1'(s(fun(X576,bool),X577))) | s(bool,X580) = s(bool,X582) | s(bool,'T_0') = s(bool,X582)) )),
  inference(superposition,[],[f1783,f2054])).
fof(f31965,plain,(
  ( ! [X573,X575,X572,X574] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.min.@_1'(s(fun('type.extreal.extreal',bool),X572))),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X572)))))) | s(bool,X573) = s(bool,X574) | s(bool,X574) = s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),X572))) | s(bool,X573) = s(bool,X575) | s(bool,'T_0') = s(bool,X575)) )),
  inference(subsumption_resolution,[],[f29368,f263])).
fof(f29368,plain,(
  ( ! [X573,X575,X572,X574] : (~p(s(bool,'T_0')) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.min.@_1'(s(fun('type.extreal.extreal',bool),X572))),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X572)))))) | s(bool,X573) = s(bool,X574) | s(bool,X574) = s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),X572))) | s(bool,X573) = s(bool,X575) | s(bool,'T_0') = s(bool,X575)) )),
  inference(superposition,[],[f1840,f2054])).
fof(f31964,plain,(
  ( ! [X553,X555,X551,X556,X552,X554] : (s(bool,X553) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X552),s('type.extreal.extreal',X551))) | s(bool,'T_0') = s(bool,X553) | s(bool,X554) = s(bool,X555) | s(bool,X555) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X551),s('type.extreal.extreal',X552))) | s(bool,X554) = s(bool,X556) | s(bool,'T_0') = s(bool,X556)) )),
  inference(subsumption_resolution,[],[f29363,f263])).
fof(f29363,plain,(
  ( ! [X553,X555,X551,X556,X552,X554] : (~p(s(bool,'T_0')) | s(bool,X553) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X552),s('type.extreal.extreal',X551))) | s(bool,'T_0') = s(bool,X553) | s(bool,X554) = s(bool,X555) | s(bool,X555) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X551),s('type.extreal.extreal',X552))) | s(bool,X554) = s(bool,X556) | s(bool,'T_0') = s(bool,X556)) )),
  inference(superposition,[],[f2635,f2054])).
fof(f31963,plain,(
  ( ! [X546,X542,X545,X541,X543,X544] : (p(s(bool,happ(s(fun(bool,bool),X541),s(bool,X542)))) | ~p(s(bool,X543)) | s(bool,'T_0') = s(bool,sK18(X541,X543,bool)) | s(bool,X544) = s(bool,X545) | s(bool,X545) = s(bool,happ(s(fun(bool,bool),X541),s(bool,'F_0'))) | s(bool,X544) = s(bool,X546) | s(bool,'T_0') = s(bool,X546)) )),
  inference(subsumption_resolution,[],[f29361,f263])).
fof(f29361,plain,(
  ( ! [X546,X542,X545,X541,X543,X544] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),X541),s(bool,X542)))) | ~p(s(bool,X543)) | s(bool,'T_0') = s(bool,sK18(X541,X543,bool)) | s(bool,X544) = s(bool,X545) | s(bool,X545) = s(bool,happ(s(fun(bool,bool),X541),s(bool,'F_0'))) | s(bool,X544) = s(bool,X546) | s(bool,'T_0') = s(bool,X546)) )),
  inference(superposition,[],[f1898,f2054])).
fof(f31962,plain,(
  ( ! [X540,X536,X538,X537,X539,X535] : (p(s(bool,happ(s(fun(bool,bool),X535),s(bool,X536)))) | ~p(s(bool,X537)) | s(bool,'T_0') = s(bool,sK19(X535,X537,bool)) | s(bool,X538) = s(bool,X539) | s(bool,X539) = s(bool,happ(s(fun(bool,bool),X535),s(bool,'F_0'))) | s(bool,X538) = s(bool,X540) | s(bool,'T_0') = s(bool,X540)) )),
  inference(subsumption_resolution,[],[f29360,f263])).
fof(f29360,plain,(
  ( ! [X540,X536,X538,X537,X539,X535] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),X535),s(bool,X536)))) | ~p(s(bool,X537)) | s(bool,'T_0') = s(bool,sK19(X535,X537,bool)) | s(bool,X538) = s(bool,X539) | s(bool,X539) = s(bool,happ(s(fun(bool,bool),X535),s(bool,'F_0'))) | s(bool,X538) = s(bool,X540) | s(bool,'T_0') = s(bool,X540)) )),
  inference(superposition,[],[f1923,f2054])).
fof(f31961,plain,(
  ( ! [X532,X534,X530,X533,X529,X531] : (p(s(bool,X530)) | p(s(bool,happ(s(fun(bool,bool),X529),s(bool,X531)))) | s(bool,'T_0') = s(bool,sK20(X529,X530,bool)) | s(bool,X532) = s(bool,X533) | s(bool,X533) = s(bool,happ(s(fun(bool,bool),X529),s(bool,'F_0'))) | s(bool,X532) = s(bool,X534) | s(bool,'T_0') = s(bool,X534)) )),
  inference(subsumption_resolution,[],[f29359,f263])).
fof(f29359,plain,(
  ( ! [X532,X534,X530,X533,X529,X531] : (~p(s(bool,'T_0')) | p(s(bool,X530)) | p(s(bool,happ(s(fun(bool,bool),X529),s(bool,X531)))) | s(bool,'T_0') = s(bool,sK20(X529,X530,bool)) | s(bool,X532) = s(bool,X533) | s(bool,X533) = s(bool,happ(s(fun(bool,bool),X529),s(bool,'F_0'))) | s(bool,X532) = s(bool,X534) | s(bool,'T_0') = s(bool,X534)) )),
  inference(superposition,[],[f1947,f2054])).
fof(f31960,plain,(
  ( ! [X528,X524,X526,X525,X527,X523] : (p(s(bool,X524)) | p(s(bool,happ(s(fun(bool,bool),X523),s(bool,X525)))) | s(bool,'T_0') = s(bool,sK21(X523,X524,bool)) | s(bool,X526) = s(bool,X527) | s(bool,X527) = s(bool,happ(s(fun(bool,bool),X523),s(bool,'F_0'))) | s(bool,X526) = s(bool,X528) | s(bool,'T_0') = s(bool,X528)) )),
  inference(subsumption_resolution,[],[f29358,f263])).
fof(f29358,plain,(
  ( ! [X528,X524,X526,X525,X527,X523] : (~p(s(bool,'T_0')) | p(s(bool,X524)) | p(s(bool,happ(s(fun(bool,bool),X523),s(bool,X525)))) | s(bool,'T_0') = s(bool,sK21(X523,X524,bool)) | s(bool,X526) = s(bool,X527) | s(bool,X527) = s(bool,happ(s(fun(bool,bool),X523),s(bool,'F_0'))) | s(bool,X526) = s(bool,X528) | s(bool,'T_0') = s(bool,X528)) )),
  inference(superposition,[],[f1959,f2054])).
fof(f31959,plain,(
  ( ! [X520,X522,X516,X518,X521,X517,X519] : (p(s(bool,X518)) | p(s(bool,happ(s(fun(X516,bool),X517),s(X516,X519)))) | s(bool,X520) = s(bool,X521) | s(bool,X521) = s(bool,happ(s(fun(X516,bool),X517),s(X516,sK21(X517,X518,X516)))) | s(bool,X520) = s(bool,X522) | s(bool,'T_0') = s(bool,X522)) )),
  inference(subsumption_resolution,[],[f29357,f263])).
fof(f29357,plain,(
  ( ! [X520,X522,X516,X518,X521,X517,X519] : (~p(s(bool,'T_0')) | p(s(bool,X518)) | p(s(bool,happ(s(fun(X516,bool),X517),s(X516,X519)))) | s(bool,X520) = s(bool,X521) | s(bool,X521) = s(bool,happ(s(fun(X516,bool),X517),s(X516,sK21(X517,X518,X516)))) | s(bool,X520) = s(bool,X522) | s(bool,'T_0') = s(bool,X522)) )),
  inference(superposition,[],[f316,f2054])).
fof(f31958,plain,(
  ( ! [X509,X511,X512,X514,X510,X513,X515] : (p(s(bool,X511)) | p(s(bool,happ(s(fun(X509,bool),X510),s(X509,X512)))) | s(bool,X513) = s(bool,X514) | s(bool,X514) = s(bool,happ(s(fun(X509,bool),X510),s(X509,sK20(X510,X511,X509)))) | s(bool,X513) = s(bool,X515) | s(bool,'T_0') = s(bool,X515)) )),
  inference(subsumption_resolution,[],[f29356,f263])).
fof(f29356,plain,(
  ( ! [X509,X511,X512,X514,X510,X513,X515] : (~p(s(bool,'T_0')) | p(s(bool,X511)) | p(s(bool,happ(s(fun(X509,bool),X510),s(X509,X512)))) | s(bool,X513) = s(bool,X514) | s(bool,X514) = s(bool,happ(s(fun(X509,bool),X510),s(X509,sK20(X510,X511,X509)))) | s(bool,X513) = s(bool,X515) | s(bool,'T_0') = s(bool,X515)) )),
  inference(superposition,[],[f314,f2054])).
fof(f31957,plain,(
  ( ! [X505,X507,X502,X508,X504,X506,X503] : (p(s(bool,happ(s(fun(X502,bool),X503),s(X502,X505)))) | ~p(s(bool,X504)) | s(bool,X506) = s(bool,X507) | s(bool,X507) = s(bool,happ(s(fun(X502,bool),X503),s(X502,sK19(X503,X504,X502)))) | s(bool,X506) = s(bool,X508) | s(bool,'T_0') = s(bool,X508)) )),
  inference(subsumption_resolution,[],[f29355,f263])).
fof(f29355,plain,(
  ( ! [X505,X507,X502,X508,X504,X506,X503] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X502,bool),X503),s(X502,X505)))) | ~p(s(bool,X504)) | s(bool,X506) = s(bool,X507) | s(bool,X507) = s(bool,happ(s(fun(X502,bool),X503),s(X502,sK19(X503,X504,X502)))) | s(bool,X506) = s(bool,X508) | s(bool,'T_0') = s(bool,X508)) )),
  inference(superposition,[],[f312,f2054])).
fof(f31956,plain,(
  ( ! [X499,X500,X496,X498,X495,X501,X497] : (p(s(bool,happ(s(fun(X495,bool),X496),s(X495,X498)))) | ~p(s(bool,X497)) | s(bool,X499) = s(bool,X500) | s(bool,X500) = s(bool,happ(s(fun(X495,bool),X496),s(X495,sK18(X496,X497,X495)))) | s(bool,X499) = s(bool,X501) | s(bool,'T_0') = s(bool,X501)) )),
  inference(subsumption_resolution,[],[f29354,f263])).
fof(f29354,plain,(
  ( ! [X499,X500,X496,X498,X495,X501,X497] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X495,bool),X496),s(X495,X498)))) | ~p(s(bool,X497)) | s(bool,X499) = s(bool,X500) | s(bool,X500) = s(bool,happ(s(fun(X495,bool),X496),s(X495,sK18(X496,X497,X495)))) | s(bool,X499) = s(bool,X501) | s(bool,'T_0') = s(bool,X501)) )),
  inference(superposition,[],[f310,f2054])).
fof(f31955,plain,(
  ( ! [X455,X457,X459,X460,X456,X458] : (p(s(bool,'|dollar|exists_1'(s(fun(X455,bool),X456)))) | s(bool,X458) = s(bool,X459) | s(bool,X459) = s(bool,happ(s(fun(X455,bool),X456),s(X455,X457))) | s(bool,X458) = s(bool,X460) | s(bool,'T_0') = s(bool,X460)) )),
  inference(subsumption_resolution,[],[f29347,f263])).
fof(f29347,plain,(
  ( ! [X455,X457,X459,X460,X456,X458] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X455,bool),X456)))) | s(bool,X458) = s(bool,X459) | s(bool,X459) = s(bool,happ(s(fun(X455,bool),X456),s(X455,X457))) | s(bool,X458) = s(bool,X460) | s(bool,'T_0') = s(bool,X460)) )),
  inference(superposition,[],[f469,f2054])).
fof(f31954,plain,(
  ( ! [X453,X449,X451,X452,X454,X450] : (p(s(bool,happ(s(fun(X449,bool),X450),s(X449,sK16(X450,X449))))) | s(bool,X452) = s(bool,X453) | s(bool,X453) = s(bool,happ(s(fun(X449,bool),X450),s(X449,X451))) | s(bool,X452) = s(bool,X454) | s(bool,'T_0') = s(bool,X454)) )),
  inference(subsumption_resolution,[],[f29346,f263])).
fof(f29346,plain,(
  ( ! [X453,X449,X451,X452,X454,X450] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X449,bool),X450),s(X449,sK16(X450,X449))))) | s(bool,X452) = s(bool,X453) | s(bool,X453) = s(bool,happ(s(fun(X449,bool),X450),s(X449,X451))) | s(bool,X452) = s(bool,X454) | s(bool,'T_0') = s(bool,X454)) )),
  inference(superposition,[],[f304,f2054])).
fof(f31953,plain,(
  ( ! [X445,X447,X443,X444,X446,X448] : (p(s(bool,happ(s(fun(X443,bool),X444),s(X443,sK17(X444,X443))))) | s(bool,X446) = s(bool,X447) | s(bool,X447) = s(bool,happ(s(fun(X443,bool),X444),s(X443,X445))) | s(bool,X446) = s(bool,X448) | s(bool,'T_0') = s(bool,X448)) )),
  inference(subsumption_resolution,[],[f29345,f263])).
fof(f29345,plain,(
  ( ! [X445,X447,X443,X444,X446,X448] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X443,bool),X444),s(X443,sK17(X444,X443))))) | s(bool,X446) = s(bool,X447) | s(bool,X447) = s(bool,happ(s(fun(X443,bool),X444),s(X443,X445))) | s(bool,X446) = s(bool,X448) | s(bool,'T_0') = s(bool,X448)) )),
  inference(superposition,[],[f305,f2054])).
fof(f31952,plain,(
  ( ! [X428,X430,X432,X429,X431] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X429),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X428)))))) | s(bool,X430) = s(bool,X431) | s(bool,X431) = s(bool,happ(s(fun('type.extreal.extreal',bool),X428),s('type.extreal.extreal',X429))) | s(bool,X430) = s(bool,X432) | s(bool,'T_0') = s(bool,X432)) )),
  inference(subsumption_resolution,[],[f29342,f263])).
fof(f29342,plain,(
  ( ! [X428,X430,X432,X429,X431] : (~p(s(bool,'T_0')) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X429),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X428)))))) | s(bool,X430) = s(bool,X431) | s(bool,X431) = s(bool,happ(s(fun('type.extreal.extreal',bool),X428),s('type.extreal.extreal',X429))) | s(bool,X430) = s(bool,X432) | s(bool,'T_0') = s(bool,X432)) )),
  inference(superposition,[],[f317,f2054])).
fof(f31948,plain,(
  ( ! [X239,X237,X235,X238,X236,X234] : (s(bool,X234) = s(bool,'|dollar|exists_1'(s(fun(X235,bool),X236))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X235,bool),X236))) | s(bool,X237) = s(bool,X238) | s(bool,X234) = s(bool,X238) | s(bool,X237) = s(bool,X239) | s(bool,'T_0') = s(bool,X239)) )),
  inference(subsumption_resolution,[],[f29299,f263])).
fof(f29299,plain,(
  ( ! [X239,X237,X235,X238,X236,X234] : (~p(s(bool,'T_0')) | s(bool,X234) = s(bool,'|dollar|exists_1'(s(fun(X235,bool),X236))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X235,bool),X236))) | s(bool,X237) = s(bool,X238) | s(bool,X234) = s(bool,X238) | s(bool,X237) = s(bool,X239) | s(bool,'T_0') = s(bool,X239)) )),
  inference(superposition,[],[f4971,f2054])).
fof(f31947,plain,(
  ( ! [X231,X229,X227,X233,X230,X228,X232] : (p(s(bool,X228)) | ~sP12(X229,X228,X227) | s(bool,X229) = s(bool,X230) | s(bool,'T_0') = s(bool,X230) | s(bool,X231) = s(bool,X232) | s(bool,X227) = s(bool,X232) | s(bool,X231) = s(bool,X233) | s(bool,'T_0') = s(bool,X233)) )),
  inference(subsumption_resolution,[],[f29298,f263])).
fof(f29298,plain,(
  ( ! [X231,X229,X227,X233,X230,X228,X232] : (~p(s(bool,'T_0')) | p(s(bool,X228)) | ~sP12(X229,X228,X227) | s(bool,X229) = s(bool,X230) | s(bool,'T_0') = s(bool,X230) | s(bool,X231) = s(bool,X232) | s(bool,X227) = s(bool,X232) | s(bool,X231) = s(bool,X233) | s(bool,'T_0') = s(bool,X233)) )),
  inference(superposition,[],[f2623,f2054])).
fof(f31946,plain,(
  ( ! [X222,X220,X225,X223,X221,X226,X224] : (sP12(X221,X220,X222) | s(bool,X221) = s(bool,X223) | s(bool,'T_0') = s(bool,X223) | s(bool,X224) = s(bool,X225) | s(bool,X220) = s(bool,X225) | s(bool,X224) = s(bool,X226) | s(bool,'T_0') = s(bool,X226)) )),
  inference(subsumption_resolution,[],[f29297,f263])).
fof(f29297,plain,(
  ( ! [X222,X220,X225,X223,X221,X226,X224] : (~p(s(bool,'T_0')) | sP12(X221,X220,X222) | s(bool,X221) = s(bool,X223) | s(bool,'T_0') = s(bool,X223) | s(bool,X224) = s(bool,X225) | s(bool,X220) = s(bool,X225) | s(bool,X224) = s(bool,X226) | s(bool,'T_0') = s(bool,X226)) )),
  inference(superposition,[],[f2622,f2054])).
fof(f31945,plain,(
  ( ! [X218,X216,X215,X213,X219,X217,X214] : (sP10(X214,X213,X215) | s(bool,X214) = s(bool,X216) | s(bool,'T_0') = s(bool,X216) | s(bool,X217) = s(bool,X218) | s(bool,X213) = s(bool,X218) | s(bool,X217) = s(bool,X219) | s(bool,'T_0') = s(bool,X219)) )),
  inference(subsumption_resolution,[],[f29296,f263])).
fof(f29296,plain,(
  ( ! [X218,X216,X215,X213,X219,X217,X214] : (~p(s(bool,'T_0')) | sP10(X214,X213,X215) | s(bool,X214) = s(bool,X216) | s(bool,'T_0') = s(bool,X216) | s(bool,X217) = s(bool,X218) | s(bool,X213) = s(bool,X218) | s(bool,X217) = s(bool,X219) | s(bool,'T_0') = s(bool,X219)) )),
  inference(superposition,[],[f2618,f2054])).
fof(f31944,plain,(
  ( ! [X212,X210,X208,X206,X211,X209,X207] : (sP10(X207,X208,X206) | s(bool,X207) = s(bool,X209) | s(bool,'T_0') = s(bool,X209) | s(bool,X210) = s(bool,X211) | s(bool,X206) = s(bool,X211) | s(bool,X210) = s(bool,X212) | s(bool,'T_0') = s(bool,X212)) )),
  inference(subsumption_resolution,[],[f29295,f263])).
fof(f29295,plain,(
  ( ! [X212,X210,X208,X206,X211,X209,X207] : (~p(s(bool,'T_0')) | sP10(X207,X208,X206) | s(bool,X207) = s(bool,X209) | s(bool,'T_0') = s(bool,X209) | s(bool,X210) = s(bool,X211) | s(bool,X206) = s(bool,X211) | s(bool,X210) = s(bool,X212) | s(bool,'T_0') = s(bool,X212)) )),
  inference(superposition,[],[f2617,f2054])).
fof(f31943,plain,(
  ( ! [X204,X202,X200,X199,X205,X203,X201] : (p(s(bool,X200)) | ~sP8(X200,X199,X201) | s(bool,X201) = s(bool,X202) | s(bool,'T_0') = s(bool,X202) | s(bool,X203) = s(bool,X204) | s(bool,X199) = s(bool,X204) | s(bool,X203) = s(bool,X205) | s(bool,'T_0') = s(bool,X205)) )),
  inference(subsumption_resolution,[],[f29294,f263])).
fof(f29294,plain,(
  ( ! [X204,X202,X200,X199,X205,X203,X201] : (~p(s(bool,'T_0')) | p(s(bool,X200)) | ~sP8(X200,X199,X201) | s(bool,X201) = s(bool,X202) | s(bool,'T_0') = s(bool,X202) | s(bool,X203) = s(bool,X204) | s(bool,X199) = s(bool,X204) | s(bool,X203) = s(bool,X205) | s(bool,'T_0') = s(bool,X205)) )),
  inference(superposition,[],[f2615,f2054])).
fof(f31942,plain,(
  ( ! [X144,X142,X140,X143,X141,X139] : (~sP9(X139,X140,X141) | sP10(X139,X140,X141) | s(bool,X142) = s(bool,X143) | s(bool,X139) = s(bool,X143) | s(bool,X142) = s(bool,X144) | s(bool,'T_0') = s(bool,X144)) )),
  inference(subsumption_resolution,[],[f29283,f263])).
fof(f29283,plain,(
  ( ! [X144,X142,X140,X143,X141,X139] : (~p(s(bool,'T_0')) | ~sP9(X139,X140,X141) | sP10(X139,X140,X141) | s(bool,X142) = s(bool,X143) | s(bool,X139) = s(bool,X143) | s(bool,X142) = s(bool,X144) | s(bool,'T_0') = s(bool,X144)) )),
  inference(superposition,[],[f453,f2054])).
fof(f31941,plain,(
  ( ! [X134,X138,X136,X135,X133,X137] : (~sP11(X133,X134,X135) | sP12(X133,X134,X135) | s(bool,X136) = s(bool,X137) | s(bool,X133) = s(bool,X137) | s(bool,X136) = s(bool,X138) | s(bool,'T_0') = s(bool,X138)) )),
  inference(subsumption_resolution,[],[f29282,f263])).
fof(f29282,plain,(
  ( ! [X134,X138,X136,X135,X133,X137] : (~p(s(bool,'T_0')) | ~sP11(X133,X134,X135) | sP12(X133,X134,X135) | s(bool,X136) = s(bool,X137) | s(bool,X133) = s(bool,X137) | s(bool,X136) = s(bool,X138) | s(bool,'T_0') = s(bool,X138)) )),
  inference(superposition,[],[f452,f2054])).
fof(f31940,plain,(
  ( ! [X127,X132,X130,X128,X131,X129] : (~p(s(bool,X128)) | p(s(bool,X129)) | ~sP12(X127,X129,X128) | s(bool,X130) = s(bool,X131) | s(bool,X127) = s(bool,X131) | s(bool,X130) = s(bool,X132) | s(bool,'T_0') = s(bool,X132)) )),
  inference(subsumption_resolution,[],[f29281,f263])).
fof(f29281,plain,(
  ( ! [X127,X132,X130,X128,X131,X129] : (~p(s(bool,'T_0')) | ~p(s(bool,X128)) | p(s(bool,X129)) | ~sP12(X127,X129,X128) | s(bool,X130) = s(bool,X131) | s(bool,X127) = s(bool,X131) | s(bool,X130) = s(bool,X132) | s(bool,'T_0') = s(bool,X132)) )),
  inference(superposition,[],[f449,f2054])).
fof(f31939,plain,(
  ( ! [X125,X123,X121,X126,X124,X122] : (~p(s(bool,X122)) | sP12(X121,X122,X123) | s(bool,X124) = s(bool,X125) | s(bool,X121) = s(bool,X125) | s(bool,X124) = s(bool,X126) | s(bool,'T_0') = s(bool,X126)) )),
  inference(subsumption_resolution,[],[f29280,f263])).
fof(f29280,plain,(
  ( ! [X125,X123,X121,X126,X124,X122] : (~p(s(bool,'T_0')) | ~p(s(bool,X122)) | sP12(X121,X122,X123) | s(bool,X124) = s(bool,X125) | s(bool,X121) = s(bool,X125) | s(bool,X124) = s(bool,X126) | s(bool,'T_0') = s(bool,X126)) )),
  inference(superposition,[],[f448,f2054])).
fof(f31938,plain,(
  ( ! [X118,X116,X120,X119,X117,X115] : (p(s(bool,X116)) | sP12(X115,X117,X116) | s(bool,X118) = s(bool,X119) | s(bool,X115) = s(bool,X119) | s(bool,X118) = s(bool,X120) | s(bool,'T_0') = s(bool,X120)) )),
  inference(subsumption_resolution,[],[f29279,f263])).
fof(f29279,plain,(
  ( ! [X118,X116,X120,X119,X117,X115] : (~p(s(bool,'T_0')) | p(s(bool,X116)) | sP12(X115,X117,X116) | s(bool,X118) = s(bool,X119) | s(bool,X115) = s(bool,X119) | s(bool,X118) = s(bool,X120) | s(bool,'T_0') = s(bool,X120)) )),
  inference(superposition,[],[f447,f2054])).
fof(f31937,plain,(
  ( ! [X111,X109,X114,X112,X110,X113] : (p(s(bool,X110)) | ~p(s(bool,X111)) | ~sP11(X111,X110,X109) | s(bool,X112) = s(bool,X113) | s(bool,X109) = s(bool,X113) | s(bool,X112) = s(bool,X114) | s(bool,'T_0') = s(bool,X114)) )),
  inference(subsumption_resolution,[],[f29278,f263])).
fof(f29278,plain,(
  ( ! [X111,X109,X114,X112,X110,X113] : (~p(s(bool,'T_0')) | p(s(bool,X110)) | ~p(s(bool,X111)) | ~sP11(X111,X110,X109) | s(bool,X112) = s(bool,X113) | s(bool,X109) = s(bool,X113) | s(bool,X112) = s(bool,X114) | s(bool,'T_0') = s(bool,X114)) )),
  inference(superposition,[],[f442,f2054])).
fof(f31936,plain,(
  ( ! [X103,X107,X105,X108,X106,X104] : (p(s(bool,X104)) | p(s(bool,X105)) | ~sP10(X103,X105,X104) | s(bool,X106) = s(bool,X107) | s(bool,X103) = s(bool,X107) | s(bool,X106) = s(bool,X108) | s(bool,'T_0') = s(bool,X108)) )),
  inference(subsumption_resolution,[],[f29277,f263])).
fof(f29277,plain,(
  ( ! [X103,X107,X105,X108,X106,X104] : (~p(s(bool,'T_0')) | p(s(bool,X104)) | p(s(bool,X105)) | ~sP10(X103,X105,X104) | s(bool,X106) = s(bool,X107) | s(bool,X103) = s(bool,X107) | s(bool,X106) = s(bool,X108) | s(bool,'T_0') = s(bool,X108)) )),
  inference(superposition,[],[f432,f2054])).
fof(f31935,plain,(
  ( ! [X101,X99,X97,X102,X100,X98] : (~p(s(bool,X98)) | sP10(X97,X98,X99) | s(bool,X100) = s(bool,X101) | s(bool,X97) = s(bool,X101) | s(bool,X100) = s(bool,X102) | s(bool,'T_0') = s(bool,X102)) )),
  inference(subsumption_resolution,[],[f29276,f263])).
fof(f29276,plain,(
  ( ! [X101,X99,X97,X102,X100,X98] : (~p(s(bool,'T_0')) | ~p(s(bool,X98)) | sP10(X97,X98,X99) | s(bool,X100) = s(bool,X101) | s(bool,X97) = s(bool,X101) | s(bool,X100) = s(bool,X102) | s(bool,'T_0') = s(bool,X102)) )),
  inference(superposition,[],[f431,f2054])).
fof(f31934,plain,(
  ( ! [X94,X92,X95,X93,X91,X96] : (~p(s(bool,X92)) | sP10(X91,X93,X92) | s(bool,X94) = s(bool,X95) | s(bool,X91) = s(bool,X95) | s(bool,X94) = s(bool,X96) | s(bool,'T_0') = s(bool,X96)) )),
  inference(subsumption_resolution,[],[f29275,f263])).
fof(f29275,plain,(
  ( ! [X94,X92,X95,X93,X91,X96] : (~p(s(bool,'T_0')) | ~p(s(bool,X92)) | sP10(X91,X93,X92) | s(bool,X94) = s(bool,X95) | s(bool,X91) = s(bool,X95) | s(bool,X94) = s(bool,X96) | s(bool,'T_0') = s(bool,X96)) )),
  inference(superposition,[],[f430,f2054])).
fof(f31933,plain,(
  ( ! [X90,X88,X87,X85,X89,X86] : (p(s(bool,X86)) | p(s(bool,X87)) | ~sP9(X85,X86,X87) | s(bool,X88) = s(bool,X89) | s(bool,X85) = s(bool,X89) | s(bool,X88) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f29274,f263])).
fof(f29274,plain,(
  ( ! [X90,X88,X87,X85,X89,X86] : (~p(s(bool,'T_0')) | p(s(bool,X86)) | p(s(bool,X87)) | ~sP9(X85,X86,X87) | s(bool,X88) = s(bool,X89) | s(bool,X85) = s(bool,X89) | s(bool,X88) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f425,f2054])).
fof(f31932,plain,(
  ( ! [X80,X83,X81,X79,X84,X82] : (~p(s(bool,X80)) | p(s(bool,X81)) | ~sP8(X81,X80,X79) | s(bool,X82) = s(bool,X83) | s(bool,X79) = s(bool,X83) | s(bool,X82) = s(bool,X84) | s(bool,'T_0') = s(bool,X84)) )),
  inference(subsumption_resolution,[],[f29273,f263])).
fof(f29273,plain,(
  ( ! [X80,X83,X81,X79,X84,X82] : (~p(s(bool,'T_0')) | ~p(s(bool,X80)) | p(s(bool,X81)) | ~sP8(X81,X80,X79) | s(bool,X82) = s(bool,X83) | s(bool,X79) = s(bool,X83) | s(bool,X82) = s(bool,X84) | s(bool,'T_0') = s(bool,X84)) )),
  inference(superposition,[],[f417,f2054])).
fof(f31931,plain,(
  ( ! [X78,X76,X74,X77,X75,X73] : (~p(s(bool,X74)) | p(s(bool,X75)) | ~sP7(X74,X73,X75) | s(bool,X76) = s(bool,X77) | s(bool,X73) = s(bool,X77) | s(bool,X76) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f29272,f263])).
fof(f29272,plain,(
  ( ! [X78,X76,X74,X77,X75,X73] : (~p(s(bool,'T_0')) | ~p(s(bool,X74)) | p(s(bool,X75)) | ~sP7(X74,X73,X75) | s(bool,X76) = s(bool,X77) | s(bool,X73) = s(bool,X77) | s(bool,X76) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f408,f2054])).
fof(f31930,plain,(
  ( ! [X70,X68,X72,X71,X69,X67] : (~p(s(bool,X68)) | p(s(bool,X69)) | ~sP4(X68,X67,X69) | s(bool,X70) = s(bool,X71) | s(bool,X67) = s(bool,X71) | s(bool,X70) = s(bool,X72) | s(bool,'T_0') = s(bool,X72)) )),
  inference(subsumption_resolution,[],[f29271,f263])).
fof(f29271,plain,(
  ( ! [X70,X68,X72,X71,X69,X67] : (~p(s(bool,'T_0')) | ~p(s(bool,X68)) | p(s(bool,X69)) | ~sP4(X68,X67,X69) | s(bool,X70) = s(bool,X71) | s(bool,X67) = s(bool,X71) | s(bool,X70) = s(bool,X72) | s(bool,'T_0') = s(bool,X72)) )),
  inference(superposition,[],[f389,f2054])).
fof(f31929,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (~p(s(bool,X62)) | p(s(bool,X63)) | ~sP5(X62,X61,X63) | s(bool,X64) = s(bool,X65) | s(bool,X61) = s(bool,X65) | s(bool,X64) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f29270,f263])).
fof(f29270,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (~p(s(bool,'T_0')) | ~p(s(bool,X62)) | p(s(bool,X63)) | ~sP5(X62,X61,X63) | s(bool,X64) = s(bool,X65) | s(bool,X61) = s(bool,X65) | s(bool,X64) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f385,f2054])).
fof(f31928,plain,(
  ( ! [X59,X57,X60,X58,X56,X55] : (sP6(X56,X57,X55) | ~sP5(X56,X55,X57) | ~sP4(X57,X55,X56) | ~sP3(X55,X57,X56) | s(bool,X58) = s(bool,X59) | s(bool,X55) = s(bool,X59) | s(bool,X58) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(subsumption_resolution,[],[f29269,f263])).
fof(f29269,plain,(
  ( ! [X59,X57,X60,X58,X56,X55] : (~p(s(bool,'T_0')) | sP6(X56,X57,X55) | ~sP5(X56,X55,X57) | ~sP4(X57,X55,X56) | ~sP3(X55,X57,X56) | s(bool,X58) = s(bool,X59) | s(bool,X55) = s(bool,X59) | s(bool,X58) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(superposition,[],[f382,f2054])).
fof(f31927,plain,(
  ( ! [X54,X52,X50,X53,X51] : (p(s(bool,X51)) | sP1(X51,X50) | s(bool,X52) = s(bool,X53) | s(bool,X50) = s(bool,X53) | s(bool,X52) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f29268,f263])).
fof(f29268,plain,(
  ( ! [X54,X52,X50,X53,X51] : (~p(s(bool,'T_0')) | p(s(bool,X51)) | sP1(X51,X50) | s(bool,X52) = s(bool,X53) | s(bool,X50) = s(bool,X53) | s(bool,X52) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f353,f2054])).
fof(f31926,plain,(
  ( ! [X47,X45,X48,X46,X49] : (p(s(bool,X46)) | sP1(X45,X46) | s(bool,X47) = s(bool,X48) | s(bool,X45) = s(bool,X48) | s(bool,X47) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(subsumption_resolution,[],[f29267,f263])).
fof(f29267,plain,(
  ( ! [X47,X45,X48,X46,X49] : (~p(s(bool,'T_0')) | p(s(bool,X46)) | sP1(X45,X46) | s(bool,X47) = s(bool,X48) | s(bool,X45) = s(bool,X48) | s(bool,X47) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(superposition,[],[f350,f2054])).
fof(f31925,plain,(
  ( ! [X43,X41,X44,X42,X40] : (sP0(X40,X41) | s(bool,X42) = s(bool,X43) | s(bool,X40) = s(bool,X43) | s(bool,X42) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(subsumption_resolution,[],[f29266,f263])).
fof(f29266,plain,(
  ( ! [X43,X41,X44,X42,X40] : (~p(s(bool,'T_0')) | sP0(X40,X41) | s(bool,X42) = s(bool,X43) | s(bool,X40) = s(bool,X43) | s(bool,X42) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f335,f2054])).
fof(f31924,plain,(
  ( ! [X39,X37,X35,X38,X36] : (sP0(X36,X35) | s(bool,X37) = s(bool,X38) | s(bool,X35) = s(bool,X38) | s(bool,X37) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(subsumption_resolution,[],[f29265,f263])).
fof(f29265,plain,(
  ( ! [X39,X37,X35,X38,X36] : (~p(s(bool,'T_0')) | sP0(X36,X35) | s(bool,X37) = s(bool,X38) | s(bool,X35) = s(bool,X38) | s(bool,X37) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f334,f2054])).
fof(f31923,plain,(
  ( ! [X30,X33,X31,X34,X32] : (~p(s(bool,X31)) | s(bool,X30) = s(bool,X31) | s(bool,X32) = s(bool,X33) | s(bool,X30) = s(bool,X33) | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f29264,f263])).
fof(f29264,plain,(
  ( ! [X30,X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X31)) | s(bool,X30) = s(bool,X31) | s(bool,X32) = s(bool,X33) | s(bool,X30) = s(bool,X33) | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f321,f2054])).
fof(f31922,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') != s(bool,X12) | s(bool,X13) = s(bool,X14) | s(bool,X12) = s(bool,X14) | s(bool,X13) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f29260,f263])).
fof(f29260,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X12) | s(bool,X13) = s(bool,X14) | s(bool,X12) = s(bool,X14) | s(bool,X13) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f280,f2054])).
fof(f29195,plain,(
  ( ! [X6,X4,X8,X7,X5] : (s(bool,X4) != s(bool,X5) | p(s(bool,X5)) | s(bool,X6) = s(bool,X7) | s(bool,X4) = s(bool,X7) | s(bool,X6) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f276,f2054])).
fof(f29187,plain,(
  ( ! [X947,X949,X950,X948] : (p(s(bool,happ(s(fun(bool,bool),X947),s(bool,X948)))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X947))) | s(bool,X948) = s(bool,X949) | s(bool,'T_0') = s(bool,sK17(X947,bool)) | s(bool,X949) = s(bool,X950) | s(bool,'T_0') = s(bool,X950)) )),
  inference(superposition,[],[f4966,f2054])).
fof(f29186,plain,(
  ( ! [X945,X942,X946,X944,X943] : (p(s(bool,happ(s(fun(bool,bool),X942),s(bool,X943)))) | s(bool,X944) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X942))) | s(bool,'T_0') = s(bool,X944) | s(bool,X943) = s(bool,X945) | s(bool,'T_0') = s(bool,sK17(X942,bool)) | s(bool,X945) = s(bool,X946) | s(bool,'T_0') = s(bool,X946)) )),
  inference(superposition,[],[f4303,f2054])).
fof(f29185,plain,(
  ( ! [X938,X940,X939,X941] : (p(s(bool,happ(s(fun(bool,bool),X938),s(bool,X939)))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X938))) | s(bool,X939) = s(bool,X940) | s(bool,'T_0') = s(bool,sK16(X938,bool)) | s(bool,X940) = s(bool,X941) | s(bool,'T_0') = s(bool,X941)) )),
  inference(superposition,[],[f4967,f2054])).
fof(f29184,plain,(
  ( ! [X936,X934,X937,X935,X933] : (p(s(bool,happ(s(fun(bool,bool),X933),s(bool,X934)))) | s(bool,X935) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X933))) | s(bool,'T_0') = s(bool,X935) | s(bool,X934) = s(bool,X936) | s(bool,'T_0') = s(bool,sK16(X933,bool)) | s(bool,X936) = s(bool,X937) | s(bool,'T_0') = s(bool,X937)) )),
  inference(superposition,[],[f3897,f2054])).
fof(f29175,plain,(
  ( ! [X895,X899,X897,X898,X896,X900] : (sP12(happ(s(fun(bool,bool),X895),s(bool,X896)),X897,X898) | ~p(s(bool,X897)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X895))) | s(bool,X896) = s(bool,X899) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X895))) | s(bool,X899) = s(bool,X900) | s(bool,'T_0') = s(bool,X900)) )),
  inference(superposition,[],[f4987,f2054])).
fof(f29174,plain,(
  ( ! [X893,X890,X894,X892,X891,X889] : (sP12(happ(s(fun(bool,bool),X889),s(bool,X890)),X891,X892) | p(s(bool,X892)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X889))) | s(bool,X890) = s(bool,X893) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X889))) | s(bool,X893) = s(bool,X894) | s(bool,'T_0') = s(bool,X894)) )),
  inference(superposition,[],[f4986,f2054])).
fof(f29173,plain,(
  ( ! [X883,X887,X885,X888,X886,X884] : (sP10(happ(s(fun(bool,bool),X883),s(bool,X884)),X885,X886) | ~p(s(bool,X885)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X883))) | s(bool,X884) = s(bool,X887) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X883))) | s(bool,X887) = s(bool,X888) | s(bool,'T_0') = s(bool,X888)) )),
  inference(superposition,[],[f4983,f2054])).
fof(f29172,plain,(
  ( ! [X881,X878,X882,X880,X879,X877] : (sP10(happ(s(fun(bool,bool),X877),s(bool,X878)),X879,X880) | ~p(s(bool,X880)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X877))) | s(bool,X878) = s(bool,X881) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X877))) | s(bool,X881) = s(bool,X882) | s(bool,'T_0') = s(bool,X882)) )),
  inference(superposition,[],[f4982,f2054])).
fof(f29171,plain,(
  ( ! [X874,X872,X876,X875,X873] : (sP1(X874,happ(s(fun(bool,bool),X872),s(bool,X873))) | p(s(bool,X874)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X872))) | s(bool,X873) = s(bool,X875) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X872))) | s(bool,X875) = s(bool,X876) | s(bool,'T_0') = s(bool,X876)) )),
  inference(superposition,[],[f4975,f2054])).
fof(f29170,plain,(
  ( ! [X870,X868,X867,X871,X869] : (sP1(happ(s(fun(bool,bool),X867),s(bool,X868)),X869) | p(s(bool,X869)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X867))) | s(bool,X868) = s(bool,X870) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X867))) | s(bool,X870) = s(bool,X871) | s(bool,'T_0') = s(bool,X871)) )),
  inference(superposition,[],[f4974,f2054])).
fof(f29169,plain,(
  ( ! [X866,X864,X863,X865,X862] : (sP0(happ(s(fun(bool,bool),X862),s(bool,X863)),X864) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X862))) | s(bool,X863) = s(bool,X865) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X862))) | s(bool,X865) = s(bool,X866) | s(bool,'T_0') = s(bool,X866)) )),
  inference(superposition,[],[f4973,f2054])).
fof(f29168,plain,(
  ( ! [X860,X859,X857,X861,X858] : (sP0(X859,happ(s(fun(bool,bool),X857),s(bool,X858))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X857))) | s(bool,X858) = s(bool,X860) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X857))) | s(bool,X860) = s(bool,X861) | s(bool,'T_0') = s(bool,X861)) )),
  inference(superposition,[],[f4972,f2054])).
fof(f29167,plain,(
  ( ! [X854,X855,X853,X856] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X853),s(bool,X854))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X853))) | s(bool,X854) = s(bool,X855) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X853))) | s(bool,X855) = s(bool,X856) | s(bool,'T_0') = s(bool,X856)) )),
  inference(superposition,[],[f4969,f2054])).
fof(f29166,plain,(
  ( ! [X850,X852,X851,X849] : (sP5('T_0',happ(s(fun(bool,bool),X849),s(bool,X850)),'|dollar|exists_1'(s(fun(bool,bool),X849))) | s(bool,X850) = s(bool,X851) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X849))) | s(bool,X851) = s(bool,X852) | s(bool,'T_0') = s(bool,X852)) )),
  inference(superposition,[],[f3862,f2054])).
fof(f29165,plain,(
  ( ! [X848,X847,X845,X846] : (sP5('T_0','|dollar|exists_1'(s(fun(bool,bool),X845)),happ(s(fun(bool,bool),X845),s(bool,X846))) | s(bool,X846) = s(bool,X847) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X845))) | s(bool,X847) = s(bool,X848) | s(bool,'T_0') = s(bool,X848)) )),
  inference(superposition,[],[f3757,f2054])).
fof(f29164,plain,(
  ( ! [X843,X841,X839,X842,X840,X844] : (sP10(happ(s(fun(bool,bool),X839),s(bool,X840)),X841,X842) | ~sP9(happ(s(fun(bool,bool),X839),s(bool,'const.min.@_1'(s(fun(bool,bool),X839)))),X841,X842) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X839)))) | s(bool,X840) = s(bool,X843) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X839))) | s(bool,X843) = s(bool,X844) | s(bool,'T_0') = s(bool,X844)) )),
  inference(superposition,[],[f1797,f2054])).
fof(f29163,plain,(
  ( ! [X835,X833,X837,X834,X838,X836] : (sP12(happ(s(fun(bool,bool),X833),s(bool,X834)),X835,X836) | ~sP11(happ(s(fun(bool,bool),X833),s(bool,'const.min.@_1'(s(fun(bool,bool),X833)))),X835,X836) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X833)))) | s(bool,X834) = s(bool,X837) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X833))) | s(bool,X837) = s(bool,X838) | s(bool,'T_0') = s(bool,X838)) )),
  inference(superposition,[],[f1796,f2054])).
fof(f29162,plain,(
  ( ! [X831,X829,X830,X828,X832,X827] : (~sP12(happ(s(fun(bool,bool),X827),s(bool,X828)),X829,X830) | ~p(s(bool,X830)) | p(s(bool,X829)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X827)))) | s(bool,X828) = s(bool,X831) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X827))) | s(bool,X831) = s(bool,X832) | s(bool,'T_0') = s(bool,X832)) )),
  inference(superposition,[],[f1795,f2054])).
fof(f29161,plain,(
  ( ! [X825,X823,X821,X826,X824,X822] : (sP12(happ(s(fun(bool,bool),X821),s(bool,X822)),X823,X824) | ~p(s(bool,X823)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X821)))) | s(bool,X822) = s(bool,X825) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X821))) | s(bool,X825) = s(bool,X826) | s(bool,'T_0') = s(bool,X826)) )),
  inference(superposition,[],[f1794,f2054])).
fof(f29160,plain,(
  ( ! [X819,X817,X818,X816,X820,X815] : (sP12(happ(s(fun(bool,bool),X815),s(bool,X816)),X817,X818) | p(s(bool,X818)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X815)))) | s(bool,X816) = s(bool,X819) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X815))) | s(bool,X819) = s(bool,X820) | s(bool,'T_0') = s(bool,X820)) )),
  inference(superposition,[],[f1793,f2054])).
fof(f29159,plain,(
  ( ! [X810,X814,X812,X811,X809,X813] : (~sP10(happ(s(fun(bool,bool),X809),s(bool,X810)),X811,X812) | p(s(bool,X812)) | p(s(bool,X811)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X809)))) | s(bool,X810) = s(bool,X813) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X809))) | s(bool,X813) = s(bool,X814) | s(bool,'T_0') = s(bool,X814)) )),
  inference(superposition,[],[f1791,f2054])).
fof(f29158,plain,(
  ( ! [X808,X806,X804,X803,X807,X805] : (sP10(happ(s(fun(bool,bool),X803),s(bool,X804)),X805,X806) | ~p(s(bool,X805)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X803)))) | s(bool,X804) = s(bool,X807) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X803))) | s(bool,X807) = s(bool,X808) | s(bool,'T_0') = s(bool,X808)) )),
  inference(superposition,[],[f1790,f2054])).
fof(f29157,plain,(
  ( ! [X798,X802,X800,X799,X797,X801] : (sP10(happ(s(fun(bool,bool),X797),s(bool,X798)),X799,X800) | ~p(s(bool,X800)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X797)))) | s(bool,X798) = s(bool,X801) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X797))) | s(bool,X801) = s(bool,X802) | s(bool,'T_0') = s(bool,X802)) )),
  inference(superposition,[],[f1789,f2054])).
fof(f29156,plain,(
  ( ! [X796,X795,X793,X794,X792] : (sP1(X794,happ(s(fun(bool,bool),X792),s(bool,X793))) | p(s(bool,X794)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X792)))) | s(bool,X793) = s(bool,X795) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X792))) | s(bool,X795) = s(bool,X796) | s(bool,'T_0') = s(bool,X796)) )),
  inference(superposition,[],[f1782,f2054])).
fof(f29155,plain,(
  ( ! [X790,X788,X787,X791,X789] : (sP1(happ(s(fun(bool,bool),X787),s(bool,X788)),X789) | p(s(bool,X789)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X787)))) | s(bool,X788) = s(bool,X790) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X787))) | s(bool,X790) = s(bool,X791) | s(bool,'T_0') = s(bool,X791)) )),
  inference(superposition,[],[f1781,f2054])).
fof(f29152,plain,(
  ( ! [X775,X773,X776,X774] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X773),s(bool,X774))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X773))) | s(bool,X774) = s(bool,X775) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X773))) | s(bool,X775) = s(bool,X776) | s(bool,'T_0') = s(bool,X776)) )),
  inference(superposition,[],[f1769,f2054])).
fof(f29145,plain,(
  ( ! [X745,X747,X743,X744,X746,X742] : (s(bool,X745) != s(bool,f(s(X742,sK23(X743,X744,bool,X742)))) | s(fun(X742,bool),X743) = s(fun(X742,bool),X744) | s(bool,X745) = s(bool,X746) | s(bool,'T_0') = s(bool,g(s(X742,sK23(X743,X744,bool,X742)))) | s(bool,X746) = s(bool,X747) | s(bool,'T_0') = s(bool,X747)) )),
  inference(superposition,[],[f366,f2054])).
fof(f31872,plain,(
  ( ! [X397,X399,X396,X398,X400] : (sP0(happ(s(fun(bool,bool),X397),s(bool,X396)),X398) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X397)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X397))) | s(bool,X396) = s(bool,X399) | s(bool,X399) = s(bool,X400) | s(bool,'T_0') = s(bool,X400)) )),
  inference(subsumption_resolution,[],[f29079,f525])).
fof(f29079,plain,(
  ( ! [X397,X399,X396,X398,X400] : (sP0(happ(s(fun(bool,bool),X397),s(bool,X396)),X398) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X397)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X397))) | s(bool,X396) = s(bool,X399) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X399) = s(bool,X400) | s(bool,'T_0') = s(bool,X400)) )),
  inference(superposition,[],[f4321,f2054])).
fof(f31871,plain,(
  ( ! [X391,X393,X395,X392,X394] : (sP0(X392,happ(s(fun(bool,bool),X393),s(bool,X391))) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X393)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X393))) | s(bool,X391) = s(bool,X394) | s(bool,X394) = s(bool,X395) | s(bool,'T_0') = s(bool,X395)) )),
  inference(subsumption_resolution,[],[f29078,f525])).
fof(f29078,plain,(
  ( ! [X391,X393,X395,X392,X394] : (sP0(X392,happ(s(fun(bool,bool),X393),s(bool,X391))) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X393)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X393))) | s(bool,X391) = s(bool,X394) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X394) = s(bool,X395) | s(bool,'T_0') = s(bool,X395)) )),
  inference(superposition,[],[f4312,f2054])).
fof(f31870,plain,(
  ( ! [X389,X387,X388,X390] : (sP6('T_0','|dollar|exists_1'(s(fun(bool,bool),X388)),happ(s(fun(bool,bool),X388),s(bool,X387))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X388))) | s(bool,X387) = s(bool,X389) | s(bool,X389) = s(bool,X390) | s(bool,'T_0') = s(bool,X390)) )),
  inference(subsumption_resolution,[],[f29077,f525])).
fof(f29077,plain,(
  ( ! [X389,X387,X388,X390] : (sP6('T_0','|dollar|exists_1'(s(fun(bool,bool),X388)),happ(s(fun(bool,bool),X388),s(bool,X387))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X388))) | s(bool,X387) = s(bool,X389) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X389) = s(bool,X390) | s(bool,'T_0') = s(bool,X390)) )),
  inference(superposition,[],[f3869,f2054])).
fof(f31869,plain,(
  ( ! [X383,X385,X384,X386] : (~sP6('F_0','|dollar|exists_1'(s(fun(bool,bool),X384)),happ(s(fun(bool,bool),X384),s(bool,X383))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X384))) | s(bool,X383) = s(bool,X385) | s(bool,X385) = s(bool,X386) | s(bool,'T_0') = s(bool,X386)) )),
  inference(subsumption_resolution,[],[f29076,f525])).
fof(f29076,plain,(
  ( ! [X383,X385,X384,X386] : (~sP6('F_0','|dollar|exists_1'(s(fun(bool,bool),X384)),happ(s(fun(bool,bool),X384),s(bool,X383))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X384))) | s(bool,X383) = s(bool,X385) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X385) = s(bool,X386) | s(bool,'T_0') = s(bool,X386)) )),
  inference(superposition,[],[f3853,f2054])).
fof(f31868,plain,(
  ( ! [X379,X380,X382,X381] : (sP6('T_0',happ(s(fun(bool,bool),X380),s(bool,X379)),'|dollar|exists_1'(s(fun(bool,bool),X380))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X380))) | s(bool,X379) = s(bool,X381) | s(bool,X381) = s(bool,X382) | s(bool,'T_0') = s(bool,X382)) )),
  inference(subsumption_resolution,[],[f29075,f525])).
fof(f29075,plain,(
  ( ! [X379,X380,X382,X381] : (sP6('T_0',happ(s(fun(bool,bool),X380),s(bool,X379)),'|dollar|exists_1'(s(fun(bool,bool),X380))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X380))) | s(bool,X379) = s(bool,X381) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X381) = s(bool,X382) | s(bool,'T_0') = s(bool,X382)) )),
  inference(superposition,[],[f3764,f2054])).
fof(f31867,plain,(
  ( ! [X377,X376,X378,X375] : (~sP6('F_0',happ(s(fun(bool,bool),X376),s(bool,X375)),'|dollar|exists_1'(s(fun(bool,bool),X376))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X376))) | s(bool,X375) = s(bool,X377) | s(bool,X377) = s(bool,X378) | s(bool,'T_0') = s(bool,X378)) )),
  inference(subsumption_resolution,[],[f29074,f525])).
fof(f29074,plain,(
  ( ! [X377,X376,X378,X375] : (~sP6('F_0',happ(s(fun(bool,bool),X376),s(bool,X375)),'|dollar|exists_1'(s(fun(bool,bool),X376))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X376))) | s(bool,X375) = s(bool,X377) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X377) = s(bool,X378) | s(bool,'T_0') = s(bool,X378)) )),
  inference(superposition,[],[f3744,f2054])).
fof(f31866,plain,(
  ( ! [X372,X374,X370,X373,X371] : (s(bool,X370) != s(bool,X371) | s(bool,X371) = s(bool,X372) | s(bool,'T_0') = s(bool,X372) | s(bool,X370) = s(bool,X373) | s(bool,X373) = s(bool,X374) | s(bool,'T_0') = s(bool,X374)) )),
  inference(subsumption_resolution,[],[f29073,f525])).
fof(f29073,plain,(
  ( ! [X372,X374,X370,X373,X371] : (s(bool,X370) != s(bool,X371) | s(bool,X371) = s(bool,X372) | s(bool,'T_0') = s(bool,X372) | s(bool,X370) = s(bool,X373) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X373) = s(bool,X374) | s(bool,'T_0') = s(bool,X374)) )),
  inference(superposition,[],[f2605,f2054])).
fof(f31865,plain,(
  ( ! [X364,X360,X362,X365,X361,X363] : (s(X361,f(s(bool,X360))) != s(X361,g(s(bool,X360))) | s(fun(bool,X361),X362) = s(fun(bool,X361),X363) | s(bool,'T_0') = s(bool,sK23(X362,X363,X361,bool)) | s(bool,X360) = s(bool,X364) | s(bool,X364) = s(bool,X365) | s(bool,'T_0') = s(bool,X365)) )),
  inference(subsumption_resolution,[],[f29072,f525])).
fof(f29072,plain,(
  ( ! [X364,X360,X362,X365,X361,X363] : (s(X361,f(s(bool,X360))) != s(X361,g(s(bool,X360))) | s(fun(bool,X361),X362) = s(fun(bool,X361),X363) | s(bool,'T_0') = s(bool,sK23(X362,X363,X361,bool)) | s(bool,X360) = s(bool,X364) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X364) = s(bool,X365) | s(bool,'T_0') = s(bool,X365)) )),
  inference(superposition,[],[f1975,f2054])).
fof(f31864,plain,(
  ( ! [X356,X358,X354,X357,X359,X355] : (~p(s(bool,happ(s(fun(bool,bool),X355),s(bool,X354)))) | p(s(bool,X356)) | p(s(bool,happ(s(fun(bool,bool),X355),s(bool,X357)))) | s(bool,'T_0') = s(bool,sK21(X355,X356,bool)) | s(bool,X354) = s(bool,X358) | s(bool,X358) = s(bool,X359) | s(bool,'T_0') = s(bool,X359)) )),
  inference(subsumption_resolution,[],[f29071,f525])).
fof(f29071,plain,(
  ( ! [X356,X358,X354,X357,X359,X355] : (~p(s(bool,happ(s(fun(bool,bool),X355),s(bool,X354)))) | p(s(bool,X356)) | p(s(bool,happ(s(fun(bool,bool),X355),s(bool,X357)))) | s(bool,'T_0') = s(bool,sK21(X355,X356,bool)) | s(bool,X354) = s(bool,X358) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X358) = s(bool,X359) | s(bool,'T_0') = s(bool,X359)) )),
  inference(superposition,[],[f1959,f2054])).
fof(f31863,plain,(
  ( ! [X350,X352,X349,X351,X353,X348] : (~p(s(bool,happ(s(fun(bool,bool),X349),s(bool,X348)))) | p(s(bool,X350)) | p(s(bool,happ(s(fun(bool,bool),X349),s(bool,X351)))) | s(bool,'T_0') = s(bool,sK20(X349,X350,bool)) | s(bool,X348) = s(bool,X352) | s(bool,X352) = s(bool,X353) | s(bool,'T_0') = s(bool,X353)) )),
  inference(subsumption_resolution,[],[f29070,f525])).
fof(f29070,plain,(
  ( ! [X350,X352,X349,X351,X353,X348] : (~p(s(bool,happ(s(fun(bool,bool),X349),s(bool,X348)))) | p(s(bool,X350)) | p(s(bool,happ(s(fun(bool,bool),X349),s(bool,X351)))) | s(bool,'T_0') = s(bool,sK20(X349,X350,bool)) | s(bool,X348) = s(bool,X352) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X352) = s(bool,X353) | s(bool,'T_0') = s(bool,X353)) )),
  inference(superposition,[],[f1947,f2054])).
fof(f31862,plain,(
  ( ! [X344,X346,X343,X345,X347,X342] : (~p(s(bool,happ(s(fun(bool,bool),X343),s(bool,X342)))) | p(s(bool,happ(s(fun(bool,bool),X343),s(bool,X344)))) | ~p(s(bool,X345)) | s(bool,'T_0') = s(bool,sK19(X343,X345,bool)) | s(bool,X342) = s(bool,X346) | s(bool,X346) = s(bool,X347) | s(bool,'T_0') = s(bool,X347)) )),
  inference(subsumption_resolution,[],[f29069,f525])).
fof(f29069,plain,(
  ( ! [X344,X346,X343,X345,X347,X342] : (~p(s(bool,happ(s(fun(bool,bool),X343),s(bool,X342)))) | p(s(bool,happ(s(fun(bool,bool),X343),s(bool,X344)))) | ~p(s(bool,X345)) | s(bool,'T_0') = s(bool,sK19(X343,X345,bool)) | s(bool,X342) = s(bool,X346) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X346) = s(bool,X347) | s(bool,'T_0') = s(bool,X347)) )),
  inference(superposition,[],[f1923,f2054])).
fof(f31861,plain,(
  ( ! [X341,X337,X339,X340,X336,X338] : (~p(s(bool,happ(s(fun(bool,bool),X337),s(bool,X336)))) | p(s(bool,happ(s(fun(bool,bool),X337),s(bool,X338)))) | ~p(s(bool,X339)) | s(bool,'T_0') = s(bool,sK18(X337,X339,bool)) | s(bool,X336) = s(bool,X340) | s(bool,X340) = s(bool,X341) | s(bool,'T_0') = s(bool,X341)) )),
  inference(subsumption_resolution,[],[f29068,f525])).
fof(f29068,plain,(
  ( ! [X341,X337,X339,X340,X336,X338] : (~p(s(bool,happ(s(fun(bool,bool),X337),s(bool,X336)))) | p(s(bool,happ(s(fun(bool,bool),X337),s(bool,X338)))) | ~p(s(bool,X339)) | s(bool,'T_0') = s(bool,sK18(X337,X339,bool)) | s(bool,X336) = s(bool,X340) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X340) = s(bool,X341) | s(bool,'T_0') = s(bool,X341)) )),
  inference(superposition,[],[f1898,f2054])).
fof(f31860,plain,(
  ( ! [X333,X335,X332,X334] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X333))) = s(bool,happ(s(fun(bool,bool),X333),s(bool,X332))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X333))) | s(bool,X332) = s(bool,X334) | s(bool,X334) = s(bool,X335) | s(bool,'T_0') = s(bool,X335)) )),
  inference(subsumption_resolution,[],[f29067,f525])).
fof(f29067,plain,(
  ( ! [X333,X335,X332,X334] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X333))) = s(bool,happ(s(fun(bool,bool),X333),s(bool,X332))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X333))) | s(bool,X332) = s(bool,X334) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X334) = s(bool,X335) | s(bool,'T_0') = s(bool,X335)) )),
  inference(superposition,[],[f1768,f2054])).
fof(f31859,plain,(
  ( ! [X329,X331,X330] : (s(bool,'T_0') != s(bool,X329) | s(bool,X329) = s(bool,X330) | s(bool,X330) = s(bool,X331) | s(bool,'T_0') = s(bool,X331)) )),
  inference(subsumption_resolution,[],[f29066,f525])).
fof(f29066,plain,(
  ( ! [X329,X331,X330] : (s(bool,'T_0') != s(bool,X329) | s(bool,X329) = s(bool,X330) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X330) = s(bool,X331) | s(bool,'T_0') = s(bool,X331)) )),
  inference(superposition,[],[f525,f2054])).
fof(f28996,plain,(
  ( ! [X990,X994,X992,X991,X995,X993] : (s(X992,f(s(bool,X993))) != s(X992,g(s(bool,X993))) | s(fun(bool,X992),X990) = s(fun(bool,X992),X991) | s(bool,X994) = s(bool,sK23(X990,X991,X992,bool)) | s(bool,'T_0') = s(bool,X993) | s(bool,X994) = s(bool,X995) | s(bool,'T_0') = s(bool,X995)) )),
  inference(superposition,[],[f366,f2054])).
fof(f28995,plain,(
  ( ! [X984,X988,X987,X985,X989,X986] : (~p(s(bool,happ(s(fun(bool,bool),X984),s(bool,X986)))) | p(s(bool,X985)) | p(s(bool,happ(s(fun(bool,bool),X984),s(bool,X987)))) | s(bool,X988) = s(bool,sK21(X984,X985,bool)) | s(bool,'T_0') = s(bool,X986) | s(bool,X988) = s(bool,X989) | s(bool,'T_0') = s(bool,X989)) )),
  inference(superposition,[],[f316,f2054])).
fof(f28994,plain,(
  ( ! [X978,X982,X980,X979,X983,X981] : (~p(s(bool,happ(s(fun(bool,bool),X978),s(bool,X980)))) | p(s(bool,X979)) | p(s(bool,happ(s(fun(bool,bool),X978),s(bool,X981)))) | s(bool,X982) = s(bool,sK20(X978,X979,bool)) | s(bool,'T_0') = s(bool,X980) | s(bool,X982) = s(bool,X983) | s(bool,'T_0') = s(bool,X983)) )),
  inference(superposition,[],[f314,f2054])).
fof(f28993,plain,(
  ( ! [X976,X975,X973,X977,X974,X972] : (~p(s(bool,happ(s(fun(bool,bool),X972),s(bool,X974)))) | p(s(bool,happ(s(fun(bool,bool),X972),s(bool,X975)))) | ~p(s(bool,X973)) | s(bool,X976) = s(bool,sK19(X972,X973,bool)) | s(bool,'T_0') = s(bool,X974) | s(bool,X976) = s(bool,X977) | s(bool,'T_0') = s(bool,X977)) )),
  inference(superposition,[],[f312,f2054])).
fof(f28992,plain,(
  ( ! [X971,X969,X967,X970,X968,X966] : (~p(s(bool,happ(s(fun(bool,bool),X966),s(bool,X968)))) | p(s(bool,happ(s(fun(bool,bool),X966),s(bool,X969)))) | ~p(s(bool,X967)) | s(bool,X970) = s(bool,sK18(X966,X967,bool)) | s(bool,'T_0') = s(bool,X968) | s(bool,X970) = s(bool,X971) | s(bool,'T_0') = s(bool,X971)) )),
  inference(superposition,[],[f310,f2054])).
fof(f28991,plain,(
  ( ! [X963,X965,X962,X964] : (p(s(bool,happ(s(fun(bool,bool),X962),s(bool,X963)))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X962))) | s(bool,X964) = s(bool,sK17(X962,bool)) | s(bool,'T_0') = s(bool,X963) | s(bool,X964) = s(bool,X965) | s(bool,'T_0') = s(bool,X965)) )),
  inference(superposition,[],[f4966,f2054])).
fof(f28990,plain,(
  ( ! [X957,X961,X958,X960,X959] : (p(s(bool,happ(s(fun(bool,bool),X957),s(bool,X958)))) | s(bool,X959) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X957))) | s(bool,'T_0') = s(bool,X959) | s(bool,X960) = s(bool,sK17(X957,bool)) | s(bool,'T_0') = s(bool,X958) | s(bool,X960) = s(bool,X961) | s(bool,'T_0') = s(bool,X961)) )),
  inference(superposition,[],[f4303,f2054])).
fof(f28989,plain,(
  ( ! [X954,X956,X955,X953] : (p(s(bool,happ(s(fun(bool,bool),X953),s(bool,X954)))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X953))) | s(bool,X955) = s(bool,sK16(X953,bool)) | s(bool,'T_0') = s(bool,X954) | s(bool,X955) = s(bool,X956) | s(bool,'T_0') = s(bool,X956)) )),
  inference(superposition,[],[f4967,f2054])).
fof(f28988,plain,(
  ( ! [X951,X949,X952,X950,X948] : (p(s(bool,happ(s(fun(bool,bool),X948),s(bool,X949)))) | s(bool,X950) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X948))) | s(bool,'T_0') = s(bool,X950) | s(bool,X951) = s(bool,sK16(X948,bool)) | s(bool,'T_0') = s(bool,X949) | s(bool,X951) = s(bool,X952) | s(bool,'T_0') = s(bool,X952)) )),
  inference(superposition,[],[f3897,f2054])).
fof(f28979,plain,(
  ( ! [X912,X914,X911,X915,X913,X910] : (sP12(happ(s(fun(bool,bool),X910),s(bool,X911)),X912,X913) | ~p(s(bool,X912)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X910))) | s(bool,X914) = s(bool,'const.min.@_1'(s(fun(bool,bool),X910))) | s(bool,'T_0') = s(bool,X911) | s(bool,X914) = s(bool,X915) | s(bool,'T_0') = s(bool,X915)) )),
  inference(superposition,[],[f4987,f2054])).
fof(f28978,plain,(
  ( ! [X907,X905,X909,X906,X904,X908] : (sP12(happ(s(fun(bool,bool),X904),s(bool,X905)),X906,X907) | p(s(bool,X907)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X904))) | s(bool,X908) = s(bool,'const.min.@_1'(s(fun(bool,bool),X904))) | s(bool,'T_0') = s(bool,X905) | s(bool,X908) = s(bool,X909) | s(bool,'T_0') = s(bool,X909)) )),
  inference(superposition,[],[f4986,f2054])).
fof(f28977,plain,(
  ( ! [X899,X903,X901,X898,X902,X900] : (sP10(happ(s(fun(bool,bool),X898),s(bool,X899)),X900,X901) | ~p(s(bool,X900)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X898))) | s(bool,X902) = s(bool,'const.min.@_1'(s(fun(bool,bool),X898))) | s(bool,'T_0') = s(bool,X899) | s(bool,X902) = s(bool,X903) | s(bool,'T_0') = s(bool,X903)) )),
  inference(superposition,[],[f4983,f2054])).
fof(f28976,plain,(
  ( ! [X895,X893,X897,X894,X892,X896] : (sP10(happ(s(fun(bool,bool),X892),s(bool,X893)),X894,X895) | ~p(s(bool,X895)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X892))) | s(bool,X896) = s(bool,'const.min.@_1'(s(fun(bool,bool),X892))) | s(bool,'T_0') = s(bool,X893) | s(bool,X896) = s(bool,X897) | s(bool,'T_0') = s(bool,X897)) )),
  inference(superposition,[],[f4982,f2054])).
fof(f28975,plain,(
  ( ! [X887,X890,X888,X891,X889] : (sP1(X889,happ(s(fun(bool,bool),X887),s(bool,X888))) | p(s(bool,X889)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X887))) | s(bool,X890) = s(bool,'const.min.@_1'(s(fun(bool,bool),X887))) | s(bool,'T_0') = s(bool,X888) | s(bool,X890) = s(bool,X891) | s(bool,'T_0') = s(bool,X891)) )),
  inference(superposition,[],[f4975,f2054])).
fof(f28974,plain,(
  ( ! [X883,X885,X882,X886,X884] : (sP1(happ(s(fun(bool,bool),X882),s(bool,X883)),X884) | p(s(bool,X884)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X882))) | s(bool,X885) = s(bool,'const.min.@_1'(s(fun(bool,bool),X882))) | s(bool,'T_0') = s(bool,X883) | s(bool,X885) = s(bool,X886) | s(bool,'T_0') = s(bool,X886)) )),
  inference(superposition,[],[f4974,f2054])).
fof(f28973,plain,(
  ( ! [X881,X878,X880,X879,X877] : (sP0(happ(s(fun(bool,bool),X877),s(bool,X878)),X879) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X877))) | s(bool,X880) = s(bool,'const.min.@_1'(s(fun(bool,bool),X877))) | s(bool,'T_0') = s(bool,X878) | s(bool,X880) = s(bool,X881) | s(bool,'T_0') = s(bool,X881)) )),
  inference(superposition,[],[f4973,f2054])).
fof(f28972,plain,(
  ( ! [X874,X872,X876,X875,X873] : (sP0(X874,happ(s(fun(bool,bool),X872),s(bool,X873))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X872))) | s(bool,X875) = s(bool,'const.min.@_1'(s(fun(bool,bool),X872))) | s(bool,'T_0') = s(bool,X873) | s(bool,X875) = s(bool,X876) | s(bool,'T_0') = s(bool,X876)) )),
  inference(superposition,[],[f4972,f2054])).
fof(f28971,plain,(
  ( ! [X870,X868,X871,X869] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X868),s(bool,X869))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X868))) | s(bool,X870) = s(bool,'const.min.@_1'(s(fun(bool,bool),X868))) | s(bool,'T_0') = s(bool,X869) | s(bool,X870) = s(bool,X871) | s(bool,'T_0') = s(bool,X871)) )),
  inference(superposition,[],[f4969,f2054])).
fof(f28970,plain,(
  ( ! [X866,X864,X867,X865] : (sP5('T_0',happ(s(fun(bool,bool),X864),s(bool,X865)),'|dollar|exists_1'(s(fun(bool,bool),X864))) | s(bool,X866) = s(bool,'const.min.@_1'(s(fun(bool,bool),X864))) | s(bool,'T_0') = s(bool,X865) | s(bool,X866) = s(bool,X867) | s(bool,'T_0') = s(bool,X867)) )),
  inference(superposition,[],[f3862,f2054])).
fof(f28969,plain,(
  ( ! [X860,X863,X861,X862] : (sP5('T_0','|dollar|exists_1'(s(fun(bool,bool),X860)),happ(s(fun(bool,bool),X860),s(bool,X861))) | s(bool,X862) = s(bool,'const.min.@_1'(s(fun(bool,bool),X860))) | s(bool,'T_0') = s(bool,X861) | s(bool,X862) = s(bool,X863) | s(bool,'T_0') = s(bool,X863)) )),
  inference(superposition,[],[f3757,f2054])).
fof(f28968,plain,(
  ( ! [X854,X859,X857,X855,X858,X856] : (sP10(happ(s(fun(bool,bool),X854),s(bool,X855)),X856,X857) | ~sP9(happ(s(fun(bool,bool),X854),s(bool,'const.min.@_1'(s(fun(bool,bool),X854)))),X856,X857) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X854)))) | s(bool,X858) = s(bool,'const.min.@_1'(s(fun(bool,bool),X854))) | s(bool,'T_0') = s(bool,X855) | s(bool,X858) = s(bool,X859) | s(bool,'T_0') = s(bool,X859)) )),
  inference(superposition,[],[f1797,f2054])).
fof(f28967,plain,(
  ( ! [X850,X848,X852,X851,X849,X853] : (sP12(happ(s(fun(bool,bool),X848),s(bool,X849)),X850,X851) | ~sP11(happ(s(fun(bool,bool),X848),s(bool,'const.min.@_1'(s(fun(bool,bool),X848)))),X850,X851) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X848)))) | s(bool,X852) = s(bool,'const.min.@_1'(s(fun(bool,bool),X848))) | s(bool,'T_0') = s(bool,X849) | s(bool,X852) = s(bool,X853) | s(bool,'T_0') = s(bool,X853)) )),
  inference(superposition,[],[f1796,f2054])).
fof(f28966,plain,(
  ( ! [X843,X847,X845,X842,X846,X844] : (~sP12(happ(s(fun(bool,bool),X842),s(bool,X843)),X844,X845) | ~p(s(bool,X845)) | p(s(bool,X844)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X842)))) | s(bool,X846) = s(bool,'const.min.@_1'(s(fun(bool,bool),X842))) | s(bool,'T_0') = s(bool,X843) | s(bool,X846) = s(bool,X847) | s(bool,'T_0') = s(bool,X847)) )),
  inference(superposition,[],[f1795,f2054])).
fof(f28965,plain,(
  ( ! [X841,X839,X837,X840,X838,X836] : (sP12(happ(s(fun(bool,bool),X836),s(bool,X837)),X838,X839) | ~p(s(bool,X838)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X836)))) | s(bool,X840) = s(bool,'const.min.@_1'(s(fun(bool,bool),X836))) | s(bool,'T_0') = s(bool,X837) | s(bool,X840) = s(bool,X841) | s(bool,'T_0') = s(bool,X841)) )),
  inference(superposition,[],[f1794,f2054])).
fof(f28964,plain,(
  ( ! [X831,X835,X833,X830,X834,X832] : (sP12(happ(s(fun(bool,bool),X830),s(bool,X831)),X832,X833) | p(s(bool,X833)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X830)))) | s(bool,X834) = s(bool,'const.min.@_1'(s(fun(bool,bool),X830))) | s(bool,'T_0') = s(bool,X831) | s(bool,X834) = s(bool,X835) | s(bool,'T_0') = s(bool,X835)) )),
  inference(superposition,[],[f1793,f2054])).
fof(f28963,plain,(
  ( ! [X825,X829,X826,X824,X828,X827] : (~sP10(happ(s(fun(bool,bool),X824),s(bool,X825)),X826,X827) | p(s(bool,X827)) | p(s(bool,X826)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X824)))) | s(bool,X828) = s(bool,'const.min.@_1'(s(fun(bool,bool),X824))) | s(bool,'T_0') = s(bool,X825) | s(bool,X828) = s(bool,X829) | s(bool,'T_0') = s(bool,X829)) )),
  inference(superposition,[],[f1791,f2054])).
fof(f28962,plain,(
  ( ! [X819,X823,X821,X818,X822,X820] : (sP10(happ(s(fun(bool,bool),X818),s(bool,X819)),X820,X821) | ~p(s(bool,X820)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X818)))) | s(bool,X822) = s(bool,'const.min.@_1'(s(fun(bool,bool),X818))) | s(bool,'T_0') = s(bool,X819) | s(bool,X822) = s(bool,X823) | s(bool,'T_0') = s(bool,X823)) )),
  inference(superposition,[],[f1790,f2054])).
fof(f28961,plain,(
  ( ! [X817,X814,X812,X816,X815,X813] : (sP10(happ(s(fun(bool,bool),X812),s(bool,X813)),X814,X815) | ~p(s(bool,X815)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X812)))) | s(bool,X816) = s(bool,'const.min.@_1'(s(fun(bool,bool),X812))) | s(bool,'T_0') = s(bool,X813) | s(bool,X816) = s(bool,X817) | s(bool,'T_0') = s(bool,X817)) )),
  inference(superposition,[],[f1789,f2054])).
fof(f28960,plain,(
  ( ! [X810,X808,X811,X809,X807] : (sP1(X809,happ(s(fun(bool,bool),X807),s(bool,X808))) | p(s(bool,X809)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X807)))) | s(bool,X810) = s(bool,'const.min.@_1'(s(fun(bool,bool),X807))) | s(bool,'T_0') = s(bool,X808) | s(bool,X810) = s(bool,X811) | s(bool,'T_0') = s(bool,X811)) )),
  inference(superposition,[],[f1782,f2054])).
fof(f28959,plain,(
  ( ! [X802,X806,X804,X803,X805] : (sP1(happ(s(fun(bool,bool),X802),s(bool,X803)),X804) | p(s(bool,X804)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X802)))) | s(bool,X805) = s(bool,'const.min.@_1'(s(fun(bool,bool),X802))) | s(bool,'T_0') = s(bool,X803) | s(bool,X805) = s(bool,X806) | s(bool,'T_0') = s(bool,X806)) )),
  inference(superposition,[],[f1781,f2054])).
fof(f28958,plain,(
  ( ! [X798,X800,X799,X797,X801] : (sP0(happ(s(fun(bool,bool),X797),s(bool,X798)),X799) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X797)))) | s(bool,X800) = s(bool,'const.min.@_1'(s(fun(bool,bool),X797))) | s(bool,'T_0') = s(bool,X798) | s(bool,X800) = s(bool,X801) | s(bool,'T_0') = s(bool,X801)) )),
  inference(superposition,[],[f1780,f2054])).
fof(f28957,plain,(
  ( ! [X796,X795,X793,X794,X792] : (sP0(X794,happ(s(fun(bool,bool),X792),s(bool,X793))) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X792)))) | s(bool,X795) = s(bool,'const.min.@_1'(s(fun(bool,bool),X792))) | s(bool,'T_0') = s(bool,X793) | s(bool,X795) = s(bool,X796) | s(bool,'T_0') = s(bool,X796)) )),
  inference(superposition,[],[f1779,f2054])).
fof(f28956,plain,(
  ( ! [X790,X788,X791,X789] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X788),s(bool,X789))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X788))) | s(bool,X790) = s(bool,'const.min.@_1'(s(fun(bool,bool),X788))) | s(bool,'T_0') = s(bool,X789) | s(bool,X790) = s(bool,X791) | s(bool,'T_0') = s(bool,X791)) )),
  inference(superposition,[],[f1769,f2054])).
fof(f28955,plain,(
  ( ! [X786,X784,X787,X785] : (sP6('T_0','|dollar|exists_1'(s(fun(bool,bool),X784)),happ(s(fun(bool,bool),X784),s(bool,X785))) | s(bool,X786) = s(bool,'const.min.@_1'(s(fun(bool,bool),X784))) | s(bool,'T_0') = s(bool,X785) | s(bool,X786) = s(bool,X787) | s(bool,'T_0') = s(bool,X787)) )),
  inference(superposition,[],[f1767,f2054])).
fof(f28954,plain,(
  ( ! [X783,X781,X782,X780] : (~sP6('F_0','|dollar|exists_1'(s(fun(bool,bool),X780)),happ(s(fun(bool,bool),X780),s(bool,X781))) | s(bool,X782) = s(bool,'const.min.@_1'(s(fun(bool,bool),X780))) | s(bool,'T_0') = s(bool,X781) | s(bool,X782) = s(bool,X783) | s(bool,'T_0') = s(bool,X783)) )),
  inference(superposition,[],[f1766,f2054])).
fof(f28953,plain,(
  ( ! [X779,X777,X778,X776] : (sP6('T_0',happ(s(fun(bool,bool),X776),s(bool,X777)),'|dollar|exists_1'(s(fun(bool,bool),X776))) | s(bool,X778) = s(bool,'const.min.@_1'(s(fun(bool,bool),X776))) | s(bool,'T_0') = s(bool,X777) | s(bool,X778) = s(bool,X779) | s(bool,'T_0') = s(bool,X779)) )),
  inference(superposition,[],[f1765,f2054])).
fof(f28952,plain,(
  ( ! [X775,X773,X774,X772] : (~sP6('F_0',happ(s(fun(bool,bool),X772),s(bool,X773)),'|dollar|exists_1'(s(fun(bool,bool),X772))) | s(bool,X774) = s(bool,'const.min.@_1'(s(fun(bool,bool),X772))) | s(bool,'T_0') = s(bool,X773) | s(bool,X774) = s(bool,X775) | s(bool,'T_0') = s(bool,X775)) )),
  inference(superposition,[],[f1764,f2054])).
fof(f28951,plain,(
  ( ! [X771,X769,X770,X768] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X768))) = s(bool,happ(s(fun(bool,bool),X768),s(bool,X769))) | s(bool,X770) = s(bool,'const.min.@_1'(s(fun(bool,bool),X768))) | s(bool,'T_0') = s(bool,X769) | s(bool,X770) = s(bool,X771) | s(bool,'T_0') = s(bool,X771)) )),
  inference(superposition,[],[f303,f2054])).
fof(f28949,plain,(
  ( ! [X761,X757,X759,X760,X762,X758] : (s(bool,X760) != s(bool,f(s(X757,sK23(X758,X759,bool,X757)))) | s(fun(X757,bool),X758) = s(fun(X757,bool),X759) | s(bool,X761) = s(bool,g(s(X757,sK23(X758,X759,bool,X757)))) | s(bool,'T_0') = s(bool,X760) | s(bool,X761) = s(bool,X762) | s(bool,'T_0') = s(bool,X762)) )),
  inference(superposition,[],[f366,f2054])).
fof(f31841,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',sK22(X1,X0))))) | s(bool,X2) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X0))),s('type.extreal.extreal',X1))) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(condensation,[],[f31840])).
fof(f31840,plain,(
  ( ! [X707,X708,X706,X709,X705] : (p(s(bool,X707)) | p(s(bool,happ(s(fun('type.extreal.extreal',bool),X705),s('type.extreal.extreal',sK22(X706,X705))))) | s(bool,X708) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X705))),s('type.extreal.extreal',X706))) | s(bool,X708) = s(bool,X709) | s(bool,'T_0') = s(bool,X709)) )),
  inference(subsumption_resolution,[],[f28938,f276])).
fof(f28938,plain,(
  ( ! [X707,X708,X706,X709,X705] : (p(s(bool,X707)) | p(s(bool,happ(s(fun('type.extreal.extreal',bool),X705),s('type.extreal.extreal',sK22(X706,X705))))) | s(bool,X708) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X705))),s('type.extreal.extreal',X706))) | s(bool,'T_0') = s(bool,X707) | s(bool,X708) = s(bool,X709) | s(bool,'T_0') = s(bool,X709)) )),
  inference(superposition,[],[f327,f2054])).
fof(f31837,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | s(bool,X2) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(condensation,[],[f31836])).
fof(f31836,plain,(
  ( ! [X685,X687,X688,X684,X686] : (p(s(bool,X686)) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X685),s('type.extreal.extreal',X684)))) | s(bool,X687) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X684),s('type.extreal.extreal',X685))) | s(bool,X687) = s(bool,X688) | s(bool,'T_0') = s(bool,X688)) )),
  inference(subsumption_resolution,[],[f28934,f276])).
fof(f28934,plain,(
  ( ! [X685,X687,X688,X684,X686] : (p(s(bool,X686)) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X685),s('type.extreal.extreal',X684)))) | s(bool,X687) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X684),s('type.extreal.extreal',X685))) | s(bool,'T_0') = s(bool,X686) | s(bool,X687) = s(bool,X688) | s(bool,'T_0') = s(bool,X688)) )),
  inference(superposition,[],[f325,f2054])).
fof(f31823,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X0))),s('type.extreal.extreal',X1)))) | s(bool,X2) = s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',sK22(X1,X0)))) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(condensation,[],[f31822])).
fof(f31822,plain,(
  ( ! [X461,X463,X460,X462,X464] : (p(s(bool,X462)) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X460))),s('type.extreal.extreal',X461)))) | s(bool,X463) = s(bool,happ(s(fun('type.extreal.extreal',bool),X460),s('type.extreal.extreal',sK22(X461,X460)))) | s(bool,X463) = s(bool,X464) | s(bool,'T_0') = s(bool,X464)) )),
  inference(subsumption_resolution,[],[f28894,f276])).
fof(f28894,plain,(
  ( ! [X461,X463,X460,X462,X464] : (p(s(bool,X462)) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X460))),s('type.extreal.extreal',X461)))) | s(bool,X463) = s(bool,happ(s(fun('type.extreal.extreal',bool),X460),s('type.extreal.extreal',sK22(X461,X460)))) | s(bool,'T_0') = s(bool,X462) | s(bool,X463) = s(bool,X464) | s(bool,'T_0') = s(bool,X464)) )),
  inference(superposition,[],[f327,f2054])).
fof(f31819,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK23(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,X3) = s(bool,f(s(X0,sK23(X1,X2,bool,X0)))) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(condensation,[],[f31818])).
fof(f31818,plain,(
  ( ! [X436,X438,X440,X437,X439,X435] : (s(fun(X435,bool),X436) = s(fun(X435,bool),X437) | s(bool,'T_0') = s(bool,g(s(X435,sK23(X436,X437,bool,X435)))) | s(bool,X439) = s(bool,f(s(X435,sK23(X436,X437,bool,X435)))) | s(bool,'T_0') = s(bool,X438) | s(bool,X439) = s(bool,X440) | s(bool,'T_0') = s(bool,X440)) )),
  inference(subsumption_resolution,[],[f28888,f270])).
fof(f28888,plain,(
  ( ! [X436,X438,X440,X437,X439,X435] : (s(bool,'F_0') != s(bool,X438) | s(fun(X435,bool),X436) = s(fun(X435,bool),X437) | s(bool,'T_0') = s(bool,g(s(X435,sK23(X436,X437,bool,X435)))) | s(bool,X439) = s(bool,f(s(X435,sK23(X436,X437,bool,X435)))) | s(bool,'T_0') = s(bool,X438) | s(bool,X439) = s(bool,X440) | s(bool,'T_0') = s(bool,X440)) )),
  inference(superposition,[],[f1976,f2054])).
fof(f31756,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X0)) | s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(condensation,[],[f31755])).
fof(f31755,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X20) = s(bool,X22) | s(bool,X20) = s(bool,X23) | s(bool,X23) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f28807,f276])).
fof(f28807,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X20) = s(bool,X22) | s(bool,X20) = s(bool,X23) | s(bool,'T_0') = s(bool,X21) | s(bool,X23) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f318,f2054])).
fof(f28604,plain,(
  ( ! [X990,X994,X992,X991,X995,X993] : (s(X992,f(s(bool,X993))) != s(X992,g(s(bool,X993))) | s(fun(bool,X992),X990) = s(fun(bool,X992),X991) | s(bool,X993) = s(bool,X994) | s(bool,'T_0') = s(bool,X994) | s(bool,X995) = s(bool,sK23(X990,X991,X992,bool)) | s(bool,'T_0') = s(bool,X995)) )),
  inference(superposition,[],[f366,f2054])).
fof(f28603,plain,(
  ( ! [X984,X988,X987,X985,X989,X986] : (~p(s(bool,happ(s(fun(bool,bool),X984),s(bool,X986)))) | p(s(bool,X985)) | p(s(bool,happ(s(fun(bool,bool),X984),s(bool,X987)))) | s(bool,X986) = s(bool,X988) | s(bool,'T_0') = s(bool,X988) | s(bool,X989) = s(bool,sK21(X984,X985,bool)) | s(bool,'T_0') = s(bool,X989)) )),
  inference(superposition,[],[f316,f2054])).
fof(f28602,plain,(
  ( ! [X978,X982,X980,X979,X983,X981] : (~p(s(bool,happ(s(fun(bool,bool),X978),s(bool,X980)))) | p(s(bool,X979)) | p(s(bool,happ(s(fun(bool,bool),X978),s(bool,X981)))) | s(bool,X980) = s(bool,X982) | s(bool,'T_0') = s(bool,X982) | s(bool,X983) = s(bool,sK20(X978,X979,bool)) | s(bool,'T_0') = s(bool,X983)) )),
  inference(superposition,[],[f314,f2054])).
fof(f28601,plain,(
  ( ! [X976,X975,X973,X977,X974,X972] : (~p(s(bool,happ(s(fun(bool,bool),X972),s(bool,X974)))) | p(s(bool,happ(s(fun(bool,bool),X972),s(bool,X975)))) | ~p(s(bool,X973)) | s(bool,X974) = s(bool,X976) | s(bool,'T_0') = s(bool,X976) | s(bool,X977) = s(bool,sK19(X972,X973,bool)) | s(bool,'T_0') = s(bool,X977)) )),
  inference(superposition,[],[f312,f2054])).
fof(f28600,plain,(
  ( ! [X971,X969,X967,X970,X968,X966] : (~p(s(bool,happ(s(fun(bool,bool),X966),s(bool,X968)))) | p(s(bool,happ(s(fun(bool,bool),X966),s(bool,X969)))) | ~p(s(bool,X967)) | s(bool,X968) = s(bool,X970) | s(bool,'T_0') = s(bool,X970) | s(bool,X971) = s(bool,sK18(X966,X967,bool)) | s(bool,'T_0') = s(bool,X971)) )),
  inference(superposition,[],[f310,f2054])).
fof(f28599,plain,(
  ( ! [X963,X965,X962,X964] : (p(s(bool,happ(s(fun(bool,bool),X962),s(bool,X963)))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X962))) | s(bool,X963) = s(bool,X964) | s(bool,'T_0') = s(bool,X964) | s(bool,X965) = s(bool,sK17(X962,bool)) | s(bool,'T_0') = s(bool,X965)) )),
  inference(superposition,[],[f4966,f2054])).
fof(f28598,plain,(
  ( ! [X957,X961,X958,X960,X959] : (p(s(bool,happ(s(fun(bool,bool),X957),s(bool,X958)))) | s(bool,X959) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X957))) | s(bool,'T_0') = s(bool,X959) | s(bool,X958) = s(bool,X960) | s(bool,'T_0') = s(bool,X960) | s(bool,X961) = s(bool,sK17(X957,bool)) | s(bool,'T_0') = s(bool,X961)) )),
  inference(superposition,[],[f4303,f2054])).
fof(f28597,plain,(
  ( ! [X954,X956,X955,X953] : (p(s(bool,happ(s(fun(bool,bool),X953),s(bool,X954)))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X953))) | s(bool,X954) = s(bool,X955) | s(bool,'T_0') = s(bool,X955) | s(bool,X956) = s(bool,sK16(X953,bool)) | s(bool,'T_0') = s(bool,X956)) )),
  inference(superposition,[],[f4967,f2054])).
fof(f28596,plain,(
  ( ! [X951,X949,X952,X950,X948] : (p(s(bool,happ(s(fun(bool,bool),X948),s(bool,X949)))) | s(bool,X950) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X948))) | s(bool,'T_0') = s(bool,X950) | s(bool,X949) = s(bool,X951) | s(bool,'T_0') = s(bool,X951) | s(bool,X952) = s(bool,sK16(X948,bool)) | s(bool,'T_0') = s(bool,X952)) )),
  inference(superposition,[],[f3897,f2054])).
fof(f28587,plain,(
  ( ! [X912,X914,X911,X915,X913,X910] : (sP12(happ(s(fun(bool,bool),X910),s(bool,X911)),X912,X913) | ~p(s(bool,X912)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X910))) | s(bool,X911) = s(bool,X914) | s(bool,'T_0') = s(bool,X914) | s(bool,X915) = s(bool,'const.min.@_1'(s(fun(bool,bool),X910))) | s(bool,'T_0') = s(bool,X915)) )),
  inference(superposition,[],[f4987,f2054])).
fof(f28586,plain,(
  ( ! [X907,X905,X909,X906,X904,X908] : (sP12(happ(s(fun(bool,bool),X904),s(bool,X905)),X906,X907) | p(s(bool,X907)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X904))) | s(bool,X905) = s(bool,X908) | s(bool,'T_0') = s(bool,X908) | s(bool,X909) = s(bool,'const.min.@_1'(s(fun(bool,bool),X904))) | s(bool,'T_0') = s(bool,X909)) )),
  inference(superposition,[],[f4986,f2054])).
fof(f28585,plain,(
  ( ! [X899,X903,X901,X898,X902,X900] : (sP10(happ(s(fun(bool,bool),X898),s(bool,X899)),X900,X901) | ~p(s(bool,X900)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X898))) | s(bool,X899) = s(bool,X902) | s(bool,'T_0') = s(bool,X902) | s(bool,X903) = s(bool,'const.min.@_1'(s(fun(bool,bool),X898))) | s(bool,'T_0') = s(bool,X903)) )),
  inference(superposition,[],[f4983,f2054])).
fof(f28584,plain,(
  ( ! [X895,X893,X897,X894,X892,X896] : (sP10(happ(s(fun(bool,bool),X892),s(bool,X893)),X894,X895) | ~p(s(bool,X895)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X892))) | s(bool,X893) = s(bool,X896) | s(bool,'T_0') = s(bool,X896) | s(bool,X897) = s(bool,'const.min.@_1'(s(fun(bool,bool),X892))) | s(bool,'T_0') = s(bool,X897)) )),
  inference(superposition,[],[f4982,f2054])).
fof(f28583,plain,(
  ( ! [X887,X890,X888,X891,X889] : (sP1(X889,happ(s(fun(bool,bool),X887),s(bool,X888))) | p(s(bool,X889)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X887))) | s(bool,X888) = s(bool,X890) | s(bool,'T_0') = s(bool,X890) | s(bool,X891) = s(bool,'const.min.@_1'(s(fun(bool,bool),X887))) | s(bool,'T_0') = s(bool,X891)) )),
  inference(superposition,[],[f4975,f2054])).
fof(f28582,plain,(
  ( ! [X883,X885,X882,X886,X884] : (sP1(happ(s(fun(bool,bool),X882),s(bool,X883)),X884) | p(s(bool,X884)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X882))) | s(bool,X883) = s(bool,X885) | s(bool,'T_0') = s(bool,X885) | s(bool,X886) = s(bool,'const.min.@_1'(s(fun(bool,bool),X882))) | s(bool,'T_0') = s(bool,X886)) )),
  inference(superposition,[],[f4974,f2054])).
fof(f28581,plain,(
  ( ! [X881,X878,X880,X879,X877] : (sP0(happ(s(fun(bool,bool),X877),s(bool,X878)),X879) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X877))) | s(bool,X878) = s(bool,X880) | s(bool,'T_0') = s(bool,X880) | s(bool,X881) = s(bool,'const.min.@_1'(s(fun(bool,bool),X877))) | s(bool,'T_0') = s(bool,X881)) )),
  inference(superposition,[],[f4973,f2054])).
fof(f28580,plain,(
  ( ! [X874,X872,X876,X875,X873] : (sP0(X874,happ(s(fun(bool,bool),X872),s(bool,X873))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X872))) | s(bool,X873) = s(bool,X875) | s(bool,'T_0') = s(bool,X875) | s(bool,X876) = s(bool,'const.min.@_1'(s(fun(bool,bool),X872))) | s(bool,'T_0') = s(bool,X876)) )),
  inference(superposition,[],[f4972,f2054])).
fof(f28579,plain,(
  ( ! [X870,X868,X871,X869] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X868),s(bool,X869))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X868))) | s(bool,X869) = s(bool,X870) | s(bool,'T_0') = s(bool,X870) | s(bool,X871) = s(bool,'const.min.@_1'(s(fun(bool,bool),X868))) | s(bool,'T_0') = s(bool,X871)) )),
  inference(superposition,[],[f4969,f2054])).
fof(f28578,plain,(
  ( ! [X866,X864,X867,X865] : (sP5('T_0',happ(s(fun(bool,bool),X864),s(bool,X865)),'|dollar|exists_1'(s(fun(bool,bool),X864))) | s(bool,X865) = s(bool,X866) | s(bool,'T_0') = s(bool,X866) | s(bool,X867) = s(bool,'const.min.@_1'(s(fun(bool,bool),X864))) | s(bool,'T_0') = s(bool,X867)) )),
  inference(superposition,[],[f3862,f2054])).
fof(f28577,plain,(
  ( ! [X860,X863,X861,X862] : (sP5('T_0','|dollar|exists_1'(s(fun(bool,bool),X860)),happ(s(fun(bool,bool),X860),s(bool,X861))) | s(bool,X861) = s(bool,X862) | s(bool,'T_0') = s(bool,X862) | s(bool,X863) = s(bool,'const.min.@_1'(s(fun(bool,bool),X860))) | s(bool,'T_0') = s(bool,X863)) )),
  inference(superposition,[],[f3757,f2054])).
fof(f28576,plain,(
  ( ! [X854,X859,X857,X855,X858,X856] : (sP10(happ(s(fun(bool,bool),X854),s(bool,X855)),X856,X857) | ~sP9(happ(s(fun(bool,bool),X854),s(bool,'const.min.@_1'(s(fun(bool,bool),X854)))),X856,X857) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X854)))) | s(bool,X855) = s(bool,X858) | s(bool,'T_0') = s(bool,X858) | s(bool,X859) = s(bool,'const.min.@_1'(s(fun(bool,bool),X854))) | s(bool,'T_0') = s(bool,X859)) )),
  inference(superposition,[],[f1797,f2054])).
fof(f28575,plain,(
  ( ! [X850,X848,X852,X851,X849,X853] : (sP12(happ(s(fun(bool,bool),X848),s(bool,X849)),X850,X851) | ~sP11(happ(s(fun(bool,bool),X848),s(bool,'const.min.@_1'(s(fun(bool,bool),X848)))),X850,X851) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X848)))) | s(bool,X849) = s(bool,X852) | s(bool,'T_0') = s(bool,X852) | s(bool,X853) = s(bool,'const.min.@_1'(s(fun(bool,bool),X848))) | s(bool,'T_0') = s(bool,X853)) )),
  inference(superposition,[],[f1796,f2054])).
fof(f28574,plain,(
  ( ! [X843,X847,X845,X842,X846,X844] : (~sP12(happ(s(fun(bool,bool),X842),s(bool,X843)),X844,X845) | ~p(s(bool,X845)) | p(s(bool,X844)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X842)))) | s(bool,X843) = s(bool,X846) | s(bool,'T_0') = s(bool,X846) | s(bool,X847) = s(bool,'const.min.@_1'(s(fun(bool,bool),X842))) | s(bool,'T_0') = s(bool,X847)) )),
  inference(superposition,[],[f1795,f2054])).
fof(f28573,plain,(
  ( ! [X841,X839,X837,X840,X838,X836] : (sP12(happ(s(fun(bool,bool),X836),s(bool,X837)),X838,X839) | ~p(s(bool,X838)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X836)))) | s(bool,X837) = s(bool,X840) | s(bool,'T_0') = s(bool,X840) | s(bool,X841) = s(bool,'const.min.@_1'(s(fun(bool,bool),X836))) | s(bool,'T_0') = s(bool,X841)) )),
  inference(superposition,[],[f1794,f2054])).
fof(f28572,plain,(
  ( ! [X831,X835,X833,X830,X834,X832] : (sP12(happ(s(fun(bool,bool),X830),s(bool,X831)),X832,X833) | p(s(bool,X833)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X830)))) | s(bool,X831) = s(bool,X834) | s(bool,'T_0') = s(bool,X834) | s(bool,X835) = s(bool,'const.min.@_1'(s(fun(bool,bool),X830))) | s(bool,'T_0') = s(bool,X835)) )),
  inference(superposition,[],[f1793,f2054])).
fof(f28571,plain,(
  ( ! [X825,X829,X826,X824,X828,X827] : (~sP10(happ(s(fun(bool,bool),X824),s(bool,X825)),X826,X827) | p(s(bool,X827)) | p(s(bool,X826)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X824)))) | s(bool,X825) = s(bool,X828) | s(bool,'T_0') = s(bool,X828) | s(bool,X829) = s(bool,'const.min.@_1'(s(fun(bool,bool),X824))) | s(bool,'T_0') = s(bool,X829)) )),
  inference(superposition,[],[f1791,f2054])).
fof(f28570,plain,(
  ( ! [X819,X823,X821,X818,X822,X820] : (sP10(happ(s(fun(bool,bool),X818),s(bool,X819)),X820,X821) | ~p(s(bool,X820)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X818)))) | s(bool,X819) = s(bool,X822) | s(bool,'T_0') = s(bool,X822) | s(bool,X823) = s(bool,'const.min.@_1'(s(fun(bool,bool),X818))) | s(bool,'T_0') = s(bool,X823)) )),
  inference(superposition,[],[f1790,f2054])).
fof(f28569,plain,(
  ( ! [X817,X814,X812,X816,X815,X813] : (sP10(happ(s(fun(bool,bool),X812),s(bool,X813)),X814,X815) | ~p(s(bool,X815)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X812)))) | s(bool,X813) = s(bool,X816) | s(bool,'T_0') = s(bool,X816) | s(bool,X817) = s(bool,'const.min.@_1'(s(fun(bool,bool),X812))) | s(bool,'T_0') = s(bool,X817)) )),
  inference(superposition,[],[f1789,f2054])).
fof(f28568,plain,(
  ( ! [X810,X808,X811,X809,X807] : (sP1(X809,happ(s(fun(bool,bool),X807),s(bool,X808))) | p(s(bool,X809)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X807)))) | s(bool,X808) = s(bool,X810) | s(bool,'T_0') = s(bool,X810) | s(bool,X811) = s(bool,'const.min.@_1'(s(fun(bool,bool),X807))) | s(bool,'T_0') = s(bool,X811)) )),
  inference(superposition,[],[f1782,f2054])).
fof(f28567,plain,(
  ( ! [X802,X806,X804,X803,X805] : (sP1(happ(s(fun(bool,bool),X802),s(bool,X803)),X804) | p(s(bool,X804)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X802)))) | s(bool,X803) = s(bool,X805) | s(bool,'T_0') = s(bool,X805) | s(bool,X806) = s(bool,'const.min.@_1'(s(fun(bool,bool),X802))) | s(bool,'T_0') = s(bool,X806)) )),
  inference(superposition,[],[f1781,f2054])).
fof(f28566,plain,(
  ( ! [X798,X800,X799,X797,X801] : (sP0(happ(s(fun(bool,bool),X797),s(bool,X798)),X799) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X797)))) | s(bool,X798) = s(bool,X800) | s(bool,'T_0') = s(bool,X800) | s(bool,X801) = s(bool,'const.min.@_1'(s(fun(bool,bool),X797))) | s(bool,'T_0') = s(bool,X801)) )),
  inference(superposition,[],[f1780,f2054])).
fof(f28565,plain,(
  ( ! [X796,X795,X793,X794,X792] : (sP0(X794,happ(s(fun(bool,bool),X792),s(bool,X793))) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X792)))) | s(bool,X793) = s(bool,X795) | s(bool,'T_0') = s(bool,X795) | s(bool,X796) = s(bool,'const.min.@_1'(s(fun(bool,bool),X792))) | s(bool,'T_0') = s(bool,X796)) )),
  inference(superposition,[],[f1779,f2054])).
fof(f28564,plain,(
  ( ! [X790,X788,X791,X789] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X788),s(bool,X789))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X788))) | s(bool,X789) = s(bool,X790) | s(bool,'T_0') = s(bool,X790) | s(bool,X791) = s(bool,'const.min.@_1'(s(fun(bool,bool),X788))) | s(bool,'T_0') = s(bool,X791)) )),
  inference(superposition,[],[f1769,f2054])).
fof(f28563,plain,(
  ( ! [X786,X784,X787,X785] : (sP6('T_0','|dollar|exists_1'(s(fun(bool,bool),X784)),happ(s(fun(bool,bool),X784),s(bool,X785))) | s(bool,X785) = s(bool,X786) | s(bool,'T_0') = s(bool,X786) | s(bool,X787) = s(bool,'const.min.@_1'(s(fun(bool,bool),X784))) | s(bool,'T_0') = s(bool,X787)) )),
  inference(superposition,[],[f1767,f2054])).
fof(f28562,plain,(
  ( ! [X783,X781,X782,X780] : (~sP6('F_0','|dollar|exists_1'(s(fun(bool,bool),X780)),happ(s(fun(bool,bool),X780),s(bool,X781))) | s(bool,X781) = s(bool,X782) | s(bool,'T_0') = s(bool,X782) | s(bool,X783) = s(bool,'const.min.@_1'(s(fun(bool,bool),X780))) | s(bool,'T_0') = s(bool,X783)) )),
  inference(superposition,[],[f1766,f2054])).
fof(f28561,plain,(
  ( ! [X779,X777,X778,X776] : (sP6('T_0',happ(s(fun(bool,bool),X776),s(bool,X777)),'|dollar|exists_1'(s(fun(bool,bool),X776))) | s(bool,X777) = s(bool,X778) | s(bool,'T_0') = s(bool,X778) | s(bool,X779) = s(bool,'const.min.@_1'(s(fun(bool,bool),X776))) | s(bool,'T_0') = s(bool,X779)) )),
  inference(superposition,[],[f1765,f2054])).
fof(f28560,plain,(
  ( ! [X775,X773,X774,X772] : (~sP6('F_0',happ(s(fun(bool,bool),X772),s(bool,X773)),'|dollar|exists_1'(s(fun(bool,bool),X772))) | s(bool,X773) = s(bool,X774) | s(bool,'T_0') = s(bool,X774) | s(bool,X775) = s(bool,'const.min.@_1'(s(fun(bool,bool),X772))) | s(bool,'T_0') = s(bool,X775)) )),
  inference(superposition,[],[f1764,f2054])).
fof(f28559,plain,(
  ( ! [X771,X769,X770,X768] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X768))) = s(bool,happ(s(fun(bool,bool),X768),s(bool,X769))) | s(bool,X769) = s(bool,X770) | s(bool,'T_0') = s(bool,X770) | s(bool,X771) = s(bool,'const.min.@_1'(s(fun(bool,bool),X768))) | s(bool,'T_0') = s(bool,X771)) )),
  inference(superposition,[],[f303,f2054])).
fof(f28557,plain,(
  ( ! [X761,X757,X759,X760,X762,X758] : (s(bool,X760) != s(bool,f(s(X757,sK23(X758,X759,bool,X757)))) | s(fun(X757,bool),X758) = s(fun(X757,bool),X759) | s(bool,X760) = s(bool,X761) | s(bool,'T_0') = s(bool,X761) | s(bool,X762) = s(bool,g(s(X757,sK23(X758,X759,bool,X757)))) | s(bool,'T_0') = s(bool,X762)) )),
  inference(superposition,[],[f366,f2054])).
fof(f28301,plain,(
  ( ! [X85,X83,X86,X84] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X83))) | s(bool,X84) = s(bool,X85) | s(bool,X85) = s(bool,happ(s(fun(bool,bool),X83),s(bool,'F_0'))) | s(bool,X84) = s(bool,X86) | s(bool,'T_0') = s(bool,X86) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X83)))) )),
  inference(superposition,[],[f2054,f1768])).
fof(f28300,plain,(
  ( ! [X80,X78,X81,X79,X82] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X78,bool),X79))) | s(bool,X80) = s(bool,X81) | s(bool,X81) = s(bool,happ(s(fun(X78,bool),X79),s(X78,'const.min.@_1'(s(fun(X78,bool),X79))))) | s(bool,X80) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f2054,f303])).
fof(f28272,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X2) = s(bool,X3) | s(bool,X1) = s(bool,X3) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4) | s(bool,X5) = s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7) | s(bool,X0) = s(bool,X5)) )),
  inference(superposition,[],[f2054,f2054])).
fof(f28262,plain,(
  ( ! [X125,X123,X126,X124] : (s(bool,X124) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X123))) | s(bool,X124) = s(bool,X125) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X123),s(bool,'F_0'))) | s(bool,X125) = s(bool,X126) | s(bool,'T_0') = s(bool,X126) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X123)))) )),
  inference(superposition,[],[f2054,f1768])).
fof(f28261,plain,(
  ( ! [X121,X118,X122,X120,X119] : (s(bool,X120) = s(bool,'|dollar|exists_1'(s(fun(X118,bool),X119))) | s(bool,X120) = s(bool,X121) | s(bool,'T_0') = s(bool,happ(s(fun(X118,bool),X119),s(X118,'const.min.@_1'(s(fun(X118,bool),X119))))) | s(bool,X121) = s(bool,X122) | s(bool,'T_0') = s(bool,X122)) )),
  inference(superposition,[],[f2054,f303])).
fof(f31209,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X2)) )),
  inference(condensation,[],[f31208])).
fof(f31208,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X0) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(duplicate_literal_removal,[],[f28237])).
fof(f28237,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X0) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2054,f270])).
fof(f28227,plain,(
  ( ! [X138,X136,X135,X137] : (s(bool,X136) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X135))) | s(bool,X137) = s(bool,happ(s(fun(bool,bool),X135),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X136) | s(bool,X137) = s(bool,X138) | s(bool,'T_0') = s(bool,X138) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X135)))) )),
  inference(superposition,[],[f2054,f1768])).
fof(f28226,plain,(
  ( ! [X134,X132,X130,X133,X131] : (s(bool,X132) = s(bool,'|dollar|exists_1'(s(fun(X130,bool),X131))) | s(bool,X133) = s(bool,happ(s(fun(X130,bool),X131),s(X130,'const.min.@_1'(s(fun(X130,bool),X131))))) | s(bool,'T_0') = s(bool,X132) | s(bool,X133) = s(bool,X134) | s(bool,'T_0') = s(bool,X134)) )),
  inference(superposition,[],[f2054,f303])).
fof(f28157,plain,(
  ( ! [X138,X136,X135,X137] : (s(bool,X136) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X135))) | s(bool,X136) = s(bool,X137) | s(bool,'T_0') = s(bool,X137) | s(bool,X138) = s(bool,happ(s(fun(bool,bool),X135),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X138) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X135)))) )),
  inference(superposition,[],[f2054,f1768])).
fof(f28156,plain,(
  ( ! [X134,X132,X130,X133,X131] : (s(bool,X132) = s(bool,'|dollar|exists_1'(s(fun(X130,bool),X131))) | s(bool,X132) = s(bool,X133) | s(bool,'T_0') = s(bool,X133) | s(bool,X134) = s(bool,happ(s(fun(X130,bool),X131),s(X130,'const.min.@_1'(s(fun(X130,bool),X131))))) | s(bool,'T_0') = s(bool,X134)) )),
  inference(superposition,[],[f2054,f303])).
fof(f31391,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X2) | s(bool,X3) = s(bool,X4) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X5) | s(bool,X1) = s(bool,X3)) )),
  inference(condensation,[],[f31390])).
fof(f31390,plain,(
  ( ! [X94,X92,X90,X88,X93,X91,X89] : (s(bool,X88) = s(bool,X89) | s(bool,X89) = s(bool,X90) | s(bool,'T_0') = s(bool,X90) | s(bool,'T_0') = s(bool,X91) | s(bool,X92) = s(bool,X93) | s(bool,X93) = s(bool,X94) | s(bool,'T_0') = s(bool,X94) | s(bool,X88) = s(bool,X92)) )),
  inference(duplicate_literal_removal,[],[f28147])).
fof(f28147,plain,(
  ( ! [X94,X92,X90,X88,X93,X91,X89] : (s(bool,X88) = s(bool,X89) | s(bool,X89) = s(bool,X90) | s(bool,'T_0') = s(bool,X90) | s(bool,'T_0') = s(bool,X91) | s(bool,'T_0') = s(bool,X91) | s(bool,X92) = s(bool,X93) | s(bool,X93) = s(bool,X94) | s(bool,'T_0') = s(bool,X94) | s(bool,X88) = s(bool,X92)) )),
  inference(superposition,[],[f2054,f2054])).
fof(f28033,plain,(
  ( ! [X10,X8,X7,X9] : (sP12(happ(s(fun(bool,bool),X7),s(bool,X8)),X9,X10) | ~p(s(bool,X9)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X7))) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X7)))) )),
  inference(superposition,[],[f4987,f783])).
fof(f28032,plain,(
  ( ! [X6,X4,X5,X3] : (sP12(happ(s(fun(bool,bool),X3),s(bool,'T_0')),X4,X5) | ~p(s(bool,X4)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) | s(bool,X6) = s(bool,'const.min.@_1'(s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f4987,f783])).
fof(f28031,plain,(
  ( ! [X2,X0,X1] : (sP12(happ(s(fun(bool,bool),X0),s(bool,'F_0')),X1,X2) | ~p(s(bool,X1)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f4987,f270])).
fof(f28007,plain,(
  ( ! [X10,X8,X7,X9] : (sP12(happ(s(fun(bool,bool),X7),s(bool,X8)),X9,X10) | p(s(bool,X10)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X7))) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X7)))) )),
  inference(superposition,[],[f4986,f783])).
fof(f28006,plain,(
  ( ! [X6,X4,X5,X3] : (sP12(happ(s(fun(bool,bool),X3),s(bool,'T_0')),X4,X5) | p(s(bool,X5)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) | s(bool,X6) = s(bool,'const.min.@_1'(s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f4986,f783])).
fof(f28005,plain,(
  ( ! [X2,X0,X1] : (sP12(happ(s(fun(bool,bool),X0),s(bool,'F_0')),X1,X2) | p(s(bool,X2)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f4986,f270])).
fof(f27993,plain,(
  ( ! [X10,X8,X7,X9] : (sP10(happ(s(fun(bool,bool),X7),s(bool,X8)),X9,X10) | ~p(s(bool,X9)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X7))) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X7)))) )),
  inference(superposition,[],[f4983,f783])).
fof(f27992,plain,(
  ( ! [X6,X4,X5,X3] : (sP10(happ(s(fun(bool,bool),X3),s(bool,'T_0')),X4,X5) | ~p(s(bool,X4)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) | s(bool,X6) = s(bool,'const.min.@_1'(s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f4983,f783])).
fof(f27991,plain,(
  ( ! [X2,X0,X1] : (sP10(happ(s(fun(bool,bool),X0),s(bool,'F_0')),X1,X2) | ~p(s(bool,X1)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f4983,f270])).
fof(f27970,plain,(
  ( ! [X10,X8,X7,X9] : (sP10(happ(s(fun(bool,bool),X7),s(bool,X8)),X9,X10) | ~p(s(bool,X10)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X7))) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X7)))) )),
  inference(superposition,[],[f4982,f783])).
fof(f27969,plain,(
  ( ! [X6,X4,X5,X3] : (sP10(happ(s(fun(bool,bool),X3),s(bool,'T_0')),X4,X5) | ~p(s(bool,X5)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) | s(bool,X6) = s(bool,'const.min.@_1'(s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f4982,f783])).
fof(f27968,plain,(
  ( ! [X2,X0,X1] : (sP10(happ(s(fun(bool,bool),X0),s(bool,'F_0')),X1,X2) | ~p(s(bool,X2)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f4982,f270])).
fof(f27946,plain,(
  ( ! [X2,X3,X1] : (s(bool,X3) = s(bool,'|dollar|exists_1'(s(fun(X1,bool),X2))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,happ(s(fun(X1,bool),X2),s(X1,sK17(X2,X1))))) )),
  inference(subsumption_resolution,[],[f27937,f468])).
fof(f27937,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,'F_0')) | s(bool,X3) = s(bool,'|dollar|exists_1'(s(fun(X1,bool),X2))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,happ(s(fun(X1,bool),X2),s(X1,sK17(X2,X1))))) )),
  inference(superposition,[],[f4303,f270])).
fof(f27931,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,happ(s(fun(bool,bool),X5),s(bool,X6)))) | s(bool,X7) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X5))) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,sK17(X5,bool))) )),
  inference(superposition,[],[f4303,f783])).
fof(f27930,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,happ(s(fun(bool,bool),X2),s(bool,'T_0')))) | s(bool,X3) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X2))) | s(bool,'T_0') = s(bool,X3) | s(bool,X4) = s(bool,sK17(X2,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f4303,f783])).
fof(f27929,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | s(bool,X1) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,sK17(X0,bool))) )),
  inference(superposition,[],[f4303,f270])).
fof(f26105,plain,(
  ( ! [X2,X3,X1] : (s(bool,X3) = s(bool,'|dollar|exists_1'(s(fun(X1,bool),X2))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,happ(s(fun(X1,bool),X2),s(X1,sK16(X2,X1))))) )),
  inference(subsumption_resolution,[],[f26096,f468])).
fof(f26096,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,'F_0')) | s(bool,X3) = s(bool,'|dollar|exists_1'(s(fun(X1,bool),X2))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,happ(s(fun(X1,bool),X2),s(X1,sK16(X2,X1))))) )),
  inference(superposition,[],[f3897,f270])).
fof(f26090,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,happ(s(fun(bool,bool),X5),s(bool,X6)))) | s(bool,X7) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X5))) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,sK16(X5,bool))) )),
  inference(superposition,[],[f3897,f783])).
fof(f26089,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,happ(s(fun(bool,bool),X2),s(bool,'T_0')))) | s(bool,X3) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X2))) | s(bool,'T_0') = s(bool,X3) | s(bool,X4) = s(bool,sK16(X2,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f3897,f783])).
fof(f26088,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | s(bool,X1) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,sK16(X0,bool))) )),
  inference(superposition,[],[f3897,f270])).
fof(f26077,plain,(
  ( ! [X6,X7,X5] : (sP1(X7,happ(s(fun(bool,bool),X5),s(bool,X6))) | p(s(bool,X7)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X5))) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X5)))) )),
  inference(superposition,[],[f4975,f783])).
fof(f26076,plain,(
  ( ! [X4,X2,X3] : (sP1(X3,happ(s(fun(bool,bool),X2),s(bool,'T_0'))) | p(s(bool,X3)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X2))) | s(bool,X4) = s(bool,'const.min.@_1'(s(fun(bool,bool),X2))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f4975,f783])).
fof(f26075,plain,(
  ( ! [X0,X1] : (sP1(X1,happ(s(fun(bool,bool),X0),s(bool,'F_0'))) | p(s(bool,X1)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f4975,f270])).
fof(f26035,plain,(
  ( ! [X6,X7,X5] : (sP1(happ(s(fun(bool,bool),X5),s(bool,X6)),X7) | p(s(bool,X7)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X5))) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X5)))) )),
  inference(superposition,[],[f4974,f783])).
fof(f26034,plain,(
  ( ! [X4,X2,X3] : (sP1(happ(s(fun(bool,bool),X2),s(bool,'T_0')),X3) | p(s(bool,X3)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X2))) | s(bool,X4) = s(bool,'const.min.@_1'(s(fun(bool,bool),X2))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f4974,f783])).
fof(f26033,plain,(
  ( ! [X0,X1] : (sP1(happ(s(fun(bool,bool),X0),s(bool,'F_0')),X1) | p(s(bool,X1)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f4974,f270])).
fof(f25994,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X3),s(bool,X4))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X3)))) )),
  inference(superposition,[],[f4969,f783])).
fof(f25993,plain,(
  ( ! [X2,X1] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1))) | s(bool,X2) = s(bool,'const.min.@_1'(s(fun(bool,bool),X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f4969,f783])).
fof(f25992,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f4969,f270])).
fof(f25973,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),X1))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.min.@_1'(s(fun('type.extreal.extreal',bool),X1))),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X1)))))) )),
  inference(subsumption_resolution,[],[f25966,f468])).
fof(f25966,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),X1))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.min.@_1'(s(fun('type.extreal.extreal',bool),X1))),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X1)))))) )),
  inference(superposition,[],[f4965,f270])).
fof(f25921,plain,(
  ( ! [X4,X2,X5,X3] : (s(bool,X4) = s(bool,happ(s(fun(X2,bool),X3),s(X2,X5))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X2,bool),X3)))) )),
  inference(subsumption_resolution,[],[f25914,f468])).
fof(f25914,plain,(
  ( ! [X4,X2,X5,X3] : (p(s(bool,'F_0')) | s(bool,X4) = s(bool,happ(s(fun(X2,bool),X3),s(X2,X5))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X2,bool),X3)))) )),
  inference(superposition,[],[f2629,f270])).
fof(f25900,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,X9) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X8),s('type.extreal.extreal',X7))) | s(bool,'T_0') = s(bool,X9) | s(bool,X10) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X8))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f25894,f263])).
fof(f25894,plain,(
  ( ! [X10,X8,X7,X9] : (~p(s(bool,'T_0')) | s(bool,X9) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X8),s('type.extreal.extreal',X7))) | s(bool,'T_0') = s(bool,X9) | s(bool,X10) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X8))) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f2635,f783])).
fof(f24674,plain,(
  ( ! [X6,X7,X5] : (sP0(happ(s(fun(bool,bool),X5),s(bool,X6)),X7) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X5))) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X5)))) )),
  inference(superposition,[],[f4973,f783])).
fof(f24673,plain,(
  ( ! [X4,X2,X3] : (sP0(happ(s(fun(bool,bool),X2),s(bool,'T_0')),X3) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X2))) | s(bool,X4) = s(bool,'const.min.@_1'(s(fun(bool,bool),X2))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f4973,f783])).
fof(f24672,plain,(
  ( ! [X0,X1] : (sP0(happ(s(fun(bool,bool),X0),s(bool,'F_0')),X1) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f4973,f270])).
fof(f24650,plain,(
  ( ! [X6,X7,X5] : (sP0(X7,happ(s(fun(bool,bool),X5),s(bool,X6))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X5))) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X5)))) )),
  inference(superposition,[],[f4972,f783])).
fof(f24649,plain,(
  ( ! [X4,X2,X3] : (sP0(X3,happ(s(fun(bool,bool),X2),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X2))) | s(bool,X4) = s(bool,'const.min.@_1'(s(fun(bool,bool),X2))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f4972,f783])).
fof(f24648,plain,(
  ( ! [X0,X1] : (sP0(X1,happ(s(fun(bool,bool),X0),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f4972,f270])).
fof(f24621,plain,(
  ( ! [X70,X72,X71,X69] : (s(bool,'|dollar|exists_1'(s(fun(X71,bool),X72))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X69),s('type.extreal.extreal',X70))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X71,bool),X72))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X70),s('type.extreal.extreal',X69)))) )),
  inference(subsumption_resolution,[],[f24593,f263])).
fof(f24593,plain,(
  ( ! [X70,X72,X71,X69] : (~p(s(bool,'T_0')) | s(bool,'|dollar|exists_1'(s(fun(X71,bool),X72))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X69),s('type.extreal.extreal',X70))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X71,bool),X72))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X70),s('type.extreal.extreal',X69)))) )),
  inference(superposition,[],[f4971,f3886])).
fof(f24585,plain,(
  ( ! [X52,X50,X51] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X50)))) | s(bool,'|dollar|exists_1'(s(fun(X51,bool),X52))) = s(bool,happ(s(fun(bool,bool),X50),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X51,bool),X52))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X50)))) )),
  inference(superposition,[],[f4971,f1768])).
fof(f24613,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,'|dollar|exists_1'(s(fun(X44,bool),X45))) = s(bool,happ(s(fun(X42,bool),X43),s(X42,'const.min.@_1'(s(fun(X42,bool),X43))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X44,bool),X45))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X42,bool),X43)))) )),
  inference(subsumption_resolution,[],[f24583,f263])).
fof(f24583,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,'T_0')) | s(bool,'|dollar|exists_1'(s(fun(X44,bool),X45))) = s(bool,happ(s(fun(X42,bool),X43),s(X42,'const.min.@_1'(s(fun(X42,bool),X43))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X44,bool),X45))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X42,bool),X43)))) )),
  inference(superposition,[],[f4971,f1769])).
fof(f24596,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,X3) = s(bool,'|dollar|exists_1'(s(fun(X4,bool),X5))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X4,bool),X5))) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f24571,f263])).
fof(f24571,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | s(bool,X3) = s(bool,'|dollar|exists_1'(s(fun(X4,bool),X5))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X4,bool),X5))) | s(bool,X3) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f4971,f783])).
fof(f24555,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK16(X1,X0))))) )),
  inference(subsumption_resolution,[],[f24546,f468])).
fof(f24546,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK16(X1,X0))))) )),
  inference(superposition,[],[f4967,f270])).
fof(f24540,plain,(
  ( ! [X4,X3] : (p(s(bool,happ(s(fun(bool,bool),X3),s(bool,X4)))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,sK16(X3,bool))) )),
  inference(superposition,[],[f4967,f783])).
fof(f24539,plain,(
  ( ! [X2,X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1))) | s(bool,X2) = s(bool,sK16(X1,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f4967,f783])).
fof(f24538,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,sK16(X0,bool))) )),
  inference(superposition,[],[f4967,f270])).
fof(f24523,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK17(X1,X0))))) )),
  inference(subsumption_resolution,[],[f24514,f468])).
fof(f24514,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK17(X1,X0))))) )),
  inference(superposition,[],[f4966,f270])).
fof(f24508,plain,(
  ( ! [X4,X3] : (p(s(bool,happ(s(fun(bool,bool),X3),s(bool,X4)))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,sK17(X3,bool))) )),
  inference(superposition,[],[f4966,f783])).
fof(f24507,plain,(
  ( ! [X2,X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1))) | s(bool,X2) = s(bool,sK17(X1,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f4966,f783])).
fof(f24506,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,sK17(X0,bool))) )),
  inference(superposition,[],[f4966,f270])).
fof(f24489,plain,(
  ( ! [X167,X165,X166,X164,X168] : (p(s(bool,X166)) | ~sP12(X167,X166,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X164),s('type.extreal.extreal',X165))) | s(bool,X167) = s(bool,X168) | s(bool,'T_0') = s(bool,X168) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X165),s('type.extreal.extreal',X164)))) )),
  inference(subsumption_resolution,[],[f24436,f263])).
fof(f24436,plain,(
  ( ! [X167,X165,X166,X164,X168] : (~p(s(bool,'T_0')) | p(s(bool,X166)) | ~sP12(X167,X166,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X164),s('type.extreal.extreal',X165))) | s(bool,X167) = s(bool,X168) | s(bool,'T_0') = s(bool,X168) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X165),s('type.extreal.extreal',X164)))) )),
  inference(superposition,[],[f2623,f3886])).
fof(f24488,plain,(
  ( ! [X163,X161,X159,X162,X160] : (sP12(X161,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X159),s('type.extreal.extreal',X160)),X162) | s(bool,X161) = s(bool,X163) | s(bool,'T_0') = s(bool,X163) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X160),s('type.extreal.extreal',X159)))) )),
  inference(subsumption_resolution,[],[f24435,f263])).
fof(f24435,plain,(
  ( ! [X163,X161,X159,X162,X160] : (~p(s(bool,'T_0')) | sP12(X161,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X159),s('type.extreal.extreal',X160)),X162) | s(bool,X161) = s(bool,X163) | s(bool,'T_0') = s(bool,X163) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X160),s('type.extreal.extreal',X159)))) )),
  inference(superposition,[],[f2622,f3886])).
fof(f24487,plain,(
  ( ! [X158,X156,X154,X157,X155] : (sP10(X156,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X154),s('type.extreal.extreal',X155)),X157) | s(bool,X156) = s(bool,X158) | s(bool,'T_0') = s(bool,X158) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X155),s('type.extreal.extreal',X154)))) )),
  inference(subsumption_resolution,[],[f24434,f263])).
fof(f24434,plain,(
  ( ! [X158,X156,X154,X157,X155] : (~p(s(bool,'T_0')) | sP10(X156,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X154),s('type.extreal.extreal',X155)),X157) | s(bool,X156) = s(bool,X158) | s(bool,'T_0') = s(bool,X158) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X155),s('type.extreal.extreal',X154)))) )),
  inference(superposition,[],[f2618,f3886])).
fof(f24486,plain,(
  ( ! [X152,X151,X149,X153,X150] : (sP10(X151,X152,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X149),s('type.extreal.extreal',X150))) | s(bool,X151) = s(bool,X153) | s(bool,'T_0') = s(bool,X153) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X150),s('type.extreal.extreal',X149)))) )),
  inference(subsumption_resolution,[],[f24433,f263])).
fof(f24433,plain,(
  ( ! [X152,X151,X149,X153,X150] : (~p(s(bool,'T_0')) | sP10(X151,X152,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X149),s('type.extreal.extreal',X150))) | s(bool,X151) = s(bool,X153) | s(bool,'T_0') = s(bool,X153) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X150),s('type.extreal.extreal',X149)))) )),
  inference(superposition,[],[f2617,f3886])).
fof(f24485,plain,(
  ( ! [X146,X144,X147,X145,X148] : (p(s(bool,X146)) | ~sP8(X146,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X144),s('type.extreal.extreal',X145)),X147) | s(bool,X147) = s(bool,X148) | s(bool,'T_0') = s(bool,X148) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X145),s('type.extreal.extreal',X144)))) )),
  inference(subsumption_resolution,[],[f24432,f263])).
fof(f24432,plain,(
  ( ! [X146,X144,X147,X145,X148] : (~p(s(bool,'T_0')) | p(s(bool,X146)) | ~sP8(X146,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X144),s('type.extreal.extreal',X145)),X147) | s(bool,X147) = s(bool,X148) | s(bool,'T_0') = s(bool,X148) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X145),s('type.extreal.extreal',X144)))) )),
  inference(superposition,[],[f2615,f3886])).
fof(f24484,plain,(
  ( ! [X134,X132,X135,X133] : (s(bool,X134) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X132),s('type.extreal.extreal',X133))) | s(bool,X134) = s(bool,X135) | s(bool,'T_0') = s(bool,X135) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X133),s('type.extreal.extreal',X132)))) )),
  inference(subsumption_resolution,[],[f24429,f263])).
fof(f24429,plain,(
  ( ! [X134,X132,X135,X133] : (~p(s(bool,'T_0')) | s(bool,X134) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X132),s('type.extreal.extreal',X133))) | s(bool,X134) = s(bool,X135) | s(bool,'T_0') = s(bool,X135) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X133),s('type.extreal.extreal',X132)))) )),
  inference(superposition,[],[f2606,f3886])).
fof(f24483,plain,(
  ( ! [X107,X105,X106,X104] : (~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',X104),s('type.extreal.extreal',X105)),X106,X107) | sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',X104),s('type.extreal.extreal',X105)),X106,X107) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X105),s('type.extreal.extreal',X104)))) )),
  inference(subsumption_resolution,[],[f24421,f263])).
fof(f24421,plain,(
  ( ! [X107,X105,X106,X104] : (~p(s(bool,'T_0')) | ~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',X104),s('type.extreal.extreal',X105)),X106,X107) | sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',X104),s('type.extreal.extreal',X105)),X106,X107) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X105),s('type.extreal.extreal',X104)))) )),
  inference(superposition,[],[f453,f3886])).
fof(f24482,plain,(
  ( ! [X103,X101,X102,X100] : (~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',X100),s('type.extreal.extreal',X101)),X102,X103) | sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal',X100),s('type.extreal.extreal',X101)),X102,X103) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X101),s('type.extreal.extreal',X100)))) )),
  inference(subsumption_resolution,[],[f24420,f263])).
fof(f24420,plain,(
  ( ! [X103,X101,X102,X100] : (~p(s(bool,'T_0')) | ~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',X100),s('type.extreal.extreal',X101)),X102,X103) | sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal',X100),s('type.extreal.extreal',X101)),X102,X103) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X101),s('type.extreal.extreal',X100)))) )),
  inference(superposition,[],[f452,f3886])).
fof(f24481,plain,(
  ( ! [X99,X97,X98,X96] : (~p(s(bool,X98)) | p(s(bool,X99)) | ~sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal',X96),s('type.extreal.extreal',X97)),X99,X98) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X97),s('type.extreal.extreal',X96)))) )),
  inference(subsumption_resolution,[],[f24419,f263])).
fof(f24419,plain,(
  ( ! [X99,X97,X98,X96] : (~p(s(bool,'T_0')) | ~p(s(bool,X98)) | p(s(bool,X99)) | ~sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal',X96),s('type.extreal.extreal',X97)),X99,X98) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X97),s('type.extreal.extreal',X96)))) )),
  inference(superposition,[],[f449,f3886])).
fof(f24480,plain,(
  ( ! [X94,X92,X95,X93] : (~p(s(bool,X94)) | sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal',X92),s('type.extreal.extreal',X93)),X94,X95) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X93),s('type.extreal.extreal',X92)))) )),
  inference(subsumption_resolution,[],[f24418,f263])).
fof(f24418,plain,(
  ( ! [X94,X92,X95,X93] : (~p(s(bool,'T_0')) | ~p(s(bool,X94)) | sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal',X92),s('type.extreal.extreal',X93)),X94,X95) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X93),s('type.extreal.extreal',X92)))) )),
  inference(superposition,[],[f448,f3886])).
fof(f24479,plain,(
  ( ! [X90,X88,X91,X89] : (p(s(bool,X90)) | sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal',X88),s('type.extreal.extreal',X89)),X91,X90) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X89),s('type.extreal.extreal',X88)))) )),
  inference(subsumption_resolution,[],[f24417,f263])).
fof(f24417,plain,(
  ( ! [X90,X88,X91,X89] : (~p(s(bool,'T_0')) | p(s(bool,X90)) | sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal',X88),s('type.extreal.extreal',X89)),X91,X90) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X89),s('type.extreal.extreal',X88)))) )),
  inference(superposition,[],[f447,f3886])).
fof(f24478,plain,(
  ( ! [X87,X85,X86,X84] : (p(s(bool,X86)) | ~p(s(bool,X87)) | ~sP11(X87,X86,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X84),s('type.extreal.extreal',X85))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X85),s('type.extreal.extreal',X84)))) )),
  inference(subsumption_resolution,[],[f24416,f263])).
fof(f24416,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,'T_0')) | p(s(bool,X86)) | ~p(s(bool,X87)) | ~sP11(X87,X86,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X84),s('type.extreal.extreal',X85))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X85),s('type.extreal.extreal',X84)))) )),
  inference(superposition,[],[f442,f3886])).
fof(f24477,plain,(
  ( ! [X80,X83,X81,X82] : (p(s(bool,X82)) | p(s(bool,X83)) | ~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',X80),s('type.extreal.extreal',X81)),X83,X82) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X81),s('type.extreal.extreal',X80)))) )),
  inference(subsumption_resolution,[],[f24415,f263])).
fof(f24415,plain,(
  ( ! [X80,X83,X81,X82] : (~p(s(bool,'T_0')) | p(s(bool,X82)) | p(s(bool,X83)) | ~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',X80),s('type.extreal.extreal',X81)),X83,X82) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X81),s('type.extreal.extreal',X80)))) )),
  inference(superposition,[],[f432,f3886])).
fof(f24476,plain,(
  ( ! [X78,X76,X79,X77] : (~p(s(bool,X78)) | sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',X76),s('type.extreal.extreal',X77)),X78,X79) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X77),s('type.extreal.extreal',X76)))) )),
  inference(subsumption_resolution,[],[f24414,f263])).
fof(f24414,plain,(
  ( ! [X78,X76,X79,X77] : (~p(s(bool,'T_0')) | ~p(s(bool,X78)) | sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',X76),s('type.extreal.extreal',X77)),X78,X79) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X77),s('type.extreal.extreal',X76)))) )),
  inference(superposition,[],[f431,f3886])).
fof(f24475,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,X74)) | sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',X72),s('type.extreal.extreal',X73)),X75,X74) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X73),s('type.extreal.extreal',X72)))) )),
  inference(subsumption_resolution,[],[f24413,f263])).
fof(f24413,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,'T_0')) | ~p(s(bool,X74)) | sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',X72),s('type.extreal.extreal',X73)),X75,X74) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X73),s('type.extreal.extreal',X72)))) )),
  inference(superposition,[],[f430,f3886])).
fof(f24474,plain,(
  ( ! [X70,X68,X71,X69] : (p(s(bool,X70)) | p(s(bool,X71)) | ~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',X68),s('type.extreal.extreal',X69)),X70,X71) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X69),s('type.extreal.extreal',X68)))) )),
  inference(subsumption_resolution,[],[f24412,f263])).
fof(f24412,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,'T_0')) | p(s(bool,X70)) | p(s(bool,X71)) | ~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',X68),s('type.extreal.extreal',X69)),X70,X71) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X69),s('type.extreal.extreal',X68)))) )),
  inference(superposition,[],[f425,f3886])).
fof(f24473,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,X66)) | p(s(bool,X67)) | ~sP8(X67,X66,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X64),s('type.extreal.extreal',X65))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X65),s('type.extreal.extreal',X64)))) )),
  inference(subsumption_resolution,[],[f24411,f263])).
fof(f24411,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,'T_0')) | ~p(s(bool,X66)) | p(s(bool,X67)) | ~sP8(X67,X66,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X64),s('type.extreal.extreal',X65))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X65),s('type.extreal.extreal',X64)))) )),
  inference(superposition,[],[f417,f3886])).
fof(f24472,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,X62)) | p(s(bool,X63)) | ~sP7(X62,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X60),s('type.extreal.extreal',X61)),X63) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X61),s('type.extreal.extreal',X60)))) )),
  inference(subsumption_resolution,[],[f24410,f263])).
fof(f24410,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,'T_0')) | ~p(s(bool,X62)) | p(s(bool,X63)) | ~sP7(X62,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X60),s('type.extreal.extreal',X61)),X63) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X61),s('type.extreal.extreal',X60)))) )),
  inference(superposition,[],[f408,f3886])).
fof(f24471,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,X58)) | p(s(bool,X59)) | ~sP4(X58,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X56),s('type.extreal.extreal',X57)),X59) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X57),s('type.extreal.extreal',X56)))) )),
  inference(subsumption_resolution,[],[f24409,f263])).
fof(f24409,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,'T_0')) | ~p(s(bool,X58)) | p(s(bool,X59)) | ~sP4(X58,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X56),s('type.extreal.extreal',X57)),X59) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X57),s('type.extreal.extreal',X56)))) )),
  inference(superposition,[],[f389,f3886])).
fof(f24470,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X54)) | p(s(bool,X55)) | ~sP5(X54,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X52),s('type.extreal.extreal',X53)),X55) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X53),s('type.extreal.extreal',X52)))) )),
  inference(subsumption_resolution,[],[f24408,f263])).
fof(f24408,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,'T_0')) | ~p(s(bool,X54)) | p(s(bool,X55)) | ~sP5(X54,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X52),s('type.extreal.extreal',X53)),X55) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X53),s('type.extreal.extreal',X52)))) )),
  inference(superposition,[],[f385,f3886])).
fof(f24469,plain,(
  ( ! [X50,X48,X51,X49] : (sP6(X50,X51,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X48),s('type.extreal.extreal',X49))) | ~sP5(X50,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X48),s('type.extreal.extreal',X49)),X51) | ~sP4(X51,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X48),s('type.extreal.extreal',X49)),X50) | ~sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal',X48),s('type.extreal.extreal',X49)),X51,X50) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X49),s('type.extreal.extreal',X48)))) )),
  inference(subsumption_resolution,[],[f24407,f263])).
fof(f24407,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,'T_0')) | sP6(X50,X51,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X48),s('type.extreal.extreal',X49))) | ~sP5(X50,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X48),s('type.extreal.extreal',X49)),X51) | ~sP4(X51,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X48),s('type.extreal.extreal',X49)),X50) | ~sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal',X48),s('type.extreal.extreal',X49)),X51,X50) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X49),s('type.extreal.extreal',X48)))) )),
  inference(superposition,[],[f382,f3886])).
fof(f24468,plain,(
  ( ! [X47,X45,X46] : (p(s(bool,X47)) | sP1(X47,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X45),s('type.extreal.extreal',X46))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X46),s('type.extreal.extreal',X45)))) )),
  inference(subsumption_resolution,[],[f24406,f263])).
fof(f24406,plain,(
  ( ! [X47,X45,X46] : (~p(s(bool,'T_0')) | p(s(bool,X47)) | sP1(X47,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X45),s('type.extreal.extreal',X46))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X46),s('type.extreal.extreal',X45)))) )),
  inference(superposition,[],[f353,f3886])).
fof(f24467,plain,(
  ( ! [X43,X44,X42] : (p(s(bool,X44)) | sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',X42),s('type.extreal.extreal',X43)),X44) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X43),s('type.extreal.extreal',X42)))) )),
  inference(subsumption_resolution,[],[f24405,f263])).
fof(f24405,plain,(
  ( ! [X43,X44,X42] : (~p(s(bool,'T_0')) | p(s(bool,X44)) | sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',X42),s('type.extreal.extreal',X43)),X44) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X43),s('type.extreal.extreal',X42)))) )),
  inference(superposition,[],[f350,f3886])).
fof(f24466,plain,(
  ( ! [X39,X41,X40] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal',X39),s('type.extreal.extreal',X40)),X41) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X40),s('type.extreal.extreal',X39)))) )),
  inference(subsumption_resolution,[],[f24404,f263])).
fof(f24404,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'T_0')) | sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal',X39),s('type.extreal.extreal',X40)),X41) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X40),s('type.extreal.extreal',X39)))) )),
  inference(superposition,[],[f335,f3886])).
fof(f24465,plain,(
  ( ! [X37,X38,X36] : (sP0(X38,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X36),s('type.extreal.extreal',X37))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X37),s('type.extreal.extreal',X36)))) )),
  inference(subsumption_resolution,[],[f24403,f263])).
fof(f24403,plain,(
  ( ! [X37,X38,X36] : (~p(s(bool,'T_0')) | sP0(X38,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X36),s('type.extreal.extreal',X37))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X37),s('type.extreal.extreal',X36)))) )),
  inference(superposition,[],[f334,f3886])).
fof(f24464,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X35)) | s(bool,X35) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X33),s('type.extreal.extreal',X34))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X34),s('type.extreal.extreal',X33)))) )),
  inference(subsumption_resolution,[],[f24402,f263])).
fof(f24402,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,'T_0')) | ~p(s(bool,X35)) | s(bool,X35) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X33),s('type.extreal.extreal',X34))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X34),s('type.extreal.extreal',X33)))) )),
  inference(superposition,[],[f321,f3886])).
fof(f24463,plain,(
  ( ! [X24,X23] : (s(bool,'F_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X23),s('type.extreal.extreal',X24))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X24),s('type.extreal.extreal',X23)))) )),
  inference(subsumption_resolution,[],[f24398,f263])).
fof(f24398,plain,(
  ( ! [X24,X23] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X23),s('type.extreal.extreal',X24))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X24),s('type.extreal.extreal',X23)))) )),
  inference(superposition,[],[f280,f3886])).
fof(f24450,plain,(
  ( ! [X19,X20] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X19),s('type.extreal.extreal',X20)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X20),s('type.extreal.extreal',X19)))) )),
  inference(trivial_inequality_removal,[],[f24396])).
fof(f24396,plain,(
  ( ! [X19,X20] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X19),s('type.extreal.extreal',X20)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X20),s('type.extreal.extreal',X19)))) )),
  inference(superposition,[],[f276,f3886])).
fof(f24462,plain,(
  ( ! [X15,X16] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X16))),s('type.extreal.extreal',X15)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X15),s('type.extreal.extreal',sK22(X15,X16))))) )),
  inference(subsumption_resolution,[],[f24394,f263])).
fof(f24394,plain,(
  ( ! [X15,X16] : (~p(s(bool,'T_0')) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X16))),s('type.extreal.extreal',X15)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X15),s('type.extreal.extreal',sK22(X15,X16))))) )),
  inference(superposition,[],[f328,f3886])).
fof(f24461,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X9),s('type.extreal.extreal',X11)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X11),s('type.extreal.extreal',X10)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X10),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X9)))))) )),
  inference(subsumption_resolution,[],[f24389,f263])).
fof(f24389,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X9),s('type.extreal.extreal',X11)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X11),s('type.extreal.extreal',X10)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X10),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X9)))))) )),
  inference(superposition,[],[f326,f3886])).
fof(f24460,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X1)))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X0)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))) )),
  inference(subsumption_resolution,[],[f24385,f263])).
fof(f24385,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X1)))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X0)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))) )),
  inference(superposition,[],[f367,f3886])).
fof(f24271,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X65)))) | p(s(bool,X66)) | ~sP12(X67,X66,happ(s(fun(bool,bool),X65),s(bool,'F_0'))) | s(bool,X67) = s(bool,X68) | s(bool,'T_0') = s(bool,X68) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X65)))) )),
  inference(superposition,[],[f2623,f1768])).
fof(f24270,plain,(
  ( ! [X61,X64,X62,X60,X63] : (~p(s(bool,'|dollar|exists_1'(s(fun(X60,bool),X61)))) | p(s(bool,X62)) | ~sP12(X63,X62,happ(s(fun(X60,bool),X61),s(X60,'const.min.@_1'(s(fun(X60,bool),X61))))) | s(bool,X63) = s(bool,X64) | s(bool,'T_0') = s(bool,X64)) )),
  inference(superposition,[],[f2623,f303])).
fof(f24323,plain,(
  ( ! [X59,X57,X58,X56,X55] : (p(s(bool,X57)) | ~sP12(X58,X57,happ(s(fun(X55,bool),X56),s(X55,'const.min.@_1'(s(fun(X55,bool),X56))))) | s(bool,X58) = s(bool,X59) | s(bool,'T_0') = s(bool,X59) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X55,bool),X56)))) )),
  inference(subsumption_resolution,[],[f24269,f263])).
fof(f24269,plain,(
  ( ! [X59,X57,X58,X56,X55] : (~p(s(bool,'T_0')) | p(s(bool,X57)) | ~sP12(X58,X57,happ(s(fun(X55,bool),X56),s(X55,'const.min.@_1'(s(fun(X55,bool),X56))))) | s(bool,X58) = s(bool,X59) | s(bool,'T_0') = s(bool,X59) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X55,bool),X56)))) )),
  inference(superposition,[],[f2623,f1769])).
fof(f24316,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X2)) | ~sP12(X0,X2,'T_0') | s(bool,'T_0') = s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f24315])).
fof(f24315,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X2)) | ~sP12(X0,X2,'T_0') | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) )),
  inference(condensation,[],[f24314])).
fof(f24314,plain,(
  ( ! [X35,X33,X36,X34,X32] : (p(s(bool,X33)) | ~sP12(X34,X33,'T_0') | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35) | s(bool,X32) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(subsumption_resolution,[],[f24263,f846])).
fof(f24263,plain,(
  ( ! [X35,X33,X36,X34,X32] : (~p(s(bool,X32)) | p(s(bool,X33)) | ~sP12(X34,X33,'T_0') | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35) | s(bool,X32) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f2623,f783])).
fof(f24311,plain,(
  ( ! [X6,X4,X8,X7,X5] : (p(s(bool,X5)) | ~sP12(X6,X5,X4) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7) | s(bool,X4) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f24257,f263])).
fof(f24257,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,X5)) | ~sP12(X6,X5,X4) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7) | s(bool,X4) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f2623,f783])).
fof(f23728,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X65)))) | p(s(bool,X66)) | ~sP8(X66,happ(s(fun(bool,bool),X65),s(bool,'F_0')),X67) | s(bool,X67) = s(bool,X68) | s(bool,'T_0') = s(bool,X68) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X65)))) )),
  inference(superposition,[],[f2615,f1768])).
fof(f23727,plain,(
  ( ! [X61,X64,X62,X60,X63] : (~p(s(bool,'|dollar|exists_1'(s(fun(X60,bool),X61)))) | p(s(bool,X62)) | ~sP8(X62,happ(s(fun(X60,bool),X61),s(X60,'const.min.@_1'(s(fun(X60,bool),X61)))),X63) | s(bool,X63) = s(bool,X64) | s(bool,'T_0') = s(bool,X64)) )),
  inference(superposition,[],[f2615,f303])).
fof(f23780,plain,(
  ( ! [X59,X57,X58,X56,X55] : (p(s(bool,X57)) | ~sP8(X57,happ(s(fun(X55,bool),X56),s(X55,'const.min.@_1'(s(fun(X55,bool),X56)))),X58) | s(bool,X58) = s(bool,X59) | s(bool,'T_0') = s(bool,X59) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X55,bool),X56)))) )),
  inference(subsumption_resolution,[],[f23726,f263])).
fof(f23726,plain,(
  ( ! [X59,X57,X58,X56,X55] : (~p(s(bool,'T_0')) | p(s(bool,X57)) | ~sP8(X57,happ(s(fun(X55,bool),X56),s(X55,'const.min.@_1'(s(fun(X55,bool),X56)))),X58) | s(bool,X58) = s(bool,X59) | s(bool,'T_0') = s(bool,X59) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X55,bool),X56)))) )),
  inference(superposition,[],[f2615,f1769])).
fof(f23773,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X2)) | ~sP8(X2,'T_0',X0) | s(bool,'T_0') = s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f23772])).
fof(f23772,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X2)) | ~sP8(X2,'T_0',X0) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) )),
  inference(condensation,[],[f23771])).
fof(f23771,plain,(
  ( ! [X35,X33,X36,X34,X32] : (p(s(bool,X33)) | ~sP8(X33,'T_0',X34) | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35) | s(bool,X32) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(subsumption_resolution,[],[f23720,f846])).
fof(f23720,plain,(
  ( ! [X35,X33,X36,X34,X32] : (~p(s(bool,X32)) | p(s(bool,X33)) | ~sP8(X33,'T_0',X34) | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35) | s(bool,X32) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f2615,f783])).
fof(f23768,plain,(
  ( ! [X6,X4,X8,X7,X5] : (p(s(bool,X5)) | ~sP8(X5,X4,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7) | s(bool,X4) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f23714,f263])).
fof(f23714,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,X5)) | ~sP8(X5,X4,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7) | s(bool,X4) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f2615,f783])).
fof(f22041,plain,(
  ( ! [X52,X50,X51] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X50)))) | s(bool,X51) = s(bool,happ(s(fun(bool,bool),X50),s(bool,'F_0'))) | s(bool,X51) = s(bool,X52) | s(bool,'T_0') = s(bool,X52) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X50)))) )),
  inference(superposition,[],[f2606,f1768])).
fof(f22080,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,X44) = s(bool,happ(s(fun(X42,bool),X43),s(X42,'const.min.@_1'(s(fun(X42,bool),X43))))) | s(bool,X44) = s(bool,X45) | s(bool,'T_0') = s(bool,X45) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X42,bool),X43)))) )),
  inference(subsumption_resolution,[],[f22039,f263])).
fof(f22039,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,'T_0')) | s(bool,X44) = s(bool,happ(s(fun(X42,bool),X43),s(X42,'const.min.@_1'(s(fun(X42,bool),X43))))) | s(bool,X44) = s(bool,X45) | s(bool,'T_0') = s(bool,X45) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X42,bool),X43)))) )),
  inference(superposition,[],[f2606,f1769])).
fof(f21931,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X50)))) | s(bool,X51) != s(bool,happ(s(fun(bool,bool),X50),s(bool,'F_0'))) | s(bool,X51) = s(bool,X52) | s(bool,'T_0') = s(bool,X52) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X50)))) )),
  inference(superposition,[],[f2204,f1768])).
fof(f21459,plain,(
  ( ! [X4,X3] : (sP5('T_0',happ(s(fun(bool,bool),X3),s(bool,X4)),'|dollar|exists_1'(s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X3)))) )),
  inference(superposition,[],[f3862,f783])).
fof(f21458,plain,(
  ( ! [X2,X1] : (sP5('T_0',happ(s(fun(bool,bool),X1),s(bool,'T_0')),'|dollar|exists_1'(s(fun(bool,bool),X1))) | s(bool,X2) = s(bool,'const.min.@_1'(s(fun(bool,bool),X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f3862,f783])).
fof(f21457,plain,(
  ( ! [X0] : (sP5('T_0',happ(s(fun(bool,bool),X0),s(bool,'F_0')),'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f3862,f270])).
fof(f21430,plain,(
  ( ! [X4,X3] : (sP5('T_0','|dollar|exists_1'(s(fun(bool,bool),X3)),happ(s(fun(bool,bool),X3),s(bool,X4))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X3)))) )),
  inference(superposition,[],[f3757,f783])).
fof(f21429,plain,(
  ( ! [X2,X1] : (sP5('T_0','|dollar|exists_1'(s(fun(bool,bool),X1)),happ(s(fun(bool,bool),X1),s(bool,'T_0'))) | s(bool,X2) = s(bool,'const.min.@_1'(s(fun(bool,bool),X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f3757,f783])).
fof(f21428,plain,(
  ( ! [X0] : (sP5('T_0','|dollar|exists_1'(s(fun(bool,bool),X0)),happ(s(fun(bool,bool),X0),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f3757,f270])).
fof(f19879,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X59)))) | sP12(X60,happ(s(fun(bool,bool),X59),s(bool,'F_0')),X61) | s(bool,X60) = s(bool,X62) | s(bool,'T_0') = s(bool,X62) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X59)))) )),
  inference(superposition,[],[f2622,f1768])).
fof(f19878,plain,(
  ( ! [X57,X54,X58,X56,X55] : (~p(s(bool,'|dollar|exists_1'(s(fun(X54,bool),X55)))) | sP12(X56,happ(s(fun(X54,bool),X55),s(X54,'const.min.@_1'(s(fun(X54,bool),X55)))),X57) | s(bool,X56) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f2622,f303])).
fof(f19916,plain,(
  ( ! [X52,X50,X53,X51,X49] : (sP12(X51,happ(s(fun(X49,bool),X50),s(X49,'const.min.@_1'(s(fun(X49,bool),X50)))),X52) | s(bool,X51) = s(bool,X53) | s(bool,'T_0') = s(bool,X53) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X49,bool),X50)))) )),
  inference(subsumption_resolution,[],[f19877,f263])).
fof(f19877,plain,(
  ( ! [X52,X50,X53,X51,X49] : (~p(s(bool,'T_0')) | sP12(X51,happ(s(fun(X49,bool),X50),s(X49,'const.min.@_1'(s(fun(X49,bool),X50)))),X52) | s(bool,X51) = s(bool,X53) | s(bool,'T_0') = s(bool,X53) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X49,bool),X50)))) )),
  inference(superposition,[],[f2622,f1769])).
fof(f19911,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | sP12(X0,'T_0',X2) | s(bool,'T_0') = s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f19910])).
fof(f19910,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | sP12(X0,'T_0',X2) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) )),
  inference(condensation,[],[f19909])).
fof(f19909,plain,(
  ( ! [X35,X33,X36,X34,X32] : (sP12(X33,'T_0',X34) | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35) | s(bool,X32) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(subsumption_resolution,[],[f19873,f846])).
fof(f19873,plain,(
  ( ! [X35,X33,X36,X34,X32] : (~p(s(bool,X32)) | sP12(X33,'T_0',X34) | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35) | s(bool,X32) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f2622,f783])).
fof(f19906,plain,(
  ( ! [X6,X4,X8,X7,X5] : (sP12(X5,X4,X6) | s(bool,X5) = s(bool,X7) | s(bool,'T_0') = s(bool,X7) | s(bool,X4) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f19867,f263])).
fof(f19867,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~p(s(bool,'T_0')) | sP12(X5,X4,X6) | s(bool,X5) = s(bool,X7) | s(bool,'T_0') = s(bool,X7) | s(bool,X4) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f2622,f783])).
fof(f19618,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X59)))) | sP10(X60,happ(s(fun(bool,bool),X59),s(bool,'F_0')),X61) | s(bool,X60) = s(bool,X62) | s(bool,'T_0') = s(bool,X62) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X59)))) )),
  inference(superposition,[],[f2618,f1768])).
fof(f19617,plain,(
  ( ! [X57,X54,X58,X56,X55] : (~p(s(bool,'|dollar|exists_1'(s(fun(X54,bool),X55)))) | sP10(X56,happ(s(fun(X54,bool),X55),s(X54,'const.min.@_1'(s(fun(X54,bool),X55)))),X57) | s(bool,X56) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f2618,f303])).
fof(f19656,plain,(
  ( ! [X52,X50,X53,X51,X49] : (sP10(X51,happ(s(fun(X49,bool),X50),s(X49,'const.min.@_1'(s(fun(X49,bool),X50)))),X52) | s(bool,X51) = s(bool,X53) | s(bool,'T_0') = s(bool,X53) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X49,bool),X50)))) )),
  inference(subsumption_resolution,[],[f19616,f263])).
fof(f19616,plain,(
  ( ! [X52,X50,X53,X51,X49] : (~p(s(bool,'T_0')) | sP10(X51,happ(s(fun(X49,bool),X50),s(X49,'const.min.@_1'(s(fun(X49,bool),X50)))),X52) | s(bool,X51) = s(bool,X53) | s(bool,'T_0') = s(bool,X53) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X49,bool),X50)))) )),
  inference(superposition,[],[f2618,f1769])).
fof(f19651,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | sP10(X0,'T_0',X2) | s(bool,'T_0') = s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f19650])).
fof(f19650,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | sP10(X0,'T_0',X2) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) )),
  inference(condensation,[],[f19649])).
fof(f19649,plain,(
  ( ! [X35,X33,X36,X34,X32] : (sP10(X33,'T_0',X34) | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35) | s(bool,X32) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(subsumption_resolution,[],[f19612,f846])).
fof(f19612,plain,(
  ( ! [X35,X33,X36,X34,X32] : (~p(s(bool,X32)) | sP10(X33,'T_0',X34) | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35) | s(bool,X32) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f2618,f783])).
fof(f19646,plain,(
  ( ! [X6,X4,X8,X7,X5] : (sP10(X5,X4,X6) | s(bool,X5) = s(bool,X7) | s(bool,'T_0') = s(bool,X7) | s(bool,X4) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f19606,f263])).
fof(f19606,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~p(s(bool,'T_0')) | sP10(X5,X4,X6) | s(bool,X5) = s(bool,X7) | s(bool,'T_0') = s(bool,X7) | s(bool,X4) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f2618,f783])).
fof(f18467,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X59)))) | sP10(X60,X61,happ(s(fun(bool,bool),X59),s(bool,'F_0'))) | s(bool,X60) = s(bool,X62) | s(bool,'T_0') = s(bool,X62) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X59)))) )),
  inference(superposition,[],[f2617,f1768])).
fof(f18466,plain,(
  ( ! [X57,X54,X58,X56,X55] : (~p(s(bool,'|dollar|exists_1'(s(fun(X54,bool),X55)))) | sP10(X56,X57,happ(s(fun(X54,bool),X55),s(X54,'const.min.@_1'(s(fun(X54,bool),X55))))) | s(bool,X56) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f2617,f303])).
fof(f18505,plain,(
  ( ! [X52,X50,X53,X51,X49] : (sP10(X51,X52,happ(s(fun(X49,bool),X50),s(X49,'const.min.@_1'(s(fun(X49,bool),X50))))) | s(bool,X51) = s(bool,X53) | s(bool,'T_0') = s(bool,X53) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X49,bool),X50)))) )),
  inference(subsumption_resolution,[],[f18465,f263])).
fof(f18465,plain,(
  ( ! [X52,X50,X53,X51,X49] : (~p(s(bool,'T_0')) | sP10(X51,X52,happ(s(fun(X49,bool),X50),s(X49,'const.min.@_1'(s(fun(X49,bool),X50))))) | s(bool,X51) = s(bool,X53) | s(bool,'T_0') = s(bool,X53) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X49,bool),X50)))) )),
  inference(superposition,[],[f2617,f1769])).
fof(f18500,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | sP10(X0,X2,'T_0') | s(bool,'T_0') = s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f18499])).
fof(f18499,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | sP10(X0,X2,'T_0') | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) )),
  inference(condensation,[],[f18498])).
fof(f18498,plain,(
  ( ! [X35,X33,X36,X34,X32] : (sP10(X33,X34,'T_0') | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35) | s(bool,X32) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(subsumption_resolution,[],[f18461,f846])).
fof(f18461,plain,(
  ( ! [X35,X33,X36,X34,X32] : (~p(s(bool,X32)) | sP10(X33,X34,'T_0') | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35) | s(bool,X32) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f2617,f783])).
fof(f18495,plain,(
  ( ! [X6,X4,X8,X7,X5] : (sP10(X5,X6,X4) | s(bool,X5) = s(bool,X7) | s(bool,'T_0') = s(bool,X7) | s(bool,X4) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f18455,f263])).
fof(f18455,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~p(s(bool,'T_0')) | sP10(X5,X6,X4) | s(bool,X5) = s(bool,X7) | s(bool,'T_0') = s(bool,X7) | s(bool,X4) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f2617,f783])).
fof(f9490,plain,(
  ( ! [X21,X19,X20] : (sP1(X20,X19) | s(bool,X20) = s(bool,X21) | s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X19)) )),
  inference(subsumption_resolution,[],[f9392,f468])).
fof(f9392,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,'F_0')) | sP1(X20,X19) | s(bool,X20) = s(bool,X21) | s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f2609,f270])).
fof(f9489,plain,(
  ( ! [X21,X19,X20] : (sP1(X19,X20) | s(bool,X20) = s(bool,X21) | s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X19)) )),
  inference(subsumption_resolution,[],[f9459,f468])).
fof(f9459,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,'F_0')) | sP1(X19,X20) | s(bool,X20) = s(bool,X21) | s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f2610,f270])).
fof(f9466,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X43)))) | sP1(happ(s(fun(bool,bool),X43),s(bool,'F_0')),X44) | s(bool,X44) = s(bool,X45) | s(bool,'T_0') = s(bool,X45) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X43)))) )),
  inference(superposition,[],[f2610,f1768])).
fof(f9465,plain,(
  ( ! [X39,X41,X42,X40] : (p(s(bool,'|dollar|exists_1'(s(fun(X39,bool),X40)))) | sP1(happ(s(fun(X39,bool),X40),s(X39,'const.min.@_1'(s(fun(X39,bool),X40)))),X41) | s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f2610,f303])).
fof(f9399,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X43)))) | sP1(X44,happ(s(fun(bool,bool),X43),s(bool,'F_0'))) | s(bool,X44) = s(bool,X45) | s(bool,'T_0') = s(bool,X45) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X43)))) )),
  inference(superposition,[],[f2609,f1768])).
fof(f9398,plain,(
  ( ! [X39,X41,X42,X40] : (p(s(bool,'|dollar|exists_1'(s(fun(X39,bool),X40)))) | sP1(X41,happ(s(fun(X39,bool),X40),s(X39,'const.min.@_1'(s(fun(X39,bool),X40))))) | s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f2609,f303])).
fof(f9325,plain,(
  ( ! [X28,X27] : (s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),X27))) | s(bool,X28) = s(bool,happ(s(fun(bool,bool),X27),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X28) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X27)))) )),
  inference(superposition,[],[f2605,f1768])).
fof(f9324,plain,(
  ( ! [X26,X24,X25] : (s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(X24,bool),X25))) | s(bool,X26) = s(bool,happ(s(fun(X24,bool),X25),s(X24,'const.min.@_1'(s(fun(X24,bool),X25))))) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f2605,f303])).
fof(f9309,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2605,f270])).
fof(f7288,plain,(
  ( ! [X6,X4,X2,X5,X3] : (sP6(X4,X5,happ(s(fun(X2,bool),X3),s(X2,'const.min.@_1'(s(fun(X2,bool),X3))))) | ~sP5(X4,happ(s(fun(X2,bool),X3),s(X2,'const.min.@_1'(s(fun(X2,bool),X3)))),X5) | ~sP4(X5,happ(s(fun(X2,bool),X3),s(X2,'const.min.@_1'(s(fun(X2,bool),X3)))),X4) | ~sP3(happ(s(fun(X2,bool),X3),s(X2,'const.min.@_1'(s(fun(X2,bool),X3)))),X5,X4) | s(bool,X6) = s(bool,'|dollar|exists_1'(s(fun(X2,bool),X3))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f7282,f263])).
fof(f7282,plain,(
  ( ! [X6,X4,X2,X5,X3] : (~p(s(bool,'T_0')) | sP6(X4,X5,happ(s(fun(X2,bool),X3),s(X2,'const.min.@_1'(s(fun(X2,bool),X3))))) | ~sP5(X4,happ(s(fun(X2,bool),X3),s(X2,'const.min.@_1'(s(fun(X2,bool),X3)))),X5) | ~sP4(X5,happ(s(fun(X2,bool),X3),s(X2,'const.min.@_1'(s(fun(X2,bool),X3)))),X4) | ~sP3(happ(s(fun(X2,bool),X3),s(X2,'const.min.@_1'(s(fun(X2,bool),X3)))),X5,X4) | s(bool,X6) = s(bool,'|dollar|exists_1'(s(fun(X2,bool),X3))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1783,f783])).
fof(f7265,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X16)))) | ~sP9(happ(s(fun(bool,bool),X16),s(bool,'const.min.@_1'(s(fun(bool,bool),X16)))),X17,X18) | sP10(happ(s(fun(bool,bool),X16),s(bool,'F_0')),X17,X18) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X16)))) )),
  inference(superposition,[],[f1797,f270])).
fof(f7263,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X8)))) | ~sP9(happ(s(fun(bool,bool),X8),s(bool,'const.min.@_1'(s(fun(bool,bool),X8)))),X10,X11) | sP10(happ(s(fun(bool,bool),X8),s(bool,X9)),X10,X11) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X8))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f1797,f783])).
fof(f7261,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0)))) | ~sP9(happ(s(fun(bool,bool),X0),s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))),X1,X2) | sP10(happ(s(fun(bool,bool),X0),s(bool,'T_0')),X1,X2) | s(bool,X3) = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1797,f783])).
fof(f7244,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X16)))) | ~sP11(happ(s(fun(bool,bool),X16),s(bool,'const.min.@_1'(s(fun(bool,bool),X16)))),X17,X18) | sP12(happ(s(fun(bool,bool),X16),s(bool,'F_0')),X17,X18) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X16)))) )),
  inference(superposition,[],[f1796,f270])).
fof(f7242,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X8)))) | ~sP11(happ(s(fun(bool,bool),X8),s(bool,'const.min.@_1'(s(fun(bool,bool),X8)))),X10,X11) | sP12(happ(s(fun(bool,bool),X8),s(bool,X9)),X10,X11) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X8))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f1796,f783])).
fof(f7240,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0)))) | ~sP11(happ(s(fun(bool,bool),X0),s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))),X1,X2) | sP12(happ(s(fun(bool,bool),X0),s(bool,'T_0')),X1,X2) | s(bool,X3) = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1796,f783])).
fof(f7223,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK23(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,f(s(X0,sK23(X1,X2,bool,X0))))) )),
  inference(condensation,[],[f7222])).
fof(f7222,plain,(
  ( ! [X10,X8,X11,X9] : (s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK23(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,f(s(X8,sK23(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(subsumption_resolution,[],[f7215,f270])).
fof(f7215,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') != s(bool,X11) | s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK23(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,f(s(X8,sK23(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f1976,f783])).
fof(f7199,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,X4)) | p(s(bool,happ(s(fun(bool,bool),X3),s(bool,X5)))) | s(bool,'T_0') = s(bool,sK21(X3,X4,bool)) | s(bool,X6) = s(bool,happ(s(fun(bool,bool),X3),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f7194,f263])).
fof(f7194,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | p(s(bool,X4)) | p(s(bool,happ(s(fun(bool,bool),X3),s(bool,X5)))) | s(bool,'T_0') = s(bool,sK21(X3,X4,bool)) | s(bool,X6) = s(bool,happ(s(fun(bool,bool),X3),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1959,f783])).
fof(f7193,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0)))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | s(bool,'T_0') = s(bool,sK21(X0,X1,bool)) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f1959,f1768])).
fof(f7179,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,X4)) | p(s(bool,happ(s(fun(bool,bool),X3),s(bool,X5)))) | s(bool,'T_0') = s(bool,sK20(X3,X4,bool)) | s(bool,X6) = s(bool,happ(s(fun(bool,bool),X3),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f7174,f263])).
fof(f7174,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | p(s(bool,X4)) | p(s(bool,happ(s(fun(bool,bool),X3),s(bool,X5)))) | s(bool,'T_0') = s(bool,sK20(X3,X4,bool)) | s(bool,X6) = s(bool,happ(s(fun(bool,bool),X3),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1947,f783])).
fof(f7173,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0)))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | s(bool,'T_0') = s(bool,sK20(X0,X1,bool)) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f1947,f1768])).
fof(f7153,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,X5)) | p(s(bool,happ(s(fun(bool,bool),X3),s(bool,X4)))) | s(bool,'T_0') = s(bool,sK19(X3,X5,bool)) | s(bool,X6) = s(bool,happ(s(fun(bool,bool),X3),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f7148,f263])).
fof(f7148,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),X3),s(bool,X4)))) | ~p(s(bool,X5)) | s(bool,'T_0') = s(bool,sK19(X3,X5,bool)) | s(bool,X6) = s(bool,happ(s(fun(bool,bool),X3),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1923,f783])).
fof(f7147,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0)))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | ~p(s(bool,X2)) | s(bool,'T_0') = s(bool,sK19(X0,X2,bool)) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f1923,f1768])).
fof(f7130,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,X5)) | p(s(bool,happ(s(fun(bool,bool),X3),s(bool,X4)))) | s(bool,'T_0') = s(bool,sK18(X3,X5,bool)) | s(bool,X6) = s(bool,happ(s(fun(bool,bool),X3),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f7125,f263])).
fof(f7125,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),X3),s(bool,X4)))) | ~p(s(bool,X5)) | s(bool,'T_0') = s(bool,sK18(X3,X5,bool)) | s(bool,X6) = s(bool,happ(s(fun(bool,bool),X3),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1898,f783])).
fof(f7124,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0)))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | ~p(s(bool,X2)) | s(bool,'T_0') = s(bool,sK18(X0,X2,bool)) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f1898,f1768])).
fof(f7071,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X8)))) | ~p(s(bool,X11)) | p(s(bool,X10)) | ~sP12(happ(s(fun(bool,bool),X8),s(bool,X9)),X10,X11) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X8))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f1795,f783])).
fof(f7069,plain,(
  ( ! [X2,X0,X3,X1] : (~sP12(happ(s(fun(bool,bool),X0),s(bool,'T_0')),X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0)))) | s(bool,X3) = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1795,f783])).
fof(f7045,plain,(
  ( ! [X6,X4,X2,X5,X3] : (~sP11(X5,X4,happ(s(fun(X2,bool),X3),s(X2,'const.min.@_1'(s(fun(X2,bool),X3))))) | ~p(s(bool,X5)) | p(s(bool,X4)) | s(bool,X6) = s(bool,'|dollar|exists_1'(s(fun(X2,bool),X3))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f7039,f263])).
fof(f7039,plain,(
  ( ! [X6,X4,X2,X5,X3] : (~p(s(bool,'T_0')) | p(s(bool,X4)) | ~p(s(bool,X5)) | ~sP11(X5,X4,happ(s(fun(X2,bool),X3),s(X2,'const.min.@_1'(s(fun(X2,bool),X3))))) | s(bool,X6) = s(bool,'|dollar|exists_1'(s(fun(X2,bool),X3))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1792,f783])).
fof(f7009,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X8)))) | p(s(bool,X11)) | p(s(bool,X10)) | ~sP10(happ(s(fun(bool,bool),X8),s(bool,X9)),X10,X11) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X8))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f1791,f783])).
fof(f7007,plain,(
  ( ! [X2,X0,X3,X1] : (~sP10(happ(s(fun(bool,bool),X0),s(bool,'T_0')),X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0)))) | s(bool,X3) = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1791,f783])).
fof(f6996,plain,(
  ( ! [X6,X4,X2,X5,X3] : (~sP9(happ(s(fun(X2,bool),X3),s(X2,'const.min.@_1'(s(fun(X2,bool),X3)))),X4,X5) | p(s(bool,X5)) | p(s(bool,X4)) | s(bool,X6) = s(bool,'|dollar|exists_1'(s(fun(X2,bool),X3))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f6990,f263])).
fof(f6990,plain,(
  ( ! [X6,X4,X2,X5,X3] : (~p(s(bool,'T_0')) | p(s(bool,X4)) | p(s(bool,X5)) | ~sP9(happ(s(fun(X2,bool),X3),s(X2,'const.min.@_1'(s(fun(X2,bool),X3)))),X4,X5) | s(bool,X6) = s(bool,'|dollar|exists_1'(s(fun(X2,bool),X3))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1788,f783])).
fof(f6967,plain,(
  ( ! [X6,X4,X2,X5,X3] : (~sP8(X5,X4,happ(s(fun(X2,bool),X3),s(X2,'const.min.@_1'(s(fun(X2,bool),X3))))) | p(s(bool,X5)) | ~p(s(bool,X4)) | s(bool,X6) = s(bool,'|dollar|exists_1'(s(fun(X2,bool),X3))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f6961,f263])).
fof(f6961,plain,(
  ( ! [X6,X4,X2,X5,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,X4)) | p(s(bool,X5)) | ~sP8(X5,X4,happ(s(fun(X2,bool),X3),s(X2,'const.min.@_1'(s(fun(X2,bool),X3))))) | s(bool,X6) = s(bool,'|dollar|exists_1'(s(fun(X2,bool),X3))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1787,f783])).
fof(f6948,plain,(
  ( ! [X6,X4,X2,X5,X3] : (~sP7(X4,happ(s(fun(X2,bool),X3),s(X2,'const.min.@_1'(s(fun(X2,bool),X3)))),X5) | p(s(bool,X5)) | ~p(s(bool,X4)) | s(bool,X6) = s(bool,'|dollar|exists_1'(s(fun(X2,bool),X3))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f6942,f263])).
fof(f6942,plain,(
  ( ! [X6,X4,X2,X5,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,X4)) | p(s(bool,X5)) | ~sP7(X4,happ(s(fun(X2,bool),X3),s(X2,'const.min.@_1'(s(fun(X2,bool),X3)))),X5) | s(bool,X6) = s(bool,'|dollar|exists_1'(s(fun(X2,bool),X3))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1786,f783])).
fof(f6931,plain,(
  ( ! [X6,X4,X2,X5,X3] : (~sP4(X4,happ(s(fun(X2,bool),X3),s(X2,'const.min.@_1'(s(fun(X2,bool),X3)))),X5) | p(s(bool,X5)) | ~p(s(bool,X4)) | s(bool,X6) = s(bool,'|dollar|exists_1'(s(fun(X2,bool),X3))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f6925,f263])).
fof(f6925,plain,(
  ( ! [X6,X4,X2,X5,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,X4)) | p(s(bool,X5)) | ~sP4(X4,happ(s(fun(X2,bool),X3),s(X2,'const.min.@_1'(s(fun(X2,bool),X3)))),X5) | s(bool,X6) = s(bool,'|dollar|exists_1'(s(fun(X2,bool),X3))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1785,f783])).
fof(f6917,plain,(
  ( ! [X6,X4,X2,X5,X3] : (~sP5(X4,happ(s(fun(X2,bool),X3),s(X2,'const.min.@_1'(s(fun(X2,bool),X3)))),X5) | p(s(bool,X5)) | ~p(s(bool,X4)) | s(bool,X6) = s(bool,'|dollar|exists_1'(s(fun(X2,bool),X3))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f6911,f263])).
fof(f6911,plain,(
  ( ! [X6,X4,X2,X5,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,X4)) | p(s(bool,X5)) | ~sP5(X4,happ(s(fun(X2,bool),X3),s(X2,'const.min.@_1'(s(fun(X2,bool),X3)))),X5) | s(bool,X6) = s(bool,'|dollar|exists_1'(s(fun(X2,bool),X3))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1784,f783])).
fof(f5007,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X12))),s('type.extreal.extreal',X13))) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X12),s('type.extreal.extreal',sK22(X13,X12))))) )),
  inference(subsumption_resolution,[],[f5003,f468])).
fof(f5003,plain,(
  ( ! [X12,X13] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X12))),s('type.extreal.extreal',X13))) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X12),s('type.extreal.extreal',sK22(X13,X12))))) )),
  inference(superposition,[],[f1866,f270])).
fof(f4990,plain,(
  ( ! [X99,X97,X98,X96] : (sP10(happ(s(fun(X96,bool),X97),s(X96,'const.min.@_1'(s(fun(X96,bool),X97)))),X98,X99) | ~sP9(happ(s(fun(X96,bool),X97),s(X96,'const.min.@_1'(s(fun(X96,bool),X97)))),X98,X99) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X96,bool),X97)))) )),
  inference(subsumption_resolution,[],[f4947,f263])).
fof(f4947,plain,(
  ( ! [X99,X97,X98,X96] : (~p(s(bool,'T_0')) | ~sP9(happ(s(fun(X96,bool),X97),s(X96,'const.min.@_1'(s(fun(X96,bool),X97)))),X98,X99) | sP10(happ(s(fun(X96,bool),X97),s(X96,'const.min.@_1'(s(fun(X96,bool),X97)))),X98,X99) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X96,bool),X97)))) )),
  inference(superposition,[],[f453,f1769])).
fof(f4989,plain,(
  ( ! [X94,X92,X95,X93] : (sP12(happ(s(fun(X92,bool),X93),s(X92,'const.min.@_1'(s(fun(X92,bool),X93)))),X94,X95) | ~sP11(happ(s(fun(X92,bool),X93),s(X92,'const.min.@_1'(s(fun(X92,bool),X93)))),X94,X95) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X92,bool),X93)))) )),
  inference(subsumption_resolution,[],[f4946,f263])).
fof(f4946,plain,(
  ( ! [X94,X92,X95,X93] : (~p(s(bool,'T_0')) | ~sP11(happ(s(fun(X92,bool),X93),s(X92,'const.min.@_1'(s(fun(X92,bool),X93)))),X94,X95) | sP12(happ(s(fun(X92,bool),X93),s(X92,'const.min.@_1'(s(fun(X92,bool),X93)))),X94,X95) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X92,bool),X93)))) )),
  inference(superposition,[],[f452,f1769])).
fof(f4988,plain,(
  ( ! [X90,X88,X91,X89] : (~sP12(happ(s(fun(X88,bool),X89),s(X88,'const.min.@_1'(s(fun(X88,bool),X89)))),X91,X90) | p(s(bool,X91)) | ~p(s(bool,X90)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X88,bool),X89)))) )),
  inference(subsumption_resolution,[],[f4945,f263])).
fof(f4945,plain,(
  ( ! [X90,X88,X91,X89] : (~p(s(bool,'T_0')) | ~p(s(bool,X90)) | p(s(bool,X91)) | ~sP12(happ(s(fun(X88,bool),X89),s(X88,'const.min.@_1'(s(fun(X88,bool),X89)))),X91,X90) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X88,bool),X89)))) )),
  inference(superposition,[],[f449,f1769])).
fof(f4987,plain,(
  ( ! [X87,X85,X86,X84] : (sP12(happ(s(fun(X84,bool),X85),s(X84,'const.min.@_1'(s(fun(X84,bool),X85)))),X86,X87) | ~p(s(bool,X86)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X84,bool),X85)))) )),
  inference(subsumption_resolution,[],[f4944,f263])).
fof(f4944,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,'T_0')) | ~p(s(bool,X86)) | sP12(happ(s(fun(X84,bool),X85),s(X84,'const.min.@_1'(s(fun(X84,bool),X85)))),X86,X87) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X84,bool),X85)))) )),
  inference(superposition,[],[f448,f1769])).
fof(f4986,plain,(
  ( ! [X80,X83,X81,X82] : (sP12(happ(s(fun(X80,bool),X81),s(X80,'const.min.@_1'(s(fun(X80,bool),X81)))),X83,X82) | p(s(bool,X82)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X80,bool),X81)))) )),
  inference(subsumption_resolution,[],[f4943,f263])).
fof(f4943,plain,(
  ( ! [X80,X83,X81,X82] : (~p(s(bool,'T_0')) | p(s(bool,X82)) | sP12(happ(s(fun(X80,bool),X81),s(X80,'const.min.@_1'(s(fun(X80,bool),X81)))),X83,X82) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X80,bool),X81)))) )),
  inference(superposition,[],[f447,f1769])).
fof(f4985,plain,(
  ( ! [X78,X76,X79,X77] : (~sP11(X79,X78,happ(s(fun(X76,bool),X77),s(X76,'const.min.@_1'(s(fun(X76,bool),X77))))) | ~p(s(bool,X79)) | p(s(bool,X78)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X76,bool),X77)))) )),
  inference(subsumption_resolution,[],[f4942,f263])).
fof(f4942,plain,(
  ( ! [X78,X76,X79,X77] : (~p(s(bool,'T_0')) | p(s(bool,X78)) | ~p(s(bool,X79)) | ~sP11(X79,X78,happ(s(fun(X76,bool),X77),s(X76,'const.min.@_1'(s(fun(X76,bool),X77))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X76,bool),X77)))) )),
  inference(superposition,[],[f442,f1769])).
fof(f4984,plain,(
  ( ! [X74,X72,X75,X73] : (~sP10(happ(s(fun(X72,bool),X73),s(X72,'const.min.@_1'(s(fun(X72,bool),X73)))),X75,X74) | p(s(bool,X75)) | p(s(bool,X74)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X72,bool),X73)))) )),
  inference(subsumption_resolution,[],[f4941,f263])).
fof(f4941,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,'T_0')) | p(s(bool,X74)) | p(s(bool,X75)) | ~sP10(happ(s(fun(X72,bool),X73),s(X72,'const.min.@_1'(s(fun(X72,bool),X73)))),X75,X74) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X72,bool),X73)))) )),
  inference(superposition,[],[f432,f1769])).
fof(f4983,plain,(
  ( ! [X70,X68,X71,X69] : (sP10(happ(s(fun(X68,bool),X69),s(X68,'const.min.@_1'(s(fun(X68,bool),X69)))),X70,X71) | ~p(s(bool,X70)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X68,bool),X69)))) )),
  inference(subsumption_resolution,[],[f4940,f263])).
fof(f4940,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,'T_0')) | ~p(s(bool,X70)) | sP10(happ(s(fun(X68,bool),X69),s(X68,'const.min.@_1'(s(fun(X68,bool),X69)))),X70,X71) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X68,bool),X69)))) )),
  inference(superposition,[],[f431,f1769])).
fof(f4982,plain,(
  ( ! [X66,X64,X67,X65] : (sP10(happ(s(fun(X64,bool),X65),s(X64,'const.min.@_1'(s(fun(X64,bool),X65)))),X67,X66) | ~p(s(bool,X66)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X64,bool),X65)))) )),
  inference(subsumption_resolution,[],[f4939,f263])).
fof(f4939,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,'T_0')) | ~p(s(bool,X66)) | sP10(happ(s(fun(X64,bool),X65),s(X64,'const.min.@_1'(s(fun(X64,bool),X65)))),X67,X66) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X64,bool),X65)))) )),
  inference(superposition,[],[f430,f1769])).
fof(f4981,plain,(
  ( ! [X61,X62,X60,X63] : (~sP9(happ(s(fun(X60,bool),X61),s(X60,'const.min.@_1'(s(fun(X60,bool),X61)))),X62,X63) | p(s(bool,X63)) | p(s(bool,X62)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X60,bool),X61)))) )),
  inference(subsumption_resolution,[],[f4938,f263])).
fof(f4938,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,'T_0')) | p(s(bool,X62)) | p(s(bool,X63)) | ~sP9(happ(s(fun(X60,bool),X61),s(X60,'const.min.@_1'(s(fun(X60,bool),X61)))),X62,X63) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X60,bool),X61)))) )),
  inference(superposition,[],[f425,f1769])).
fof(f4980,plain,(
  ( ! [X59,X57,X58,X56] : (~sP8(X59,X58,happ(s(fun(X56,bool),X57),s(X56,'const.min.@_1'(s(fun(X56,bool),X57))))) | p(s(bool,X59)) | ~p(s(bool,X58)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X56,bool),X57)))) )),
  inference(subsumption_resolution,[],[f4937,f263])).
fof(f4937,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,'T_0')) | ~p(s(bool,X58)) | p(s(bool,X59)) | ~sP8(X59,X58,happ(s(fun(X56,bool),X57),s(X56,'const.min.@_1'(s(fun(X56,bool),X57))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X56,bool),X57)))) )),
  inference(superposition,[],[f417,f1769])).
fof(f4979,plain,(
  ( ! [X54,X52,X55,X53] : (~sP7(X54,happ(s(fun(X52,bool),X53),s(X52,'const.min.@_1'(s(fun(X52,bool),X53)))),X55) | p(s(bool,X55)) | ~p(s(bool,X54)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X52,bool),X53)))) )),
  inference(subsumption_resolution,[],[f4936,f263])).
fof(f4936,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,'T_0')) | ~p(s(bool,X54)) | p(s(bool,X55)) | ~sP7(X54,happ(s(fun(X52,bool),X53),s(X52,'const.min.@_1'(s(fun(X52,bool),X53)))),X55) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X52,bool),X53)))) )),
  inference(superposition,[],[f408,f1769])).
fof(f4978,plain,(
  ( ! [X50,X48,X51,X49] : (~sP4(X50,happ(s(fun(X48,bool),X49),s(X48,'const.min.@_1'(s(fun(X48,bool),X49)))),X51) | p(s(bool,X51)) | ~p(s(bool,X50)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X48,bool),X49)))) )),
  inference(subsumption_resolution,[],[f4935,f263])).
fof(f4935,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X50)) | p(s(bool,X51)) | ~sP4(X50,happ(s(fun(X48,bool),X49),s(X48,'const.min.@_1'(s(fun(X48,bool),X49)))),X51) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X48,bool),X49)))) )),
  inference(superposition,[],[f389,f1769])).
fof(f4977,plain,(
  ( ! [X47,X45,X46,X44] : (~sP5(X46,happ(s(fun(X44,bool),X45),s(X44,'const.min.@_1'(s(fun(X44,bool),X45)))),X47) | p(s(bool,X47)) | ~p(s(bool,X46)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X44,bool),X45)))) )),
  inference(subsumption_resolution,[],[f4934,f263])).
fof(f4934,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X46)) | p(s(bool,X47)) | ~sP5(X46,happ(s(fun(X44,bool),X45),s(X44,'const.min.@_1'(s(fun(X44,bool),X45)))),X47) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X44,bool),X45)))) )),
  inference(superposition,[],[f385,f1769])).
fof(f4976,plain,(
  ( ! [X43,X41,X42,X40] : (sP6(X42,X43,happ(s(fun(X40,bool),X41),s(X40,'const.min.@_1'(s(fun(X40,bool),X41))))) | ~sP5(X42,happ(s(fun(X40,bool),X41),s(X40,'const.min.@_1'(s(fun(X40,bool),X41)))),X43) | ~sP4(X43,happ(s(fun(X40,bool),X41),s(X40,'const.min.@_1'(s(fun(X40,bool),X41)))),X42) | ~sP3(happ(s(fun(X40,bool),X41),s(X40,'const.min.@_1'(s(fun(X40,bool),X41)))),X43,X42) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X40,bool),X41)))) )),
  inference(subsumption_resolution,[],[f4933,f263])).
fof(f4933,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,'T_0')) | sP6(X42,X43,happ(s(fun(X40,bool),X41),s(X40,'const.min.@_1'(s(fun(X40,bool),X41))))) | ~sP5(X42,happ(s(fun(X40,bool),X41),s(X40,'const.min.@_1'(s(fun(X40,bool),X41)))),X43) | ~sP4(X43,happ(s(fun(X40,bool),X41),s(X40,'const.min.@_1'(s(fun(X40,bool),X41)))),X42) | ~sP3(happ(s(fun(X40,bool),X41),s(X40,'const.min.@_1'(s(fun(X40,bool),X41)))),X43,X42) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X40,bool),X41)))) )),
  inference(superposition,[],[f382,f1769])).
fof(f4975,plain,(
  ( ! [X39,X37,X38] : (sP1(X39,happ(s(fun(X37,bool),X38),s(X37,'const.min.@_1'(s(fun(X37,bool),X38))))) | p(s(bool,X39)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X37,bool),X38)))) )),
  inference(subsumption_resolution,[],[f4932,f263])).
fof(f4932,plain,(
  ( ! [X39,X37,X38] : (~p(s(bool,'T_0')) | p(s(bool,X39)) | sP1(X39,happ(s(fun(X37,bool),X38),s(X37,'const.min.@_1'(s(fun(X37,bool),X38))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X37,bool),X38)))) )),
  inference(superposition,[],[f353,f1769])).
fof(f4974,plain,(
  ( ! [X35,X36,X34] : (sP1(happ(s(fun(X34,bool),X35),s(X34,'const.min.@_1'(s(fun(X34,bool),X35)))),X36) | p(s(bool,X36)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X34,bool),X35)))) )),
  inference(subsumption_resolution,[],[f4931,f263])).
fof(f4931,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,'T_0')) | p(s(bool,X36)) | sP1(happ(s(fun(X34,bool),X35),s(X34,'const.min.@_1'(s(fun(X34,bool),X35)))),X36) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X34,bool),X35)))) )),
  inference(superposition,[],[f350,f1769])).
fof(f4973,plain,(
  ( ! [X33,X31,X32] : (sP0(happ(s(fun(X31,bool),X32),s(X31,'const.min.@_1'(s(fun(X31,bool),X32)))),X33) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X31,bool),X32)))) )),
  inference(subsumption_resolution,[],[f4930,f263])).
fof(f4930,plain,(
  ( ! [X33,X31,X32] : (~p(s(bool,'T_0')) | sP0(happ(s(fun(X31,bool),X32),s(X31,'const.min.@_1'(s(fun(X31,bool),X32)))),X33) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X31,bool),X32)))) )),
  inference(superposition,[],[f335,f1769])).
fof(f4972,plain,(
  ( ! [X30,X28,X29] : (sP0(X30,happ(s(fun(X28,bool),X29),s(X28,'const.min.@_1'(s(fun(X28,bool),X29))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X28,bool),X29)))) )),
  inference(subsumption_resolution,[],[f4929,f263])).
fof(f4929,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | sP0(X30,happ(s(fun(X28,bool),X29),s(X28,'const.min.@_1'(s(fun(X28,bool),X29))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X28,bool),X29)))) )),
  inference(superposition,[],[f334,f1769])).
fof(f4971,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,X27)) | s(bool,X27) = s(bool,'|dollar|exists_1'(s(fun(X25,bool),X26))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X25,bool),X26)))) )),
  inference(forward_demodulation,[],[f4970,f303])).
fof(f4970,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,X27)) | s(bool,X27) = s(bool,happ(s(fun(X25,bool),X26),s(X25,'const.min.@_1'(s(fun(X25,bool),X26))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X25,bool),X26)))) )),
  inference(subsumption_resolution,[],[f4928,f263])).
fof(f4928,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | ~p(s(bool,X27)) | s(bool,X27) = s(bool,happ(s(fun(X25,bool),X26),s(X25,'const.min.@_1'(s(fun(X25,bool),X26))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X25,bool),X26)))) )),
  inference(superposition,[],[f321,f1769])).
fof(f4969,plain,(
  ( ! [X15,X16] : (s(bool,'F_0') != s(bool,happ(s(fun(X15,bool),X16),s(X15,'const.min.@_1'(s(fun(X15,bool),X16))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X15,bool),X16)))) )),
  inference(subsumption_resolution,[],[f4924,f263])).
fof(f4924,plain,(
  ( ! [X15,X16] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,happ(s(fun(X15,bool),X16),s(X15,'const.min.@_1'(s(fun(X15,bool),X16))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X15,bool),X16)))) )),
  inference(superposition,[],[f280,f1769])).
fof(f4967,plain,(
  ( ! [X6,X5] : (p(s(bool,happ(s(fun(X5,bool),X6),s(X5,sK16(X6,X5))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X5,bool),X6)))) )),
  inference(subsumption_resolution,[],[f4919,f263])).
fof(f4919,plain,(
  ( ! [X6,X5] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X5,bool),X6),s(X5,sK16(X6,X5))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X5,bool),X6)))) )),
  inference(superposition,[],[f304,f1769])).
fof(f4966,plain,(
  ( ! [X4,X3] : (p(s(bool,happ(s(fun(X3,bool),X4),s(X3,sK17(X4,X3))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X3,bool),X4)))) )),
  inference(subsumption_resolution,[],[f4918,f263])).
fof(f4918,plain,(
  ( ! [X4,X3] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X3,bool),X4),s(X3,sK17(X4,X3))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X3,bool),X4)))) )),
  inference(superposition,[],[f305,f1769])).
fof(f4965,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.min.@_1'(s(fun('type.extreal.extreal',bool),X0))),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X0)))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),X0)))) )),
  inference(subsumption_resolution,[],[f4916,f263])).
fof(f4916,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.min.@_1'(s(fun('type.extreal.extreal',bool),X0))),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X0)))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),X0)))) )),
  inference(superposition,[],[f317,f1769])).
fof(f4909,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X4),s(bool,X5))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X4))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X4))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1769,f783])).
fof(f4894,plain,(
  ( ! [X66,X64,X65] : (sP10(happ(s(fun(bool,bool),X64),s(bool,'F_0')),X65,X66) | ~sP9(happ(s(fun(bool,bool),X64),s(bool,'F_0')),X65,X66) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X64)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X64)))) )),
  inference(superposition,[],[f453,f1768])).
fof(f4893,plain,(
  ( ! [X61,X62,X63] : (sP12(happ(s(fun(bool,bool),X61),s(bool,'F_0')),X62,X63) | ~sP11(happ(s(fun(bool,bool),X61),s(bool,'F_0')),X62,X63) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X61)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X61)))) )),
  inference(superposition,[],[f452,f1768])).
fof(f4892,plain,(
  ( ! [X59,X60,X58] : (~sP12(happ(s(fun(bool,bool),X58),s(bool,'F_0')),X60,X59) | ~p(s(bool,X59)) | p(s(bool,X60)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X58)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X58)))) )),
  inference(superposition,[],[f449,f1768])).
fof(f4889,plain,(
  ( ! [X50,X51,X49] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X49)))) | p(s(bool,X50)) | ~p(s(bool,X51)) | ~sP11(X51,X50,happ(s(fun(bool,bool),X49),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X49)))) )),
  inference(superposition,[],[f442,f1768])).
fof(f4888,plain,(
  ( ! [X47,X48,X46] : (~sP10(happ(s(fun(bool,bool),X46),s(bool,'F_0')),X48,X47) | p(s(bool,X47)) | p(s(bool,X48)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X46)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X46)))) )),
  inference(superposition,[],[f432,f1768])).
fof(f4885,plain,(
  ( ! [X39,X37,X38] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X37)))) | p(s(bool,X38)) | p(s(bool,X39)) | ~sP9(happ(s(fun(bool,bool),X37),s(bool,'F_0')),X38,X39) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X37)))) )),
  inference(superposition,[],[f425,f1768])).
fof(f4884,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X34)))) | ~p(s(bool,X35)) | p(s(bool,X36)) | ~sP8(X36,X35,happ(s(fun(bool,bool),X34),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X34)))) )),
  inference(superposition,[],[f417,f1768])).
fof(f4883,plain,(
  ( ! [X33,X31,X32] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X31)))) | ~p(s(bool,X32)) | p(s(bool,X33)) | ~sP7(X32,happ(s(fun(bool,bool),X31),s(bool,'F_0')),X33) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X31)))) )),
  inference(superposition,[],[f408,f1768])).
fof(f4882,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X28)))) | ~p(s(bool,X29)) | p(s(bool,X30)) | ~sP4(X29,happ(s(fun(bool,bool),X28),s(bool,'F_0')),X30) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X28)))) )),
  inference(superposition,[],[f389,f1768])).
fof(f4881,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X25)))) | ~p(s(bool,X26)) | p(s(bool,X27)) | ~sP5(X26,happ(s(fun(bool,bool),X25),s(bool,'F_0')),X27) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X25)))) )),
  inference(superposition,[],[f385,f1768])).
fof(f4880,plain,(
  ( ! [X24,X23,X22] : (sP6(X23,X24,happ(s(fun(bool,bool),X22),s(bool,'F_0'))) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X22)))) | ~sP5(X23,happ(s(fun(bool,bool),X22),s(bool,'F_0')),X24) | ~sP4(X24,happ(s(fun(bool,bool),X22),s(bool,'F_0')),X23) | ~sP3(happ(s(fun(bool,bool),X22),s(bool,'F_0')),X24,X23) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X22)))) )),
  inference(superposition,[],[f382,f1768])).
fof(f4875,plain,(
  ( ! [X12,X13] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X12)))) | ~p(s(bool,X13)) | s(bool,X13) = s(bool,happ(s(fun(bool,bool),X12),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X12)))) )),
  inference(superposition,[],[f321,f1768])).
fof(f4873,plain,(
  ( ! [X8,X9] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X8)))) | p(s(bool,X9)) | s(bool,X9) = s(bool,happ(s(fun(bool,bool),X8),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X8)))) )),
  inference(superposition,[],[f318,f1768])).
fof(f4872,plain,(
  ( ! [X7] : (p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X7)))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X7),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X7)))) )),
  inference(superposition,[],[f281,f1768])).
fof(f4871,plain,(
  ( ! [X6] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X6)))) | s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X6),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X6)))) )),
  inference(superposition,[],[f280,f1768])).
fof(f4870,plain,(
  ( ! [X5] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X5)))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X5),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X5)))) )),
  inference(superposition,[],[f277,f1768])).
fof(f4869,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(bool,bool),X4))) | p(s(bool,happ(s(fun(bool,bool),X4),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X4)))) )),
  inference(superposition,[],[f276,f1768])).
fof(f4864,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X9),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X9))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X9)))) )),
  inference(superposition,[],[f1768,f270])).
fof(f4842,plain,(
  ( ! [X17,X18,X16] : (sP12(happ(s(fun(bool,bool),X16),s(bool,'F_0')),X17,X18) | ~p(s(bool,X17)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X16)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X16)))) )),
  inference(superposition,[],[f1794,f270])).
fof(f4840,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X8)))) | ~p(s(bool,X10)) | sP12(happ(s(fun(bool,bool),X8),s(bool,X9)),X10,X11) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X8))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f1794,f783])).
fof(f4838,plain,(
  ( ! [X2,X0,X3,X1] : (sP12(happ(s(fun(bool,bool),X0),s(bool,'T_0')),X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0)))) | s(bool,X3) = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1794,f783])).
fof(f4825,plain,(
  ( ! [X17,X18,X16] : (sP12(happ(s(fun(bool,bool),X16),s(bool,'F_0')),X17,X18) | p(s(bool,X18)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X16)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X16)))) )),
  inference(superposition,[],[f1793,f270])).
fof(f4823,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X8)))) | p(s(bool,X11)) | sP12(happ(s(fun(bool,bool),X8),s(bool,X9)),X10,X11) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X8))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f1793,f783])).
fof(f4821,plain,(
  ( ! [X2,X0,X3,X1] : (sP12(happ(s(fun(bool,bool),X0),s(bool,'T_0')),X1,X2) | p(s(bool,X2)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0)))) | s(bool,X3) = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1793,f783])).
fof(f4811,plain,(
  ( ! [X17,X18,X16] : (sP10(happ(s(fun(bool,bool),X16),s(bool,'F_0')),X17,X18) | ~p(s(bool,X17)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X16)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X16)))) )),
  inference(superposition,[],[f1790,f270])).
fof(f4809,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X8)))) | ~p(s(bool,X10)) | sP10(happ(s(fun(bool,bool),X8),s(bool,X9)),X10,X11) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X8))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f1790,f783])).
fof(f4807,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(happ(s(fun(bool,bool),X0),s(bool,'T_0')),X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0)))) | s(bool,X3) = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1790,f783])).
fof(f4756,plain,(
  ( ! [X17,X18,X16] : (sP10(happ(s(fun(bool,bool),X16),s(bool,'F_0')),X17,X18) | ~p(s(bool,X18)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X16)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X16)))) )),
  inference(superposition,[],[f1789,f270])).
fof(f4754,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X8)))) | ~p(s(bool,X11)) | sP10(happ(s(fun(bool,bool),X8),s(bool,X9)),X10,X11) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X8))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f1789,f783])).
fof(f4752,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(happ(s(fun(bool,bool),X0),s(bool,'T_0')),X1,X2) | ~p(s(bool,X2)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0)))) | s(bool,X3) = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1789,f783])).
fof(f4735,plain,(
  ( ! [X12,X13] : (sP1(X13,happ(s(fun(bool,bool),X12),s(bool,'F_0'))) | p(s(bool,X13)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X12)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X12)))) )),
  inference(superposition,[],[f1782,f270])).
fof(f4733,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X6)))) | p(s(bool,X8)) | sP1(X8,happ(s(fun(bool,bool),X6),s(bool,X7))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f1782,f783])).
fof(f4731,plain,(
  ( ! [X2,X0,X1] : (sP1(X1,happ(s(fun(bool,bool),X0),s(bool,'T_0'))) | p(s(bool,X1)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0)))) | s(bool,X2) = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1782,f783])).
fof(f4726,plain,(
  ( ! [X12,X13] : (sP1(happ(s(fun(bool,bool),X12),s(bool,'F_0')),X13) | p(s(bool,X13)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X12)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X12)))) )),
  inference(superposition,[],[f1781,f270])).
fof(f4724,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X6)))) | p(s(bool,X8)) | sP1(happ(s(fun(bool,bool),X6),s(bool,X7)),X8) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f1781,f783])).
fof(f4722,plain,(
  ( ! [X2,X0,X1] : (sP1(happ(s(fun(bool,bool),X0),s(bool,'T_0')),X1) | p(s(bool,X1)) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0)))) | s(bool,X2) = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1781,f783])).
fof(f4321,plain,(
  ( ! [X12,X13] : (sP0(happ(s(fun(bool,bool),X12),s(bool,'F_0')),X13) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X12)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X12)))) )),
  inference(superposition,[],[f1780,f270])).
fof(f4319,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X6)))) | sP0(happ(s(fun(bool,bool),X6),s(bool,X7)),X8) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f1780,f783])).
fof(f4317,plain,(
  ( ! [X2,X0,X1] : (sP0(happ(s(fun(bool,bool),X0),s(bool,'T_0')),X1) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0)))) | s(bool,X2) = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1780,f783])).
fof(f4312,plain,(
  ( ! [X12,X13] : (sP0(X13,happ(s(fun(bool,bool),X12),s(bool,'F_0'))) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X12)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X12)))) )),
  inference(superposition,[],[f1779,f270])).
fof(f4310,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X6)))) | sP0(X8,happ(s(fun(bool,bool),X6),s(bool,X7))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f1779,f783])).
fof(f4308,plain,(
  ( ! [X2,X0,X1] : (sP0(X1,happ(s(fun(bool,bool),X0),s(bool,'T_0'))) | ~p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0)))) | s(bool,X2) = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1779,f783])).
fof(f4303,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK17(X1,X0))))) | s(bool,X2) = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f4298,f263])).
fof(f4298,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK17(X1,X0))))) | s(bool,X2) = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1831,f783])).
fof(f3897,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK16(X1,X0))))) | s(bool,X2) = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f3892,f263])).
fof(f3892,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK16(X1,X0))))) | s(bool,X2) = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1817,f783])).
fof(f3886,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X13),s('type.extreal.extreal',X12))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X12),s('type.extreal.extreal',X13)))) )),
  inference(subsumption_resolution,[],[f3880,f468])).
fof(f3880,plain,(
  ( ! [X12,X13] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X13),s('type.extreal.extreal',X12))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X12),s('type.extreal.extreal',X13)))) )),
  inference(superposition,[],[f1296,f270])).
fof(f3869,plain,(
  ( ! [X8] : (sP6('T_0','|dollar|exists_1'(s(fun(bool,bool),X8)),happ(s(fun(bool,bool),X8),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X8)))) )),
  inference(superposition,[],[f1767,f270])).
fof(f3867,plain,(
  ( ! [X4,X5] : (sP6('T_0','|dollar|exists_1'(s(fun(bool,bool),X4)),happ(s(fun(bool,bool),X4),s(bool,X5))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X4))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1767,f783])).
fof(f3865,plain,(
  ( ! [X0,X1] : (sP6('T_0','|dollar|exists_1'(s(fun(bool,bool),X0)),happ(s(fun(bool,bool),X0),s(bool,'T_0'))) | s(bool,X1) = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f1767,f783])).
fof(f3862,plain,(
  ( ! [X0,X1] : (sP5('T_0',happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(unit_resulting_resolution,[],[f1767,f380])).
fof(f3853,plain,(
  ( ! [X8] : (~sP6('F_0','|dollar|exists_1'(s(fun(bool,bool),X8)),happ(s(fun(bool,bool),X8),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X8)))) )),
  inference(superposition,[],[f1766,f270])).
fof(f3851,plain,(
  ( ! [X4,X5] : (~sP6('F_0','|dollar|exists_1'(s(fun(bool,bool),X4)),happ(s(fun(bool,bool),X4),s(bool,X5))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X4))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1766,f783])).
fof(f3849,plain,(
  ( ! [X0,X1] : (~sP6('F_0','|dollar|exists_1'(s(fun(bool,bool),X0)),happ(s(fun(bool,bool),X0),s(bool,'T_0'))) | s(bool,X1) = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f1766,f783])).
fof(f3764,plain,(
  ( ! [X8] : (sP6('T_0',happ(s(fun(bool,bool),X8),s(bool,'F_0')),'|dollar|exists_1'(s(fun(bool,bool),X8))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X8)))) )),
  inference(superposition,[],[f1765,f270])).
fof(f3762,plain,(
  ( ! [X4,X5] : (sP6('T_0',happ(s(fun(bool,bool),X4),s(bool,X5)),'|dollar|exists_1'(s(fun(bool,bool),X4))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X4))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1765,f783])).
fof(f3760,plain,(
  ( ! [X0,X1] : (sP6('T_0',happ(s(fun(bool,bool),X0),s(bool,'T_0')),'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,X1) = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f1765,f783])).
fof(f3757,plain,(
  ( ! [X0,X1] : (sP5('T_0','|dollar|exists_1'(s(fun(X0,bool),X1)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f1765,f380])).
fof(f3744,plain,(
  ( ! [X8] : (~sP6('F_0',happ(s(fun(bool,bool),X8),s(bool,'F_0')),'|dollar|exists_1'(s(fun(bool,bool),X8))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X8)))) )),
  inference(superposition,[],[f1764,f270])).
fof(f3742,plain,(
  ( ! [X4,X5] : (~sP6('F_0',happ(s(fun(bool,bool),X4),s(bool,X5)),'|dollar|exists_1'(s(fun(bool,bool),X4))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X4))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1764,f783])).
fof(f3740,plain,(
  ( ! [X0,X1] : (~sP6('F_0',happ(s(fun(bool,bool),X0),s(bool,'T_0')),'|dollar|exists_1'(s(fun(bool,bool),X0))) | s(bool,X1) = s(bool,'const.min.@_1'(s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f1764,f783])).
fof(f280,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f57])).
fof(f57,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.extreal.sup_lt','thm.bool.EQ_CLAUSES|split|3')).
fof(f281,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f173])).
fof(f1898,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | ~p(s(bool,X1)) | s(bool,'T_0') = s(bool,sK18(X0,X1,bool))) )),
  inference(superposition,[],[f310,f270])).
fof(f1923,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | ~p(s(bool,X1)) | s(bool,'T_0') = s(bool,sK19(X0,X1,bool))) )),
  inference(superposition,[],[f312,f270])).
fof(f1947,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | s(bool,'T_0') = s(bool,sK20(X0,X1,bool))) )),
  inference(superposition,[],[f314,f270])).
fof(f1959,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | s(bool,'T_0') = s(bool,sK21(X0,X1,bool))) )),
  inference(superposition,[],[f316,f270])).
fof(f1975,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK23(X0,X1,X2,bool))) )),
  inference(superposition,[],[f366,f270])).
fof(f2134,plain,(
  ( ! [X182,X184,X183,X181] : (s(bool,X184) != s(bool,f(s(X181,sK23(X182,X183,bool,X181)))) | s(fun(X181,bool),X182) = s(fun(X181,bool),X183) | s(bool,'T_0') = s(bool,X184) | s(bool,'T_0') = s(bool,g(s(X181,sK23(X182,X183,bool,X181))))) )),
  inference(superposition,[],[f366,f783])).
fof(f304,plain,(
  ( ! [X4,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK16(X1,X0)))))) )),
  inference(cnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0,X1] : ((p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK16(X1,X0))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK17(X1,X0))))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16,sK17])],[f188])).
fof(f188,plain,(
  ! [X0,X1] : ((? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f187])).
fof(f187,plain,(
  ! [X0,X1] : ((? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1] : (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(ennf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1] : (~? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(flattening,[],[f73])).
fof(f73,plain,(
  ! [X0] : ! [X1] : (~? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X9] : ! [X6] : (~? [X8] : p(s(bool,happ(s(fun(X9,bool),X6),s(X9,X8)))) <=> ! [X8] : ~p(s(bool,happ(s(fun(X9,bool),X6),s(X9,X8)))))),
  file('thm.extreal.sup_lt','thm.bool.NOT_EXISTS_THM')).
fof(f305,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK17(X1,X0)))))) )),
  inference(cnf_transformation,[],[f189])).
fof(f310,plain,(
  ( ! [X2,X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK18(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK18(X2,X1,X0))))) | ! [X5] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & ((~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK19(X2,X1,X0)))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK18,sK19])],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,X1)) | ? [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | ! [X5] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & (? [X4] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,X1)) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (? [X4] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (? [X4] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X9] : ! [X6,X13] : ((p(s(bool,X6)) & ! [X8] : p(s(bool,happ(s(fun(X9,bool),X13),s(X9,X8))))) <=> ! [X8] : (p(s(bool,X6)) & p(s(bool,happ(s(fun(X9,bool),X13),s(X9,X8))))))),
  file('thm.extreal.sup_lt','thm.bool.RIGHT_AND_FORALL_THM')).
fof(f312,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK19(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f194])).
fof(f314,plain,(
  ( ! [X2,X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK20(X2,X1,X0))))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK20(X2,X1,X0)))))) | p(s(bool,X1)) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5))))) & ((~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK21(X2,X1,X0)))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20,sK21])],[f197])).
fof(f197,plain,(
  ! [X0,X1,X2] : ((? [X6] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | p(s(bool,X1)) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5))))) & ((~p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f196])).
fof(f196,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) & ((~p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((~p(s(bool,X1)) & ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f81])).
fof(f81,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X9] : ! [X6,X13] : (! [X8] : (p(s(bool,X6)) | p(s(bool,happ(s(fun(X9,bool),X13),s(X9,X8))))) <=> (p(s(bool,X6)) | ! [X8] : p(s(bool,happ(s(fun(X9,bool),X13),s(X9,X8))))))),
  file('thm.extreal.sup_lt','thm.bool.RIGHT_FORALL_OR_THM')).
fof(f316,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK21(X2,X1,X0))))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f198])).
fof(f2627,plain,(
  ( ! [X109,X107,X110,X108] : (p(s(bool,happ(s(fun(X107,bool),X108),s(X107,sK17(X108,X107))))) | s(bool,X110) = s(bool,happ(s(fun(X107,bool),X108),s(X107,X109))) | s(bool,'T_0') = s(bool,X110)) )),
  inference(subsumption_resolution,[],[f2271,f263])).
fof(f2271,plain,(
  ( ! [X109,X107,X110,X108] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X107,bool),X108),s(X107,sK17(X108,X107))))) | s(bool,X110) = s(bool,happ(s(fun(X107,bool),X108),s(X107,X109))) | s(bool,'T_0') = s(bool,X110)) )),
  inference(superposition,[],[f305,f783])).
fof(f2628,plain,(
  ( ! [X111,X114,X112,X113] : (p(s(bool,happ(s(fun(X111,bool),X112),s(X111,sK16(X112,X111))))) | s(bool,X114) = s(bool,happ(s(fun(X111,bool),X112),s(X111,X113))) | s(bool,'T_0') = s(bool,X114)) )),
  inference(subsumption_resolution,[],[f2272,f263])).
fof(f2272,plain,(
  ( ! [X111,X114,X112,X113] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X111,bool),X112),s(X111,sK16(X112,X111))))) | s(bool,X114) = s(bool,happ(s(fun(X111,bool),X112),s(X111,X113))) | s(bool,'T_0') = s(bool,X114)) )),
  inference(superposition,[],[f304,f783])).
fof(f2630,plain,(
  ( ! [X125,X123,X126,X124,X122] : (~p(s(bool,X124)) | p(s(bool,happ(s(fun(X122,bool),X123),s(X122,X125)))) | s(bool,X126) = s(bool,happ(s(fun(X122,bool),X123),s(X122,sK18(X123,X124,X122)))) | s(bool,'T_0') = s(bool,X126)) )),
  inference(subsumption_resolution,[],[f2275,f263])).
fof(f2275,plain,(
  ( ! [X125,X123,X126,X124,X122] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X122,bool),X123),s(X122,X125)))) | ~p(s(bool,X124)) | s(bool,X126) = s(bool,happ(s(fun(X122,bool),X123),s(X122,sK18(X123,X124,X122)))) | s(bool,'T_0') = s(bool,X126)) )),
  inference(superposition,[],[f310,f783])).
fof(f2631,plain,(
  ( ! [X127,X130,X128,X131,X129] : (~p(s(bool,X129)) | p(s(bool,happ(s(fun(X127,bool),X128),s(X127,X130)))) | s(bool,X131) = s(bool,happ(s(fun(X127,bool),X128),s(X127,sK19(X128,X129,X127)))) | s(bool,'T_0') = s(bool,X131)) )),
  inference(subsumption_resolution,[],[f2276,f263])).
fof(f2276,plain,(
  ( ! [X127,X130,X128,X131,X129] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X127,bool),X128),s(X127,X130)))) | ~p(s(bool,X129)) | s(bool,X131) = s(bool,happ(s(fun(X127,bool),X128),s(X127,sK19(X128,X129,X127)))) | s(bool,'T_0') = s(bool,X131)) )),
  inference(superposition,[],[f312,f783])).
fof(f2632,plain,(
  ( ! [X134,X132,X136,X135,X133] : (p(s(bool,X134)) | p(s(bool,happ(s(fun(X132,bool),X133),s(X132,X135)))) | s(bool,X136) = s(bool,happ(s(fun(X132,bool),X133),s(X132,sK20(X133,X134,X132)))) | s(bool,'T_0') = s(bool,X136)) )),
  inference(subsumption_resolution,[],[f2277,f263])).
fof(f2277,plain,(
  ( ! [X134,X132,X136,X135,X133] : (~p(s(bool,'T_0')) | p(s(bool,X134)) | p(s(bool,happ(s(fun(X132,bool),X133),s(X132,X135)))) | s(bool,X136) = s(bool,happ(s(fun(X132,bool),X133),s(X132,sK20(X133,X134,X132)))) | s(bool,'T_0') = s(bool,X136)) )),
  inference(superposition,[],[f314,f783])).
fof(f2633,plain,(
  ( ! [X140,X138,X141,X139,X137] : (p(s(bool,X139)) | p(s(bool,happ(s(fun(X137,bool),X138),s(X137,X140)))) | s(bool,X141) = s(bool,happ(s(fun(X137,bool),X138),s(X137,sK21(X138,X139,X137)))) | s(bool,'T_0') = s(bool,X141)) )),
  inference(subsumption_resolution,[],[f2278,f263])).
fof(f2278,plain,(
  ( ! [X140,X138,X141,X139,X137] : (~p(s(bool,'T_0')) | p(s(bool,X139)) | p(s(bool,happ(s(fun(X137,bool),X138),s(X137,X140)))) | s(bool,X141) = s(bool,happ(s(fun(X137,bool),X138),s(X137,sK21(X138,X139,X137)))) | s(bool,'T_0') = s(bool,X141)) )),
  inference(superposition,[],[f316,f783])).
fof(f1817,plain,(
  ( ! [X0,X1] : (~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK16(X1,X0)))))) )),
  inference(superposition,[],[f304,f303])).
fof(f1831,plain,(
  ( ! [X0,X1] : (~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK17(X1,X0)))))) )),
  inference(superposition,[],[f305,f303])).
fof(f2637,plain,(
  ( ! [X158,X159,X157,X160] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X157),s('type.extreal.extreal',X159)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X159),s('type.extreal.extreal',X158)))) | s(bool,X160) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X157))),s('type.extreal.extreal',X158))) | s(bool,'T_0') = s(bool,X160)) )),
  inference(subsumption_resolution,[],[f2285,f263])).
fof(f2285,plain,(
  ( ! [X158,X159,X157,X160] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X157),s('type.extreal.extreal',X159)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X159),s('type.extreal.extreal',X158)))) | s(bool,X160) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X157))),s('type.extreal.extreal',X158))) | s(bool,'T_0') = s(bool,X160)) )),
  inference(superposition,[],[f326,f783])).
fof(f2626,plain,(
  ( ! [X103,X101,X102] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X102),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X101)))))) | s(bool,X103) = s(bool,happ(s(fun('type.extreal.extreal',bool),X101),s('type.extreal.extreal',X102))) | s(bool,'T_0') = s(bool,X103)) )),
  inference(subsumption_resolution,[],[f2269,f263])).
fof(f2269,plain,(
  ( ! [X103,X101,X102] : (~p(s(bool,'T_0')) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X102),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X101)))))) | s(bool,X103) = s(bool,happ(s(fun('type.extreal.extreal',bool),X101),s('type.extreal.extreal',X102))) | s(bool,'T_0') = s(bool,X103)) )),
  inference(superposition,[],[f317,f783])).
fof(f1867,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X0))),s('type.extreal.extreal',X1)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',sK22(X1,X0))))) )),
  inference(subsumption_resolution,[],[f1865,f468])).
fof(f1865,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X0))),s('type.extreal.extreal',X1)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',sK22(X1,X0))))) )),
  inference(superposition,[],[f327,f270])).
fof(f1866,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',sK22(X1,X0))))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X0))),s('type.extreal.extreal',X1)))) )),
  inference(subsumption_resolution,[],[f1864,f468])).
fof(f1864,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',sK22(X1,X0))))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X0))),s('type.extreal.extreal',X1)))) )),
  inference(superposition,[],[f327,f270])).
fof(f327,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X0))),s('type.extreal.extreal',X1)))) | p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',sK22(X1,X0)))))) )),
  inference(cnf_transformation,[],[f202])).
fof(f2135,plain,(
  ( ! [X185,X186] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X185))) = s(bool,happ(s(fun(bool,bool),X185),s(bool,X186))) | s(bool,'T_0') = s(bool,X186) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X185)))) )),
  inference(superposition,[],[f303,f783])).
fof(f2136,plain,(
  ( ! [X189,X187,X190,X188] : (~p(s(bool,X188)) | p(s(bool,happ(s(fun(bool,bool),X187),s(bool,X190)))) | ~p(s(bool,happ(s(fun(bool,bool),X187),s(bool,X189)))) | s(bool,'T_0') = s(bool,X189) | s(bool,'T_0') = s(bool,sK18(X187,X188,bool))) )),
  inference(superposition,[],[f310,f783])).
fof(f2137,plain,(
  ( ! [X191,X194,X192,X193] : (~p(s(bool,X192)) | p(s(bool,happ(s(fun(bool,bool),X191),s(bool,X194)))) | ~p(s(bool,happ(s(fun(bool,bool),X191),s(bool,X193)))) | s(bool,'T_0') = s(bool,X193) | s(bool,'T_0') = s(bool,sK19(X191,X192,bool))) )),
  inference(superposition,[],[f312,f783])).
fof(f2138,plain,(
  ( ! [X198,X196,X197,X195] : (~p(s(bool,happ(s(fun(bool,bool),X195),s(bool,X197)))) | p(s(bool,X196)) | p(s(bool,happ(s(fun(bool,bool),X195),s(bool,X198)))) | s(bool,'T_0') = s(bool,X197) | s(bool,'T_0') = s(bool,sK20(X195,X196,bool))) )),
  inference(superposition,[],[f314,f783])).
fof(f2139,plain,(
  ( ! [X202,X200,X199,X201] : (~p(s(bool,happ(s(fun(bool,bool),X199),s(bool,X201)))) | p(s(bool,X200)) | p(s(bool,happ(s(fun(bool,bool),X199),s(bool,X202)))) | s(bool,'T_0') = s(bool,X201) | s(bool,'T_0') = s(bool,sK21(X199,X200,bool))) )),
  inference(superposition,[],[f316,f783])).
fof(f2292,plain,(
  ( ! [X185,X182,X184,X183] : (~p(s(bool,happ(s(fun(bool,bool),X182),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X182),s(bool,X184)))) | ~p(s(bool,X183)) | s(bool,X185) = s(bool,sK18(X182,X183,bool)) | s(bool,'T_0') = s(bool,X185)) )),
  inference(superposition,[],[f310,f783])).
fof(f2293,plain,(
  ( ! [X189,X187,X188,X186] : (~p(s(bool,happ(s(fun(bool,bool),X186),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X186),s(bool,X188)))) | ~p(s(bool,X187)) | s(bool,X189) = s(bool,sK19(X186,X187,bool)) | s(bool,'T_0') = s(bool,X189)) )),
  inference(superposition,[],[f312,f783])).
fof(f2294,plain,(
  ( ! [X191,X192,X190,X193] : (~p(s(bool,happ(s(fun(bool,bool),X190),s(bool,'T_0')))) | p(s(bool,X191)) | p(s(bool,happ(s(fun(bool,bool),X190),s(bool,X192)))) | s(bool,X193) = s(bool,sK20(X190,X191,bool)) | s(bool,'T_0') = s(bool,X193)) )),
  inference(superposition,[],[f314,f783])).
fof(f2295,plain,(
  ( ! [X196,X194,X197,X195] : (~p(s(bool,happ(s(fun(bool,bool),X194),s(bool,'T_0')))) | p(s(bool,X195)) | p(s(bool,happ(s(fun(bool,bool),X194),s(bool,X196)))) | s(bool,X197) = s(bool,sK21(X194,X195,bool)) | s(bool,'T_0') = s(bool,X197)) )),
  inference(superposition,[],[f316,f783])).
fof(f2634,plain,(
  ( ! [X146,X144,X147,X145] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X146),s('type.extreal.extreal',X144)))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X146),s('type.extreal.extreal',X145)))) | s(bool,X147) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X144),s('type.extreal.extreal',X145))) | s(bool,'T_0') = s(bool,X147)) )),
  inference(subsumption_resolution,[],[f2281,f263])).
fof(f2281,plain,(
  ( ! [X146,X144,X147,X145] : (~p(s(bool,'T_0')) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X146),s('type.extreal.extreal',X145)))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X146),s('type.extreal.extreal',X144)))) | s(bool,X147) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X144),s('type.extreal.extreal',X145))) | s(bool,'T_0') = s(bool,X147)) )),
  inference(superposition,[],[f367,f783])).
fof(f2635,plain,(
  ( ! [X152,X151,X153] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X152),s('type.extreal.extreal',X151)))) | s(bool,X153) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X151),s('type.extreal.extreal',X152))) | s(bool,'T_0') = s(bool,X153)) )),
  inference(subsumption_resolution,[],[f2283,f263])).
fof(f2283,plain,(
  ( ! [X152,X151,X153] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X152),s('type.extreal.extreal',X151)))) | s(bool,X153) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X151),s('type.extreal.extreal',X152))) | s(bool,'T_0') = s(bool,X153)) )),
  inference(superposition,[],[f324,f783])).
fof(f1779,plain,(
  ( ! [X23,X21,X22] : (sP0(X23,happ(s(fun(X21,bool),X22),s(X21,'const.min.@_1'(s(fun(X21,bool),X22))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X21,bool),X22))))) )),
  inference(superposition,[],[f334,f303])).
fof(f1780,plain,(
  ( ! [X26,X24,X25] : (sP0(happ(s(fun(X24,bool),X25),s(X24,'const.min.@_1'(s(fun(X24,bool),X25)))),X26) | ~p(s(bool,'|dollar|exists_1'(s(fun(X24,bool),X25))))) )),
  inference(superposition,[],[f335,f303])).
fof(f1781,plain,(
  ( ! [X28,X29,X27] : (sP1(happ(s(fun(X27,bool),X28),s(X27,'const.min.@_1'(s(fun(X27,bool),X28)))),X29) | p(s(bool,X29)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X27,bool),X28))))) )),
  inference(superposition,[],[f350,f303])).
fof(f1782,plain,(
  ( ! [X30,X31,X32] : (sP1(X32,happ(s(fun(X30,bool),X31),s(X30,'const.min.@_1'(s(fun(X30,bool),X31))))) | p(s(bool,X32)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X30,bool),X31))))) )),
  inference(superposition,[],[f353,f303])).
fof(f1783,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,'|dollar|exists_1'(s(fun(X33,bool),X34)))) | sP6(X35,X36,happ(s(fun(X33,bool),X34),s(X33,'const.min.@_1'(s(fun(X33,bool),X34))))) | ~sP5(X35,happ(s(fun(X33,bool),X34),s(X33,'const.min.@_1'(s(fun(X33,bool),X34)))),X36) | ~sP4(X36,happ(s(fun(X33,bool),X34),s(X33,'const.min.@_1'(s(fun(X33,bool),X34)))),X35) | ~sP3(happ(s(fun(X33,bool),X34),s(X33,'const.min.@_1'(s(fun(X33,bool),X34)))),X36,X35)) )),
  inference(superposition,[],[f382,f303])).
fof(f1784,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,'|dollar|exists_1'(s(fun(X37,bool),X38)))) | ~p(s(bool,X39)) | p(s(bool,X40)) | ~sP5(X39,happ(s(fun(X37,bool),X38),s(X37,'const.min.@_1'(s(fun(X37,bool),X38)))),X40)) )),
  inference(superposition,[],[f385,f303])).
fof(f1785,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,'|dollar|exists_1'(s(fun(X41,bool),X42)))) | ~p(s(bool,X43)) | p(s(bool,X44)) | ~sP4(X43,happ(s(fun(X41,bool),X42),s(X41,'const.min.@_1'(s(fun(X41,bool),X42)))),X44)) )),
  inference(superposition,[],[f389,f303])).
fof(f1786,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,'|dollar|exists_1'(s(fun(X45,bool),X46)))) | ~p(s(bool,X47)) | p(s(bool,X48)) | ~sP7(X47,happ(s(fun(X45,bool),X46),s(X45,'const.min.@_1'(s(fun(X45,bool),X46)))),X48)) )),
  inference(superposition,[],[f408,f303])).
fof(f1787,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,'|dollar|exists_1'(s(fun(X49,bool),X50)))) | ~p(s(bool,X51)) | p(s(bool,X52)) | ~sP8(X52,X51,happ(s(fun(X49,bool),X50),s(X49,'const.min.@_1'(s(fun(X49,bool),X50)))))) )),
  inference(superposition,[],[f417,f303])).
fof(f1788,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,'|dollar|exists_1'(s(fun(X53,bool),X54)))) | p(s(bool,X55)) | p(s(bool,X56)) | ~sP9(happ(s(fun(X53,bool),X54),s(X53,'const.min.@_1'(s(fun(X53,bool),X54)))),X55,X56)) )),
  inference(superposition,[],[f425,f303])).
fof(f1789,plain,(
  ( ! [X59,X57,X60,X58] : (sP10(happ(s(fun(X57,bool),X58),s(X57,'const.min.@_1'(s(fun(X57,bool),X58)))),X60,X59) | ~p(s(bool,X59)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X57,bool),X58))))) )),
  inference(superposition,[],[f430,f303])).
fof(f1790,plain,(
  ( ! [X61,X64,X62,X63] : (sP10(happ(s(fun(X61,bool),X62),s(X61,'const.min.@_1'(s(fun(X61,bool),X62)))),X63,X64) | ~p(s(bool,X63)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X61,bool),X62))))) )),
  inference(superposition,[],[f431,f303])).
fof(f1791,plain,(
  ( ! [X68,X66,X67,X65] : (~sP10(happ(s(fun(X65,bool),X66),s(X65,'const.min.@_1'(s(fun(X65,bool),X66)))),X68,X67) | p(s(bool,X67)) | p(s(bool,X68)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X65,bool),X66))))) )),
  inference(superposition,[],[f432,f303])).
fof(f1792,plain,(
  ( ! [X70,X72,X71,X69] : (~p(s(bool,'|dollar|exists_1'(s(fun(X69,bool),X70)))) | p(s(bool,X71)) | ~p(s(bool,X72)) | ~sP11(X72,X71,happ(s(fun(X69,bool),X70),s(X69,'const.min.@_1'(s(fun(X69,bool),X70)))))) )),
  inference(superposition,[],[f442,f303])).
fof(f1793,plain,(
  ( ! [X76,X74,X75,X73] : (sP12(happ(s(fun(X73,bool),X74),s(X73,'const.min.@_1'(s(fun(X73,bool),X74)))),X76,X75) | p(s(bool,X75)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X73,bool),X74))))) )),
  inference(superposition,[],[f447,f303])).
fof(f1794,plain,(
  ( ! [X80,X78,X79,X77] : (sP12(happ(s(fun(X77,bool),X78),s(X77,'const.min.@_1'(s(fun(X77,bool),X78)))),X79,X80) | ~p(s(bool,X79)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X77,bool),X78))))) )),
  inference(superposition,[],[f448,f303])).
fof(f1795,plain,(
  ( ! [X83,X81,X84,X82] : (~sP12(happ(s(fun(X81,bool),X82),s(X81,'const.min.@_1'(s(fun(X81,bool),X82)))),X84,X83) | ~p(s(bool,X83)) | p(s(bool,X84)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X81,bool),X82))))) )),
  inference(superposition,[],[f449,f303])).
fof(f1796,plain,(
  ( ! [X88,X87,X85,X86] : (sP12(happ(s(fun(X85,bool),X86),s(X85,'const.min.@_1'(s(fun(X85,bool),X86)))),X87,X88) | ~sP11(happ(s(fun(X85,bool),X86),s(X85,'const.min.@_1'(s(fun(X85,bool),X86)))),X87,X88) | ~p(s(bool,'|dollar|exists_1'(s(fun(X85,bool),X86))))) )),
  inference(superposition,[],[f452,f303])).
fof(f1797,plain,(
  ( ! [X92,X90,X91,X89] : (sP10(happ(s(fun(X89,bool),X90),s(X89,'const.min.@_1'(s(fun(X89,bool),X90)))),X91,X92) | ~sP9(happ(s(fun(X89,bool),X90),s(X89,'const.min.@_1'(s(fun(X89,bool),X90)))),X91,X92) | ~p(s(bool,'|dollar|exists_1'(s(fun(X89,bool),X90))))) )),
  inference(superposition,[],[f453,f303])).
fof(f2636,plain,(
  ( ! [X156,X154,X155] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X155))),s('type.extreal.extreal',X154)))) | s(bool,X156) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK22(X154,X155)),s('type.extreal.extreal',X154))) | s(bool,'T_0') = s(bool,X156)) )),
  inference(subsumption_resolution,[],[f2284,f263])).
fof(f2284,plain,(
  ( ! [X156,X154,X155] : (~p(s(bool,'T_0')) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X155))),s('type.extreal.extreal',X154)))) | s(bool,X156) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK22(X154,X155)),s('type.extreal.extreal',X154))) | s(bool,'T_0') = s(bool,X156)) )),
  inference(superposition,[],[f328,f783])).
fof(f2625,plain,(
  ( ! [X90,X88,X87,X89] : (~sP9(X87,X88,X89) | sP10(X87,X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f2262,f263])).
fof(f2262,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,'T_0')) | ~sP9(X87,X88,X89) | sP10(X87,X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f453,f783])).
fof(f2624,plain,(
  ( ! [X85,X83,X86,X84] : (~sP11(X83,X84,X85) | sP12(X83,X84,X85) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f2261,f263])).
fof(f2261,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,'T_0')) | ~sP11(X83,X84,X85) | sP12(X83,X84,X85) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f452,f783])).
fof(f2623,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X80)) | p(s(bool,X81)) | ~sP12(X79,X81,X80) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f2260,f263])).
fof(f2260,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,'T_0')) | ~p(s(bool,X80)) | p(s(bool,X81)) | ~sP12(X79,X81,X80) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f449,f783])).
fof(f2622,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X76)) | sP12(X75,X76,X77) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f2259,f263])).
fof(f2259,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,'T_0')) | ~p(s(bool,X76)) | sP12(X75,X76,X77) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f448,f783])).
fof(f2621,plain,(
  ( ! [X74,X72,X71,X73] : (sP12(X71,X73,X72) | p(s(bool,X72)) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f2258,f263])).
fof(f2258,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,'T_0')) | p(s(bool,X72)) | sP12(X71,X73,X72) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f447,f783])).
fof(f2620,plain,(
  ( ! [X70,X68,X69,X67] : (~sP11(X69,X68,X67) | ~p(s(bool,X69)) | p(s(bool,X68)) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f2257,f263])).
fof(f2257,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | p(s(bool,X68)) | ~p(s(bool,X69)) | ~sP11(X69,X68,X67) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f442,f783])).
fof(f2619,plain,(
  ( ! [X66,X64,X65,X63] : (~sP10(X63,X65,X64) | p(s(bool,X65)) | p(s(bool,X64)) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f2256,f263])).
fof(f2256,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'T_0')) | p(s(bool,X64)) | p(s(bool,X65)) | ~sP10(X63,X65,X64) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f432,f783])).
fof(f2618,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X60)) | sP10(X59,X60,X61) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f2255,f263])).
fof(f2255,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | ~p(s(bool,X60)) | sP10(X59,X60,X61) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f431,f783])).
fof(f2617,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X56)) | sP10(X55,X57,X56) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f2254,f263])).
fof(f2254,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X56)) | sP10(X55,X57,X56) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f430,f783])).
fof(f2616,plain,(
  ( ! [X54,X52,X53,X51] : (~sP9(X51,X52,X53) | p(s(bool,X53)) | p(s(bool,X52)) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f2253,f263])).
fof(f2253,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | p(s(bool,X52)) | p(s(bool,X53)) | ~sP9(X51,X52,X53) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f425,f783])).
fof(f2615,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X48)) | p(s(bool,X49)) | ~sP8(X49,X48,X47) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f2252,f263])).
fof(f2252,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X48)) | p(s(bool,X49)) | ~sP8(X49,X48,X47) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f417,f783])).
fof(f2614,plain,(
  ( ! [X45,X43,X46,X44] : (~sP7(X44,X43,X45) | p(s(bool,X45)) | ~p(s(bool,X44)) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f2251,f263])).
fof(f2251,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X44)) | p(s(bool,X45)) | ~sP7(X44,X43,X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f408,f783])).
fof(f2613,plain,(
  ( ! [X39,X41,X42,X40] : (~sP4(X40,X39,X41) | p(s(bool,X41)) | ~p(s(bool,X40)) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f2250,f263])).
fof(f2250,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X40)) | p(s(bool,X41)) | ~sP4(X40,X39,X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f389,f783])).
fof(f2612,plain,(
  ( ! [X37,X35,X38,X36] : (~sP5(X36,X35,X37) | p(s(bool,X37)) | ~p(s(bool,X36)) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f2249,f263])).
fof(f2249,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X36)) | p(s(bool,X37)) | ~sP5(X36,X35,X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f385,f783])).
fof(f2611,plain,(
  ( ! [X33,X31,X34,X32] : (~sP4(X33,X31,X32) | ~sP5(X32,X31,X33) | sP6(X32,X33,X31) | ~sP3(X31,X33,X32) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f2248,f263])).
fof(f2248,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | sP6(X32,X33,X31) | ~sP5(X32,X31,X33) | ~sP4(X33,X31,X32) | ~sP3(X31,X33,X32) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f382,f783])).
fof(f2608,plain,(
  ( ! [X24,X23,X22] : (sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f2245,f263])).
fof(f2245,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f335,f783])).
fof(f2607,plain,(
  ( ! [X21,X19,X20] : (sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f2244,f263])).
fof(f2244,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f334,f783])).
fof(f2606,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f2243,f263])).
fof(f2243,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f321,f783])).
fof(f2605,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f2239,f263])).
fof(f2239,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f280,f783])).
fof(f2140,plain,(
  ( ! [X206,X204,X205,X203] : (s(X205,f(s(bool,X206))) != s(X205,g(s(bool,X206))) | s(fun(bool,X205),X203) = s(fun(bool,X205),X204) | s(bool,'T_0') = s(bool,X206) | s(bool,'T_0') = s(bool,sK23(X203,X204,X205,bool))) )),
  inference(superposition,[],[f366,f783])).
fof(f1764,plain,(
  ( ! [X0,X1] : (~sP6('F_0',happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(unit_resulting_resolution,[],[f468,f303,f400])).
fof(f1765,plain,(
  ( ! [X0,X1] : (sP6('T_0',happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(unit_resulting_resolution,[],[f263,f303,f398])).
fof(f1766,plain,(
  ( ! [X0,X1] : (~sP6('F_0','|dollar|exists_1'(s(fun(X0,bool),X1)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f468,f303,f400])).
fof(f1767,plain,(
  ( ! [X0,X1] : (sP6('T_0','|dollar|exists_1'(s(fun(X0,bool),X1)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f263,f303,f398])).
fof(f1361,plain,(
  sP6('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f491,f484,f480,f468,f383])).
fof(f1313,plain,(
  sP6('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f497,f484,f477,f263,f382])).
fof(f1233,plain,(
  ~sP12('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1218,f437])).
fof(f1218,plain,(
  ~sP11('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f263,f468,f263,f442])).
fof(f1201,plain,(
  sP12('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f520,f263,f468,f441])).
fof(f1142,plain,(
  ~sP10('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1126,f420])).
fof(f1126,plain,(
  ~sP9('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f468,f468,f263,f425])).
fof(f1111,plain,(
  sP10('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f513,f468,f468,f424])).
fof(f1041,plain,(
  ~sP8('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1026,f401])).
fof(f1026,plain,(
  ~sP7('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f468,f263,f263,f408])).
fof(f1012,plain,(
  sP8('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f500,f263,f263,f404])).
fof(f1008,plain,(
  ~sP6('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f995,f378])).
fof(f995,plain,(
  ~sP3('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f468,f468,f468,f393])).
fof(f987,plain,(
  ~sP6('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f972,f379])).
fof(f972,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f468,f263,f263,f389])).
fof(f947,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f468,f263,f263,f385])).
fof(f901,plain,(
  ~sP6('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f263,f525,f399])).
fof(f904,plain,(
  ~sP6('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f263,f525,f399])).
fof(f876,plain,(
  sP6('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f468,f525,f397])).
fof(f878,plain,(
  sP6('F_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f468,f525,f397])).
fof(f710,plain,(
  ( ! [X0] : (sP12('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f263,f263,f448])).
fof(f698,plain,(
  ( ! [X0] : (sP12('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f468,f263,f447])).
fof(f690,plain,(
  ( ! [X0] : (~sP12('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f263,f468,f436])).
fof(f688,plain,(
  ( ! [X0] : (~sP12('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f468,f468,f435])).
fof(f674,plain,(
  ( ! [X0] : (sP10('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f263,f263,f431])).
fof(f662,plain,(
  ( ! [X0] : (sP10('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f263,f263,f430])).
fof(f659,plain,(
  ( ! [X0] : (~sP10('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f263,f468,f419])).
fof(f658,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f263,f468,f418])).
fof(f654,plain,(
  ( ! [X0] : (sP8('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f468,f468,f413])).
fof(f647,plain,(
  ( ! [X0] : (sP8('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f468,f468,f412])).
fof(f645,plain,(
  ( ! [X0] : (~sP8('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f263,f468,f403])).
fof(f644,plain,(
  ( ! [X0] : (~sP8('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f263,f468,f402])).
fof(f639,plain,(
  ( ! [X0] : (sP2('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f263,f263,f373])).
fof(f622,plain,(
  ( ! [X0] : (~sP2('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f468,f468,f369])).
fof(f613,plain,(
  ( ! [X0] : (~sP2(X0,'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f468,f468,f368])).
fof(f602,plain,(
  ~sP0('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f468,f468,f470])).
fof(f586,plain,(
  sP1('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f468,f263,f353])).
fof(f578,plain,(
  sP1('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f468,f263,f350])).
fof(f575,plain,(
  ~sP1('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f263,f263,f349])).
fof(f565,plain,(
  ~sP1('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f468,f468,f348])).
fof(f522,plain,(
  ( ! [X0,X1] : (sP2(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f263,f455])).
fof(f520,plain,(
  ( ! [X0,X1] : (sP11('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f468,f445])).
fof(f517,plain,(
  ( ! [X0,X1] : (sP11(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f263,f444])).
fof(f515,plain,(
  ( ! [X0,X1] : (sP11(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f468,f443])).
fof(f513,plain,(
  ( ! [X0,X1] : (sP9('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f468,f428])).
fof(f510,plain,(
  ( ! [X0,X1] : (sP9(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f263,f427])).
fof(f507,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f263,f426])).
fof(f505,plain,(
  ( ! [X0,X1] : (sP7('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f468,f411])).
fof(f503,plain,(
  ( ! [X0,X1] : (sP7(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f468,f410])).
fof(f500,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f263,f409])).
fof(f497,plain,(
  ( ! [X0,X1] : (sP3('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f263,f396])).
fof(f494,plain,(
  ( ! [X0,X1] : (sP3(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f263,f395])).
fof(f491,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f263,f394])).
fof(f489,plain,(
  ( ! [X0,X1] : (sP4('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f468,f392])).
fof(f487,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f468,f391])).
fof(f484,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f263,f390])).
fof(f482,plain,(
  ( ! [X0,X1] : (sP5('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f468,f388])).
fof(f480,plain,(
  ( ! [X0,X1] : (sP5(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f468,f387])).
fof(f477,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f263,f386])).
fof(f474,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f263,f335])).
fof(f471,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f263,f334])).
fof(f446,plain,(
  ( ! [X2,X0,X1] : (sP12(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f259])).
fof(f259,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP12(X2,X0,X1)) & (~sP12(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f258])).
fof(f258,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP12(X2,X0,X1)) & (~sP12(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP12(X2,X0,X1))),
  inference(definition_folding,[],[f145,f164,f163])).
fof(f163,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f164,plain,(
  ! [X2,X0,X1] : (sP12(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP11(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP12])])).
fof(f145,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f113])).
fof(f113,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X18,X19,X17] : ((p(s(bool,X17)) <=> (p(s(bool,X19)) => p(s(bool,X18)))) <=> ((p(s(bool,X17)) | p(s(bool,X19))) & (p(s(bool,X17)) | ~p(s(bool,X18))) & (~p(s(bool,X19)) | p(s(bool,X18)) | ~p(s(bool,X17)))))),
  file('thm.extreal.sup_lt','thm.sat.sat_27_prove_ltb')).
fof(f447,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP12(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f259])).
fof(f448,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP12(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f259])).
fof(f449,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP12(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f259])).
fof(f442,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP11(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f257])).
fof(f257,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f256])).
fof(f256,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f255])).
fof(f255,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f163])).
fof(f443,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f257])).
fof(f444,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f257])).
fof(f445,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f257])).
fof(f435,plain,(
  ( ! [X2,X0,X1] : (~sP12(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f254])).
fof(f254,plain,(
  ! [X0,X1,X2] : ((~sP12(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP11(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP11(X0,X1,X2) | sP12(X0,X1,X2)))),
  inference(rectify,[],[f253])).
fof(f253,plain,(
  ! [X2,X0,X1] : ((~sP12(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP11(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP11(X2,X0,X1) | sP12(X2,X0,X1)))),
  inference(flattening,[],[f252])).
fof(f252,plain,(
  ! [X2,X0,X1] : ((~sP12(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP11(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP11(X2,X0,X1)) | sP12(X2,X0,X1)))),
  inference(nnf_transformation,[],[f164])).
fof(f436,plain,(
  ( ! [X2,X0,X1] : (~sP12(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f254])).
fof(f437,plain,(
  ( ! [X2,X0,X1] : (~sP12(X0,X1,X2) | sP11(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f254])).
fof(f452,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP11(X0,X1,X2) | sP12(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f438])).
fof(f438,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~sP11(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f254])).
fof(f439,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~sP11(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f254])).
fof(f440,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~sP11(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f254])).
fof(f441,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~sP11(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f254])).
fof(f429,plain,(
  ( ! [X2,X0,X1] : (sP10(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f251])).
fof(f251,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP10(X2,X0,X1)) & (~sP10(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f250])).
fof(f250,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP10(X2,X0,X1)) & (~sP10(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP10(X2,X0,X1))),
  inference(definition_folding,[],[f112,f161,f160])).
fof(f160,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f161,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP9(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f112,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f111])).
fof(f111,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X18,X19,X17] : ((p(s(bool,X17)) <=> (p(s(bool,X19)) | p(s(bool,X18)))) <=> ((p(s(bool,X17)) | ~p(s(bool,X19))) & (p(s(bool,X17)) | ~p(s(bool,X18))) & (p(s(bool,X19)) | p(s(bool,X18)) | ~p(s(bool,X17)))))),
  file('thm.extreal.sup_lt','thm.sat.sat_25_prove_ltb')).
fof(f430,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP10(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f251])).
fof(f431,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP10(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f251])).
fof(f432,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | ~sP10(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f251])).
fof(f425,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f249])).
fof(f249,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f248])).
fof(f248,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f247])).
fof(f247,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f160])).
fof(f426,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f249])).
fof(f427,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f249])).
fof(f428,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f249])).
fof(f418,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f246])).
fof(f246,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP9(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP9(X0,X1,X2) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f245])).
fof(f245,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP9(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP9(X2,X0,X1) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f244])).
fof(f244,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP9(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP9(X2,X0,X1)) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f161])).
fof(f419,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f246])).
fof(f420,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f246])).
fof(f453,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP9(X0,X1,X2) | sP10(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f421])).
fof(f421,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~sP9(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f246])).
fof(f422,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~sP9(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f246])).
fof(f423,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~sP9(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f246])).
fof(f424,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~sP9(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f246])).
fof(f412,plain,(
  ( ! [X2,X0,X1] : (sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f243])).
fof(f243,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP8(X2,X0,X1)) & (~sP8(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f242])).
fof(f242,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP8(X2,X0,X1)) & (~sP8(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP8(X2,X0,X1))),
  inference(definition_folding,[],[f110,f158,f157])).
fof(f157,plain,(
  ! [X0,X1,X2] : (sP7(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f158,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (sP7(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f110,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f109])).
fof(f109,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X18,X19,X17] : ((p(s(bool,X17)) <=> (p(s(bool,X19)) & p(s(bool,X18)))) <=> ((p(s(bool,X17)) | ~p(s(bool,X19)) | ~p(s(bool,X18))) & (p(s(bool,X19)) | ~p(s(bool,X17))) & (p(s(bool,X18)) | ~p(s(bool,X17)))))),
  file('thm.extreal.sup_lt','thm.sat.sat_23_prove_ltb')).
fof(f413,plain,(
  ( ! [X2,X0,X1] : (sP8(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f243])).
fof(f414,plain,(
  ( ! [X2,X0,X1] : (sP8(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f243])).
fof(f417,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP8(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f243])).
fof(f408,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP7(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(flattening,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(nnf_transformation,[],[f157])).
fof(f409,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f241])).
fof(f410,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f241])).
fof(f411,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f241])).
fof(f401,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | sP7(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | (sP7(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP7(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f238])).
fof(f238,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (sP7(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP7(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f237])).
fof(f237,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (sP7(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP7(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f158])).
fof(f402,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f239])).
fof(f403,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f239])).
fof(f404,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP7(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f239])).
fof(f454,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0)) | ~sP7(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f407])).
fof(f407,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP7(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f239])).
fof(f397,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f236])).
fof(f236,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP6(X2,X1,X0)) & (~sP6(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP6(X2,X1,X0))),
  inference(definition_folding,[],[f108,f155,f154,f153,f152])).
fof(f152,plain,(
  ! [X0,X1,X2] : (sP3(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f153,plain,(
  ! [X1,X0,X2] : (sP4(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f154,plain,(
  ! [X2,X0,X1] : (sP5(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f155,plain,(
  ! [X2,X1,X0] : (sP6(X2,X1,X0) <=> (sP3(X0,X1,X2) & sP4(X1,X0,X2) & sP5(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f108,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f107])).
fof(f107,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X18,X19,X17] : ((p(s(bool,X17)) <=> s(bool,X18) = s(bool,X19)) <=> ((p(s(bool,X17)) | p(s(bool,X19)) | p(s(bool,X18))) & (p(s(bool,X17)) | ~p(s(bool,X18)) | ~p(s(bool,X19))) & (p(s(bool,X19)) | ~p(s(bool,X18)) | ~p(s(bool,X17))) & (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X17)))))),
  file('thm.extreal.sup_lt','thm.sat.sat_2|49|_prove_ltb')).
fof(f398,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f236])).
fof(f399,plain,(
  ( ! [X2,X0,X1] : (~sP6(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f236])).
fof(f400,plain,(
  ( ! [X2,X0,X1] : (~sP6(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f236])).
fof(f393,plain,(
  ( ! [X2,X0,X1] : (~sP3(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f235])).
fof(f235,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(flattening,[],[f234])).
fof(f234,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(nnf_transformation,[],[f152])).
fof(f394,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f235])).
fof(f395,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f235])).
fof(f396,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f235])).
fof(f389,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f233])).
fof(f233,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f232])).
fof(f232,plain,(
  ! [X1,X0,X2] : ((~sP4(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP4(X1,X0,X2)))),
  inference(flattening,[],[f231])).
fof(f231,plain,(
  ! [X1,X0,X2] : ((~sP4(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP4(X1,X0,X2)))),
  inference(nnf_transformation,[],[f153])).
fof(f390,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f233])).
fof(f391,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f233])).
fof(f392,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f233])).
fof(f385,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f230])).
fof(f230,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f229])).
fof(f229,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP5(X2,X0,X1)))),
  inference(flattening,[],[f228])).
fof(f228,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP5(X2,X0,X1)))),
  inference(nnf_transformation,[],[f154])).
fof(f386,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f230])).
fof(f387,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f230])).
fof(f388,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f230])).
fof(f378,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f227])).
fof(f227,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (sP3(X2,X1,X0) & sP4(X1,X2,X0) & sP5(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP3(X2,X1,X0) | ~sP4(X1,X2,X0) | ~sP5(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(rectify,[],[f226])).
fof(f226,plain,(
  ! [X2,X1,X0] : ((~sP6(X2,X1,X0) | (sP3(X0,X1,X2) & sP4(X1,X0,X2) & sP5(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP3(X0,X1,X2) | ~sP4(X1,X0,X2) | ~sP5(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP6(X2,X1,X0)))),
  inference(flattening,[],[f225])).
fof(f225,plain,(
  ! [X2,X1,X0] : ((~sP6(X2,X1,X0) | (sP3(X0,X1,X2) & sP4(X1,X0,X2) & sP5(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP3(X0,X1,X2) | ~sP4(X1,X0,X2) | ~sP5(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP6(X2,X1,X0)))),
  inference(nnf_transformation,[],[f155])).
fof(f379,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | sP4(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f227])).
fof(f380,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | sP5(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f227])).
fof(f381,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f227])).
fof(f382,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP6(X0,X1,X2) | ~sP5(X0,X2,X1) | ~sP4(X1,X2,X0) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f227])).
fof(f383,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X2,X1) | p(s(bool,X1)) | sP6(X0,X1,X2) | ~sP4(X1,X2,X0) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f227])).
fof(f384,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X2,X1) | p(s(bool,X0)) | sP6(X0,X1,X2) | ~sP4(X1,X2,X0) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f227])).
fof(f368,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f222])).
fof(f222,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | ((p(s(bool,X1)) | p(s(bool,X2))) & (p(s(bool,X1)) | p(s(bool,X0))))) & ((~p(s(bool,X1)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f221])).
fof(f221,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f220])).
fof(f220,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2))))) & (((~p(s(bool,X0)) & ~p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X2)))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f369,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f222])).
fof(f455,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(duplicate_literal_removal,[],[f370])).
fof(f370,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f222])).
fof(f371,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f222])).
fof(f372,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f222])).
fof(f373,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f222])).
fof(f348,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f213])).
fof(f349,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f213])).
fof(f470,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f333,f468])).
fof(f333,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f206])).
fof(f206,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f205])).
fof(f205,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f334,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f206])).
fof(f335,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f206])).
fof(f37120,plain,(
  ( ! [X471,X470] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X471)))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X470))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X471))) = s(bool,happ(s(fun(bool,bool),X471),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X470)))) )),
  inference(subsumption_resolution,[],[f36727,f33951])).
fof(f36727,plain,(
  ( ! [X471,X470] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X471)))) | s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X470),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X470))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X471))) = s(bool,happ(s(fun(bool,bool),X471),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X470)))) )),
  inference(superposition,[],[f4871,f2291])).
fof(f37119,plain,(
  ( ! [X457,X456] : (p(s(bool,happ(s(fun(X456,bool),X457),s(X456,sK17(X457,X456))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X456,bool),X457)))) ) | $spl206),
  inference(cnf_transformation,[],[f37119_D])).
fof(f37119_D,plain,(
  ( ! [X457,X456] : (p(s(bool,happ(s(fun(X456,bool),X457),s(X456,sK17(X457,X456))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X456,bool),X457)))) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f36723,plain,(
  ( ! [X457,X456,X458] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X458)))) | p(s(bool,happ(s(fun(X456,bool),X457),s(X456,sK17(X457,X456))))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X458))) = s(bool,happ(s(fun(bool,bool),X458),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X456,bool),X457)))) )),
  inference(superposition,[],[f1831,f2291])).
fof(f37118,plain,(
  ( ! [X453,X454] : (p(s(bool,happ(s(fun(X453,bool),X454),s(X453,sK16(X454,X453))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X453,bool),X454)))) ) | $spl204),
  inference(cnf_transformation,[],[f37118_D])).
fof(f37118_D,plain,(
  ( ! [X453,X454] : (p(s(bool,happ(s(fun(X453,bool),X454),s(X453,sK16(X454,X453))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X453,bool),X454)))) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f36722,plain,(
  ( ! [X453,X455,X454] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X455)))) | p(s(bool,happ(s(fun(X453,bool),X454),s(X453,sK16(X454,X453))))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X455))) = s(bool,happ(s(fun(bool,bool),X455),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X453,bool),X454)))) )),
  inference(superposition,[],[f1817,f2291])).
fof(f37117,plain,(
  ( ! [X449,X451,X452,X448] : (p(s(bool,X451)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X448,bool),X449))) | ~sP11(X452,X451,happ(s(fun(X448,bool),X449),s(X448,'const.min.@_1'(s(fun(X448,bool),X449))))) | ~p(s(bool,X452))) ) | $spl202),
  inference(cnf_transformation,[],[f37117_D])).
fof(f37117_D,plain,(
  ( ! [X449,X451,X452,X448] : (p(s(bool,X451)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X448,bool),X449))) | ~sP11(X452,X451,happ(s(fun(X448,bool),X449),s(X448,'const.min.@_1'(s(fun(X448,bool),X449))))) | ~p(s(bool,X452))) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f36721,plain,(
  ( ! [X449,X451,X452,X448,X450] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X450)))) | p(s(bool,X451)) | ~p(s(bool,X452)) | ~sP11(X452,X451,happ(s(fun(X448,bool),X449),s(X448,'const.min.@_1'(s(fun(X448,bool),X449))))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X450))) = s(bool,happ(s(fun(bool,bool),X450),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X448,bool),X449)))) )),
  inference(superposition,[],[f1792,f2291])).
fof(f37116,plain,(
  ( ! [X447,X443,X444,X446] : (p(s(bool,X446)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X443,bool),X444))) | ~sP9(happ(s(fun(X443,bool),X444),s(X443,'const.min.@_1'(s(fun(X443,bool),X444)))),X446,X447) | p(s(bool,X447))) ) | $spl200),
  inference(cnf_transformation,[],[f37116_D])).
fof(f37116_D,plain,(
  ( ! [X447,X443,X444,X446] : (p(s(bool,X446)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X443,bool),X444))) | ~sP9(happ(s(fun(X443,bool),X444),s(X443,'const.min.@_1'(s(fun(X443,bool),X444)))),X446,X447) | p(s(bool,X447))) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f36720,plain,(
  ( ! [X445,X447,X443,X444,X446] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X445)))) | p(s(bool,X446)) | p(s(bool,X447)) | ~sP9(happ(s(fun(X443,bool),X444),s(X443,'const.min.@_1'(s(fun(X443,bool),X444)))),X446,X447) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X445))) = s(bool,happ(s(fun(bool,bool),X445),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X443,bool),X444)))) )),
  inference(superposition,[],[f1788,f2291])).
fof(f37115,plain,(
  ( ! [X441,X438,X442,X439] : (~p(s(bool,X441)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X438,bool),X439))) | ~sP8(X442,X441,happ(s(fun(X438,bool),X439),s(X438,'const.min.@_1'(s(fun(X438,bool),X439))))) | p(s(bool,X442))) ) | $spl198),
  inference(cnf_transformation,[],[f37115_D])).
fof(f37115_D,plain,(
  ( ! [X441,X438,X442,X439] : (~p(s(bool,X441)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X438,bool),X439))) | ~sP8(X442,X441,happ(s(fun(X438,bool),X439),s(X438,'const.min.@_1'(s(fun(X438,bool),X439))))) | p(s(bool,X442))) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f36719,plain,(
  ( ! [X441,X438,X440,X442,X439] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X440)))) | ~p(s(bool,X441)) | p(s(bool,X442)) | ~sP8(X442,X441,happ(s(fun(X438,bool),X439),s(X438,'const.min.@_1'(s(fun(X438,bool),X439))))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X440))) = s(bool,happ(s(fun(bool,bool),X440),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X438,bool),X439)))) )),
  inference(superposition,[],[f1787,f2291])).
fof(f37114,plain,(
  ( ! [X436,X434,X437,X433] : (~p(s(bool,X436)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X433,bool),X434))) | ~sP7(X436,happ(s(fun(X433,bool),X434),s(X433,'const.min.@_1'(s(fun(X433,bool),X434)))),X437) | p(s(bool,X437))) ) | $spl196),
  inference(cnf_transformation,[],[f37114_D])).
fof(f37114_D,plain,(
  ( ! [X436,X434,X437,X433] : (~p(s(bool,X436)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X433,bool),X434))) | ~sP7(X436,happ(s(fun(X433,bool),X434),s(X433,'const.min.@_1'(s(fun(X433,bool),X434)))),X437) | p(s(bool,X437))) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f36718,plain,(
  ( ! [X436,X434,X437,X433,X435] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X435)))) | ~p(s(bool,X436)) | p(s(bool,X437)) | ~sP7(X436,happ(s(fun(X433,bool),X434),s(X433,'const.min.@_1'(s(fun(X433,bool),X434)))),X437) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X435))) = s(bool,happ(s(fun(bool,bool),X435),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X433,bool),X434)))) )),
  inference(superposition,[],[f1786,f2291])).
fof(f37113,plain,(
  ( ! [X428,X432,X429,X431] : (~p(s(bool,X431)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X428,bool),X429))) | ~sP4(X431,happ(s(fun(X428,bool),X429),s(X428,'const.min.@_1'(s(fun(X428,bool),X429)))),X432) | p(s(bool,X432))) ) | $spl194),
  inference(cnf_transformation,[],[f37113_D])).
fof(f37113_D,plain,(
  ( ! [X428,X432,X429,X431] : (~p(s(bool,X431)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X428,bool),X429))) | ~sP4(X431,happ(s(fun(X428,bool),X429),s(X428,'const.min.@_1'(s(fun(X428,bool),X429)))),X432) | p(s(bool,X432))) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f36717,plain,(
  ( ! [X428,X430,X432,X429,X431] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X430)))) | ~p(s(bool,X431)) | p(s(bool,X432)) | ~sP4(X431,happ(s(fun(X428,bool),X429),s(X428,'const.min.@_1'(s(fun(X428,bool),X429)))),X432) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X430))) = s(bool,happ(s(fun(bool,bool),X430),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X428,bool),X429)))) )),
  inference(superposition,[],[f1785,f2291])).
fof(f37112,plain,(
  ( ! [X424,X426,X423,X427] : (~p(s(bool,X426)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X423,bool),X424))) | ~sP5(X426,happ(s(fun(X423,bool),X424),s(X423,'const.min.@_1'(s(fun(X423,bool),X424)))),X427) | p(s(bool,X427))) ) | $spl192),
  inference(cnf_transformation,[],[f37112_D])).
fof(f37112_D,plain,(
  ( ! [X424,X426,X423,X427] : (~p(s(bool,X426)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X423,bool),X424))) | ~sP5(X426,happ(s(fun(X423,bool),X424),s(X423,'const.min.@_1'(s(fun(X423,bool),X424)))),X427) | p(s(bool,X427))) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f36716,plain,(
  ( ! [X424,X426,X423,X425,X427] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X425)))) | ~p(s(bool,X426)) | p(s(bool,X427)) | ~sP5(X426,happ(s(fun(X423,bool),X424),s(X423,'const.min.@_1'(s(fun(X423,bool),X424)))),X427) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X425))) = s(bool,happ(s(fun(bool,bool),X425),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X423,bool),X424)))) )),
  inference(superposition,[],[f1784,f2291])).
fof(f37111,plain,(
  ( ! [X422,X418,X421,X419] : (sP6(X421,X422,happ(s(fun(X418,bool),X419),s(X418,'const.min.@_1'(s(fun(X418,bool),X419))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X418,bool),X419))) | ~sP3(happ(s(fun(X418,bool),X419),s(X418,'const.min.@_1'(s(fun(X418,bool),X419)))),X422,X421) | ~sP4(X422,happ(s(fun(X418,bool),X419),s(X418,'const.min.@_1'(s(fun(X418,bool),X419)))),X421) | ~sP5(X421,happ(s(fun(X418,bool),X419),s(X418,'const.min.@_1'(s(fun(X418,bool),X419)))),X422)) ) | $spl190),
  inference(cnf_transformation,[],[f37111_D])).
fof(f37111_D,plain,(
  ( ! [X422,X418,X421,X419] : (sP6(X421,X422,happ(s(fun(X418,bool),X419),s(X418,'const.min.@_1'(s(fun(X418,bool),X419))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X418,bool),X419))) | ~sP3(happ(s(fun(X418,bool),X419),s(X418,'const.min.@_1'(s(fun(X418,bool),X419)))),X422,X421) | ~sP4(X422,happ(s(fun(X418,bool),X419),s(X418,'const.min.@_1'(s(fun(X418,bool),X419)))),X421) | ~sP5(X421,happ(s(fun(X418,bool),X419),s(X418,'const.min.@_1'(s(fun(X418,bool),X419)))),X422)) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f36715,plain,(
  ( ! [X420,X422,X418,X421,X419] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X420)))) | sP6(X421,X422,happ(s(fun(X418,bool),X419),s(X418,'const.min.@_1'(s(fun(X418,bool),X419))))) | ~sP5(X421,happ(s(fun(X418,bool),X419),s(X418,'const.min.@_1'(s(fun(X418,bool),X419)))),X422) | ~sP4(X422,happ(s(fun(X418,bool),X419),s(X418,'const.min.@_1'(s(fun(X418,bool),X419)))),X421) | ~sP3(happ(s(fun(X418,bool),X419),s(X418,'const.min.@_1'(s(fun(X418,bool),X419)))),X422,X421) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X420))) = s(bool,happ(s(fun(bool,bool),X420),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X418,bool),X419)))) )),
  inference(superposition,[],[f1783,f2291])).
fof(f37110,plain,(
  ( ! [X416] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.min.@_1'(s(fun('type.extreal.extreal',bool),X416))),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X416)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),X416)))) ) | $spl188),
  inference(cnf_transformation,[],[f37110_D])).
fof(f37110_D,plain,(
  ( ! [X416] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.min.@_1'(s(fun('type.extreal.extreal',bool),X416))),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X416)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),X416)))) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f36714,plain,(
  ( ! [X416,X417] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X417)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.min.@_1'(s(fun('type.extreal.extreal',bool),X416))),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X416)))))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X417))) = s(bool,happ(s(fun(bool,bool),X417),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),X416)))) )),
  inference(superposition,[],[f1840,f2291])).
fof(f37108,plain,(
  ( ! [X405,X406] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X405)))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X405))) = s(bool,happ(s(fun(bool,bool),X405),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X406)) )),
  inference(global_subsumption,[],[f262,f261,f260,f263,f270,f277,f276,f321,f318,f335,f334,f470,f353,f350,f349,f348,f366,f373,f372,f371,f455,f369,f368,f384,f383,f382,f381,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f396,f395,f394,f393,f400,f399,f398,f397,f454,f404,f403,f402,f401,f411,f410,f409,f408,f417,f414,f413,f412,f424,f423,f422,f453,f420,f419,f418,f428,f427,f426,f425,f432,f431,f430,f429,f441,f440,f439,f452,f437,f436,f435,f445,f444,f443,f442,f449,f448,f447,f446,f471,f474,f477,f480,f482,f484,f487,f489,f491,f494,f497,f500,f503,f505,f507,f510,f513,f515,f517,f520,f522,f525,f565,f575,f578,f586,f602,f613,f622,f639,f644,f645,f647,f654,f658,f659,f662,f674,f688,f690,f698,f710,f783,f846,f878,f876,f904,f901,f947,f972,f987,f995,f1008,f1012,f1026,f1041,f1111,f1126,f1142,f1201,f1218,f1233,f1313,f1361,f1767,f1766,f1765,f1764,f2054,f2140,f2204,f2553,f2557,f2561,f2564,f2567,f2570,f2573,f2576,f2579,f2582,f2585,f2588,f2591,f2594,f2597,f2600,f2603,f2605,f2606,f2607,f2608,f2609,f2610,f2611,f2612,f2613,f2614,f2615,f2616,f2617,f2618,f2619,f2620,f2621,f2622,f2623,f2624,f2625,f2296,f2643,f2645,f2648,f2650,f2652,f2654,f2656,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2636,f328,f1840,f2064,f2039,f1797,f1796,f1795,f1794,f1793,f1792,f1791,f1790,f1789,f1788,f1787,f1786,f1785,f1784,f1783,f1782,f1781,f1780,f1779,f2635,f2634,f1296,f367,f325,f324,f2291,f2295,f2294,f2293,f2292,f2139,f2138,f2137,f2136,f2135,f327,f1866,f1867,f317,f326,f2626,f2637,f1831,f1817,f303,f2633,f2632,f2631,f2630,f2629,f2628,f2627,f469,f316,f314,f312,f310,f305,f304,f2134,f2290,f1976,f1975,f1959,f1947,f1923,f1898,f1769,f1768,f468,f281,f280,f3740,f3742,f3744,f3757,f3760,f3762,f3764,f3849,f3851,f3853,f3862,f3865,f3867,f3869,f3882,f3885,f3886,f3897,f4303,f4308,f4310,f4312,f4317,f4319,f4321,f4722,f4724,f4726,f4731,f4733,f4735,f4747,f4752,f4754,f4756,f4807,f4809,f4811,f4821,f4823,f4825,f4838,f4840,f4842,f4860,f4862,f4864,f4869,f4870,f4871,f4872,f4873,f4875,f4880,f4881,f4882,f4883,f4884,f4885,f4888,f4889,f4892,f4893,f4894,f4907,f4909,f4965,f4966,f4967,f4969,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f5004,f5006,f5007,f6917,f6931,f6948,f6967,f6996,f7007,f7009,f7045,f7069,f7071,f7090,f7091,f7092,f7093,f7094,f7124,f7130,f7147,f7153,f7173,f7179,f7193,f7199,f7208,f7209,f7223,f7240,f7242,f7244,f7261,f7263,f7265,f7288,f9309,f9324,f9325,f9417,f9398,f9399,f9484,f9465,f9466,f9489,f9490,f18495,f18500,f18505,f18466,f18467,f19646,f19651,f19656,f19617,f19618,f19906,f19911,f19916,f19878,f19879,f21428,f21429,f21430,f21457,f21458,f21459,f21931,f22080,f22041,f23768,f23773,f23780,f23727,f23728,f24311,f24316,f24323,f24270,f24271,f24460,f24461,f24462,f24450,f24463,f24464,f24465,f24466,f24467,f24468,f24469,f24470,f24471,f24472,f24473,f24474,f24475,f24476,f24477,f24478,f24479,f24480,f24481,f24482,f24483,f24484,f24485,f24486,f24487,f24488,f24489,f24506,f24507,f24508,f24523,f24538,f24539,f24540,f24555,f24596,f24613,f24585,f24621,f24648,f24649,f24650,f24672,f24673,f24674,f25900,f25921,f25973,f25992,f25993,f25994,f26033,f26034,f26035,f26075,f26076,f26077,f26088,f26089,f26090,f26105,f27929,f27930,f27931,f27946,f27968,f27969,f27970,f27991,f27992,f27993,f28005,f28006,f28007,f28031,f28032,f28033,f31391,f28156,f28157,f28226,f28227,f31209,f28261,f28262,f28272,f28300,f28301,f28557,f28559,f28560,f28561,f28562,f28563,f28564,f28565,f28566,f28567,f28568,f28569,f28570,f28571,f28572,f28573,f28574,f28575,f28576,f28577,f28578,f28579,f28580,f28581,f28582,f28583,f28584,f28585,f28586,f28587,f28596,f28597,f28598,f28599,f28600,f28601,f28602,f28603,f28604,f31756,f31760,f31766,f31768,f31819,f31820,f31823,f31824,f31826,f31828,f31830,f31832,f31834,f31837,f31838,f31841,f31842,f28949,f28951,f28952,f28953,f28954,f28955,f28956,f28957,f28958,f28959,f28960,f28961,f28962,f28963,f28964,f28965,f28966,f28967,f28968,f28969,f28970,f28971,f28972,f28973,f28974,f28975,f28976,f28977,f28978,f28979,f28988,f28989,f28990,f28991,f28992,f28993,f28994,f28995,f28996,f31844,f31859,f31860,f31861,f31862,f31863,f31864,f31865,f31866,f31867,f31868,f31869,f31870,f31871,f31872,f29145,f29152,f29155,f29156,f29157,f29158,f29159,f29160,f29161,f29162,f29163,f29164,f29165,f29166,f29167,f29168,f29169,f29170,f29171,f29172,f29173,f29174,f29175,f29184,f29185,f29186,f29187,f29195,f31922,f31923,f31924,f31925,f31926,f31927,f31928,f31929,f31930,f31931,f31932,f31933,f31934,f31935,f31936,f31937,f31938,f31939,f31940,f31941,f31942,f31943,f31944,f31945,f31946,f31947,f31948,f31952,f31953,f31954,f31955,f31956,f31957,f31958,f31959,f31960,f31961,f31962,f31963,f31964,f31965,f31966,f31967,f31968,f31969,f31970,f31971,f31972,f31973,f31974,f31975,f31976,f31977,f31978,f29393,f29395,f29396,f29397,f29398,f29399,f29400,f29401,f29402,f29403,f29404,f29405,f29406,f29407,f29408,f29409,f29410,f29411,f29412,f29413,f29414,f29415,f29416,f29417,f29418,f29419,f29420,f29421,f29422,f29423,f29432,f29433,f29434,f29435,f29436,f29437,f29438,f29439,f29440,f31989,f31990,f32041,f32044,f32045,f32046,f32047,f32048,f32049,f32052,f32055,f32125,f30120,f32469,f32470,f32471,f32472,f32473,f32474,f32475,f32476,f32477,f32478,f32479,f32480,f32481,f32482,f32483,f32484,f32485,f32486,f32487,f32488,f32489,f32490,f32491,f32492,f32493,f32494,f32495,f32496,f32497,f32498,f32499,f32500,f32501,f32502,f32503,f32504,f32505,f32506,f32507,f32508,f32509,f32510,f32511,f32512,f32513,f32514,f32515,f32516,f32517,f32518,f32519,f32520,f32521,f32522,f32523,f32524,f32525,f32526,f32527,f32528,f32529,f32691,f32692,f32693,f32696,f32698,f32699,f32700,f32732,f32735,f32740,f32824,f32825,f32826,f32829,f32831,f32832,f32833,f32865,f32868,f32873,f32932,f33004,f33066,f33132,f33183,f33335,f33469,f33623,f33852,f33850,f33916,f33923,f33951,f33954,f33959,f33989,f33992,f33997,f34045,f34048,f34053,f34210,f34211,f34212,f34215,f34217,f34218,f34219,f34272,f34273,f34274,f34277,f34279,f34280,f34281,f34411,f34412,f34413,f34416,f34418,f34419,f34420,f34493,f34494,f34495,f34498,f34500,f34501,f34502,f34520,f34521,f34522,f34525,f34527,f34528,f34529,f34549,f34550,f34551,f34554,f34556,f34557,f34558,f34578,f34579,f34580,f34583,f34585,f34586,f34587,f34664,f34665,f34666,f34669,f34671,f34672,f34673,f35269,f34731,f34762,f34763,f34766,f34767,f34769,f34770,f34771,f34924,f34925,f34930,f34945,f34980,f35440,f35443,f35444,f35445,f35446,f35447,f35448,f35449,f35450,f35451,f35452,f35453,f35454,f35455,f35456,f35457,f35458,f35459,f35460,f36038,f36032,f36018,f35621,f35789,f36202,f36204,f35923,f36266,f36267,f36287,f36290,f36291,f36292,f36296,f36321,f36322,f36323,f36324,f36325,f36327,f36332,f36333,f36334,f36335,f36336,f36337,f36341,f36345,f36346,f36362,f36365,f36366,f36367,f36368,f36369,f36370,f36371,f36372,f36373,f36374,f36375,f36376,f36377,f36378,f36445,f36466,f36469,f36470,f36471,f36475,f36476,f37031,f37026,f37012,f37010,f36505,f36506,f36542,f36543,f36544,f36545,f36546,f36548,f36553,f36554,f36555,f36556,f36557,f36558,f36562,f36566,f36567,f36583,f36584,f36587,f36588,f36589,f36590,f36591,f36592,f36593,f36594,f36595,f36596,f36597,f36598,f36599,f36600,f37057,f36989,f37060,f37062,f36618,f37063,f36619,f37064,f36620,f36621,f36622,f37067,f36623,f36624,f36625,f36626,f36627,f36628,f36629,f36630,f36631,f36632,f36633,f36634,f36635,f37080,f36636,f37081,f36968,f36966,f36660,f36661,f36662,f36663,f36664,f36963,f36961,f36959,f36957,f37087,f37088,f36680,f36684,f36685,f36686,f36698,f37096,f36699,f37098,f36700,f37100,f36701,f37101,f36702,f37102,f36703,f37103,f36704,f37104,f36705,f37106,f36708])).
fof(f37107,plain,(
  ( ! [X397,X396,X394] : (s(bool,'T_0') = s(bool,sK18(X394,X397,bool)) | p(s(bool,happ(s(fun(bool,bool),X394),s(bool,X396)))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X394),s(bool,'F_0')))) ) | $spl186),
  inference(cnf_transformation,[],[f37107_D])).
fof(f37107_D,plain,(
  ( ! [X397,X396,X394] : (s(bool,'T_0') = s(bool,sK18(X394,X397,bool)) | p(s(bool,happ(s(fun(bool,bool),X394),s(bool,X396)))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X394),s(bool,'F_0')))) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f37106,plain,(
  ( ! [X397,X395,X396,X394] : (s(bool,'T_0') = s(bool,sK18(X394,X397,bool)) | ~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X395)))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X394),s(bool,'F_0'))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X395))) = s(bool,happ(s(fun(bool,bool),X395),s(bool,'T_0'))) | p(s(bool,happ(s(fun(bool,bool),X394),s(bool,X396))))) )),
  inference(global_subsumption,[],[f262,f261,f260,f263,f270,f277,f276,f321,f318,f335,f334,f470,f353,f350,f349,f348,f366,f373,f372,f371,f455,f369,f368,f384,f383,f382,f381,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f396,f395,f394,f393,f400,f399,f398,f397,f454,f404,f403,f402,f401,f411,f410,f409,f408,f417,f414,f413,f412,f424,f423,f422,f453,f420,f419,f418,f428,f427,f426,f425,f432,f431,f430,f429,f441,f440,f439,f452,f437,f436,f435,f445,f444,f443,f442,f449,f448,f447,f446,f471,f474,f477,f480,f482,f484,f487,f489,f491,f494,f497,f500,f503,f505,f507,f510,f513,f515,f517,f520,f522,f525,f565,f575,f578,f586,f602,f613,f622,f639,f644,f645,f647,f654,f658,f659,f662,f674,f688,f690,f698,f710,f783,f846,f878,f876,f904,f901,f947,f972,f987,f995,f1008,f1012,f1026,f1041,f1111,f1126,f1142,f1201,f1218,f1233,f1313,f1361,f1767,f1766,f1765,f1764,f2054,f2140,f2204,f2553,f2557,f2561,f2564,f2567,f2570,f2573,f2576,f2579,f2582,f2585,f2588,f2591,f2594,f2597,f2600,f2603,f2605,f2606,f2607,f2608,f2609,f2610,f2611,f2612,f2613,f2614,f2615,f2616,f2617,f2618,f2619,f2620,f2621,f2622,f2623,f2624,f2625,f2296,f2643,f2645,f2648,f2650,f2652,f2654,f2656,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2636,f328,f1840,f2064,f2039,f1797,f1796,f1795,f1794,f1793,f1792,f1791,f1790,f1789,f1788,f1787,f1786,f1785,f1784,f1783,f1782,f1781,f1780,f1779,f2635,f2634,f1296,f367,f325,f324,f2291,f2295,f2294,f2293,f2292,f2139,f2138,f2137,f2136,f2135,f327,f1866,f1867,f317,f326,f2626,f2637,f1831,f1817,f303,f2633,f2632,f2631,f2630,f2629,f2628,f2627,f469,f316,f314,f312,f310,f305,f304,f2134,f2290,f1976,f1975,f1959,f1947,f1923,f1898,f1769,f1768,f468,f281,f280,f3740,f3742,f3744,f3757,f3760,f3762,f3764,f3849,f3851,f3853,f3862,f3865,f3867,f3869,f3882,f3885,f3886,f3897,f4303,f4308,f4310,f4312,f4317,f4319,f4321,f4722,f4724,f4726,f4731,f4733,f4735,f4747,f4752,f4754,f4756,f4807,f4809,f4811,f4821,f4823,f4825,f4838,f4840,f4842,f4860,f4862,f4864,f4869,f4870,f4871,f4872,f4873,f4875,f4880,f4881,f4882,f4883,f4884,f4885,f4888,f4889,f4892,f4893,f4894,f4907,f4909,f4965,f4966,f4967,f4969,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f5004,f5006,f5007,f6917,f6931,f6948,f6967,f6996,f7007,f7009,f7045,f7069,f7071,f7090,f7091,f7092,f7093,f7094,f7124,f7130,f7147,f7153,f7173,f7179,f7193,f7199,f7208,f7209,f7223,f7240,f7242,f7244,f7261,f7263,f7265,f7288,f9309,f9324,f9325,f9417,f9398,f9399,f9484,f9465,f9466,f9489,f9490,f18495,f18500,f18505,f18466,f18467,f19646,f19651,f19656,f19617,f19618,f19906,f19911,f19916,f19878,f19879,f21428,f21429,f21430,f21457,f21458,f21459,f21931,f22080,f22041,f23768,f23773,f23780,f23727,f23728,f24311,f24316,f24323,f24270,f24271,f24460,f24461,f24462,f24450,f24463,f24464,f24465,f24466,f24467,f24468,f24469,f24470,f24471,f24472,f24473,f24474,f24475,f24476,f24477,f24478,f24479,f24480,f24481,f24482,f24483,f24484,f24485,f24486,f24487,f24488,f24489,f24506,f24507,f24508,f24523,f24538,f24539,f24540,f24555,f24596,f24613,f24585,f24621,f24648,f24649,f24650,f24672,f24673,f24674,f25900,f25921,f25973,f25992,f25993,f25994,f26033,f26034,f26035,f26075,f26076,f26077,f26088,f26089,f26090,f26105,f27929,f27930,f27931,f27946,f27968,f27969,f27970,f27991,f27992,f27993,f28005,f28006,f28007,f28031,f28032,f28033,f31391,f28156,f28157,f28226,f28227,f31209,f28261,f28262,f28272,f28300,f28301,f28557,f28559,f28560,f28561,f28562,f28563,f28564,f28565,f28566,f28567,f28568,f28569,f28570,f28571,f28572,f28573,f28574,f28575,f28576,f28577,f28578,f28579,f28580,f28581,f28582,f28583,f28584,f28585,f28586,f28587,f28596,f28597,f28598,f28599,f28600,f28601,f28602,f28603,f28604,f31756,f31760,f31766,f31768,f31819,f31820,f31823,f31824,f31826,f31828,f31830,f31832,f31834,f31837,f31838,f31841,f31842,f28949,f28951,f28952,f28953,f28954,f28955,f28956,f28957,f28958,f28959,f28960,f28961,f28962,f28963,f28964,f28965,f28966,f28967,f28968,f28969,f28970,f28971,f28972,f28973,f28974,f28975,f28976,f28977,f28978,f28979,f28988,f28989,f28990,f28991,f28992,f28993,f28994,f28995,f28996,f31844,f31859,f31860,f31861,f31862,f31863,f31864,f31865,f31866,f31867,f31868,f31869,f31870,f31871,f31872,f29145,f29152,f29155,f29156,f29157,f29158,f29159,f29160,f29161,f29162,f29163,f29164,f29165,f29166,f29167,f29168,f29169,f29170,f29171,f29172,f29173,f29174,f29175,f29184,f29185,f29186,f29187,f29195,f31922,f31923,f31924,f31925,f31926,f31927,f31928,f31929,f31930,f31931,f31932,f31933,f31934,f31935,f31936,f31937,f31938,f31939,f31940,f31941,f31942,f31943,f31944,f31945,f31946,f31947,f31948,f31952,f31953,f31954,f31955,f31956,f31957,f31958,f31959,f31960,f31961,f31962,f31963,f31964,f31965,f31966,f31967,f31968,f31969,f31970,f31971,f31972,f31973,f31974,f31975,f31976,f31977,f31978,f29393,f29395,f29396,f29397,f29398,f29399,f29400,f29401,f29402,f29403,f29404,f29405,f29406,f29407,f29408,f29409,f29410,f29411,f29412,f29413,f29414,f29415,f29416,f29417,f29418,f29419,f29420,f29421,f29422,f29423,f29432,f29433,f29434,f29435,f29436,f29437,f29438,f29439,f29440,f31989,f31990,f32041,f32044,f32045,f32046,f32047,f32048,f32049,f32052,f32055,f32125,f30120,f32469,f32470,f32471,f32472,f32473,f32474,f32475,f32476,f32477,f32478,f32479,f32480,f32481,f32482,f32483,f32484,f32485,f32486,f32487,f32488,f32489,f32490,f32491,f32492,f32493,f32494,f32495,f32496,f32497,f32498,f32499,f32500,f32501,f32502,f32503,f32504,f32505,f32506,f32507,f32508,f32509,f32510,f32511,f32512,f32513,f32514,f32515,f32516,f32517,f32518,f32519,f32520,f32521,f32522,f32523,f32524,f32525,f32526,f32527,f32528,f32529,f32691,f32692,f32693,f32696,f32698,f32699,f32700,f32732,f32735,f32740,f32824,f32825,f32826,f32829,f32831,f32832,f32833,f32865,f32868,f32873,f32932,f33004,f33066,f33132,f33183,f33335,f33469,f33623,f33852,f33850,f33916,f33923,f33951,f33954,f33959,f33989,f33992,f33997,f34045,f34048,f34053,f34210,f34211,f34212,f34215,f34217,f34218,f34219,f34272,f34273,f34274,f34277,f34279,f34280,f34281,f34411,f34412,f34413,f34416,f34418,f34419,f34420,f34493,f34494,f34495,f34498,f34500,f34501,f34502,f34520,f34521,f34522,f34525,f34527,f34528,f34529,f34549,f34550,f34551,f34554,f34556,f34557,f34558,f34578,f34579,f34580,f34583,f34585,f34586,f34587,f34664,f34665,f34666,f34669,f34671,f34672,f34673,f35269,f34731,f34762,f34763,f34766,f34767,f34769,f34770,f34771,f34924,f34925,f34930,f34945,f34980,f35440,f35443,f35444,f35445,f35446,f35447,f35448,f35449,f35450,f35451,f35452,f35453,f35454,f35455,f35456,f35457,f35458,f35459,f35460,f36038,f36032,f36018,f35621,f35789,f36202,f36204,f35923,f36266,f36267,f36287,f36290,f36291,f36292,f36296,f36321,f36322,f36323,f36324,f36325,f36327,f36332,f36333,f36334,f36335,f36336,f36337,f36341,f36345,f36346,f36362,f36365,f36366,f36367,f36368,f36369,f36370,f36371,f36372,f36373,f36374,f36375,f36376,f36377,f36378,f36445,f36466,f36469,f36470,f36471,f36475,f36476,f37031,f37026,f37012,f37010,f36505,f36506,f36542,f36543,f36544,f36545,f36546,f36548,f36553,f36554,f36555,f36556,f36557,f36558,f36562,f36566,f36567,f36583,f36584,f36587,f36588,f36589,f36590,f36591,f36592,f36593,f36594,f36595,f36596,f36597,f36598,f36599,f36600,f37057,f36989,f37060,f37062,f36618,f37063,f36619,f37064,f36620,f36621,f36622,f37067,f36623,f36624,f36625,f36626,f36627,f36628,f36629,f36630,f36631,f36632,f36633,f36634,f36635,f37080,f36636,f37081,f36968,f36966,f36660,f36661,f36662,f36663,f36664,f36963,f36961,f36959,f36957,f37087,f37088,f36680,f36684,f36685,f36686,f36698,f37096,f36699,f37098,f36700,f37100,f36701,f37101,f36702,f37102,f36703,f37103,f36704,f37104,f36705])).
fof(f37105,plain,(
  ( ! [X393,X390,X392] : (s(bool,'T_0') = s(bool,sK19(X390,X393,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X390),s(bool,'F_0'))) | p(s(bool,happ(s(fun(bool,bool),X390),s(bool,X392))))) ) | $spl184),
  inference(cnf_transformation,[],[f37105_D])).
fof(f37105_D,plain,(
  ( ! [X393,X390,X392] : (s(bool,'T_0') = s(bool,sK19(X390,X393,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X390),s(bool,'F_0'))) | p(s(bool,happ(s(fun(bool,bool),X390),s(bool,X392))))) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f37104,plain,(
  ( ! [X391,X393,X390,X392] : (s(bool,'T_0') = s(bool,sK19(X390,X393,bool)) | p(s(bool,happ(s(fun(bool,bool),X390),s(bool,X392)))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X391))) = s(bool,happ(s(fun(bool,bool),X391),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X390),s(bool,'F_0'))) | ~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X391))))) )),
  inference(global_subsumption,[],[f262,f261,f260,f263,f270,f277,f276,f321,f318,f335,f334,f470,f353,f350,f349,f348,f366,f373,f372,f371,f455,f369,f368,f384,f383,f382,f381,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f396,f395,f394,f393,f400,f399,f398,f397,f454,f404,f403,f402,f401,f411,f410,f409,f408,f417,f414,f413,f412,f424,f423,f422,f453,f420,f419,f418,f428,f427,f426,f425,f432,f431,f430,f429,f441,f440,f439,f452,f437,f436,f435,f445,f444,f443,f442,f449,f448,f447,f446,f471,f474,f477,f480,f482,f484,f487,f489,f491,f494,f497,f500,f503,f505,f507,f510,f513,f515,f517,f520,f522,f525,f565,f575,f578,f586,f602,f613,f622,f639,f644,f645,f647,f654,f658,f659,f662,f674,f688,f690,f698,f710,f783,f846,f878,f876,f904,f901,f947,f972,f987,f995,f1008,f1012,f1026,f1041,f1111,f1126,f1142,f1201,f1218,f1233,f1313,f1361,f1767,f1766,f1765,f1764,f2054,f2140,f2204,f2553,f2557,f2561,f2564,f2567,f2570,f2573,f2576,f2579,f2582,f2585,f2588,f2591,f2594,f2597,f2600,f2603,f2605,f2606,f2607,f2608,f2609,f2610,f2611,f2612,f2613,f2614,f2615,f2616,f2617,f2618,f2619,f2620,f2621,f2622,f2623,f2624,f2625,f2296,f2643,f2645,f2648,f2650,f2652,f2654,f2656,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2636,f328,f1840,f2064,f2039,f1797,f1796,f1795,f1794,f1793,f1792,f1791,f1790,f1789,f1788,f1787,f1786,f1785,f1784,f1783,f1782,f1781,f1780,f1779,f2635,f2634,f1296,f367,f325,f324,f2291,f2295,f2294,f2293,f2292,f2139,f2138,f2137,f2136,f2135,f327,f1866,f1867,f317,f326,f2626,f2637,f1831,f1817,f303,f2633,f2632,f2631,f2630,f2629,f2628,f2627,f469,f316,f314,f312,f310,f305,f304,f2134,f2290,f1976,f1975,f1959,f1947,f1923,f1898,f1769,f1768,f468,f281,f280,f3740,f3742,f3744,f3757,f3760,f3762,f3764,f3849,f3851,f3853,f3862,f3865,f3867,f3869,f3882,f3885,f3886,f3897,f4303,f4308,f4310,f4312,f4317,f4319,f4321,f4722,f4724,f4726,f4731,f4733,f4735,f4747,f4752,f4754,f4756,f4807,f4809,f4811,f4821,f4823,f4825,f4838,f4840,f4842,f4860,f4862,f4864,f4869,f4870,f4871,f4872,f4873,f4875,f4880,f4881,f4882,f4883,f4884,f4885,f4888,f4889,f4892,f4893,f4894,f4907,f4909,f4965,f4966,f4967,f4969,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f5004,f5006,f5007,f6917,f6931,f6948,f6967,f6996,f7007,f7009,f7045,f7069,f7071,f7090,f7091,f7092,f7093,f7094,f7124,f7130,f7147,f7153,f7173,f7179,f7193,f7199,f7208,f7209,f7223,f7240,f7242,f7244,f7261,f7263,f7265,f7288,f9309,f9324,f9325,f9417,f9398,f9399,f9484,f9465,f9466,f9489,f9490,f18495,f18500,f18505,f18466,f18467,f19646,f19651,f19656,f19617,f19618,f19906,f19911,f19916,f19878,f19879,f21428,f21429,f21430,f21457,f21458,f21459,f21931,f22080,f22041,f23768,f23773,f23780,f23727,f23728,f24311,f24316,f24323,f24270,f24271,f24460,f24461,f24462,f24450,f24463,f24464,f24465,f24466,f24467,f24468,f24469,f24470,f24471,f24472,f24473,f24474,f24475,f24476,f24477,f24478,f24479,f24480,f24481,f24482,f24483,f24484,f24485,f24486,f24487,f24488,f24489,f24506,f24507,f24508,f24523,f24538,f24539,f24540,f24555,f24596,f24613,f24585,f24621,f24648,f24649,f24650,f24672,f24673,f24674,f25900,f25921,f25973,f25992,f25993,f25994,f26033,f26034,f26035,f26075,f26076,f26077,f26088,f26089,f26090,f26105,f27929,f27930,f27931,f27946,f27968,f27969,f27970,f27991,f27992,f27993,f28005,f28006,f28007,f28031,f28032,f28033,f31391,f28156,f28157,f28226,f28227,f31209,f28261,f28262,f28272,f28300,f28301,f28557,f28559,f28560,f28561,f28562,f28563,f28564,f28565,f28566,f28567,f28568,f28569,f28570,f28571,f28572,f28573,f28574,f28575,f28576,f28577,f28578,f28579,f28580,f28581,f28582,f28583,f28584,f28585,f28586,f28587,f28596,f28597,f28598,f28599,f28600,f28601,f28602,f28603,f28604,f31756,f31760,f31766,f31768,f31819,f31820,f31823,f31824,f31826,f31828,f31830,f31832,f31834,f31837,f31838,f31841,f31842,f28949,f28951,f28952,f28953,f28954,f28955,f28956,f28957,f28958,f28959,f28960,f28961,f28962,f28963,f28964,f28965,f28966,f28967,f28968,f28969,f28970,f28971,f28972,f28973,f28974,f28975,f28976,f28977,f28978,f28979,f28988,f28989,f28990,f28991,f28992,f28993,f28994,f28995,f28996,f31844,f31859,f31860,f31861,f31862,f31863,f31864,f31865,f31866,f31867,f31868,f31869,f31870,f31871,f31872,f29145,f29152,f29155,f29156,f29157,f29158,f29159,f29160,f29161,f29162,f29163,f29164,f29165,f29166,f29167,f29168,f29169,f29170,f29171,f29172,f29173,f29174,f29175,f29184,f29185,f29186,f29187,f29195,f31922,f31923,f31924,f31925,f31926,f31927,f31928,f31929,f31930,f31931,f31932,f31933,f31934,f31935,f31936,f31937,f31938,f31939,f31940,f31941,f31942,f31943,f31944,f31945,f31946,f31947,f31948,f31952,f31953,f31954,f31955,f31956,f31957,f31958,f31959,f31960,f31961,f31962,f31963,f31964,f31965,f31966,f31967,f31968,f31969,f31970,f31971,f31972,f31973,f31974,f31975,f31976,f31977,f31978,f29393,f29395,f29396,f29397,f29398,f29399,f29400,f29401,f29402,f29403,f29404,f29405,f29406,f29407,f29408,f29409,f29410,f29411,f29412,f29413,f29414,f29415,f29416,f29417,f29418,f29419,f29420,f29421,f29422,f29423,f29432,f29433,f29434,f29435,f29436,f29437,f29438,f29439,f29440,f31989,f31990,f32041,f32044,f32045,f32046,f32047,f32048,f32049,f32052,f32055,f32125,f30120,f32469,f32470,f32471,f32472,f32473,f32474,f32475,f32476,f32477,f32478,f32479,f32480,f32481,f32482,f32483,f32484,f32485,f32486,f32487,f32488,f32489,f32490,f32491,f32492,f32493,f32494,f32495,f32496,f32497,f32498,f32499,f32500,f32501,f32502,f32503,f32504,f32505,f32506,f32507,f32508,f32509,f32510,f32511,f32512,f32513,f32514,f32515,f32516,f32517,f32518,f32519,f32520,f32521,f32522,f32523,f32524,f32525,f32526,f32527,f32528,f32529,f32691,f32692,f32693,f32696,f32698,f32699,f32700,f32732,f32735,f32740,f32824,f32825,f32826,f32829,f32831,f32832,f32833,f32865,f32868,f32873,f32932,f33004,f33066,f33132,f33183,f33335,f33469,f33623,f33852,f33850,f33916,f33923,f33951,f33954,f33959,f33989,f33992,f33997,f34045,f34048,f34053,f34210,f34211,f34212,f34215,f34217,f34218,f34219,f34272,f34273,f34274,f34277,f34279,f34280,f34281,f34411,f34412,f34413,f34416,f34418,f34419,f34420,f34493,f34494,f34495,f34498,f34500,f34501,f34502,f34520,f34521,f34522,f34525,f34527,f34528,f34529,f34549,f34550,f34551,f34554,f34556,f34557,f34558,f34578,f34579,f34580,f34583,f34585,f34586,f34587,f34664,f34665,f34666,f34669,f34671,f34672,f34673,f35269,f34731,f34762,f34763,f34766,f34767,f34769,f34770,f34771,f34924,f34925,f34930,f34945,f34980,f35440,f35443,f35444,f35445,f35446,f35447,f35448,f35449,f35450,f35451,f35452,f35453,f35454,f35455,f35456,f35457,f35458,f35459,f35460,f36038,f36032,f36018,f35621,f35789,f36202,f36204,f35923,f36266,f36267,f36287,f36290,f36291,f36292,f36296,f36321,f36322,f36323,f36324,f36325,f36327,f36332,f36333,f36334,f36335,f36336,f36337,f36341,f36345,f36346,f36362,f36365,f36366,f36367,f36368,f36369,f36370,f36371,f36372,f36373,f36374,f36375,f36376,f36377,f36378,f36445,f36466,f36469,f36470,f36471,f36475,f36476,f37031,f37026,f37012,f37010,f36505,f36506,f36542,f36543,f36544,f36545,f36546,f36548,f36553,f36554,f36555,f36556,f36557,f36558,f36562,f36566,f36567,f36583,f36584,f36587,f36588,f36589,f36590,f36591,f36592,f36593,f36594,f36595,f36596,f36597,f36598,f36599,f36600,f37057,f36989,f37060,f37062,f36618,f37063,f36619,f37064,f36620,f36621,f36622,f37067,f36623,f36624,f36625,f36626,f36627,f36628,f36629,f36630,f36631,f36632,f36633,f36634,f36635,f37080,f36636,f37081,f36968,f36966,f36660,f36661,f36662,f36663,f36664,f36963,f36961,f36959,f36957,f37087,f37088,f36680,f36684,f36685,f36686,f36698,f37096,f36699,f37098,f36700,f37100,f36701,f37101,f36702,f37102,f36703,f37103,f36704])).
fof(f37103,plain,(
  ( ! [X387,X388] : (p(s(bool,X388)) | ~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X387)))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X387))) = s(bool,happ(s(fun(bool,bool),X387),s(bool,'T_0')))) )),
  inference(global_subsumption,[],[f262,f261,f260,f263,f270,f277,f276,f321,f318,f335,f334,f470,f353,f350,f349,f348,f366,f373,f372,f371,f455,f369,f368,f384,f383,f382,f381,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f396,f395,f394,f393,f400,f399,f398,f397,f454,f404,f403,f402,f401,f411,f410,f409,f408,f417,f414,f413,f412,f424,f423,f422,f453,f420,f419,f418,f428,f427,f426,f425,f432,f431,f430,f429,f441,f440,f439,f452,f437,f436,f435,f445,f444,f443,f442,f449,f448,f447,f446,f471,f474,f477,f480,f482,f484,f487,f489,f491,f494,f497,f500,f503,f505,f507,f510,f513,f515,f517,f520,f522,f525,f565,f575,f578,f586,f602,f613,f622,f639,f644,f645,f647,f654,f658,f659,f662,f674,f688,f690,f698,f710,f783,f846,f878,f876,f904,f901,f947,f972,f987,f995,f1008,f1012,f1026,f1041,f1111,f1126,f1142,f1201,f1218,f1233,f1313,f1361,f1767,f1766,f1765,f1764,f2054,f2140,f2204,f2553,f2557,f2561,f2564,f2567,f2570,f2573,f2576,f2579,f2582,f2585,f2588,f2591,f2594,f2597,f2600,f2603,f2605,f2606,f2607,f2608,f2609,f2610,f2611,f2612,f2613,f2614,f2615,f2616,f2617,f2618,f2619,f2620,f2621,f2622,f2623,f2624,f2625,f2296,f2643,f2645,f2648,f2650,f2652,f2654,f2656,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2636,f328,f1840,f2064,f2039,f1797,f1796,f1795,f1794,f1793,f1792,f1791,f1790,f1789,f1788,f1787,f1786,f1785,f1784,f1783,f1782,f1781,f1780,f1779,f2635,f2634,f1296,f367,f325,f324,f2291,f2295,f2294,f2293,f2292,f2139,f2138,f2137,f2136,f2135,f327,f1866,f1867,f317,f326,f2626,f2637,f1831,f1817,f303,f2633,f2632,f2631,f2630,f2629,f2628,f2627,f469,f316,f314,f312,f310,f305,f304,f2134,f2290,f1976,f1975,f1959,f1947,f1923,f1898,f1769,f1768,f468,f281,f280,f3740,f3742,f3744,f3757,f3760,f3762,f3764,f3849,f3851,f3853,f3862,f3865,f3867,f3869,f3882,f3885,f3886,f3897,f4303,f4308,f4310,f4312,f4317,f4319,f4321,f4722,f4724,f4726,f4731,f4733,f4735,f4747,f4752,f4754,f4756,f4807,f4809,f4811,f4821,f4823,f4825,f4838,f4840,f4842,f4860,f4862,f4864,f4869,f4870,f4871,f4872,f4873,f4875,f4880,f4881,f4882,f4883,f4884,f4885,f4888,f4889,f4892,f4893,f4894,f4907,f4909,f4965,f4966,f4967,f4969,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f5004,f5006,f5007,f6917,f6931,f6948,f6967,f6996,f7007,f7009,f7045,f7069,f7071,f7090,f7091,f7092,f7093,f7094,f7124,f7130,f7147,f7153,f7173,f7179,f7193,f7199,f7208,f7209,f7223,f7240,f7242,f7244,f7261,f7263,f7265,f7288,f9309,f9324,f9325,f9417,f9398,f9399,f9484,f9465,f9466,f9489,f9490,f18495,f18500,f18505,f18466,f18467,f19646,f19651,f19656,f19617,f19618,f19906,f19911,f19916,f19878,f19879,f21428,f21429,f21430,f21457,f21458,f21459,f21931,f22080,f22041,f23768,f23773,f23780,f23727,f23728,f24311,f24316,f24323,f24270,f24271,f24460,f24461,f24462,f24450,f24463,f24464,f24465,f24466,f24467,f24468,f24469,f24470,f24471,f24472,f24473,f24474,f24475,f24476,f24477,f24478,f24479,f24480,f24481,f24482,f24483,f24484,f24485,f24486,f24487,f24488,f24489,f24506,f24507,f24508,f24523,f24538,f24539,f24540,f24555,f24596,f24613,f24585,f24621,f24648,f24649,f24650,f24672,f24673,f24674,f25900,f25921,f25973,f25992,f25993,f25994,f26033,f26034,f26035,f26075,f26076,f26077,f26088,f26089,f26090,f26105,f27929,f27930,f27931,f27946,f27968,f27969,f27970,f27991,f27992,f27993,f28005,f28006,f28007,f28031,f28032,f28033,f31391,f28156,f28157,f28226,f28227,f31209,f28261,f28262,f28272,f28300,f28301,f28557,f28559,f28560,f28561,f28562,f28563,f28564,f28565,f28566,f28567,f28568,f28569,f28570,f28571,f28572,f28573,f28574,f28575,f28576,f28577,f28578,f28579,f28580,f28581,f28582,f28583,f28584,f28585,f28586,f28587,f28596,f28597,f28598,f28599,f28600,f28601,f28602,f28603,f28604,f31756,f31760,f31766,f31768,f31819,f31820,f31823,f31824,f31826,f31828,f31830,f31832,f31834,f31837,f31838,f31841,f31842,f28949,f28951,f28952,f28953,f28954,f28955,f28956,f28957,f28958,f28959,f28960,f28961,f28962,f28963,f28964,f28965,f28966,f28967,f28968,f28969,f28970,f28971,f28972,f28973,f28974,f28975,f28976,f28977,f28978,f28979,f28988,f28989,f28990,f28991,f28992,f28993,f28994,f28995,f28996,f31844,f31859,f31860,f31861,f31862,f31863,f31864,f31865,f31866,f31867,f31868,f31869,f31870,f31871,f31872,f29145,f29152,f29155,f29156,f29157,f29158,f29159,f29160,f29161,f29162,f29163,f29164,f29165,f29166,f29167,f29168,f29169,f29170,f29171,f29172,f29173,f29174,f29175,f29184,f29185,f29186,f29187,f29195,f31922,f31923,f31924,f31925,f31926,f31927,f31928,f31929,f31930,f31931,f31932,f31933,f31934,f31935,f31936,f31937,f31938,f31939,f31940,f31941,f31942,f31943,f31944,f31945,f31946,f31947,f31948,f31952,f31953,f31954,f31955,f31956,f31957,f31958,f31959,f31960,f31961,f31962,f31963,f31964,f31965,f31966,f31967,f31968,f31969,f31970,f31971,f31972,f31973,f31974,f31975,f31976,f31977,f31978,f29393,f29395,f29396,f29397,f29398,f29399,f29400,f29401,f29402,f29403,f29404,f29405,f29406,f29407,f29408,f29409,f29410,f29411,f29412,f29413,f29414,f29415,f29416,f29417,f29418,f29419,f29420,f29421,f29422,f29423,f29432,f29433,f29434,f29435,f29436,f29437,f29438,f29439,f29440,f31989,f31990,f32041,f32044,f32045,f32046,f32047,f32048,f32049,f32052,f32055,f32125,f30120,f32469,f32470,f32471,f32472,f32473,f32474,f32475,f32476,f32477,f32478,f32479,f32480,f32481,f32482,f32483,f32484,f32485,f32486,f32487,f32488,f32489,f32490,f32491,f32492,f32493,f32494,f32495,f32496,f32497,f32498,f32499,f32500,f32501,f32502,f32503,f32504,f32505,f32506,f32507,f32508,f32509,f32510,f32511,f32512,f32513,f32514,f32515,f32516,f32517,f32518,f32519,f32520,f32521,f32522,f32523,f32524,f32525,f32526,f32527,f32528,f32529,f32691,f32692,f32693,f32696,f32698,f32699,f32700,f32732,f32735,f32740,f32824,f32825,f32826,f32829,f32831,f32832,f32833,f32865,f32868,f32873,f32932,f33004,f33066,f33132,f33183,f33335,f33469,f33623,f33852,f33850,f33916,f33923,f33951,f33954,f33959,f33989,f33992,f33997,f34045,f34048,f34053,f34210,f34211,f34212,f34215,f34217,f34218,f34219,f34272,f34273,f34274,f34277,f34279,f34280,f34281,f34411,f34412,f34413,f34416,f34418,f34419,f34420,f34493,f34494,f34495,f34498,f34500,f34501,f34502,f34520,f34521,f34522,f34525,f34527,f34528,f34529,f34549,f34550,f34551,f34554,f34556,f34557,f34558,f34578,f34579,f34580,f34583,f34585,f34586,f34587,f34664,f34665,f34666,f34669,f34671,f34672,f34673,f35269,f34731,f34762,f34763,f34766,f34767,f34769,f34770,f34771,f34924,f34925,f34930,f34945,f34980,f35440,f35443,f35444,f35445,f35446,f35447,f35448,f35449,f35450,f35451,f35452,f35453,f35454,f35455,f35456,f35457,f35458,f35459,f35460,f36038,f36032,f36018,f35621,f35789,f36202,f36204,f35923,f36266,f36267,f36287,f36290,f36291,f36292,f36296,f36321,f36322,f36323,f36324,f36325,f36327,f36332,f36333,f36334,f36335,f36336,f36337,f36341,f36345,f36346,f36362,f36365,f36366,f36367,f36368,f36369,f36370,f36371,f36372,f36373,f36374,f36375,f36376,f36377,f36378,f36445,f36466,f36469,f36470,f36471,f36475,f36476,f37031,f37026,f37012,f37010,f36505,f36506,f36542,f36543,f36544,f36545,f36546,f36548,f36553,f36554,f36555,f36556,f36557,f36558,f36562,f36566,f36567,f36583,f36584,f36587,f36588,f36589,f36590,f36591,f36592,f36593,f36594,f36595,f36596,f36597,f36598,f36599,f36600,f37057,f36989,f37060,f37062,f36618,f37063,f36619,f37064,f36620,f36621,f36622,f37067,f36623,f36624,f36625,f36626,f36627,f36628,f36629,f36630,f36631,f36632,f36633,f36634,f36635,f37080,f36636,f37081,f36968,f36966,f36660,f36661,f36662,f36663,f36664,f36963,f36961,f36959,f36957,f37087,f37088,f36680,f36684,f36685,f36686,f36698,f37096,f36699,f37098,f36700,f37100,f36701,f37101,f36702,f37102,f36703])).
fof(f37102,plain,(
  ( ! [X383,X384] : (p(s(bool,X384)) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X383))) = s(bool,happ(s(fun(bool,bool),X383),s(bool,'T_0'))) | ~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X383))))) )),
  inference(global_subsumption,[],[f262,f261,f260,f263,f270,f277,f276,f321,f318,f335,f334,f470,f353,f350,f349,f348,f366,f373,f372,f371,f455,f369,f368,f384,f383,f382,f381,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f396,f395,f394,f393,f400,f399,f398,f397,f454,f404,f403,f402,f401,f411,f410,f409,f408,f417,f414,f413,f412,f424,f423,f422,f453,f420,f419,f418,f428,f427,f426,f425,f432,f431,f430,f429,f441,f440,f439,f452,f437,f436,f435,f445,f444,f443,f442,f449,f448,f447,f446,f471,f474,f477,f480,f482,f484,f487,f489,f491,f494,f497,f500,f503,f505,f507,f510,f513,f515,f517,f520,f522,f525,f565,f575,f578,f586,f602,f613,f622,f639,f644,f645,f647,f654,f658,f659,f662,f674,f688,f690,f698,f710,f783,f846,f878,f876,f904,f901,f947,f972,f987,f995,f1008,f1012,f1026,f1041,f1111,f1126,f1142,f1201,f1218,f1233,f1313,f1361,f1767,f1766,f1765,f1764,f2054,f2140,f2204,f2553,f2557,f2561,f2564,f2567,f2570,f2573,f2576,f2579,f2582,f2585,f2588,f2591,f2594,f2597,f2600,f2603,f2605,f2606,f2607,f2608,f2609,f2610,f2611,f2612,f2613,f2614,f2615,f2616,f2617,f2618,f2619,f2620,f2621,f2622,f2623,f2624,f2625,f2296,f2643,f2645,f2648,f2650,f2652,f2654,f2656,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2636,f328,f1840,f2064,f2039,f1797,f1796,f1795,f1794,f1793,f1792,f1791,f1790,f1789,f1788,f1787,f1786,f1785,f1784,f1783,f1782,f1781,f1780,f1779,f2635,f2634,f1296,f367,f325,f324,f2291,f2295,f2294,f2293,f2292,f2139,f2138,f2137,f2136,f2135,f327,f1866,f1867,f317,f326,f2626,f2637,f1831,f1817,f303,f2633,f2632,f2631,f2630,f2629,f2628,f2627,f469,f316,f314,f312,f310,f305,f304,f2134,f2290,f1976,f1975,f1959,f1947,f1923,f1898,f1769,f1768,f468,f281,f280,f3740,f3742,f3744,f3757,f3760,f3762,f3764,f3849,f3851,f3853,f3862,f3865,f3867,f3869,f3882,f3885,f3886,f3897,f4303,f4308,f4310,f4312,f4317,f4319,f4321,f4722,f4724,f4726,f4731,f4733,f4735,f4747,f4752,f4754,f4756,f4807,f4809,f4811,f4821,f4823,f4825,f4838,f4840,f4842,f4860,f4862,f4864,f4869,f4870,f4871,f4872,f4873,f4875,f4880,f4881,f4882,f4883,f4884,f4885,f4888,f4889,f4892,f4893,f4894,f4907,f4909,f4965,f4966,f4967,f4969,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f5004,f5006,f5007,f6917,f6931,f6948,f6967,f6996,f7007,f7009,f7045,f7069,f7071,f7090,f7091,f7092,f7093,f7094,f7124,f7130,f7147,f7153,f7173,f7179,f7193,f7199,f7208,f7209,f7223,f7240,f7242,f7244,f7261,f7263,f7265,f7288,f9309,f9324,f9325,f9417,f9398,f9399,f9484,f9465,f9466,f9489,f9490,f18495,f18500,f18505,f18466,f18467,f19646,f19651,f19656,f19617,f19618,f19906,f19911,f19916,f19878,f19879,f21428,f21429,f21430,f21457,f21458,f21459,f21931,f22080,f22041,f23768,f23773,f23780,f23727,f23728,f24311,f24316,f24323,f24270,f24271,f24460,f24461,f24462,f24450,f24463,f24464,f24465,f24466,f24467,f24468,f24469,f24470,f24471,f24472,f24473,f24474,f24475,f24476,f24477,f24478,f24479,f24480,f24481,f24482,f24483,f24484,f24485,f24486,f24487,f24488,f24489,f24506,f24507,f24508,f24523,f24538,f24539,f24540,f24555,f24596,f24613,f24585,f24621,f24648,f24649,f24650,f24672,f24673,f24674,f25900,f25921,f25973,f25992,f25993,f25994,f26033,f26034,f26035,f26075,f26076,f26077,f26088,f26089,f26090,f26105,f27929,f27930,f27931,f27946,f27968,f27969,f27970,f27991,f27992,f27993,f28005,f28006,f28007,f28031,f28032,f28033,f31391,f28156,f28157,f28226,f28227,f31209,f28261,f28262,f28272,f28300,f28301,f28557,f28559,f28560,f28561,f28562,f28563,f28564,f28565,f28566,f28567,f28568,f28569,f28570,f28571,f28572,f28573,f28574,f28575,f28576,f28577,f28578,f28579,f28580,f28581,f28582,f28583,f28584,f28585,f28586,f28587,f28596,f28597,f28598,f28599,f28600,f28601,f28602,f28603,f28604,f31756,f31760,f31766,f31768,f31819,f31820,f31823,f31824,f31826,f31828,f31830,f31832,f31834,f31837,f31838,f31841,f31842,f28949,f28951,f28952,f28953,f28954,f28955,f28956,f28957,f28958,f28959,f28960,f28961,f28962,f28963,f28964,f28965,f28966,f28967,f28968,f28969,f28970,f28971,f28972,f28973,f28974,f28975,f28976,f28977,f28978,f28979,f28988,f28989,f28990,f28991,f28992,f28993,f28994,f28995,f28996,f31844,f31859,f31860,f31861,f31862,f31863,f31864,f31865,f31866,f31867,f31868,f31869,f31870,f31871,f31872,f29145,f29152,f29155,f29156,f29157,f29158,f29159,f29160,f29161,f29162,f29163,f29164,f29165,f29166,f29167,f29168,f29169,f29170,f29171,f29172,f29173,f29174,f29175,f29184,f29185,f29186,f29187,f29195,f31922,f31923,f31924,f31925,f31926,f31927,f31928,f31929,f31930,f31931,f31932,f31933,f31934,f31935,f31936,f31937,f31938,f31939,f31940,f31941,f31942,f31943,f31944,f31945,f31946,f31947,f31948,f31952,f31953,f31954,f31955,f31956,f31957,f31958,f31959,f31960,f31961,f31962,f31963,f31964,f31965,f31966,f31967,f31968,f31969,f31970,f31971,f31972,f31973,f31974,f31975,f31976,f31977,f31978,f29393,f29395,f29396,f29397,f29398,f29399,f29400,f29401,f29402,f29403,f29404,f29405,f29406,f29407,f29408,f29409,f29410,f29411,f29412,f29413,f29414,f29415,f29416,f29417,f29418,f29419,f29420,f29421,f29422,f29423,f29432,f29433,f29434,f29435,f29436,f29437,f29438,f29439,f29440,f31989,f31990,f32041,f32044,f32045,f32046,f32047,f32048,f32049,f32052,f32055,f32125,f30120,f32469,f32470,f32471,f32472,f32473,f32474,f32475,f32476,f32477,f32478,f32479,f32480,f32481,f32482,f32483,f32484,f32485,f32486,f32487,f32488,f32489,f32490,f32491,f32492,f32493,f32494,f32495,f32496,f32497,f32498,f32499,f32500,f32501,f32502,f32503,f32504,f32505,f32506,f32507,f32508,f32509,f32510,f32511,f32512,f32513,f32514,f32515,f32516,f32517,f32518,f32519,f32520,f32521,f32522,f32523,f32524,f32525,f32526,f32527,f32528,f32529,f32691,f32692,f32693,f32696,f32698,f32699,f32700,f32732,f32735,f32740,f32824,f32825,f32826,f32829,f32831,f32832,f32833,f32865,f32868,f32873,f32932,f33004,f33066,f33132,f33183,f33335,f33469,f33623,f33852,f33850,f33916,f33923,f33951,f33954,f33959,f33989,f33992,f33997,f34045,f34048,f34053,f34210,f34211,f34212,f34215,f34217,f34218,f34219,f34272,f34273,f34274,f34277,f34279,f34280,f34281,f34411,f34412,f34413,f34416,f34418,f34419,f34420,f34493,f34494,f34495,f34498,f34500,f34501,f34502,f34520,f34521,f34522,f34525,f34527,f34528,f34529,f34549,f34550,f34551,f34554,f34556,f34557,f34558,f34578,f34579,f34580,f34583,f34585,f34586,f34587,f34664,f34665,f34666,f34669,f34671,f34672,f34673,f35269,f34731,f34762,f34763,f34766,f34767,f34769,f34770,f34771,f34924,f34925,f34930,f34945,f34980,f35440,f35443,f35444,f35445,f35446,f35447,f35448,f35449,f35450,f35451,f35452,f35453,f35454,f35455,f35456,f35457,f35458,f35459,f35460,f36038,f36032,f36018,f35621,f35789,f36202,f36204,f35923,f36266,f36267,f36287,f36290,f36291,f36292,f36296,f36321,f36322,f36323,f36324,f36325,f36327,f36332,f36333,f36334,f36335,f36336,f36337,f36341,f36345,f36346,f36362,f36365,f36366,f36367,f36368,f36369,f36370,f36371,f36372,f36373,f36374,f36375,f36376,f36377,f36378,f36445,f36466,f36469,f36470,f36471,f36475,f36476,f37031,f37026,f37012,f37010,f36505,f36506,f36542,f36543,f36544,f36545,f36546,f36548,f36553,f36554,f36555,f36556,f36557,f36558,f36562,f36566,f36567,f36583,f36584,f36587,f36588,f36589,f36590,f36591,f36592,f36593,f36594,f36595,f36596,f36597,f36598,f36599,f36600,f37057,f36989,f37060,f37062,f36618,f37063,f36619,f37064,f36620,f36621,f36622,f37067,f36623,f36624,f36625,f36626,f36627,f36628,f36629,f36630,f36631,f36632,f36633,f36634,f36635,f37080,f36636,f37081,f36968,f36966,f36660,f36661,f36662,f36663,f36664,f36963,f36961,f36959,f36957,f37087,f37088,f36680,f36684,f36685,f36686,f36698,f37096,f36699,f37098,f36700,f37100,f36701,f37101,f36702])).
fof(f37101,plain,(
  ( ! [X379,X380] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X380)))) | p(s(bool,X379)) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X380))) = s(bool,happ(s(fun(bool,bool),X380),s(bool,'T_0')))) )),
  inference(global_subsumption,[],[f262,f261,f260,f263,f270,f277,f276,f321,f318,f335,f334,f470,f353,f350,f349,f348,f366,f373,f372,f371,f455,f369,f368,f384,f383,f382,f381,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f396,f395,f394,f393,f400,f399,f398,f397,f454,f404,f403,f402,f401,f411,f410,f409,f408,f417,f414,f413,f412,f424,f423,f422,f453,f420,f419,f418,f428,f427,f426,f425,f432,f431,f430,f429,f441,f440,f439,f452,f437,f436,f435,f445,f444,f443,f442,f449,f448,f447,f446,f471,f474,f477,f480,f482,f484,f487,f489,f491,f494,f497,f500,f503,f505,f507,f510,f513,f515,f517,f520,f522,f525,f565,f575,f578,f586,f602,f613,f622,f639,f644,f645,f647,f654,f658,f659,f662,f674,f688,f690,f698,f710,f783,f846,f878,f876,f904,f901,f947,f972,f987,f995,f1008,f1012,f1026,f1041,f1111,f1126,f1142,f1201,f1218,f1233,f1313,f1361,f1767,f1766,f1765,f1764,f2054,f2140,f2204,f2553,f2557,f2561,f2564,f2567,f2570,f2573,f2576,f2579,f2582,f2585,f2588,f2591,f2594,f2597,f2600,f2603,f2605,f2606,f2607,f2608,f2609,f2610,f2611,f2612,f2613,f2614,f2615,f2616,f2617,f2618,f2619,f2620,f2621,f2622,f2623,f2624,f2625,f2296,f2643,f2645,f2648,f2650,f2652,f2654,f2656,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2636,f328,f1840,f2064,f2039,f1797,f1796,f1795,f1794,f1793,f1792,f1791,f1790,f1789,f1788,f1787,f1786,f1785,f1784,f1783,f1782,f1781,f1780,f1779,f2635,f2634,f1296,f367,f325,f324,f2291,f2295,f2294,f2293,f2292,f2139,f2138,f2137,f2136,f2135,f327,f1866,f1867,f317,f326,f2626,f2637,f1831,f1817,f303,f2633,f2632,f2631,f2630,f2629,f2628,f2627,f469,f316,f314,f312,f310,f305,f304,f2134,f2290,f1976,f1975,f1959,f1947,f1923,f1898,f1769,f1768,f468,f281,f280,f3740,f3742,f3744,f3757,f3760,f3762,f3764,f3849,f3851,f3853,f3862,f3865,f3867,f3869,f3882,f3885,f3886,f3897,f4303,f4308,f4310,f4312,f4317,f4319,f4321,f4722,f4724,f4726,f4731,f4733,f4735,f4747,f4752,f4754,f4756,f4807,f4809,f4811,f4821,f4823,f4825,f4838,f4840,f4842,f4860,f4862,f4864,f4869,f4870,f4871,f4872,f4873,f4875,f4880,f4881,f4882,f4883,f4884,f4885,f4888,f4889,f4892,f4893,f4894,f4907,f4909,f4965,f4966,f4967,f4969,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f5004,f5006,f5007,f6917,f6931,f6948,f6967,f6996,f7007,f7009,f7045,f7069,f7071,f7090,f7091,f7092,f7093,f7094,f7124,f7130,f7147,f7153,f7173,f7179,f7193,f7199,f7208,f7209,f7223,f7240,f7242,f7244,f7261,f7263,f7265,f7288,f9309,f9324,f9325,f9417,f9398,f9399,f9484,f9465,f9466,f9489,f9490,f18495,f18500,f18505,f18466,f18467,f19646,f19651,f19656,f19617,f19618,f19906,f19911,f19916,f19878,f19879,f21428,f21429,f21430,f21457,f21458,f21459,f21931,f22080,f22041,f23768,f23773,f23780,f23727,f23728,f24311,f24316,f24323,f24270,f24271,f24460,f24461,f24462,f24450,f24463,f24464,f24465,f24466,f24467,f24468,f24469,f24470,f24471,f24472,f24473,f24474,f24475,f24476,f24477,f24478,f24479,f24480,f24481,f24482,f24483,f24484,f24485,f24486,f24487,f24488,f24489,f24506,f24507,f24508,f24523,f24538,f24539,f24540,f24555,f24596,f24613,f24585,f24621,f24648,f24649,f24650,f24672,f24673,f24674,f25900,f25921,f25973,f25992,f25993,f25994,f26033,f26034,f26035,f26075,f26076,f26077,f26088,f26089,f26090,f26105,f27929,f27930,f27931,f27946,f27968,f27969,f27970,f27991,f27992,f27993,f28005,f28006,f28007,f28031,f28032,f28033,f31391,f28156,f28157,f28226,f28227,f31209,f28261,f28262,f28272,f28300,f28301,f28557,f28559,f28560,f28561,f28562,f28563,f28564,f28565,f28566,f28567,f28568,f28569,f28570,f28571,f28572,f28573,f28574,f28575,f28576,f28577,f28578,f28579,f28580,f28581,f28582,f28583,f28584,f28585,f28586,f28587,f28596,f28597,f28598,f28599,f28600,f28601,f28602,f28603,f28604,f31756,f31760,f31766,f31768,f31819,f31820,f31823,f31824,f31826,f31828,f31830,f31832,f31834,f31837,f31838,f31841,f31842,f28949,f28951,f28952,f28953,f28954,f28955,f28956,f28957,f28958,f28959,f28960,f28961,f28962,f28963,f28964,f28965,f28966,f28967,f28968,f28969,f28970,f28971,f28972,f28973,f28974,f28975,f28976,f28977,f28978,f28979,f28988,f28989,f28990,f28991,f28992,f28993,f28994,f28995,f28996,f31844,f31859,f31860,f31861,f31862,f31863,f31864,f31865,f31866,f31867,f31868,f31869,f31870,f31871,f31872,f29145,f29152,f29155,f29156,f29157,f29158,f29159,f29160,f29161,f29162,f29163,f29164,f29165,f29166,f29167,f29168,f29169,f29170,f29171,f29172,f29173,f29174,f29175,f29184,f29185,f29186,f29187,f29195,f31922,f31923,f31924,f31925,f31926,f31927,f31928,f31929,f31930,f31931,f31932,f31933,f31934,f31935,f31936,f31937,f31938,f31939,f31940,f31941,f31942,f31943,f31944,f31945,f31946,f31947,f31948,f31952,f31953,f31954,f31955,f31956,f31957,f31958,f31959,f31960,f31961,f31962,f31963,f31964,f31965,f31966,f31967,f31968,f31969,f31970,f31971,f31972,f31973,f31974,f31975,f31976,f31977,f31978,f29393,f29395,f29396,f29397,f29398,f29399,f29400,f29401,f29402,f29403,f29404,f29405,f29406,f29407,f29408,f29409,f29410,f29411,f29412,f29413,f29414,f29415,f29416,f29417,f29418,f29419,f29420,f29421,f29422,f29423,f29432,f29433,f29434,f29435,f29436,f29437,f29438,f29439,f29440,f31989,f31990,f32041,f32044,f32045,f32046,f32047,f32048,f32049,f32052,f32055,f32125,f30120,f32469,f32470,f32471,f32472,f32473,f32474,f32475,f32476,f32477,f32478,f32479,f32480,f32481,f32482,f32483,f32484,f32485,f32486,f32487,f32488,f32489,f32490,f32491,f32492,f32493,f32494,f32495,f32496,f32497,f32498,f32499,f32500,f32501,f32502,f32503,f32504,f32505,f32506,f32507,f32508,f32509,f32510,f32511,f32512,f32513,f32514,f32515,f32516,f32517,f32518,f32519,f32520,f32521,f32522,f32523,f32524,f32525,f32526,f32527,f32528,f32529,f32691,f32692,f32693,f32696,f32698,f32699,f32700,f32732,f32735,f32740,f32824,f32825,f32826,f32829,f32831,f32832,f32833,f32865,f32868,f32873,f32932,f33004,f33066,f33132,f33183,f33335,f33469,f33623,f33852,f33850,f33916,f33923,f33951,f33954,f33959,f33989,f33992,f33997,f34045,f34048,f34053,f34210,f34211,f34212,f34215,f34217,f34218,f34219,f34272,f34273,f34274,f34277,f34279,f34280,f34281,f34411,f34412,f34413,f34416,f34418,f34419,f34420,f34493,f34494,f34495,f34498,f34500,f34501,f34502,f34520,f34521,f34522,f34525,f34527,f34528,f34529,f34549,f34550,f34551,f34554,f34556,f34557,f34558,f34578,f34579,f34580,f34583,f34585,f34586,f34587,f34664,f34665,f34666,f34669,f34671,f34672,f34673,f35269,f34731,f34762,f34763,f34766,f34767,f34769,f34770,f34771,f34924,f34925,f34930,f34945,f34980,f35440,f35443,f35444,f35445,f35446,f35447,f35448,f35449,f35450,f35451,f35452,f35453,f35454,f35455,f35456,f35457,f35458,f35459,f35460,f36038,f36032,f36018,f35621,f35789,f36202,f36204,f35923,f36266,f36267,f36287,f36290,f36291,f36292,f36296,f36321,f36322,f36323,f36324,f36325,f36327,f36332,f36333,f36334,f36335,f36336,f36337,f36341,f36345,f36346,f36362,f36365,f36366,f36367,f36368,f36369,f36370,f36371,f36372,f36373,f36374,f36375,f36376,f36377,f36378,f36445,f36466,f36469,f36470,f36471,f36475,f36476,f37031,f37026,f37012,f37010,f36505,f36506,f36542,f36543,f36544,f36545,f36546,f36548,f36553,f36554,f36555,f36556,f36557,f36558,f36562,f36566,f36567,f36583,f36584,f36587,f36588,f36589,f36590,f36591,f36592,f36593,f36594,f36595,f36596,f36597,f36598,f36599,f36600,f37057,f36989,f37060,f37062,f36618,f37063,f36619,f37064,f36620,f36621,f36622,f37067,f36623,f36624,f36625,f36626,f36627,f36628,f36629,f36630,f36631,f36632,f36633,f36634,f36635,f37080,f36636,f37081,f36968,f36966,f36660,f36661,f36662,f36663,f36664,f36963,f36961,f36959,f36957,f37087,f37088,f36680,f36684,f36685,f36686,f36698,f37096,f36699,f37098,f36700,f37100,f36701])).
fof(f37100,plain,(
  ( ! [X374,X375] : (p(s(bool,X374)) | ~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X375)))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X375))) = s(bool,happ(s(fun(bool,bool),X375),s(bool,'T_0')))) )),
  inference(global_subsumption,[],[f262,f261,f260,f263,f270,f277,f276,f321,f318,f335,f334,f470,f353,f350,f349,f348,f366,f373,f372,f371,f455,f369,f368,f384,f383,f382,f381,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f396,f395,f394,f393,f400,f399,f398,f397,f454,f404,f403,f402,f401,f411,f410,f409,f408,f417,f414,f413,f412,f424,f423,f422,f453,f420,f419,f418,f428,f427,f426,f425,f432,f431,f430,f429,f441,f440,f439,f452,f437,f436,f435,f445,f444,f443,f442,f449,f448,f447,f446,f471,f474,f477,f480,f482,f484,f487,f489,f491,f494,f497,f500,f503,f505,f507,f510,f513,f515,f517,f520,f522,f525,f565,f575,f578,f586,f602,f613,f622,f639,f644,f645,f647,f654,f658,f659,f662,f674,f688,f690,f698,f710,f783,f846,f878,f876,f904,f901,f947,f972,f987,f995,f1008,f1012,f1026,f1041,f1111,f1126,f1142,f1201,f1218,f1233,f1313,f1361,f1767,f1766,f1765,f1764,f2054,f2140,f2204,f2553,f2557,f2561,f2564,f2567,f2570,f2573,f2576,f2579,f2582,f2585,f2588,f2591,f2594,f2597,f2600,f2603,f2605,f2606,f2607,f2608,f2609,f2610,f2611,f2612,f2613,f2614,f2615,f2616,f2617,f2618,f2619,f2620,f2621,f2622,f2623,f2624,f2625,f2296,f2643,f2645,f2648,f2650,f2652,f2654,f2656,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2636,f328,f1840,f2064,f2039,f1797,f1796,f1795,f1794,f1793,f1792,f1791,f1790,f1789,f1788,f1787,f1786,f1785,f1784,f1783,f1782,f1781,f1780,f1779,f2635,f2634,f1296,f367,f325,f324,f2291,f2295,f2294,f2293,f2292,f2139,f2138,f2137,f2136,f2135,f327,f1866,f1867,f317,f326,f2626,f2637,f1831,f1817,f303,f2633,f2632,f2631,f2630,f2629,f2628,f2627,f469,f316,f314,f312,f310,f305,f304,f2134,f2290,f1976,f1975,f1959,f1947,f1923,f1898,f1769,f1768,f468,f281,f280,f3740,f3742,f3744,f3757,f3760,f3762,f3764,f3849,f3851,f3853,f3862,f3865,f3867,f3869,f3882,f3885,f3886,f3897,f4303,f4308,f4310,f4312,f4317,f4319,f4321,f4722,f4724,f4726,f4731,f4733,f4735,f4747,f4752,f4754,f4756,f4807,f4809,f4811,f4821,f4823,f4825,f4838,f4840,f4842,f4860,f4862,f4864,f4869,f4870,f4871,f4872,f4873,f4875,f4880,f4881,f4882,f4883,f4884,f4885,f4888,f4889,f4892,f4893,f4894,f4907,f4909,f4965,f4966,f4967,f4969,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f5004,f5006,f5007,f6917,f6931,f6948,f6967,f6996,f7007,f7009,f7045,f7069,f7071,f7090,f7091,f7092,f7093,f7094,f7124,f7130,f7147,f7153,f7173,f7179,f7193,f7199,f7208,f7209,f7223,f7240,f7242,f7244,f7261,f7263,f7265,f7288,f9309,f9324,f9325,f9417,f9398,f9399,f9484,f9465,f9466,f9489,f9490,f18495,f18500,f18505,f18466,f18467,f19646,f19651,f19656,f19617,f19618,f19906,f19911,f19916,f19878,f19879,f21428,f21429,f21430,f21457,f21458,f21459,f21931,f22080,f22041,f23768,f23773,f23780,f23727,f23728,f24311,f24316,f24323,f24270,f24271,f24460,f24461,f24462,f24450,f24463,f24464,f24465,f24466,f24467,f24468,f24469,f24470,f24471,f24472,f24473,f24474,f24475,f24476,f24477,f24478,f24479,f24480,f24481,f24482,f24483,f24484,f24485,f24486,f24487,f24488,f24489,f24506,f24507,f24508,f24523,f24538,f24539,f24540,f24555,f24596,f24613,f24585,f24621,f24648,f24649,f24650,f24672,f24673,f24674,f25900,f25921,f25973,f25992,f25993,f25994,f26033,f26034,f26035,f26075,f26076,f26077,f26088,f26089,f26090,f26105,f27929,f27930,f27931,f27946,f27968,f27969,f27970,f27991,f27992,f27993,f28005,f28006,f28007,f28031,f28032,f28033,f31391,f28156,f28157,f28226,f28227,f31209,f28261,f28262,f28272,f28300,f28301,f28557,f28559,f28560,f28561,f28562,f28563,f28564,f28565,f28566,f28567,f28568,f28569,f28570,f28571,f28572,f28573,f28574,f28575,f28576,f28577,f28578,f28579,f28580,f28581,f28582,f28583,f28584,f28585,f28586,f28587,f28596,f28597,f28598,f28599,f28600,f28601,f28602,f28603,f28604,f31756,f31760,f31766,f31768,f31819,f31820,f31823,f31824,f31826,f31828,f31830,f31832,f31834,f31837,f31838,f31841,f31842,f28949,f28951,f28952,f28953,f28954,f28955,f28956,f28957,f28958,f28959,f28960,f28961,f28962,f28963,f28964,f28965,f28966,f28967,f28968,f28969,f28970,f28971,f28972,f28973,f28974,f28975,f28976,f28977,f28978,f28979,f28988,f28989,f28990,f28991,f28992,f28993,f28994,f28995,f28996,f31844,f31859,f31860,f31861,f31862,f31863,f31864,f31865,f31866,f31867,f31868,f31869,f31870,f31871,f31872,f29145,f29152,f29155,f29156,f29157,f29158,f29159,f29160,f29161,f29162,f29163,f29164,f29165,f29166,f29167,f29168,f29169,f29170,f29171,f29172,f29173,f29174,f29175,f29184,f29185,f29186,f29187,f29195,f31922,f31923,f31924,f31925,f31926,f31927,f31928,f31929,f31930,f31931,f31932,f31933,f31934,f31935,f31936,f31937,f31938,f31939,f31940,f31941,f31942,f31943,f31944,f31945,f31946,f31947,f31948,f31952,f31953,f31954,f31955,f31956,f31957,f31958,f31959,f31960,f31961,f31962,f31963,f31964,f31965,f31966,f31967,f31968,f31969,f31970,f31971,f31972,f31973,f31974,f31975,f31976,f31977,f31978,f29393,f29395,f29396,f29397,f29398,f29399,f29400,f29401,f29402,f29403,f29404,f29405,f29406,f29407,f29408,f29409,f29410,f29411,f29412,f29413,f29414,f29415,f29416,f29417,f29418,f29419,f29420,f29421,f29422,f29423,f29432,f29433,f29434,f29435,f29436,f29437,f29438,f29439,f29440,f31989,f31990,f32041,f32044,f32045,f32046,f32047,f32048,f32049,f32052,f32055,f32125,f30120,f32469,f32470,f32471,f32472,f32473,f32474,f32475,f32476,f32477,f32478,f32479,f32480,f32481,f32482,f32483,f32484,f32485,f32486,f32487,f32488,f32489,f32490,f32491,f32492,f32493,f32494,f32495,f32496,f32497,f32498,f32499,f32500,f32501,f32502,f32503,f32504,f32505,f32506,f32507,f32508,f32509,f32510,f32511,f32512,f32513,f32514,f32515,f32516,f32517,f32518,f32519,f32520,f32521,f32522,f32523,f32524,f32525,f32526,f32527,f32528,f32529,f32691,f32692,f32693,f32696,f32698,f32699,f32700,f32732,f32735,f32740,f32824,f32825,f32826,f32829,f32831,f32832,f32833,f32865,f32868,f32873,f32932,f33004,f33066,f33132,f33183,f33335,f33469,f33623,f33852,f33850,f33916,f33923,f33951,f33954,f33959,f33989,f33992,f33997,f34045,f34048,f34053,f34210,f34211,f34212,f34215,f34217,f34218,f34219,f34272,f34273,f34274,f34277,f34279,f34280,f34281,f34411,f34412,f34413,f34416,f34418,f34419,f34420,f34493,f34494,f34495,f34498,f34500,f34501,f34502,f34520,f34521,f34522,f34525,f34527,f34528,f34529,f34549,f34550,f34551,f34554,f34556,f34557,f34558,f34578,f34579,f34580,f34583,f34585,f34586,f34587,f34664,f34665,f34666,f34669,f34671,f34672,f34673,f35269,f34731,f34762,f34763,f34766,f34767,f34769,f34770,f34771,f34924,f34925,f34930,f34945,f34980,f35440,f35443,f35444,f35445,f35446,f35447,f35448,f35449,f35450,f35451,f35452,f35453,f35454,f35455,f35456,f35457,f35458,f35459,f35460,f36038,f36032,f36018,f35621,f35789,f36202,f36204,f35923,f36266,f36267,f36287,f36290,f36291,f36292,f36296,f36321,f36322,f36323,f36324,f36325,f36327,f36332,f36333,f36334,f36335,f36336,f36337,f36341,f36345,f36346,f36362,f36365,f36366,f36367,f36368,f36369,f36370,f36371,f36372,f36373,f36374,f36375,f36376,f36377,f36378,f36445,f36466,f36469,f36470,f36471,f36475,f36476,f37031,f37026,f37012,f37010,f36505,f36506,f36542,f36543,f36544,f36545,f36546,f36548,f36553,f36554,f36555,f36556,f36557,f36558,f36562,f36566,f36567,f36583,f36584,f36587,f36588,f36589,f36590,f36591,f36592,f36593,f36594,f36595,f36596,f36597,f36598,f36599,f36600,f37057,f36989,f37060,f37062,f36618,f37063,f36619,f37064,f36620,f36621,f36622,f37067,f36623,f36624,f36625,f36626,f36627,f36628,f36629,f36630,f36631,f36632,f36633,f36634,f36635,f37080,f36636,f37081,f36968,f36966,f36660,f36661,f36662,f36663,f36664,f36963,f36961,f36959,f36957,f37087,f37088,f36680,f36684,f36685,f36686,f36698,f37096,f36699,f37098,f36700])).
fof(f37099,plain,(
  ( ! [X368,X367,X369,X371] : (p(s(bool,happ(s(fun(X367,bool),X368),s(X367,X371)))) | s(bool,'T_0') = s(bool,happ(s(fun(X367,bool),X368),s(X367,sK19(X368,X369,X367))))) ) | $spl182),
  inference(cnf_transformation,[],[f37099_D])).
fof(f37099_D,plain,(
  ( ! [X368,X367,X369,X371] : (p(s(bool,happ(s(fun(X367,bool),X368),s(X367,X371)))) | s(bool,'T_0') = s(bool,happ(s(fun(X367,bool),X368),s(X367,sK19(X368,X369,X367))))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f37098,plain,(
  ( ! [X368,X370,X367,X369,X371] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X370))) = s(bool,happ(s(fun(bool,bool),X370),s(bool,'T_0'))) | p(s(bool,happ(s(fun(X367,bool),X368),s(X367,X371)))) | ~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X370)))) | s(bool,'T_0') = s(bool,happ(s(fun(X367,bool),X368),s(X367,sK19(X368,X369,X367))))) )),
  inference(global_subsumption,[],[f262,f261,f260,f263,f270,f277,f276,f321,f318,f335,f334,f470,f353,f350,f349,f348,f366,f373,f372,f371,f455,f369,f368,f384,f383,f382,f381,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f396,f395,f394,f393,f400,f399,f398,f397,f454,f404,f403,f402,f401,f411,f410,f409,f408,f417,f414,f413,f412,f424,f423,f422,f453,f420,f419,f418,f428,f427,f426,f425,f432,f431,f430,f429,f441,f440,f439,f452,f437,f436,f435,f445,f444,f443,f442,f449,f448,f447,f446,f471,f474,f477,f480,f482,f484,f487,f489,f491,f494,f497,f500,f503,f505,f507,f510,f513,f515,f517,f520,f522,f525,f565,f575,f578,f586,f602,f613,f622,f639,f644,f645,f647,f654,f658,f659,f662,f674,f688,f690,f698,f710,f783,f846,f878,f876,f904,f901,f947,f972,f987,f995,f1008,f1012,f1026,f1041,f1111,f1126,f1142,f1201,f1218,f1233,f1313,f1361,f1767,f1766,f1765,f1764,f2054,f2140,f2204,f2553,f2557,f2561,f2564,f2567,f2570,f2573,f2576,f2579,f2582,f2585,f2588,f2591,f2594,f2597,f2600,f2603,f2605,f2606,f2607,f2608,f2609,f2610,f2611,f2612,f2613,f2614,f2615,f2616,f2617,f2618,f2619,f2620,f2621,f2622,f2623,f2624,f2625,f2296,f2643,f2645,f2648,f2650,f2652,f2654,f2656,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2636,f328,f1840,f2064,f2039,f1797,f1796,f1795,f1794,f1793,f1792,f1791,f1790,f1789,f1788,f1787,f1786,f1785,f1784,f1783,f1782,f1781,f1780,f1779,f2635,f2634,f1296,f367,f325,f324,f2291,f2295,f2294,f2293,f2292,f2139,f2138,f2137,f2136,f2135,f327,f1866,f1867,f317,f326,f2626,f2637,f1831,f1817,f303,f2633,f2632,f2631,f2630,f2629,f2628,f2627,f469,f316,f314,f312,f310,f305,f304,f2134,f2290,f1976,f1975,f1959,f1947,f1923,f1898,f1769,f1768,f468,f281,f280,f3740,f3742,f3744,f3757,f3760,f3762,f3764,f3849,f3851,f3853,f3862,f3865,f3867,f3869,f3882,f3885,f3886,f3897,f4303,f4308,f4310,f4312,f4317,f4319,f4321,f4722,f4724,f4726,f4731,f4733,f4735,f4747,f4752,f4754,f4756,f4807,f4809,f4811,f4821,f4823,f4825,f4838,f4840,f4842,f4860,f4862,f4864,f4869,f4870,f4871,f4872,f4873,f4875,f4880,f4881,f4882,f4883,f4884,f4885,f4888,f4889,f4892,f4893,f4894,f4907,f4909,f4965,f4966,f4967,f4969,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f5004,f5006,f5007,f6917,f6931,f6948,f6967,f6996,f7007,f7009,f7045,f7069,f7071,f7090,f7091,f7092,f7093,f7094,f7124,f7130,f7147,f7153,f7173,f7179,f7193,f7199,f7208,f7209,f7223,f7240,f7242,f7244,f7261,f7263,f7265,f7288,f9309,f9324,f9325,f9417,f9398,f9399,f9484,f9465,f9466,f9489,f9490,f18495,f18500,f18505,f18466,f18467,f19646,f19651,f19656,f19617,f19618,f19906,f19911,f19916,f19878,f19879,f21428,f21429,f21430,f21457,f21458,f21459,f21931,f22080,f22041,f23768,f23773,f23780,f23727,f23728,f24311,f24316,f24323,f24270,f24271,f24460,f24461,f24462,f24450,f24463,f24464,f24465,f24466,f24467,f24468,f24469,f24470,f24471,f24472,f24473,f24474,f24475,f24476,f24477,f24478,f24479,f24480,f24481,f24482,f24483,f24484,f24485,f24486,f24487,f24488,f24489,f24506,f24507,f24508,f24523,f24538,f24539,f24540,f24555,f24596,f24613,f24585,f24621,f24648,f24649,f24650,f24672,f24673,f24674,f25900,f25921,f25973,f25992,f25993,f25994,f26033,f26034,f26035,f26075,f26076,f26077,f26088,f26089,f26090,f26105,f27929,f27930,f27931,f27946,f27968,f27969,f27970,f27991,f27992,f27993,f28005,f28006,f28007,f28031,f28032,f28033,f31391,f28156,f28157,f28226,f28227,f31209,f28261,f28262,f28272,f28300,f28301,f28557,f28559,f28560,f28561,f28562,f28563,f28564,f28565,f28566,f28567,f28568,f28569,f28570,f28571,f28572,f28573,f28574,f28575,f28576,f28577,f28578,f28579,f28580,f28581,f28582,f28583,f28584,f28585,f28586,f28587,f28596,f28597,f28598,f28599,f28600,f28601,f28602,f28603,f28604,f31756,f31760,f31766,f31768,f31819,f31820,f31823,f31824,f31826,f31828,f31830,f31832,f31834,f31837,f31838,f31841,f31842,f28949,f28951,f28952,f28953,f28954,f28955,f28956,f28957,f28958,f28959,f28960,f28961,f28962,f28963,f28964,f28965,f28966,f28967,f28968,f28969,f28970,f28971,f28972,f28973,f28974,f28975,f28976,f28977,f28978,f28979,f28988,f28989,f28990,f28991,f28992,f28993,f28994,f28995,f28996,f31844,f31859,f31860,f31861,f31862,f31863,f31864,f31865,f31866,f31867,f31868,f31869,f31870,f31871,f31872,f29145,f29152,f29155,f29156,f29157,f29158,f29159,f29160,f29161,f29162,f29163,f29164,f29165,f29166,f29167,f29168,f29169,f29170,f29171,f29172,f29173,f29174,f29175,f29184,f29185,f29186,f29187,f29195,f31922,f31923,f31924,f31925,f31926,f31927,f31928,f31929,f31930,f31931,f31932,f31933,f31934,f31935,f31936,f31937,f31938,f31939,f31940,f31941,f31942,f31943,f31944,f31945,f31946,f31947,f31948,f31952,f31953,f31954,f31955,f31956,f31957,f31958,f31959,f31960,f31961,f31962,f31963,f31964,f31965,f31966,f31967,f31968,f31969,f31970,f31971,f31972,f31973,f31974,f31975,f31976,f31977,f31978,f29393,f29395,f29396,f29397,f29398,f29399,f29400,f29401,f29402,f29403,f29404,f29405,f29406,f29407,f29408,f29409,f29410,f29411,f29412,f29413,f29414,f29415,f29416,f29417,f29418,f29419,f29420,f29421,f29422,f29423,f29432,f29433,f29434,f29435,f29436,f29437,f29438,f29439,f29440,f31989,f31990,f32041,f32044,f32045,f32046,f32047,f32048,f32049,f32052,f32055,f32125,f30120,f32469,f32470,f32471,f32472,f32473,f32474,f32475,f32476,f32477,f32478,f32479,f32480,f32481,f32482,f32483,f32484,f32485,f32486,f32487,f32488,f32489,f32490,f32491,f32492,f32493,f32494,f32495,f32496,f32497,f32498,f32499,f32500,f32501,f32502,f32503,f32504,f32505,f32506,f32507,f32508,f32509,f32510,f32511,f32512,f32513,f32514,f32515,f32516,f32517,f32518,f32519,f32520,f32521,f32522,f32523,f32524,f32525,f32526,f32527,f32528,f32529,f32691,f32692,f32693,f32696,f32698,f32699,f32700,f32732,f32735,f32740,f32824,f32825,f32826,f32829,f32831,f32832,f32833,f32865,f32868,f32873,f32932,f33004,f33066,f33132,f33183,f33335,f33469,f33623,f33852,f33850,f33916,f33923,f33951,f33954,f33959,f33989,f33992,f33997,f34045,f34048,f34053,f34210,f34211,f34212,f34215,f34217,f34218,f34219,f34272,f34273,f34274,f34277,f34279,f34280,f34281,f34411,f34412,f34413,f34416,f34418,f34419,f34420,f34493,f34494,f34495,f34498,f34500,f34501,f34502,f34520,f34521,f34522,f34525,f34527,f34528,f34529,f34549,f34550,f34551,f34554,f34556,f34557,f34558,f34578,f34579,f34580,f34583,f34585,f34586,f34587,f34664,f34665,f34666,f34669,f34671,f34672,f34673,f35269,f34731,f34762,f34763,f34766,f34767,f34769,f34770,f34771,f34924,f34925,f34930,f34945,f34980,f35440,f35443,f35444,f35445,f35446,f35447,f35448,f35449,f35450,f35451,f35452,f35453,f35454,f35455,f35456,f35457,f35458,f35459,f35460,f36038,f36032,f36018,f35621,f35789,f36202,f36204,f35923,f36266,f36267,f36287,f36290,f36291,f36292,f36296,f36321,f36322,f36323,f36324,f36325,f36327,f36332,f36333,f36334,f36335,f36336,f36337,f36341,f36345,f36346,f36362,f36365,f36366,f36367,f36368,f36369,f36370,f36371,f36372,f36373,f36374,f36375,f36376,f36377,f36378,f36445,f36466,f36469,f36470,f36471,f36475,f36476,f37031,f37026,f37012,f37010,f36505,f36506,f36542,f36543,f36544,f36545,f36546,f36548,f36553,f36554,f36555,f36556,f36557,f36558,f36562,f36566,f36567,f36583,f36584,f36587,f36588,f36589,f36590,f36591,f36592,f36593,f36594,f36595,f36596,f36597,f36598,f36599,f36600,f37057,f36989,f37060,f37062,f36618,f37063,f36619,f37064,f36620,f36621,f36622,f37067,f36623,f36624,f36625,f36626,f36627,f36628,f36629,f36630,f36631,f36632,f36633,f36634,f36635,f37080,f36636,f37081,f36968,f36966,f36660,f36661,f36662,f36663,f36664,f36963,f36961,f36959,f36957,f37087,f37088,f36680,f36684,f36685,f36686,f36698,f37096,f36699])).
fof(f37097,plain,(
  ( ! [X364,X366,X362,X363] : (p(s(bool,happ(s(fun(X362,bool),X363),s(X362,X366)))) | s(bool,'T_0') = s(bool,happ(s(fun(X362,bool),X363),s(X362,sK18(X363,X364,X362))))) ) | $spl180),
  inference(cnf_transformation,[],[f37097_D])).
fof(f37097_D,plain,(
  ( ! [X364,X366,X362,X363] : (p(s(bool,happ(s(fun(X362,bool),X363),s(X362,X366)))) | s(bool,'T_0') = s(bool,happ(s(fun(X362,bool),X363),s(X362,sK18(X363,X364,X362))))) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f37096,plain,(
  ( ! [X364,X366,X362,X365,X363] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X365)))) | p(s(bool,happ(s(fun(X362,bool),X363),s(X362,X366)))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X365))) = s(bool,happ(s(fun(bool,bool),X365),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(X362,bool),X363),s(X362,sK18(X363,X364,X362))))) )),
  inference(global_subsumption,[],[f262,f261,f260,f263,f270,f277,f276,f321,f318,f335,f334,f470,f353,f350,f349,f348,f366,f373,f372,f371,f455,f369,f368,f384,f383,f382,f381,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f396,f395,f394,f393,f400,f399,f398,f397,f454,f404,f403,f402,f401,f411,f410,f409,f408,f417,f414,f413,f412,f424,f423,f422,f453,f420,f419,f418,f428,f427,f426,f425,f432,f431,f430,f429,f441,f440,f439,f452,f437,f436,f435,f445,f444,f443,f442,f449,f448,f447,f446,f471,f474,f477,f480,f482,f484,f487,f489,f491,f494,f497,f500,f503,f505,f507,f510,f513,f515,f517,f520,f522,f525,f565,f575,f578,f586,f602,f613,f622,f639,f644,f645,f647,f654,f658,f659,f662,f674,f688,f690,f698,f710,f783,f846,f878,f876,f904,f901,f947,f972,f987,f995,f1008,f1012,f1026,f1041,f1111,f1126,f1142,f1201,f1218,f1233,f1313,f1361,f1767,f1766,f1765,f1764,f2054,f2140,f2204,f2553,f2557,f2561,f2564,f2567,f2570,f2573,f2576,f2579,f2582,f2585,f2588,f2591,f2594,f2597,f2600,f2603,f2605,f2606,f2607,f2608,f2609,f2610,f2611,f2612,f2613,f2614,f2615,f2616,f2617,f2618,f2619,f2620,f2621,f2622,f2623,f2624,f2625,f2296,f2643,f2645,f2648,f2650,f2652,f2654,f2656,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2636,f328,f1840,f2064,f2039,f1797,f1796,f1795,f1794,f1793,f1792,f1791,f1790,f1789,f1788,f1787,f1786,f1785,f1784,f1783,f1782,f1781,f1780,f1779,f2635,f2634,f1296,f367,f325,f324,f2291,f2295,f2294,f2293,f2292,f2139,f2138,f2137,f2136,f2135,f327,f1866,f1867,f317,f326,f2626,f2637,f1831,f1817,f303,f2633,f2632,f2631,f2630,f2629,f2628,f2627,f469,f316,f314,f312,f310,f305,f304,f2134,f2290,f1976,f1975,f1959,f1947,f1923,f1898,f1769,f1768,f468,f281,f280,f3740,f3742,f3744,f3757,f3760,f3762,f3764,f3849,f3851,f3853,f3862,f3865,f3867,f3869,f3882,f3885,f3886,f3897,f4303,f4308,f4310,f4312,f4317,f4319,f4321,f4722,f4724,f4726,f4731,f4733,f4735,f4747,f4752,f4754,f4756,f4807,f4809,f4811,f4821,f4823,f4825,f4838,f4840,f4842,f4860,f4862,f4864,f4869,f4870,f4871,f4872,f4873,f4875,f4880,f4881,f4882,f4883,f4884,f4885,f4888,f4889,f4892,f4893,f4894,f4907,f4909,f4965,f4966,f4967,f4969,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f5004,f5006,f5007,f6917,f6931,f6948,f6967,f6996,f7007,f7009,f7045,f7069,f7071,f7090,f7091,f7092,f7093,f7094,f7124,f7130,f7147,f7153,f7173,f7179,f7193,f7199,f7208,f7209,f7223,f7240,f7242,f7244,f7261,f7263,f7265,f7288,f9309,f9324,f9325,f9417,f9398,f9399,f9484,f9465,f9466,f9489,f9490,f18495,f18500,f18505,f18466,f18467,f19646,f19651,f19656,f19617,f19618,f19906,f19911,f19916,f19878,f19879,f21428,f21429,f21430,f21457,f21458,f21459,f21931,f22080,f22041,f23768,f23773,f23780,f23727,f23728,f24311,f24316,f24323,f24270,f24271,f24460,f24461,f24462,f24450,f24463,f24464,f24465,f24466,f24467,f24468,f24469,f24470,f24471,f24472,f24473,f24474,f24475,f24476,f24477,f24478,f24479,f24480,f24481,f24482,f24483,f24484,f24485,f24486,f24487,f24488,f24489,f24506,f24507,f24508,f24523,f24538,f24539,f24540,f24555,f24596,f24613,f24585,f24621,f24648,f24649,f24650,f24672,f24673,f24674,f25900,f25921,f25973,f25992,f25993,f25994,f26033,f26034,f26035,f26075,f26076,f26077,f26088,f26089,f26090,f26105,f27929,f27930,f27931,f27946,f27968,f27969,f27970,f27991,f27992,f27993,f28005,f28006,f28007,f28031,f28032,f28033,f31391,f28156,f28157,f28226,f28227,f31209,f28261,f28262,f28272,f28300,f28301,f28557,f28559,f28560,f28561,f28562,f28563,f28564,f28565,f28566,f28567,f28568,f28569,f28570,f28571,f28572,f28573,f28574,f28575,f28576,f28577,f28578,f28579,f28580,f28581,f28582,f28583,f28584,f28585,f28586,f28587,f28596,f28597,f28598,f28599,f28600,f28601,f28602,f28603,f28604,f31756,f31760,f31766,f31768,f31819,f31820,f31823,f31824,f31826,f31828,f31830,f31832,f31834,f31837,f31838,f31841,f31842,f28949,f28951,f28952,f28953,f28954,f28955,f28956,f28957,f28958,f28959,f28960,f28961,f28962,f28963,f28964,f28965,f28966,f28967,f28968,f28969,f28970,f28971,f28972,f28973,f28974,f28975,f28976,f28977,f28978,f28979,f28988,f28989,f28990,f28991,f28992,f28993,f28994,f28995,f28996,f31844,f31859,f31860,f31861,f31862,f31863,f31864,f31865,f31866,f31867,f31868,f31869,f31870,f31871,f31872,f29145,f29152,f29155,f29156,f29157,f29158,f29159,f29160,f29161,f29162,f29163,f29164,f29165,f29166,f29167,f29168,f29169,f29170,f29171,f29172,f29173,f29174,f29175,f29184,f29185,f29186,f29187,f29195,f31922,f31923,f31924,f31925,f31926,f31927,f31928,f31929,f31930,f31931,f31932,f31933,f31934,f31935,f31936,f31937,f31938,f31939,f31940,f31941,f31942,f31943,f31944,f31945,f31946,f31947,f31948,f31952,f31953,f31954,f31955,f31956,f31957,f31958,f31959,f31960,f31961,f31962,f31963,f31964,f31965,f31966,f31967,f31968,f31969,f31970,f31971,f31972,f31973,f31974,f31975,f31976,f31977,f31978,f29393,f29395,f29396,f29397,f29398,f29399,f29400,f29401,f29402,f29403,f29404,f29405,f29406,f29407,f29408,f29409,f29410,f29411,f29412,f29413,f29414,f29415,f29416,f29417,f29418,f29419,f29420,f29421,f29422,f29423,f29432,f29433,f29434,f29435,f29436,f29437,f29438,f29439,f29440,f31989,f31990,f32041,f32044,f32045,f32046,f32047,f32048,f32049,f32052,f32055,f32125,f30120,f32469,f32470,f32471,f32472,f32473,f32474,f32475,f32476,f32477,f32478,f32479,f32480,f32481,f32482,f32483,f32484,f32485,f32486,f32487,f32488,f32489,f32490,f32491,f32492,f32493,f32494,f32495,f32496,f32497,f32498,f32499,f32500,f32501,f32502,f32503,f32504,f32505,f32506,f32507,f32508,f32509,f32510,f32511,f32512,f32513,f32514,f32515,f32516,f32517,f32518,f32519,f32520,f32521,f32522,f32523,f32524,f32525,f32526,f32527,f32528,f32529,f32691,f32692,f32693,f32696,f32698,f32699,f32700,f32732,f32735,f32740,f32824,f32825,f32826,f32829,f32831,f32832,f32833,f32865,f32868,f32873,f32932,f33004,f33066,f33132,f33183,f33335,f33469,f33623,f33852,f33850,f33916,f33923,f33951,f33954,f33959,f33989,f33992,f33997,f34045,f34048,f34053,f34210,f34211,f34212,f34215,f34217,f34218,f34219,f34272,f34273,f34274,f34277,f34279,f34280,f34281,f34411,f34412,f34413,f34416,f34418,f34419,f34420,f34493,f34494,f34495,f34498,f34500,f34501,f34502,f34520,f34521,f34522,f34525,f34527,f34528,f34529,f34549,f34550,f34551,f34554,f34556,f34557,f34558,f34578,f34579,f34580,f34583,f34585,f34586,f34587,f34664,f34665,f34666,f34669,f34671,f34672,f34673,f35269,f34731,f34762,f34763,f34766,f34767,f34769,f34770,f34771,f34924,f34925,f34930,f34945,f34980,f35440,f35443,f35444,f35445,f35446,f35447,f35448,f35449,f35450,f35451,f35452,f35453,f35454,f35455,f35456,f35457,f35458,f35459,f35460,f36038,f36032,f36018,f35621,f35789,f36202,f36204,f35923,f36266,f36267,f36287,f36290,f36291,f36292,f36296,f36321,f36322,f36323,f36324,f36325,f36327,f36332,f36333,f36334,f36335,f36336,f36337,f36341,f36345,f36346,f36362,f36365,f36366,f36367,f36368,f36369,f36370,f36371,f36372,f36373,f36374,f36375,f36376,f36377,f36378,f36445,f36466,f36469,f36470,f36471,f36475,f36476,f37031,f37026,f37012,f37010,f36505,f36506,f36542,f36543,f36544,f36545,f36546,f36548,f36553,f36554,f36555,f36556,f36557,f36558,f36562,f36566,f36567,f36583,f36584,f36587,f36588,f36589,f36590,f36591,f36592,f36593,f36594,f36595,f36596,f36597,f36598,f36599,f36600,f37057,f36989,f37060,f37062,f36618,f37063,f36619,f37064,f36620,f36621,f36622,f37067,f36623,f36624,f36625,f36626,f36627,f36628,f36629,f36630,f36631,f36632,f36633,f36634,f36635,f37080,f36636,f37081,f36968,f36966,f36660,f36661,f36662,f36663,f36664,f36963,f36961,f36959,f36957,f37087,f37088,f36680,f36684,f36685,f36686,f36698])).
fof(f37094,plain,(
  ( ! [X318,X317,X319] : (p(s(bool,'|dollar|exists_1'(s(fun(X317,bool),X318)))) | s(bool,'T_0') = s(bool,happ(s(fun(X317,bool),X318),s(X317,X319)))) ) | $spl178),
  inference(cnf_transformation,[],[f37094_D])).
fof(f37094_D,plain,(
  ( ! [X318,X317,X319] : (p(s(bool,'|dollar|exists_1'(s(fun(X317,bool),X318)))) | s(bool,'T_0') = s(bool,happ(s(fun(X317,bool),X318),s(X317,X319)))) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f36686,plain,(
  ( ! [X318,X320,X317,X319] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X320)))) | p(s(bool,'|dollar|exists_1'(s(fun(X317,bool),X318)))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X320))) = s(bool,happ(s(fun(bool,bool),X320),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(X317,bool),X318),s(X317,X319)))) )),
  inference(superposition,[],[f469,f2291])).
fof(f37093,plain,(
  ( ! [X313,X315,X314] : (p(s(bool,happ(s(fun(X313,bool),X314),s(X313,sK16(X314,X313))))) | s(bool,'T_0') = s(bool,happ(s(fun(X313,bool),X314),s(X313,X315)))) ) | $spl176),
  inference(cnf_transformation,[],[f37093_D])).
fof(f37093_D,plain,(
  ( ! [X313,X315,X314] : (p(s(bool,happ(s(fun(X313,bool),X314),s(X313,sK16(X314,X313))))) | s(bool,'T_0') = s(bool,happ(s(fun(X313,bool),X314),s(X313,X315)))) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f36685,plain,(
  ( ! [X313,X315,X316,X314] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X316)))) | p(s(bool,happ(s(fun(X313,bool),X314),s(X313,sK16(X314,X313))))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X316))) = s(bool,happ(s(fun(bool,bool),X316),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(X313,bool),X314),s(X313,X315)))) )),
  inference(superposition,[],[f304,f2291])).
fof(f37092,plain,(
  ( ! [X310,X309,X311] : (p(s(bool,happ(s(fun(X309,bool),X310),s(X309,sK17(X310,X309))))) | s(bool,'T_0') = s(bool,happ(s(fun(X309,bool),X310),s(X309,X311)))) ) | $spl174),
  inference(cnf_transformation,[],[f37092_D])).
fof(f37092_D,plain,(
  ( ! [X310,X309,X311] : (p(s(bool,happ(s(fun(X309,bool),X310),s(X309,sK17(X310,X309))))) | s(bool,'T_0') = s(bool,happ(s(fun(X309,bool),X310),s(X309,X311)))) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f36684,plain,(
  ( ! [X310,X312,X309,X311] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X312)))) | p(s(bool,happ(s(fun(X309,bool),X310),s(X309,sK17(X310,X309))))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X312))) = s(bool,happ(s(fun(bool,bool),X312),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(X309,bool),X310),s(X309,X311)))) )),
  inference(superposition,[],[f305,f2291])).
fof(f37091,plain,(
  ( ! [X298,X297] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X298),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X297)))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X297),s('type.extreal.extreal',X298)))) ) | $spl172),
  inference(cnf_transformation,[],[f37091_D])).
fof(f37091_D,plain,(
  ( ! [X298,X297] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X298),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X297)))))) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X297),s('type.extreal.extreal',X298)))) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f36680,plain,(
  ( ! [X298,X297,X299] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X299)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X298),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X297)))))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X299))) = s(bool,happ(s(fun(bool,bool),X299),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X297),s('type.extreal.extreal',X298)))) )),
  inference(superposition,[],[f317,f2291])).
fof(f37088,plain,(
  ( ! [X218,X217] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X218))) = s(bool,happ(s(fun(bool,bool),X218),s(bool,'T_0'))) | ~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X218)))) | s(bool,'T_0') = s(bool,X217)) )),
  inference(global_subsumption,[],[f262,f261,f260,f263,f270,f277,f276,f321,f318,f335,f334,f470,f353,f350,f349,f348,f366,f373,f372,f371,f455,f369,f368,f384,f383,f382,f381,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f396,f395,f394,f393,f400,f399,f398,f397,f454,f404,f403,f402,f401,f411,f410,f409,f408,f417,f414,f413,f412,f424,f423,f422,f453,f420,f419,f418,f428,f427,f426,f425,f432,f431,f430,f429,f441,f440,f439,f452,f437,f436,f435,f445,f444,f443,f442,f449,f448,f447,f446,f471,f474,f477,f480,f482,f484,f487,f489,f491,f494,f497,f500,f503,f505,f507,f510,f513,f515,f517,f520,f522,f525,f565,f575,f578,f586,f602,f613,f622,f639,f644,f645,f647,f654,f658,f659,f662,f674,f688,f690,f698,f710,f783,f846,f878,f876,f904,f901,f947,f972,f987,f995,f1008,f1012,f1026,f1041,f1111,f1126,f1142,f1201,f1218,f1233,f1313,f1361,f1767,f1766,f1765,f1764,f2054,f2140,f2204,f2553,f2557,f2561,f2564,f2567,f2570,f2573,f2576,f2579,f2582,f2585,f2588,f2591,f2594,f2597,f2600,f2603,f2605,f2606,f2607,f2608,f2609,f2610,f2611,f2612,f2613,f2614,f2615,f2616,f2617,f2618,f2619,f2620,f2621,f2622,f2623,f2624,f2625,f2296,f2643,f2645,f2648,f2650,f2652,f2654,f2656,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2636,f328,f1840,f2064,f2039,f1797,f1796,f1795,f1794,f1793,f1792,f1791,f1790,f1789,f1788,f1787,f1786,f1785,f1784,f1783,f1782,f1781,f1780,f1779,f2635,f2634,f1296,f367,f325,f324,f2291,f2295,f2294,f2293,f2292,f2139,f2138,f2137,f2136,f2135,f327,f1866,f1867,f317,f326,f2626,f2637,f1831,f1817,f303,f2633,f2632,f2631,f2630,f2629,f2628,f2627,f469,f316,f314,f312,f310,f305,f304,f2134,f2290,f1976,f1975,f1959,f1947,f1923,f1898,f1769,f1768,f468,f281,f280,f3740,f3742,f3744,f3757,f3760,f3762,f3764,f3849,f3851,f3853,f3862,f3865,f3867,f3869,f3882,f3885,f3886,f3897,f4303,f4308,f4310,f4312,f4317,f4319,f4321,f4722,f4724,f4726,f4731,f4733,f4735,f4747,f4752,f4754,f4756,f4807,f4809,f4811,f4821,f4823,f4825,f4838,f4840,f4842,f4860,f4862,f4864,f4869,f4870,f4871,f4872,f4873,f4875,f4880,f4881,f4882,f4883,f4884,f4885,f4888,f4889,f4892,f4893,f4894,f4907,f4909,f4965,f4966,f4967,f4969,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f5004,f5006,f5007,f6917,f6931,f6948,f6967,f6996,f7007,f7009,f7045,f7069,f7071,f7090,f7091,f7092,f7093,f7094,f7124,f7130,f7147,f7153,f7173,f7179,f7193,f7199,f7208,f7209,f7223,f7240,f7242,f7244,f7261,f7263,f7265,f7288,f9309,f9324,f9325,f9417,f9398,f9399,f9484,f9465,f9466,f9489,f9490,f18495,f18500,f18505,f18466,f18467,f19646,f19651,f19656,f19617,f19618,f19906,f19911,f19916,f19878,f19879,f21428,f21429,f21430,f21457,f21458,f21459,f21931,f22080,f22041,f23768,f23773,f23780,f23727,f23728,f24311,f24316,f24323,f24270,f24271,f24460,f24461,f24462,f24450,f24463,f24464,f24465,f24466,f24467,f24468,f24469,f24470,f24471,f24472,f24473,f24474,f24475,f24476,f24477,f24478,f24479,f24480,f24481,f24482,f24483,f24484,f24485,f24486,f24487,f24488,f24489,f24506,f24507,f24508,f24523,f24538,f24539,f24540,f24555,f24596,f24613,f24585,f24621,f24648,f24649,f24650,f24672,f24673,f24674,f25900,f25921,f25973,f25992,f25993,f25994,f26033,f26034,f26035,f26075,f26076,f26077,f26088,f26089,f26090,f26105,f27929,f27930,f27931,f27946,f27968,f27969,f27970,f27991,f27992,f27993,f28005,f28006,f28007,f28031,f28032,f28033,f31391,f28156,f28157,f28226,f28227,f31209,f28261,f28262,f28272,f28300,f28301,f28557,f28559,f28560,f28561,f28562,f28563,f28564,f28565,f28566,f28567,f28568,f28569,f28570,f28571,f28572,f28573,f28574,f28575,f28576,f28577,f28578,f28579,f28580,f28581,f28582,f28583,f28584,f28585,f28586,f28587,f28596,f28597,f28598,f28599,f28600,f28601,f28602,f28603,f28604,f31756,f31760,f31766,f31768,f31819,f31820,f31823,f31824,f31826,f31828,f31830,f31832,f31834,f31837,f31838,f31841,f31842,f28949,f28951,f28952,f28953,f28954,f28955,f28956,f28957,f28958,f28959,f28960,f28961,f28962,f28963,f28964,f28965,f28966,f28967,f28968,f28969,f28970,f28971,f28972,f28973,f28974,f28975,f28976,f28977,f28978,f28979,f28988,f28989,f28990,f28991,f28992,f28993,f28994,f28995,f28996,f31844,f31859,f31860,f31861,f31862,f31863,f31864,f31865,f31866,f31867,f31868,f31869,f31870,f31871,f31872,f29145,f29152,f29155,f29156,f29157,f29158,f29159,f29160,f29161,f29162,f29163,f29164,f29165,f29166,f29167,f29168,f29169,f29170,f29171,f29172,f29173,f29174,f29175,f29184,f29185,f29186,f29187,f29195,f31922,f31923,f31924,f31925,f31926,f31927,f31928,f31929,f31930,f31931,f31932,f31933,f31934,f31935,f31936,f31937,f31938,f31939,f31940,f31941,f31942,f31943,f31944,f31945,f31946,f31947,f31948,f31952,f31953,f31954,f31955,f31956,f31957,f31958,f31959,f31960,f31961,f31962,f31963,f31964,f31965,f31966,f31967,f31968,f31969,f31970,f31971,f31972,f31973,f31974,f31975,f31976,f31977,f31978,f29393,f29395,f29396,f29397,f29398,f29399,f29400,f29401,f29402,f29403,f29404,f29405,f29406,f29407,f29408,f29409,f29410,f29411,f29412,f29413,f29414,f29415,f29416,f29417,f29418,f29419,f29420,f29421,f29422,f29423,f29432,f29433,f29434,f29435,f29436,f29437,f29438,f29439,f29440,f31989,f31990,f32041,f32044,f32045,f32046,f32047,f32048,f32049,f32052,f32055,f32125,f30120,f32469,f32470,f32471,f32472,f32473,f32474,f32475,f32476,f32477,f32478,f32479,f32480,f32481,f32482,f32483,f32484,f32485,f32486,f32487,f32488,f32489,f32490,f32491,f32492,f32493,f32494,f32495,f32496,f32497,f32498,f32499,f32500,f32501,f32502,f32503,f32504,f32505,f32506,f32507,f32508,f32509,f32510,f32511,f32512,f32513,f32514,f32515,f32516,f32517,f32518,f32519,f32520,f32521,f32522,f32523,f32524,f32525,f32526,f32527,f32528,f32529,f32691,f32692,f32693,f32696,f32698,f32699,f32700,f32732,f32735,f32740,f32824,f32825,f32826,f32829,f32831,f32832,f32833,f32865,f32868,f32873,f32932,f33004,f33066,f33132,f33183,f33335,f33469,f33623,f33852,f33850,f33916,f33923,f33951,f33954,f33959,f33989,f33992,f33997,f34045,f34048,f34053,f34210,f34211,f34212,f34215,f34217,f34218,f34219,f34272,f34273,f34274,f34277,f34279,f34280,f34281,f34411,f34412,f34413,f34416,f34418,f34419,f34420,f34493,f34494,f34495,f34498,f34500,f34501,f34502,f34520,f34521,f34522,f34525,f34527,f34528,f34529,f34549,f34550,f34551,f34554,f34556,f34557,f34558,f34578,f34579,f34580,f34583,f34585,f34586,f34587,f34664,f34665,f34666,f34669,f34671,f34672,f34673,f35269,f34731,f34762,f34763,f34766,f34767,f34769,f34770,f34771,f34924,f34925,f34930,f34945,f34980,f35440,f35443,f35444,f35445,f35446,f35447,f35448,f35449,f35450,f35451,f35452,f35453,f35454,f35455,f35456,f35457,f35458,f35459,f35460,f36038,f36032,f36018,f35621,f35789,f36202,f36204,f35923,f36266,f36267,f36287,f36290,f36291,f36292,f36296,f36321,f36322,f36323,f36324,f36325,f36327,f36332,f36333,f36334,f36335,f36336,f36337,f36341,f36345,f36346,f36362,f36365,f36366,f36367,f36368,f36369,f36370,f36371,f36372,f36373,f36374,f36375,f36376,f36377,f36378,f36445,f36466,f36469,f36470,f36471,f36475,f36476,f37031,f37026,f37012,f37010,f36505,f36506,f36542,f36543,f36544,f36545,f36546,f36548,f36553,f36554,f36555,f36556,f36557,f36558,f36562,f36566,f36567,f36583,f36584,f36587,f36588,f36589,f36590,f36591,f36592,f36593,f36594,f36595,f36596,f36597,f36598,f36599,f36600,f37057,f36989,f37060,f37062,f36618,f37063,f36619,f37064,f36620,f36621,f36622,f37067,f36623,f36624,f36625,f36626,f36627,f36628,f36629,f36630,f36631,f36632,f36633,f36634,f36635,f37080,f36636,f37081,f36968,f36966,f36660,f36661,f36662,f36663,f36664,f36963,f36961,f36959,f36957,f37087])).
fof(f37086,plain,(
  ( ! [X204,X202,X200,X203] : (p(s(bool,X202)) | s(bool,'T_0') = s(bool,X200) | s(bool,'T_0') = s(bool,X204) | s(bool,X203) = s(bool,X204) | ~sP12(X203,X202,X200)) ) | $spl170),
  inference(cnf_transformation,[],[f37086_D])).
fof(f37086_D,plain,(
  ( ! [X204,X202,X200,X203] : (p(s(bool,X202)) | s(bool,'T_0') = s(bool,X200) | s(bool,'T_0') = s(bool,X204) | s(bool,X203) = s(bool,X204) | ~sP12(X203,X202,X200)) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f36664,plain,(
  ( ! [X204,X202,X200,X203,X201] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X201)))) | p(s(bool,X202)) | ~sP12(X203,X202,X200) | s(bool,X203) = s(bool,X204) | s(bool,'T_0') = s(bool,X204) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X201))) = s(bool,happ(s(fun(bool,bool),X201),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X200)) )),
  inference(superposition,[],[f2623,f2291])).
fof(f37085,plain,(
  ( ! [X198,X199,X197,X195] : (sP12(X197,X195,X198) | s(bool,'T_0') = s(bool,X195) | s(bool,'T_0') = s(bool,X199) | s(bool,X197) = s(bool,X199)) ) | $spl168),
  inference(cnf_transformation,[],[f37085_D])).
fof(f37085_D,plain,(
  ( ! [X198,X199,X197,X195] : (sP12(X197,X195,X198) | s(bool,'T_0') = s(bool,X195) | s(bool,'T_0') = s(bool,X199) | s(bool,X197) = s(bool,X199)) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f36663,plain,(
  ( ! [X198,X196,X199,X197,X195] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X196)))) | sP12(X197,X195,X198) | s(bool,X197) = s(bool,X199) | s(bool,'T_0') = s(bool,X199) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X196))) = s(bool,happ(s(fun(bool,bool),X196),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X195)) )),
  inference(superposition,[],[f2622,f2291])).
fof(f37084,plain,(
  ( ! [X194,X192,X190,X193] : (sP10(X192,X190,X193) | s(bool,'T_0') = s(bool,X190) | s(bool,'T_0') = s(bool,X194) | s(bool,X192) = s(bool,X194)) ) | $spl166),
  inference(cnf_transformation,[],[f37084_D])).
fof(f37084_D,plain,(
  ( ! [X194,X192,X190,X193] : (sP10(X192,X190,X193) | s(bool,'T_0') = s(bool,X190) | s(bool,'T_0') = s(bool,X194) | s(bool,X192) = s(bool,X194)) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f36662,plain,(
  ( ! [X191,X194,X192,X190,X193] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X191)))) | sP10(X192,X190,X193) | s(bool,X192) = s(bool,X194) | s(bool,'T_0') = s(bool,X194) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X191))) = s(bool,happ(s(fun(bool,bool),X191),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X190)) )),
  inference(superposition,[],[f2618,f2291])).
fof(f37083,plain,(
  ( ! [X189,X187,X185,X188] : (sP10(X187,X188,X185) | s(bool,'T_0') = s(bool,X185) | s(bool,'T_0') = s(bool,X189) | s(bool,X187) = s(bool,X189)) ) | $spl164),
  inference(cnf_transformation,[],[f37083_D])).
fof(f37083_D,plain,(
  ( ! [X189,X187,X185,X188] : (sP10(X187,X188,X185) | s(bool,'T_0') = s(bool,X185) | s(bool,'T_0') = s(bool,X189) | s(bool,X187) = s(bool,X189)) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f36661,plain,(
  ( ! [X189,X187,X185,X188,X186] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X186)))) | sP10(X187,X188,X185) | s(bool,X187) = s(bool,X189) | s(bool,'T_0') = s(bool,X189) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X186))) = s(bool,happ(s(fun(bool,bool),X186),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X185)) )),
  inference(superposition,[],[f2617,f2291])).
fof(f37082,plain,(
  ( ! [X182,X180,X184,X183] : (p(s(bool,X182)) | s(bool,'T_0') = s(bool,X180) | s(bool,'T_0') = s(bool,X184) | s(bool,X183) = s(bool,X184) | ~sP8(X182,X180,X183)) ) | $spl162),
  inference(cnf_transformation,[],[f37082_D])).
fof(f37082_D,plain,(
  ( ! [X182,X180,X184,X183] : (p(s(bool,X182)) | s(bool,'T_0') = s(bool,X180) | s(bool,'T_0') = s(bool,X184) | s(bool,X183) = s(bool,X184) | ~sP8(X182,X180,X183)) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f36660,plain,(
  ( ! [X182,X180,X184,X183,X181] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X181)))) | p(s(bool,X182)) | ~sP8(X182,X180,X183) | s(bool,X183) = s(bool,X184) | s(bool,'T_0') = s(bool,X184) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X181))) = s(bool,happ(s(fun(bool,bool),X181),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X180)) )),
  inference(superposition,[],[f2615,f2291])).
fof(f36966,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X1)))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1))) = s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) )),
  inference(condensation,[],[f36965])).
fof(f36965,plain,(
  ( ! [X2,X3,X1] : (s(bool,X1) = s(bool,X2) | ~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X3)))) | s(bool,'T_0') = s(bool,X1) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) = s(bool,happ(s(fun(bool,bool),X3),s(bool,'T_0')))) )),
  inference(duplicate_literal_removal,[],[f36964])).
fof(f36964,plain,(
  ( ! [X2,X3,X1] : (s(bool,X1) = s(bool,X2) | ~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X3)))) | s(bool,'T_0') = s(bool,X1) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) = s(bool,happ(s(fun(bool,bool),X3),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(condensation,[],[f36657])).
fof(f36657,plain,(
  ( ! [X171,X169,X170,X168] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X169)))) | s(bool,X168) = s(bool,X170) | s(bool,X170) = s(bool,X171) | s(bool,'T_0') = s(bool,X171) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X169))) = s(bool,happ(s(fun(bool,bool),X169),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X168)) )),
  inference(superposition,[],[f2606,f2291])).
fof(f37081,plain,(
  ( ! [X88,X87] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X88)))) | s(bool,'T_0') = s(bool,X87) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X88))) = s(bool,happ(s(fun(bool,bool),X88),s(bool,'T_0')))) )),
  inference(global_subsumption,[],[f262,f261,f260,f263,f270,f277,f276,f321,f318,f335,f334,f470,f353,f350,f349,f348,f366,f373,f372,f371,f455,f369,f368,f384,f383,f382,f381,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f396,f395,f394,f393,f400,f399,f398,f397,f454,f404,f403,f402,f401,f411,f410,f409,f408,f417,f414,f413,f412,f424,f423,f422,f453,f420,f419,f418,f428,f427,f426,f425,f432,f431,f430,f429,f441,f440,f439,f452,f437,f436,f435,f445,f444,f443,f442,f449,f448,f447,f446,f471,f474,f477,f480,f482,f484,f487,f489,f491,f494,f497,f500,f503,f505,f507,f510,f513,f515,f517,f520,f522,f525,f565,f575,f578,f586,f602,f613,f622,f639,f644,f645,f647,f654,f658,f659,f662,f674,f688,f690,f698,f710,f783,f846,f878,f876,f904,f901,f947,f972,f987,f995,f1008,f1012,f1026,f1041,f1111,f1126,f1142,f1201,f1218,f1233,f1313,f1361,f1767,f1766,f1765,f1764,f2054,f2140,f2204,f2553,f2557,f2561,f2564,f2567,f2570,f2573,f2576,f2579,f2582,f2585,f2588,f2591,f2594,f2597,f2600,f2603,f2605,f2606,f2607,f2608,f2609,f2610,f2611,f2612,f2613,f2614,f2615,f2616,f2617,f2618,f2619,f2620,f2621,f2622,f2623,f2624,f2625,f2296,f2643,f2645,f2648,f2650,f2652,f2654,f2656,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2636,f328,f1840,f2064,f2039,f1797,f1796,f1795,f1794,f1793,f1792,f1791,f1790,f1789,f1788,f1787,f1786,f1785,f1784,f1783,f1782,f1781,f1780,f1779,f2635,f2634,f1296,f367,f325,f324,f2291,f2295,f2294,f2293,f2292,f2139,f2138,f2137,f2136,f2135,f327,f1866,f1867,f317,f326,f2626,f2637,f1831,f1817,f303,f2633,f2632,f2631,f2630,f2629,f2628,f2627,f469,f316,f314,f312,f310,f305,f304,f2134,f2290,f1976,f1975,f1959,f1947,f1923,f1898,f1769,f1768,f468,f281,f280,f3740,f3742,f3744,f3757,f3760,f3762,f3764,f3849,f3851,f3853,f3862,f3865,f3867,f3869,f3882,f3885,f3886,f3897,f4303,f4308,f4310,f4312,f4317,f4319,f4321,f4722,f4724,f4726,f4731,f4733,f4735,f4747,f4752,f4754,f4756,f4807,f4809,f4811,f4821,f4823,f4825,f4838,f4840,f4842,f4860,f4862,f4864,f4869,f4870,f4871,f4872,f4873,f4875,f4880,f4881,f4882,f4883,f4884,f4885,f4888,f4889,f4892,f4893,f4894,f4907,f4909,f4965,f4966,f4967,f4969,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f5004,f5006,f5007,f6917,f6931,f6948,f6967,f6996,f7007,f7009,f7045,f7069,f7071,f7090,f7091,f7092,f7093,f7094,f7124,f7130,f7147,f7153,f7173,f7179,f7193,f7199,f7208,f7209,f7223,f7240,f7242,f7244,f7261,f7263,f7265,f7288,f9309,f9324,f9325,f9417,f9398,f9399,f9484,f9465,f9466,f9489,f9490,f18495,f18500,f18505,f18466,f18467,f19646,f19651,f19656,f19617,f19618,f19906,f19911,f19916,f19878,f19879,f21428,f21429,f21430,f21457,f21458,f21459,f21931,f22080,f22041,f23768,f23773,f23780,f23727,f23728,f24311,f24316,f24323,f24270,f24271,f24460,f24461,f24462,f24450,f24463,f24464,f24465,f24466,f24467,f24468,f24469,f24470,f24471,f24472,f24473,f24474,f24475,f24476,f24477,f24478,f24479,f24480,f24481,f24482,f24483,f24484,f24485,f24486,f24487,f24488,f24489,f24506,f24507,f24508,f24523,f24538,f24539,f24540,f24555,f24596,f24613,f24585,f24621,f24648,f24649,f24650,f24672,f24673,f24674,f25900,f25921,f25973,f25992,f25993,f25994,f26033,f26034,f26035,f26075,f26076,f26077,f26088,f26089,f26090,f26105,f27929,f27930,f27931,f27946,f27968,f27969,f27970,f27991,f27992,f27993,f28005,f28006,f28007,f28031,f28032,f28033,f31391,f28156,f28157,f28226,f28227,f31209,f28261,f28262,f28272,f28300,f28301,f28557,f28559,f28560,f28561,f28562,f28563,f28564,f28565,f28566,f28567,f28568,f28569,f28570,f28571,f28572,f28573,f28574,f28575,f28576,f28577,f28578,f28579,f28580,f28581,f28582,f28583,f28584,f28585,f28586,f28587,f28596,f28597,f28598,f28599,f28600,f28601,f28602,f28603,f28604,f31756,f31760,f31766,f31768,f31819,f31820,f31823,f31824,f31826,f31828,f31830,f31832,f31834,f31837,f31838,f31841,f31842,f28949,f28951,f28952,f28953,f28954,f28955,f28956,f28957,f28958,f28959,f28960,f28961,f28962,f28963,f28964,f28965,f28966,f28967,f28968,f28969,f28970,f28971,f28972,f28973,f28974,f28975,f28976,f28977,f28978,f28979,f28988,f28989,f28990,f28991,f28992,f28993,f28994,f28995,f28996,f31844,f31859,f31860,f31861,f31862,f31863,f31864,f31865,f31866,f31867,f31868,f31869,f31870,f31871,f31872,f29145,f29152,f29155,f29156,f29157,f29158,f29159,f29160,f29161,f29162,f29163,f29164,f29165,f29166,f29167,f29168,f29169,f29170,f29171,f29172,f29173,f29174,f29175,f29184,f29185,f29186,f29187,f29195,f31922,f31923,f31924,f31925,f31926,f31927,f31928,f31929,f31930,f31931,f31932,f31933,f31934,f31935,f31936,f31937,f31938,f31939,f31940,f31941,f31942,f31943,f31944,f31945,f31946,f31947,f31948,f31952,f31953,f31954,f31955,f31956,f31957,f31958,f31959,f31960,f31961,f31962,f31963,f31964,f31965,f31966,f31967,f31968,f31969,f31970,f31971,f31972,f31973,f31974,f31975,f31976,f31977,f31978,f29393,f29395,f29396,f29397,f29398,f29399,f29400,f29401,f29402,f29403,f29404,f29405,f29406,f29407,f29408,f29409,f29410,f29411,f29412,f29413,f29414,f29415,f29416,f29417,f29418,f29419,f29420,f29421,f29422,f29423,f29432,f29433,f29434,f29435,f29436,f29437,f29438,f29439,f29440,f31989,f31990,f32041,f32044,f32045,f32046,f32047,f32048,f32049,f32052,f32055,f32125,f30120,f32469,f32470,f32471,f32472,f32473,f32474,f32475,f32476,f32477,f32478,f32479,f32480,f32481,f32482,f32483,f32484,f32485,f32486,f32487,f32488,f32489,f32490,f32491,f32492,f32493,f32494,f32495,f32496,f32497,f32498,f32499,f32500,f32501,f32502,f32503,f32504,f32505,f32506,f32507,f32508,f32509,f32510,f32511,f32512,f32513,f32514,f32515,f32516,f32517,f32518,f32519,f32520,f32521,f32522,f32523,f32524,f32525,f32526,f32527,f32528,f32529,f32691,f32692,f32693,f32696,f32698,f32699,f32700,f32732,f32735,f32740,f32824,f32825,f32826,f32829,f32831,f32832,f32833,f32865,f32868,f32873,f32932,f33004,f33066,f33132,f33183,f33335,f33469,f33623,f33852,f33850,f33916,f33923,f33951,f33954,f33959,f33989,f33992,f33997,f34045,f34048,f34053,f34210,f34211,f34212,f34215,f34217,f34218,f34219,f34272,f34273,f34274,f34277,f34279,f34280,f34281,f34411,f34412,f34413,f34416,f34418,f34419,f34420,f34493,f34494,f34495,f34498,f34500,f34501,f34502,f34520,f34521,f34522,f34525,f34527,f34528,f34529,f34549,f34550,f34551,f34554,f34556,f34557,f34558,f34578,f34579,f34580,f34583,f34585,f34586,f34587,f34664,f34665,f34666,f34669,f34671,f34672,f34673,f35269,f34731,f34762,f34763,f34766,f34767,f34769,f34770,f34771,f34924,f34925,f34930,f34945,f34980,f35440,f35443,f35444,f35445,f35446,f35447,f35448,f35449,f35450,f35451,f35452,f35453,f35454,f35455,f35456,f35457,f35458,f35459,f35460,f36038,f36032,f36018,f35621,f35789,f36202,f36204,f35923,f36266,f36267,f36287,f36290,f36291,f36292,f36296,f36321,f36322,f36323,f36324,f36325,f36327,f36332,f36333,f36334,f36335,f36336,f36337,f36341,f36345,f36346,f36362,f36365,f36366,f36367,f36368,f36369,f36370,f36371,f36372,f36373,f36374,f36375,f36376,f36377,f36378,f36445,f36466,f36469,f36470,f36471,f36475,f36476,f37031,f37026,f37012,f37010,f36505,f36506,f36542,f36543,f36544,f36545,f36546,f36548,f36553,f36554,f36555,f36556,f36557,f36558,f36562,f36566,f36567,f36583,f36584,f36587,f36588,f36589,f36590,f36591,f36592,f36593,f36594,f36595,f36596,f36597,f36598,f36599,f36600,f37057,f36989,f37060,f37062,f36618,f37063,f36619,f37064,f36620,f36621,f36622,f37067,f36623,f36624,f36625,f36626,f36627,f36628,f36629,f36630,f36631,f36632,f36633,f36634,f36635,f37080,f36636])).
fof(f37080,plain,(
  ( ! [X83,X84] : (s(bool,'T_0') = s(bool,X83) | ~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X84)))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X84))) = s(bool,happ(s(fun(bool,bool),X84),s(bool,'T_0')))) )),
  inference(global_subsumption,[],[f262,f261,f260,f263,f270,f277,f276,f321,f318,f335,f334,f470,f353,f350,f349,f348,f366,f373,f372,f371,f455,f369,f368,f384,f383,f382,f381,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f396,f395,f394,f393,f400,f399,f398,f397,f454,f404,f403,f402,f401,f411,f410,f409,f408,f417,f414,f413,f412,f424,f423,f422,f453,f420,f419,f418,f428,f427,f426,f425,f432,f431,f430,f429,f441,f440,f439,f452,f437,f436,f435,f445,f444,f443,f442,f449,f448,f447,f446,f471,f474,f477,f480,f482,f484,f487,f489,f491,f494,f497,f500,f503,f505,f507,f510,f513,f515,f517,f520,f522,f525,f565,f575,f578,f586,f602,f613,f622,f639,f644,f645,f647,f654,f658,f659,f662,f674,f688,f690,f698,f710,f783,f846,f878,f876,f904,f901,f947,f972,f987,f995,f1008,f1012,f1026,f1041,f1111,f1126,f1142,f1201,f1218,f1233,f1313,f1361,f1767,f1766,f1765,f1764,f2054,f2140,f2204,f2553,f2557,f2561,f2564,f2567,f2570,f2573,f2576,f2579,f2582,f2585,f2588,f2591,f2594,f2597,f2600,f2603,f2605,f2606,f2607,f2608,f2609,f2610,f2611,f2612,f2613,f2614,f2615,f2616,f2617,f2618,f2619,f2620,f2621,f2622,f2623,f2624,f2625,f2296,f2643,f2645,f2648,f2650,f2652,f2654,f2656,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2636,f328,f1840,f2064,f2039,f1797,f1796,f1795,f1794,f1793,f1792,f1791,f1790,f1789,f1788,f1787,f1786,f1785,f1784,f1783,f1782,f1781,f1780,f1779,f2635,f2634,f1296,f367,f325,f324,f2291,f2295,f2294,f2293,f2292,f2139,f2138,f2137,f2136,f2135,f327,f1866,f1867,f317,f326,f2626,f2637,f1831,f1817,f303,f2633,f2632,f2631,f2630,f2629,f2628,f2627,f469,f316,f314,f312,f310,f305,f304,f2134,f2290,f1976,f1975,f1959,f1947,f1923,f1898,f1769,f1768,f468,f281,f280,f3740,f3742,f3744,f3757,f3760,f3762,f3764,f3849,f3851,f3853,f3862,f3865,f3867,f3869,f3882,f3885,f3886,f3897,f4303,f4308,f4310,f4312,f4317,f4319,f4321,f4722,f4724,f4726,f4731,f4733,f4735,f4747,f4752,f4754,f4756,f4807,f4809,f4811,f4821,f4823,f4825,f4838,f4840,f4842,f4860,f4862,f4864,f4869,f4870,f4871,f4872,f4873,f4875,f4880,f4881,f4882,f4883,f4884,f4885,f4888,f4889,f4892,f4893,f4894,f4907,f4909,f4965,f4966,f4967,f4969,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f5004,f5006,f5007,f6917,f6931,f6948,f6967,f6996,f7007,f7009,f7045,f7069,f7071,f7090,f7091,f7092,f7093,f7094,f7124,f7130,f7147,f7153,f7173,f7179,f7193,f7199,f7208,f7209,f7223,f7240,f7242,f7244,f7261,f7263,f7265,f7288,f9309,f9324,f9325,f9417,f9398,f9399,f9484,f9465,f9466,f9489,f9490,f18495,f18500,f18505,f18466,f18467,f19646,f19651,f19656,f19617,f19618,f19906,f19911,f19916,f19878,f19879,f21428,f21429,f21430,f21457,f21458,f21459,f21931,f22080,f22041,f23768,f23773,f23780,f23727,f23728,f24311,f24316,f24323,f24270,f24271,f24460,f24461,f24462,f24450,f24463,f24464,f24465,f24466,f24467,f24468,f24469,f24470,f24471,f24472,f24473,f24474,f24475,f24476,f24477,f24478,f24479,f24480,f24481,f24482,f24483,f24484,f24485,f24486,f24487,f24488,f24489,f24506,f24507,f24508,f24523,f24538,f24539,f24540,f24555,f24596,f24613,f24585,f24621,f24648,f24649,f24650,f24672,f24673,f24674,f25900,f25921,f25973,f25992,f25993,f25994,f26033,f26034,f26035,f26075,f26076,f26077,f26088,f26089,f26090,f26105,f27929,f27930,f27931,f27946,f27968,f27969,f27970,f27991,f27992,f27993,f28005,f28006,f28007,f28031,f28032,f28033,f31391,f28156,f28157,f28226,f28227,f31209,f28261,f28262,f28272,f28300,f28301,f28557,f28559,f28560,f28561,f28562,f28563,f28564,f28565,f28566,f28567,f28568,f28569,f28570,f28571,f28572,f28573,f28574,f28575,f28576,f28577,f28578,f28579,f28580,f28581,f28582,f28583,f28584,f28585,f28586,f28587,f28596,f28597,f28598,f28599,f28600,f28601,f28602,f28603,f28604,f31756,f31760,f31766,f31768,f31819,f31820,f31823,f31824,f31826,f31828,f31830,f31832,f31834,f31837,f31838,f31841,f31842,f28949,f28951,f28952,f28953,f28954,f28955,f28956,f28957,f28958,f28959,f28960,f28961,f28962,f28963,f28964,f28965,f28966,f28967,f28968,f28969,f28970,f28971,f28972,f28973,f28974,f28975,f28976,f28977,f28978,f28979,f28988,f28989,f28990,f28991,f28992,f28993,f28994,f28995,f28996,f31844,f31859,f31860,f31861,f31862,f31863,f31864,f31865,f31866,f31867,f31868,f31869,f31870,f31871,f31872,f29145,f29152,f29155,f29156,f29157,f29158,f29159,f29160,f29161,f29162,f29163,f29164,f29165,f29166,f29167,f29168,f29169,f29170,f29171,f29172,f29173,f29174,f29175,f29184,f29185,f29186,f29187,f29195,f31922,f31923,f31924,f31925,f31926,f31927,f31928,f31929,f31930,f31931,f31932,f31933,f31934,f31935,f31936,f31937,f31938,f31939,f31940,f31941,f31942,f31943,f31944,f31945,f31946,f31947,f31948,f31952,f31953,f31954,f31955,f31956,f31957,f31958,f31959,f31960,f31961,f31962,f31963,f31964,f31965,f31966,f31967,f31968,f31969,f31970,f31971,f31972,f31973,f31974,f31975,f31976,f31977,f31978,f29393,f29395,f29396,f29397,f29398,f29399,f29400,f29401,f29402,f29403,f29404,f29405,f29406,f29407,f29408,f29409,f29410,f29411,f29412,f29413,f29414,f29415,f29416,f29417,f29418,f29419,f29420,f29421,f29422,f29423,f29432,f29433,f29434,f29435,f29436,f29437,f29438,f29439,f29440,f31989,f31990,f32041,f32044,f32045,f32046,f32047,f32048,f32049,f32052,f32055,f32125,f30120,f32469,f32470,f32471,f32472,f32473,f32474,f32475,f32476,f32477,f32478,f32479,f32480,f32481,f32482,f32483,f32484,f32485,f32486,f32487,f32488,f32489,f32490,f32491,f32492,f32493,f32494,f32495,f32496,f32497,f32498,f32499,f32500,f32501,f32502,f32503,f32504,f32505,f32506,f32507,f32508,f32509,f32510,f32511,f32512,f32513,f32514,f32515,f32516,f32517,f32518,f32519,f32520,f32521,f32522,f32523,f32524,f32525,f32526,f32527,f32528,f32529,f32691,f32692,f32693,f32696,f32698,f32699,f32700,f32732,f32735,f32740,f32824,f32825,f32826,f32829,f32831,f32832,f32833,f32865,f32868,f32873,f32932,f33004,f33066,f33132,f33183,f33335,f33469,f33623,f33852,f33850,f33916,f33923,f33951,f33954,f33959,f33989,f33992,f33997,f34045,f34048,f34053,f34210,f34211,f34212,f34215,f34217,f34218,f34219,f34272,f34273,f34274,f34277,f34279,f34280,f34281,f34411,f34412,f34413,f34416,f34418,f34419,f34420,f34493,f34494,f34495,f34498,f34500,f34501,f34502,f34520,f34521,f34522,f34525,f34527,f34528,f34529,f34549,f34550,f34551,f34554,f34556,f34557,f34558,f34578,f34579,f34580,f34583,f34585,f34586,f34587,f34664,f34665,f34666,f34669,f34671,f34672,f34673,f35269,f34731,f34762,f34763,f34766,f34767,f34769,f34770,f34771,f34924,f34925,f34930,f34945,f34980,f35440,f35443,f35444,f35445,f35446,f35447,f35448,f35449,f35450,f35451,f35452,f35453,f35454,f35455,f35456,f35457,f35458,f35459,f35460,f36038,f36032,f36018,f35621,f35789,f36202,f36204,f35923,f36266,f36267,f36287,f36290,f36291,f36292,f36296,f36321,f36322,f36323,f36324,f36325,f36327,f36332,f36333,f36334,f36335,f36336,f36337,f36341,f36345,f36346,f36362,f36365,f36366,f36367,f36368,f36369,f36370,f36371,f36372,f36373,f36374,f36375,f36376,f36377,f36378,f36445,f36466,f36469,f36470,f36471,f36475,f36476,f37031,f37026,f37012,f37010,f36505,f36506,f36542,f36543,f36544,f36545,f36546,f36548,f36553,f36554,f36555,f36556,f36557,f36558,f36562,f36566,f36567,f36583,f36584,f36587,f36588,f36589,f36590,f36591,f36592,f36593,f36594,f36595,f36596,f36597,f36598,f36599,f36600,f37057,f36989,f37060,f37062,f36618,f37063,f36619,f37064,f36620,f36621,f36622,f37067,f36623,f36624,f36625,f36626,f36627,f36628,f36629,f36630,f36631,f36632,f36633,f36634,f36635])).
fof(f37079,plain,(
  ( ! [X81,X79,X82] : (~p(s(bool,X81)) | s(bool,'T_0') = s(bool,X79) | ~sP12(X79,X82,X81) | p(s(bool,X82))) ) | $spl160),
  inference(cnf_transformation,[],[f37079_D])).
fof(f37079_D,plain,(
  ( ! [X81,X79,X82] : (~p(s(bool,X81)) | s(bool,'T_0') = s(bool,X79) | ~sP12(X79,X82,X81) | p(s(bool,X82))) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f36634,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X80)))) | ~p(s(bool,X81)) | p(s(bool,X82)) | ~sP12(X79,X82,X81) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X80))) = s(bool,happ(s(fun(bool,bool),X80),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X79)) )),
  inference(superposition,[],[f449,f2291])).
fof(f37078,plain,(
  ( ! [X78,X77,X75] : (~p(s(bool,X77)) | s(bool,'T_0') = s(bool,X75) | sP12(X75,X77,X78)) ) | $spl158),
  inference(cnf_transformation,[],[f37078_D])).
fof(f37078_D,plain,(
  ( ! [X78,X77,X75] : (~p(s(bool,X77)) | s(bool,'T_0') = s(bool,X75) | sP12(X75,X77,X78)) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f36633,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X76)))) | ~p(s(bool,X77)) | sP12(X75,X77,X78) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X76))) = s(bool,happ(s(fun(bool,bool),X76),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f448,f2291])).
fof(f37077,plain,(
  ( ! [X74,X71,X73] : (p(s(bool,X73)) | s(bool,'T_0') = s(bool,X71) | sP12(X71,X74,X73)) ) | $spl156),
  inference(cnf_transformation,[],[f37077_D])).
fof(f37077_D,plain,(
  ( ! [X74,X71,X73] : (p(s(bool,X73)) | s(bool,'T_0') = s(bool,X71) | sP12(X71,X74,X73)) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f36632,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X72)))) | p(s(bool,X73)) | sP12(X71,X74,X73) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X72))) = s(bool,happ(s(fun(bool,bool),X72),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X71)) )),
  inference(superposition,[],[f447,f2291])).
fof(f37076,plain,(
  ( ! [X70,X69,X67] : (p(s(bool,X69)) | s(bool,'T_0') = s(bool,X67) | ~sP11(X70,X69,X67) | ~p(s(bool,X70))) ) | $spl154),
  inference(cnf_transformation,[],[f37076_D])).
fof(f37076_D,plain,(
  ( ! [X70,X69,X67] : (p(s(bool,X69)) | s(bool,'T_0') = s(bool,X67) | ~sP11(X70,X69,X67) | ~p(s(bool,X70))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f36631,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X68)))) | p(s(bool,X69)) | ~p(s(bool,X70)) | ~sP11(X70,X69,X67) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X68))) = s(bool,happ(s(fun(bool,bool),X68),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X67)) )),
  inference(superposition,[],[f442,f2291])).
fof(f37075,plain,(
  ( ! [X66,X65,X63] : (p(s(bool,X65)) | s(bool,'T_0') = s(bool,X63) | ~sP10(X63,X66,X65) | p(s(bool,X66))) ) | $spl152),
  inference(cnf_transformation,[],[f37075_D])).
fof(f37075_D,plain,(
  ( ! [X66,X65,X63] : (p(s(bool,X65)) | s(bool,'T_0') = s(bool,X63) | ~sP10(X63,X66,X65) | p(s(bool,X66))) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f36630,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X64)))) | p(s(bool,X65)) | p(s(bool,X66)) | ~sP10(X63,X66,X65) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X64))) = s(bool,happ(s(fun(bool,bool),X64),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X63)) )),
  inference(superposition,[],[f432,f2291])).
fof(f37074,plain,(
  ( ! [X61,X59,X62] : (~p(s(bool,X61)) | s(bool,'T_0') = s(bool,X59) | sP10(X59,X61,X62)) ) | $spl150),
  inference(cnf_transformation,[],[f37074_D])).
fof(f37074_D,plain,(
  ( ! [X61,X59,X62] : (~p(s(bool,X61)) | s(bool,'T_0') = s(bool,X59) | sP10(X59,X61,X62)) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f36629,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X60)))) | ~p(s(bool,X61)) | sP10(X59,X61,X62) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X60))) = s(bool,happ(s(fun(bool,bool),X60),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f431,f2291])).
fof(f37073,plain,(
  ( ! [X57,X58,X55] : (~p(s(bool,X57)) | s(bool,'T_0') = s(bool,X55) | sP10(X55,X58,X57)) ) | $spl148),
  inference(cnf_transformation,[],[f37073_D])).
fof(f37073_D,plain,(
  ( ! [X57,X58,X55] : (~p(s(bool,X57)) | s(bool,'T_0') = s(bool,X55) | sP10(X55,X58,X57)) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f36628,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X56)))) | ~p(s(bool,X57)) | sP10(X55,X58,X57) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X56))) = s(bool,happ(s(fun(bool,bool),X56),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f430,f2291])).
fof(f37072,plain,(
  ( ! [X54,X53,X51] : (p(s(bool,X53)) | s(bool,'T_0') = s(bool,X51) | ~sP9(X51,X53,X54) | p(s(bool,X54))) ) | $spl146),
  inference(cnf_transformation,[],[f37072_D])).
fof(f37072_D,plain,(
  ( ! [X54,X53,X51] : (p(s(bool,X53)) | s(bool,'T_0') = s(bool,X51) | ~sP9(X51,X53,X54) | p(s(bool,X54))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f36627,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X52)))) | p(s(bool,X53)) | p(s(bool,X54)) | ~sP9(X51,X53,X54) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X52))) = s(bool,happ(s(fun(bool,bool),X52),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f425,f2291])).
fof(f37071,plain,(
  ( ! [X47,X50,X49] : (~p(s(bool,X49)) | s(bool,'T_0') = s(bool,X47) | ~sP8(X50,X49,X47) | p(s(bool,X50))) ) | $spl144),
  inference(cnf_transformation,[],[f37071_D])).
fof(f37071_D,plain,(
  ( ! [X47,X50,X49] : (~p(s(bool,X49)) | s(bool,'T_0') = s(bool,X47) | ~sP8(X50,X49,X47) | p(s(bool,X50))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f36626,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X48)))) | ~p(s(bool,X49)) | p(s(bool,X50)) | ~sP8(X50,X49,X47) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X48))) = s(bool,happ(s(fun(bool,bool),X48),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f417,f2291])).
fof(f37070,plain,(
  ( ! [X45,X43,X46] : (~p(s(bool,X45)) | s(bool,'T_0') = s(bool,X43) | ~sP7(X45,X43,X46) | p(s(bool,X46))) ) | $spl142),
  inference(cnf_transformation,[],[f37070_D])).
fof(f37070_D,plain,(
  ( ! [X45,X43,X46] : (~p(s(bool,X45)) | s(bool,'T_0') = s(bool,X43) | ~sP7(X45,X43,X46) | p(s(bool,X46))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f36625,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X44)))) | ~p(s(bool,X45)) | p(s(bool,X46)) | ~sP7(X45,X43,X46) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X44))) = s(bool,happ(s(fun(bool,bool),X44),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X43)) )),
  inference(superposition,[],[f408,f2291])).
fof(f37069,plain,(
  ( ! [X39,X41,X42] : (~p(s(bool,X41)) | s(bool,'T_0') = s(bool,X39) | ~sP4(X41,X39,X42) | p(s(bool,X42))) ) | $spl140),
  inference(cnf_transformation,[],[f37069_D])).
fof(f37069_D,plain,(
  ( ! [X39,X41,X42] : (~p(s(bool,X41)) | s(bool,'T_0') = s(bool,X39) | ~sP4(X41,X39,X42) | p(s(bool,X42))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f36624,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X40)))) | ~p(s(bool,X41)) | p(s(bool,X42)) | ~sP4(X41,X39,X42) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X40))) = s(bool,happ(s(fun(bool,bool),X40),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f389,f2291])).
fof(f37068,plain,(
  ( ! [X37,X35,X38] : (~p(s(bool,X37)) | s(bool,'T_0') = s(bool,X35) | ~sP5(X37,X35,X38) | p(s(bool,X38))) ) | $spl138),
  inference(cnf_transformation,[],[f37068_D])).
fof(f37068_D,plain,(
  ( ! [X37,X35,X38] : (~p(s(bool,X37)) | s(bool,'T_0') = s(bool,X35) | ~sP5(X37,X35,X38) | p(s(bool,X38))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f36623,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X36)))) | ~p(s(bool,X37)) | p(s(bool,X38)) | ~sP5(X37,X35,X38) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X36))) = s(bool,happ(s(fun(bool,bool),X36),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f385,f2291])).
fof(f37067,plain,(
  ( ! [X31,X32] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X32)))) | s(bool,'T_0') = s(bool,X31) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X32))) = s(bool,happ(s(fun(bool,bool),X32),s(bool,'T_0')))) )),
  inference(global_subsumption,[],[f262,f261,f260,f263,f270,f277,f276,f321,f318,f335,f334,f470,f353,f350,f349,f348,f366,f373,f372,f371,f455,f369,f368,f384,f383,f382,f381,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f396,f395,f394,f393,f400,f399,f398,f397,f454,f404,f403,f402,f401,f411,f410,f409,f408,f417,f414,f413,f412,f424,f423,f422,f453,f420,f419,f418,f428,f427,f426,f425,f432,f431,f430,f429,f441,f440,f439,f452,f437,f436,f435,f445,f444,f443,f442,f449,f448,f447,f446,f471,f474,f477,f480,f482,f484,f487,f489,f491,f494,f497,f500,f503,f505,f507,f510,f513,f515,f517,f520,f522,f525,f565,f575,f578,f586,f602,f613,f622,f639,f644,f645,f647,f654,f658,f659,f662,f674,f688,f690,f698,f710,f783,f846,f878,f876,f904,f901,f947,f972,f987,f995,f1008,f1012,f1026,f1041,f1111,f1126,f1142,f1201,f1218,f1233,f1313,f1361,f1767,f1766,f1765,f1764,f2054,f2140,f2204,f2553,f2557,f2561,f2564,f2567,f2570,f2573,f2576,f2579,f2582,f2585,f2588,f2591,f2594,f2597,f2600,f2603,f2605,f2606,f2607,f2608,f2609,f2610,f2611,f2612,f2613,f2614,f2615,f2616,f2617,f2618,f2619,f2620,f2621,f2622,f2623,f2624,f2625,f2296,f2643,f2645,f2648,f2650,f2652,f2654,f2656,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2636,f328,f1840,f2064,f2039,f1797,f1796,f1795,f1794,f1793,f1792,f1791,f1790,f1789,f1788,f1787,f1786,f1785,f1784,f1783,f1782,f1781,f1780,f1779,f2635,f2634,f1296,f367,f325,f324,f2291,f2295,f2294,f2293,f2292,f2139,f2138,f2137,f2136,f2135,f327,f1866,f1867,f317,f326,f2626,f2637,f1831,f1817,f303,f2633,f2632,f2631,f2630,f2629,f2628,f2627,f469,f316,f314,f312,f310,f305,f304,f2134,f2290,f1976,f1975,f1959,f1947,f1923,f1898,f1769,f1768,f468,f281,f280,f3740,f3742,f3744,f3757,f3760,f3762,f3764,f3849,f3851,f3853,f3862,f3865,f3867,f3869,f3882,f3885,f3886,f3897,f4303,f4308,f4310,f4312,f4317,f4319,f4321,f4722,f4724,f4726,f4731,f4733,f4735,f4747,f4752,f4754,f4756,f4807,f4809,f4811,f4821,f4823,f4825,f4838,f4840,f4842,f4860,f4862,f4864,f4869,f4870,f4871,f4872,f4873,f4875,f4880,f4881,f4882,f4883,f4884,f4885,f4888,f4889,f4892,f4893,f4894,f4907,f4909,f4965,f4966,f4967,f4969,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f5004,f5006,f5007,f6917,f6931,f6948,f6967,f6996,f7007,f7009,f7045,f7069,f7071,f7090,f7091,f7092,f7093,f7094,f7124,f7130,f7147,f7153,f7173,f7179,f7193,f7199,f7208,f7209,f7223,f7240,f7242,f7244,f7261,f7263,f7265,f7288,f9309,f9324,f9325,f9417,f9398,f9399,f9484,f9465,f9466,f9489,f9490,f18495,f18500,f18505,f18466,f18467,f19646,f19651,f19656,f19617,f19618,f19906,f19911,f19916,f19878,f19879,f21428,f21429,f21430,f21457,f21458,f21459,f21931,f22080,f22041,f23768,f23773,f23780,f23727,f23728,f24311,f24316,f24323,f24270,f24271,f24460,f24461,f24462,f24450,f24463,f24464,f24465,f24466,f24467,f24468,f24469,f24470,f24471,f24472,f24473,f24474,f24475,f24476,f24477,f24478,f24479,f24480,f24481,f24482,f24483,f24484,f24485,f24486,f24487,f24488,f24489,f24506,f24507,f24508,f24523,f24538,f24539,f24540,f24555,f24596,f24613,f24585,f24621,f24648,f24649,f24650,f24672,f24673,f24674,f25900,f25921,f25973,f25992,f25993,f25994,f26033,f26034,f26035,f26075,f26076,f26077,f26088,f26089,f26090,f26105,f27929,f27930,f27931,f27946,f27968,f27969,f27970,f27991,f27992,f27993,f28005,f28006,f28007,f28031,f28032,f28033,f31391,f28156,f28157,f28226,f28227,f31209,f28261,f28262,f28272,f28300,f28301,f28557,f28559,f28560,f28561,f28562,f28563,f28564,f28565,f28566,f28567,f28568,f28569,f28570,f28571,f28572,f28573,f28574,f28575,f28576,f28577,f28578,f28579,f28580,f28581,f28582,f28583,f28584,f28585,f28586,f28587,f28596,f28597,f28598,f28599,f28600,f28601,f28602,f28603,f28604,f31756,f31760,f31766,f31768,f31819,f31820,f31823,f31824,f31826,f31828,f31830,f31832,f31834,f31837,f31838,f31841,f31842,f28949,f28951,f28952,f28953,f28954,f28955,f28956,f28957,f28958,f28959,f28960,f28961,f28962,f28963,f28964,f28965,f28966,f28967,f28968,f28969,f28970,f28971,f28972,f28973,f28974,f28975,f28976,f28977,f28978,f28979,f28988,f28989,f28990,f28991,f28992,f28993,f28994,f28995,f28996,f31844,f31859,f31860,f31861,f31862,f31863,f31864,f31865,f31866,f31867,f31868,f31869,f31870,f31871,f31872,f29145,f29152,f29155,f29156,f29157,f29158,f29159,f29160,f29161,f29162,f29163,f29164,f29165,f29166,f29167,f29168,f29169,f29170,f29171,f29172,f29173,f29174,f29175,f29184,f29185,f29186,f29187,f29195,f31922,f31923,f31924,f31925,f31926,f31927,f31928,f31929,f31930,f31931,f31932,f31933,f31934,f31935,f31936,f31937,f31938,f31939,f31940,f31941,f31942,f31943,f31944,f31945,f31946,f31947,f31948,f31952,f31953,f31954,f31955,f31956,f31957,f31958,f31959,f31960,f31961,f31962,f31963,f31964,f31965,f31966,f31967,f31968,f31969,f31970,f31971,f31972,f31973,f31974,f31975,f31976,f31977,f31978,f29393,f29395,f29396,f29397,f29398,f29399,f29400,f29401,f29402,f29403,f29404,f29405,f29406,f29407,f29408,f29409,f29410,f29411,f29412,f29413,f29414,f29415,f29416,f29417,f29418,f29419,f29420,f29421,f29422,f29423,f29432,f29433,f29434,f29435,f29436,f29437,f29438,f29439,f29440,f31989,f31990,f32041,f32044,f32045,f32046,f32047,f32048,f32049,f32052,f32055,f32125,f30120,f32469,f32470,f32471,f32472,f32473,f32474,f32475,f32476,f32477,f32478,f32479,f32480,f32481,f32482,f32483,f32484,f32485,f32486,f32487,f32488,f32489,f32490,f32491,f32492,f32493,f32494,f32495,f32496,f32497,f32498,f32499,f32500,f32501,f32502,f32503,f32504,f32505,f32506,f32507,f32508,f32509,f32510,f32511,f32512,f32513,f32514,f32515,f32516,f32517,f32518,f32519,f32520,f32521,f32522,f32523,f32524,f32525,f32526,f32527,f32528,f32529,f32691,f32692,f32693,f32696,f32698,f32699,f32700,f32732,f32735,f32740,f32824,f32825,f32826,f32829,f32831,f32832,f32833,f32865,f32868,f32873,f32932,f33004,f33066,f33132,f33183,f33335,f33469,f33623,f33852,f33850,f33916,f33923,f33951,f33954,f33959,f33989,f33992,f33997,f34045,f34048,f34053,f34210,f34211,f34212,f34215,f34217,f34218,f34219,f34272,f34273,f34274,f34277,f34279,f34280,f34281,f34411,f34412,f34413,f34416,f34418,f34419,f34420,f34493,f34494,f34495,f34498,f34500,f34501,f34502,f34520,f34521,f34522,f34525,f34527,f34528,f34529,f34549,f34550,f34551,f34554,f34556,f34557,f34558,f34578,f34579,f34580,f34583,f34585,f34586,f34587,f34664,f34665,f34666,f34669,f34671,f34672,f34673,f35269,f34731,f34762,f34763,f34766,f34767,f34769,f34770,f34771,f34924,f34925,f34930,f34945,f34980,f35440,f35443,f35444,f35445,f35446,f35447,f35448,f35449,f35450,f35451,f35452,f35453,f35454,f35455,f35456,f35457,f35458,f35459,f35460,f36038,f36032,f36018,f35621,f35789,f36202,f36204,f35923,f36266,f36267,f36287,f36290,f36291,f36292,f36296,f36321,f36322,f36323,f36324,f36325,f36327,f36332,f36333,f36334,f36335,f36336,f36337,f36341,f36345,f36346,f36362,f36365,f36366,f36367,f36368,f36369,f36370,f36371,f36372,f36373,f36374,f36375,f36376,f36377,f36378,f36445,f36466,f36469,f36470,f36471,f36475,f36476,f37031,f37026,f37012,f37010,f36505,f36506,f36542,f36543,f36544,f36545,f36546,f36548,f36553,f36554,f36555,f36556,f36557,f36558,f36562,f36566,f36567,f36583,f36584,f36587,f36588,f36589,f36590,f36591,f36592,f36593,f36594,f36595,f36596,f36597,f36598,f36599,f36600,f37057,f36989,f37060,f37062,f36618,f37063,f36619,f37064,f36620,f36621,f36622])).
fof(f37066,plain,(
  ( ! [X30,X28] : (p(s(bool,X30)) | s(bool,'T_0') = s(bool,X28) | sP1(X30,X28)) ) | $spl136),
  inference(cnf_transformation,[],[f37066_D])).
fof(f37066_D,plain,(
  ( ! [X30,X28] : (p(s(bool,X30)) | s(bool,'T_0') = s(bool,X28) | sP1(X30,X28)) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f36621,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X29)))) | p(s(bool,X30)) | sP1(X30,X28) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X29))) = s(bool,happ(s(fun(bool,bool),X29),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f353,f2291])).
fof(f37065,plain,(
  ( ! [X27,X25] : (p(s(bool,X27)) | s(bool,'T_0') = s(bool,X25) | sP1(X25,X27)) ) | $spl134),
  inference(cnf_transformation,[],[f37065_D])).
fof(f37065_D,plain,(
  ( ! [X27,X25] : (p(s(bool,X27)) | s(bool,'T_0') = s(bool,X25) | sP1(X25,X27)) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f36620,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X26)))) | p(s(bool,X27)) | sP1(X25,X27) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X26))) = s(bool,happ(s(fun(bool,bool),X26),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f350,f2291])).
fof(f37064,plain,(
  ( ! [X23,X22] : (s(bool,'T_0') = s(bool,X22) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X23))) = s(bool,happ(s(fun(bool,bool),X23),s(bool,'T_0'))) | ~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X23))))) )),
  inference(global_subsumption,[],[f262,f261,f260,f263,f270,f277,f276,f321,f318,f335,f334,f470,f353,f350,f349,f348,f366,f373,f372,f371,f455,f369,f368,f384,f383,f382,f381,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f396,f395,f394,f393,f400,f399,f398,f397,f454,f404,f403,f402,f401,f411,f410,f409,f408,f417,f414,f413,f412,f424,f423,f422,f453,f420,f419,f418,f428,f427,f426,f425,f432,f431,f430,f429,f441,f440,f439,f452,f437,f436,f435,f445,f444,f443,f442,f449,f448,f447,f446,f471,f474,f477,f480,f482,f484,f487,f489,f491,f494,f497,f500,f503,f505,f507,f510,f513,f515,f517,f520,f522,f525,f565,f575,f578,f586,f602,f613,f622,f639,f644,f645,f647,f654,f658,f659,f662,f674,f688,f690,f698,f710,f783,f846,f878,f876,f904,f901,f947,f972,f987,f995,f1008,f1012,f1026,f1041,f1111,f1126,f1142,f1201,f1218,f1233,f1313,f1361,f1767,f1766,f1765,f1764,f2054,f2140,f2204,f2553,f2557,f2561,f2564,f2567,f2570,f2573,f2576,f2579,f2582,f2585,f2588,f2591,f2594,f2597,f2600,f2603,f2605,f2606,f2607,f2608,f2609,f2610,f2611,f2612,f2613,f2614,f2615,f2616,f2617,f2618,f2619,f2620,f2621,f2622,f2623,f2624,f2625,f2296,f2643,f2645,f2648,f2650,f2652,f2654,f2656,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2636,f328,f1840,f2064,f2039,f1797,f1796,f1795,f1794,f1793,f1792,f1791,f1790,f1789,f1788,f1787,f1786,f1785,f1784,f1783,f1782,f1781,f1780,f1779,f2635,f2634,f1296,f367,f325,f324,f2291,f2295,f2294,f2293,f2292,f2139,f2138,f2137,f2136,f2135,f327,f1866,f1867,f317,f326,f2626,f2637,f1831,f1817,f303,f2633,f2632,f2631,f2630,f2629,f2628,f2627,f469,f316,f314,f312,f310,f305,f304,f2134,f2290,f1976,f1975,f1959,f1947,f1923,f1898,f1769,f1768,f468,f281,f280,f3740,f3742,f3744,f3757,f3760,f3762,f3764,f3849,f3851,f3853,f3862,f3865,f3867,f3869,f3882,f3885,f3886,f3897,f4303,f4308,f4310,f4312,f4317,f4319,f4321,f4722,f4724,f4726,f4731,f4733,f4735,f4747,f4752,f4754,f4756,f4807,f4809,f4811,f4821,f4823,f4825,f4838,f4840,f4842,f4860,f4862,f4864,f4869,f4870,f4871,f4872,f4873,f4875,f4880,f4881,f4882,f4883,f4884,f4885,f4888,f4889,f4892,f4893,f4894,f4907,f4909,f4965,f4966,f4967,f4969,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f5004,f5006,f5007,f6917,f6931,f6948,f6967,f6996,f7007,f7009,f7045,f7069,f7071,f7090,f7091,f7092,f7093,f7094,f7124,f7130,f7147,f7153,f7173,f7179,f7193,f7199,f7208,f7209,f7223,f7240,f7242,f7244,f7261,f7263,f7265,f7288,f9309,f9324,f9325,f9417,f9398,f9399,f9484,f9465,f9466,f9489,f9490,f18495,f18500,f18505,f18466,f18467,f19646,f19651,f19656,f19617,f19618,f19906,f19911,f19916,f19878,f19879,f21428,f21429,f21430,f21457,f21458,f21459,f21931,f22080,f22041,f23768,f23773,f23780,f23727,f23728,f24311,f24316,f24323,f24270,f24271,f24460,f24461,f24462,f24450,f24463,f24464,f24465,f24466,f24467,f24468,f24469,f24470,f24471,f24472,f24473,f24474,f24475,f24476,f24477,f24478,f24479,f24480,f24481,f24482,f24483,f24484,f24485,f24486,f24487,f24488,f24489,f24506,f24507,f24508,f24523,f24538,f24539,f24540,f24555,f24596,f24613,f24585,f24621,f24648,f24649,f24650,f24672,f24673,f24674,f25900,f25921,f25973,f25992,f25993,f25994,f26033,f26034,f26035,f26075,f26076,f26077,f26088,f26089,f26090,f26105,f27929,f27930,f27931,f27946,f27968,f27969,f27970,f27991,f27992,f27993,f28005,f28006,f28007,f28031,f28032,f28033,f31391,f28156,f28157,f28226,f28227,f31209,f28261,f28262,f28272,f28300,f28301,f28557,f28559,f28560,f28561,f28562,f28563,f28564,f28565,f28566,f28567,f28568,f28569,f28570,f28571,f28572,f28573,f28574,f28575,f28576,f28577,f28578,f28579,f28580,f28581,f28582,f28583,f28584,f28585,f28586,f28587,f28596,f28597,f28598,f28599,f28600,f28601,f28602,f28603,f28604,f31756,f31760,f31766,f31768,f31819,f31820,f31823,f31824,f31826,f31828,f31830,f31832,f31834,f31837,f31838,f31841,f31842,f28949,f28951,f28952,f28953,f28954,f28955,f28956,f28957,f28958,f28959,f28960,f28961,f28962,f28963,f28964,f28965,f28966,f28967,f28968,f28969,f28970,f28971,f28972,f28973,f28974,f28975,f28976,f28977,f28978,f28979,f28988,f28989,f28990,f28991,f28992,f28993,f28994,f28995,f28996,f31844,f31859,f31860,f31861,f31862,f31863,f31864,f31865,f31866,f31867,f31868,f31869,f31870,f31871,f31872,f29145,f29152,f29155,f29156,f29157,f29158,f29159,f29160,f29161,f29162,f29163,f29164,f29165,f29166,f29167,f29168,f29169,f29170,f29171,f29172,f29173,f29174,f29175,f29184,f29185,f29186,f29187,f29195,f31922,f31923,f31924,f31925,f31926,f31927,f31928,f31929,f31930,f31931,f31932,f31933,f31934,f31935,f31936,f31937,f31938,f31939,f31940,f31941,f31942,f31943,f31944,f31945,f31946,f31947,f31948,f31952,f31953,f31954,f31955,f31956,f31957,f31958,f31959,f31960,f31961,f31962,f31963,f31964,f31965,f31966,f31967,f31968,f31969,f31970,f31971,f31972,f31973,f31974,f31975,f31976,f31977,f31978,f29393,f29395,f29396,f29397,f29398,f29399,f29400,f29401,f29402,f29403,f29404,f29405,f29406,f29407,f29408,f29409,f29410,f29411,f29412,f29413,f29414,f29415,f29416,f29417,f29418,f29419,f29420,f29421,f29422,f29423,f29432,f29433,f29434,f29435,f29436,f29437,f29438,f29439,f29440,f31989,f31990,f32041,f32044,f32045,f32046,f32047,f32048,f32049,f32052,f32055,f32125,f30120,f32469,f32470,f32471,f32472,f32473,f32474,f32475,f32476,f32477,f32478,f32479,f32480,f32481,f32482,f32483,f32484,f32485,f32486,f32487,f32488,f32489,f32490,f32491,f32492,f32493,f32494,f32495,f32496,f32497,f32498,f32499,f32500,f32501,f32502,f32503,f32504,f32505,f32506,f32507,f32508,f32509,f32510,f32511,f32512,f32513,f32514,f32515,f32516,f32517,f32518,f32519,f32520,f32521,f32522,f32523,f32524,f32525,f32526,f32527,f32528,f32529,f32691,f32692,f32693,f32696,f32698,f32699,f32700,f32732,f32735,f32740,f32824,f32825,f32826,f32829,f32831,f32832,f32833,f32865,f32868,f32873,f32932,f33004,f33066,f33132,f33183,f33335,f33469,f33623,f33852,f33850,f33916,f33923,f33951,f33954,f33959,f33989,f33992,f33997,f34045,f34048,f34053,f34210,f34211,f34212,f34215,f34217,f34218,f34219,f34272,f34273,f34274,f34277,f34279,f34280,f34281,f34411,f34412,f34413,f34416,f34418,f34419,f34420,f34493,f34494,f34495,f34498,f34500,f34501,f34502,f34520,f34521,f34522,f34525,f34527,f34528,f34529,f34549,f34550,f34551,f34554,f34556,f34557,f34558,f34578,f34579,f34580,f34583,f34585,f34586,f34587,f34664,f34665,f34666,f34669,f34671,f34672,f34673,f35269,f34731,f34762,f34763,f34766,f34767,f34769,f34770,f34771,f34924,f34925,f34930,f34945,f34980,f35440,f35443,f35444,f35445,f35446,f35447,f35448,f35449,f35450,f35451,f35452,f35453,f35454,f35455,f35456,f35457,f35458,f35459,f35460,f36038,f36032,f36018,f35621,f35789,f36202,f36204,f35923,f36266,f36267,f36287,f36290,f36291,f36292,f36296,f36321,f36322,f36323,f36324,f36325,f36327,f36332,f36333,f36334,f36335,f36336,f36337,f36341,f36345,f36346,f36362,f36365,f36366,f36367,f36368,f36369,f36370,f36371,f36372,f36373,f36374,f36375,f36376,f36377,f36378,f36445,f36466,f36469,f36470,f36471,f36475,f36476,f37031,f37026,f37012,f37010,f36505,f36506,f36542,f36543,f36544,f36545,f36546,f36548,f36553,f36554,f36555,f36556,f36557,f36558,f36562,f36566,f36567,f36583,f36584,f36587,f36588,f36589,f36590,f36591,f36592,f36593,f36594,f36595,f36596,f36597,f36598,f36599,f36600,f37057,f36989,f37060,f37062,f36618,f37063,f36619])).
fof(f37063,plain,(
  ( ! [X19,X20] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X20))) = s(bool,happ(s(fun(bool,bool),X20),s(bool,'T_0'))) | ~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X20)))) | s(bool,'T_0') = s(bool,X19)) )),
  inference(global_subsumption,[],[f262,f261,f260,f263,f270,f277,f276,f321,f318,f335,f334,f470,f353,f350,f349,f348,f366,f373,f372,f371,f455,f369,f368,f384,f383,f382,f381,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f396,f395,f394,f393,f400,f399,f398,f397,f454,f404,f403,f402,f401,f411,f410,f409,f408,f417,f414,f413,f412,f424,f423,f422,f453,f420,f419,f418,f428,f427,f426,f425,f432,f431,f430,f429,f441,f440,f439,f452,f437,f436,f435,f445,f444,f443,f442,f449,f448,f447,f446,f471,f474,f477,f480,f482,f484,f487,f489,f491,f494,f497,f500,f503,f505,f507,f510,f513,f515,f517,f520,f522,f525,f565,f575,f578,f586,f602,f613,f622,f639,f644,f645,f647,f654,f658,f659,f662,f674,f688,f690,f698,f710,f783,f846,f878,f876,f904,f901,f947,f972,f987,f995,f1008,f1012,f1026,f1041,f1111,f1126,f1142,f1201,f1218,f1233,f1313,f1361,f1767,f1766,f1765,f1764,f2054,f2140,f2204,f2553,f2557,f2561,f2564,f2567,f2570,f2573,f2576,f2579,f2582,f2585,f2588,f2591,f2594,f2597,f2600,f2603,f2605,f2606,f2607,f2608,f2609,f2610,f2611,f2612,f2613,f2614,f2615,f2616,f2617,f2618,f2619,f2620,f2621,f2622,f2623,f2624,f2625,f2296,f2643,f2645,f2648,f2650,f2652,f2654,f2656,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2636,f328,f1840,f2064,f2039,f1797,f1796,f1795,f1794,f1793,f1792,f1791,f1790,f1789,f1788,f1787,f1786,f1785,f1784,f1783,f1782,f1781,f1780,f1779,f2635,f2634,f1296,f367,f325,f324,f2291,f2295,f2294,f2293,f2292,f2139,f2138,f2137,f2136,f2135,f327,f1866,f1867,f317,f326,f2626,f2637,f1831,f1817,f303,f2633,f2632,f2631,f2630,f2629,f2628,f2627,f469,f316,f314,f312,f310,f305,f304,f2134,f2290,f1976,f1975,f1959,f1947,f1923,f1898,f1769,f1768,f468,f281,f280,f3740,f3742,f3744,f3757,f3760,f3762,f3764,f3849,f3851,f3853,f3862,f3865,f3867,f3869,f3882,f3885,f3886,f3897,f4303,f4308,f4310,f4312,f4317,f4319,f4321,f4722,f4724,f4726,f4731,f4733,f4735,f4747,f4752,f4754,f4756,f4807,f4809,f4811,f4821,f4823,f4825,f4838,f4840,f4842,f4860,f4862,f4864,f4869,f4870,f4871,f4872,f4873,f4875,f4880,f4881,f4882,f4883,f4884,f4885,f4888,f4889,f4892,f4893,f4894,f4907,f4909,f4965,f4966,f4967,f4969,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f5004,f5006,f5007,f6917,f6931,f6948,f6967,f6996,f7007,f7009,f7045,f7069,f7071,f7090,f7091,f7092,f7093,f7094,f7124,f7130,f7147,f7153,f7173,f7179,f7193,f7199,f7208,f7209,f7223,f7240,f7242,f7244,f7261,f7263,f7265,f7288,f9309,f9324,f9325,f9417,f9398,f9399,f9484,f9465,f9466,f9489,f9490,f18495,f18500,f18505,f18466,f18467,f19646,f19651,f19656,f19617,f19618,f19906,f19911,f19916,f19878,f19879,f21428,f21429,f21430,f21457,f21458,f21459,f21931,f22080,f22041,f23768,f23773,f23780,f23727,f23728,f24311,f24316,f24323,f24270,f24271,f24460,f24461,f24462,f24450,f24463,f24464,f24465,f24466,f24467,f24468,f24469,f24470,f24471,f24472,f24473,f24474,f24475,f24476,f24477,f24478,f24479,f24480,f24481,f24482,f24483,f24484,f24485,f24486,f24487,f24488,f24489,f24506,f24507,f24508,f24523,f24538,f24539,f24540,f24555,f24596,f24613,f24585,f24621,f24648,f24649,f24650,f24672,f24673,f24674,f25900,f25921,f25973,f25992,f25993,f25994,f26033,f26034,f26035,f26075,f26076,f26077,f26088,f26089,f26090,f26105,f27929,f27930,f27931,f27946,f27968,f27969,f27970,f27991,f27992,f27993,f28005,f28006,f28007,f28031,f28032,f28033,f31391,f28156,f28157,f28226,f28227,f31209,f28261,f28262,f28272,f28300,f28301,f28557,f28559,f28560,f28561,f28562,f28563,f28564,f28565,f28566,f28567,f28568,f28569,f28570,f28571,f28572,f28573,f28574,f28575,f28576,f28577,f28578,f28579,f28580,f28581,f28582,f28583,f28584,f28585,f28586,f28587,f28596,f28597,f28598,f28599,f28600,f28601,f28602,f28603,f28604,f31756,f31760,f31766,f31768,f31819,f31820,f31823,f31824,f31826,f31828,f31830,f31832,f31834,f31837,f31838,f31841,f31842,f28949,f28951,f28952,f28953,f28954,f28955,f28956,f28957,f28958,f28959,f28960,f28961,f28962,f28963,f28964,f28965,f28966,f28967,f28968,f28969,f28970,f28971,f28972,f28973,f28974,f28975,f28976,f28977,f28978,f28979,f28988,f28989,f28990,f28991,f28992,f28993,f28994,f28995,f28996,f31844,f31859,f31860,f31861,f31862,f31863,f31864,f31865,f31866,f31867,f31868,f31869,f31870,f31871,f31872,f29145,f29152,f29155,f29156,f29157,f29158,f29159,f29160,f29161,f29162,f29163,f29164,f29165,f29166,f29167,f29168,f29169,f29170,f29171,f29172,f29173,f29174,f29175,f29184,f29185,f29186,f29187,f29195,f31922,f31923,f31924,f31925,f31926,f31927,f31928,f31929,f31930,f31931,f31932,f31933,f31934,f31935,f31936,f31937,f31938,f31939,f31940,f31941,f31942,f31943,f31944,f31945,f31946,f31947,f31948,f31952,f31953,f31954,f31955,f31956,f31957,f31958,f31959,f31960,f31961,f31962,f31963,f31964,f31965,f31966,f31967,f31968,f31969,f31970,f31971,f31972,f31973,f31974,f31975,f31976,f31977,f31978,f29393,f29395,f29396,f29397,f29398,f29399,f29400,f29401,f29402,f29403,f29404,f29405,f29406,f29407,f29408,f29409,f29410,f29411,f29412,f29413,f29414,f29415,f29416,f29417,f29418,f29419,f29420,f29421,f29422,f29423,f29432,f29433,f29434,f29435,f29436,f29437,f29438,f29439,f29440,f31989,f31990,f32041,f32044,f32045,f32046,f32047,f32048,f32049,f32052,f32055,f32125,f30120,f32469,f32470,f32471,f32472,f32473,f32474,f32475,f32476,f32477,f32478,f32479,f32480,f32481,f32482,f32483,f32484,f32485,f32486,f32487,f32488,f32489,f32490,f32491,f32492,f32493,f32494,f32495,f32496,f32497,f32498,f32499,f32500,f32501,f32502,f32503,f32504,f32505,f32506,f32507,f32508,f32509,f32510,f32511,f32512,f32513,f32514,f32515,f32516,f32517,f32518,f32519,f32520,f32521,f32522,f32523,f32524,f32525,f32526,f32527,f32528,f32529,f32691,f32692,f32693,f32696,f32698,f32699,f32700,f32732,f32735,f32740,f32824,f32825,f32826,f32829,f32831,f32832,f32833,f32865,f32868,f32873,f32932,f33004,f33066,f33132,f33183,f33335,f33469,f33623,f33852,f33850,f33916,f33923,f33951,f33954,f33959,f33989,f33992,f33997,f34045,f34048,f34053,f34210,f34211,f34212,f34215,f34217,f34218,f34219,f34272,f34273,f34274,f34277,f34279,f34280,f34281,f34411,f34412,f34413,f34416,f34418,f34419,f34420,f34493,f34494,f34495,f34498,f34500,f34501,f34502,f34520,f34521,f34522,f34525,f34527,f34528,f34529,f34549,f34550,f34551,f34554,f34556,f34557,f34558,f34578,f34579,f34580,f34583,f34585,f34586,f34587,f34664,f34665,f34666,f34669,f34671,f34672,f34673,f35269,f34731,f34762,f34763,f34766,f34767,f34769,f34770,f34771,f34924,f34925,f34930,f34945,f34980,f35440,f35443,f35444,f35445,f35446,f35447,f35448,f35449,f35450,f35451,f35452,f35453,f35454,f35455,f35456,f35457,f35458,f35459,f35460,f36038,f36032,f36018,f35621,f35789,f36202,f36204,f35923,f36266,f36267,f36287,f36290,f36291,f36292,f36296,f36321,f36322,f36323,f36324,f36325,f36327,f36332,f36333,f36334,f36335,f36336,f36337,f36341,f36345,f36346,f36362,f36365,f36366,f36367,f36368,f36369,f36370,f36371,f36372,f36373,f36374,f36375,f36376,f36377,f36378,f36445,f36466,f36469,f36470,f36471,f36475,f36476,f37031,f37026,f37012,f37010,f36505,f36506,f36542,f36543,f36544,f36545,f36546,f36548,f36553,f36554,f36555,f36556,f36557,f36558,f36562,f36566,f36567,f36583,f36584,f36587,f36588,f36589,f36590,f36591,f36592,f36593,f36594,f36595,f36596,f36597,f36598,f36599,f36600,f37057,f36989,f37060,f37062,f36618])).
fof(f37062,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X1)))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1))) = s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) )),
  inference(condensation,[],[f37061])).
fof(f37061,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X17)))) | s(bool,X16) = s(bool,X18) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X17))) = s(bool,happ(s(fun(bool,bool),X17),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(subsumption_resolution,[],[f36617,f846])).
fof(f36617,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X17)))) | ~p(s(bool,X18)) | s(bool,X16) = s(bool,X18) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X17))) = s(bool,happ(s(fun(bool,bool),X17),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f321,f2291])).
fof(f37060,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X7)))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X7))) = s(bool,happ(s(fun(bool,bool),X7),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f36613,f270])).
fof(f36613,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X7)))) | s(bool,'F_0') != s(bool,X6) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X7))) = s(bool,happ(s(fun(bool,bool),X7),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f280,f2291])).
fof(f36989,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X5)))) | s(bool,'T_0') = s(bool,X4) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X5))) = s(bool,happ(s(fun(bool,bool),X5),s(bool,'T_0')))) )),
  inference(duplicate_literal_removal,[],[f36612])).
fof(f36612,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.min.@_1'(s(fun(bool,bool),X5)))) | s(bool,'T_0') = s(bool,X4) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X5))) = s(bool,happ(s(fun(bool,bool),X5),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f277,f2291])).
fof(f37058,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'const.min.@_1'(s(fun(bool,bool),X3))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) = s(bool,happ(s(fun(bool,bool),X3),s(bool,'T_0')))) ) | $spl130),
  inference(cnf_transformation,[],[f37058_D])).
fof(f37058_D,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'const.min.@_1'(s(fun(bool,bool),X3))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) = s(bool,happ(s(fun(bool,bool),X3),s(bool,'T_0')))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f37057,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'const.min.@_1'(s(fun(bool,bool),X3))) | p(s(bool,X2)) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) = s(bool,happ(s(fun(bool,bool),X3),s(bool,'T_0')))) )),
  inference(subsumption_resolution,[],[f36611,f276])).
fof(f36611,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'const.min.@_1'(s(fun(bool,bool),X3))) | p(s(bool,X2)) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) = s(bool,happ(s(fun(bool,bool),X3),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f276,f2291])).
fof(f37050,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X1))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1))) = s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) ) | $spl126),
  inference(cnf_transformation,[],[f37050_D])).
fof(f37050_D,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X1))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1))) = s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f37053,plain,(
  ( ! [X68,X67] : (s(bool,'F_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X68))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X68))) = s(bool,happ(s(fun(bool,bool),X68),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X67)) ) | $spl4),
  inference(forward_demodulation,[],[f37008,f13097])).
fof(f37008,plain,(
  ( ! [X68,X67] : (s(bool,'const.min.@_1'(s(fun(bool,bool),X68))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal',sK14))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X68))) = s(bool,happ(s(fun(bool,bool),X68),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X67)) ) | $spl4),
  inference(duplicate_literal_removal,[],[f36496])).
fof(f36496,plain,(
  ( ! [X68,X67] : (s(bool,'const.min.@_1'(s(fun(bool,bool),X68))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal',sK14))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X68))) = s(bool,happ(s(fun(bool,bool),X68),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) ) | $spl4),
  inference(superposition,[],[f2291,f13304])).
fof(f13304,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal',sK14))) | s(bool,'T_0') = s(bool,X0)) ) | $spl4),
  inference(subsumption_resolution,[],[f13176,f263])).
fof(f13176,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal',sK14))) | s(bool,'T_0') = s(bool,X0)) ) | $spl4),
  inference(superposition,[],[f11742,f783])).
fof(f37052,plain,(
  ( ! [X66,X65] : (s(bool,'F_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X66))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X66))) = s(bool,happ(s(fun(bool,bool),X66),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X65)) ) | $spl3),
  inference(forward_demodulation,[],[f37009,f12565])).
fof(f12565,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13))))) | $spl3),
  inference(unit_resulting_resolution,[],[f525,f463,f846])).
fof(f37009,plain,(
  ( ! [X66,X65] : (s(bool,'const.min.@_1'(s(fun(bool,bool),X66))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13))))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X66))) = s(bool,happ(s(fun(bool,bool),X66),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X65)) ) | $spl3),
  inference(duplicate_literal_removal,[],[f36495])).
fof(f36495,plain,(
  ( ! [X66,X65] : (s(bool,'const.min.@_1'(s(fun(bool,bool),X66))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13))))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X66))) = s(bool,happ(s(fun(bool,bool),X66),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X65) | s(bool,'T_0') = s(bool,X65)) ) | $spl3),
  inference(superposition,[],[f2291,f12772])).
fof(f12772,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(subsumption_resolution,[],[f12644,f263])).
fof(f12644,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f463,f783])).
fof(f37051,plain,(
  ( ! [X64,X63] : (s(bool,'const.min.@_1'(s(fun(bool,bool),X63))) = s(bool,'const.min.@_1'(s(fun(bool,bool),X64))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X63))) = s(bool,happ(s(fun(bool,bool),X63),s(bool,'T_0'))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X64))) = s(bool,happ(s(fun(bool,bool),X64),s(bool,'T_0')))) ) | $spl128),
  inference(cnf_transformation,[],[f37051_D])).
fof(f37051_D,plain,(
  ( ! [X64,X63] : (s(bool,'const.min.@_1'(s(fun(bool,bool),X63))) = s(bool,'const.min.@_1'(s(fun(bool,bool),X64))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X63))) = s(bool,happ(s(fun(bool,bool),X63),s(bool,'T_0'))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X64))) = s(bool,happ(s(fun(bool,bool),X64),s(bool,'T_0')))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f37010,plain,(
  ( ! [X64,X62,X63] : (s(bool,'const.min.@_1'(s(fun(bool,bool),X63))) = s(bool,'const.min.@_1'(s(fun(bool,bool),X64))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X64))) = s(bool,happ(s(fun(bool,bool),X64),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X62) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X63))) = s(bool,happ(s(fun(bool,bool),X63),s(bool,'T_0')))) )),
  inference(duplicate_literal_removal,[],[f36494])).
fof(f36494,plain,(
  ( ! [X64,X62,X63] : (s(bool,'const.min.@_1'(s(fun(bool,bool),X63))) = s(bool,'const.min.@_1'(s(fun(bool,bool),X64))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X64))) = s(bool,happ(s(fun(bool,bool),X64),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X62) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X63))) = s(bool,happ(s(fun(bool,bool),X63),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f2291,f2291])).
fof(f37031,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X1))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1))) = s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f36479])).
fof(f36479,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X1))) | s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1))) = s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2291,f270])).
fof(f36205,plain,(
  ( ! [X243,X241,X244,X240] : (sP1(happ(s(fun(X240,bool),X241),s(X240,'const.min.@_1'(s(fun(X240,bool),X241)))),X243) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X240,bool),X241))) | s(bool,'T_0') = s(bool,X244) | s(bool,X243) = s(bool,X244)) ) | $spl124),
  inference(cnf_transformation,[],[f36205_D])).
fof(f36205_D,plain,(
  ( ! [X243,X241,X244,X240] : (sP1(happ(s(fun(X240,bool),X241),s(X240,'const.min.@_1'(s(fun(X240,bool),X241)))),X243) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X240,bool),X241))) | s(bool,'T_0') = s(bool,X244) | s(bool,X243) = s(bool,X244)) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f36204,plain,(
  ( ! [X243,X241,X244,X242,X240] : (p(s(bool,X242)) | sP1(happ(s(fun(X240,bool),X241),s(X240,'const.min.@_1'(s(fun(X240,bool),X241)))),X243) | s(bool,X243) = s(bool,X244) | s(bool,'T_0') = s(bool,X244) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X240,bool),X241)))) )),
  inference(subsumption_resolution,[],[f35903,f276])).
fof(f35903,plain,(
  ( ! [X243,X241,X244,X242,X240] : (p(s(bool,X242)) | sP1(happ(s(fun(X240,bool),X241),s(X240,'const.min.@_1'(s(fun(X240,bool),X241)))),X243) | s(bool,X243) = s(bool,X244) | s(bool,'T_0') = s(bool,X244) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X240,bool),X241))) | s(bool,'T_0') = s(bool,X242)) )),
  inference(superposition,[],[f2610,f2064])).
fof(f36203,plain,(
  ( ! [X239,X235,X238,X236] : (sP1(X238,happ(s(fun(X235,bool),X236),s(X235,'const.min.@_1'(s(fun(X235,bool),X236))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X235,bool),X236))) | s(bool,'T_0') = s(bool,X239) | s(bool,X238) = s(bool,X239)) ) | $spl122),
  inference(cnf_transformation,[],[f36203_D])).
fof(f36203_D,plain,(
  ( ! [X239,X235,X238,X236] : (sP1(X238,happ(s(fun(X235,bool),X236),s(X235,'const.min.@_1'(s(fun(X235,bool),X236))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X235,bool),X236))) | s(bool,'T_0') = s(bool,X239) | s(bool,X238) = s(bool,X239)) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f36202,plain,(
  ( ! [X239,X237,X235,X238,X236] : (p(s(bool,X237)) | sP1(X238,happ(s(fun(X235,bool),X236),s(X235,'const.min.@_1'(s(fun(X235,bool),X236))))) | s(bool,X238) = s(bool,X239) | s(bool,'T_0') = s(bool,X239) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X235,bool),X236)))) )),
  inference(subsumption_resolution,[],[f35902,f276])).
fof(f35902,plain,(
  ( ! [X239,X237,X235,X238,X236] : (p(s(bool,X237)) | sP1(X238,happ(s(fun(X235,bool),X236),s(X235,'const.min.@_1'(s(fun(X235,bool),X236))))) | s(bool,X238) = s(bool,X239) | s(bool,'T_0') = s(bool,X239) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X235,bool),X236))) | s(bool,'T_0') = s(bool,X237)) )),
  inference(superposition,[],[f2609,f2064])).
fof(f36201,plain,(
  ( ! [X223,X224] : (sP1(happ(s(fun(X223,bool),X224),s(X223,'const.min.@_1'(s(fun(X223,bool),X224)))),'T_0') | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X223,bool),X224)))) ) | $spl120),
  inference(cnf_transformation,[],[f36201_D])).
fof(f36201_D,plain,(
  ( ! [X223,X224] : (sP1(happ(s(fun(X223,bool),X224),s(X223,'const.min.@_1'(s(fun(X223,bool),X224)))),'T_0') | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X223,bool),X224)))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f36200,plain,(
  ( ! [X225,X223,X224] : (p(s(bool,X225)) | sP1(happ(s(fun(X223,bool),X224),s(X223,'const.min.@_1'(s(fun(X223,bool),X224)))),'T_0') | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X223,bool),X224)))) ) | $spl12),
  inference(subsumption_resolution,[],[f35899,f276])).
fof(f35899,plain,(
  ( ! [X225,X223,X224] : (p(s(bool,X225)) | sP1(happ(s(fun(X223,bool),X224),s(X223,'const.min.@_1'(s(fun(X223,bool),X224)))),'T_0') | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X223,bool),X224))) | s(bool,'T_0') = s(bool,X225)) ) | $spl12),
  inference(superposition,[],[f2558,f2064])).
fof(f36199,plain,(
  ( ! [X220,X221] : (sP1('T_0',happ(s(fun(X220,bool),X221),s(X220,'const.min.@_1'(s(fun(X220,bool),X221))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X220,bool),X221)))) ) | $spl118),
  inference(cnf_transformation,[],[f36199_D])).
fof(f36199_D,plain,(
  ( ! [X220,X221] : (sP1('T_0',happ(s(fun(X220,bool),X221),s(X220,'const.min.@_1'(s(fun(X220,bool),X221))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X220,bool),X221)))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f36198,plain,(
  ( ! [X222,X220,X221] : (p(s(bool,X222)) | sP1('T_0',happ(s(fun(X220,bool),X221),s(X220,'const.min.@_1'(s(fun(X220,bool),X221))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X220,bool),X221)))) ) | $spl8),
  inference(subsumption_resolution,[],[f35898,f276])).
fof(f35898,plain,(
  ( ! [X222,X220,X221] : (p(s(bool,X222)) | sP1('T_0',happ(s(fun(X220,bool),X221),s(X220,'const.min.@_1'(s(fun(X220,bool),X221))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X220,bool),X221))) | s(bool,'T_0') = s(bool,X222)) ) | $spl8),
  inference(superposition,[],[f2554,f2064])).
fof(f34049,plain,(
  ( ! [X39,X43,X41,X38,X42] : (s(bool,X41) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK22(X39,X38)),s('type.extreal.extreal',X39))) | s(bool,'T_0') = s(bool,X43) | s(bool,X42) = s(bool,X43) | s(bool,X42) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X38))),s('type.extreal.extreal',X39))) | s(bool,'T_0') = s(bool,X41)) ) | $spl116),
  inference(cnf_transformation,[],[f34049_D])).
fof(f34049_D,plain,(
  ( ! [X39,X43,X41,X38,X42] : (s(bool,X41) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK22(X39,X38)),s('type.extreal.extreal',X39))) | s(bool,'T_0') = s(bool,X43) | s(bool,X42) = s(bool,X43) | s(bool,X42) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X38))),s('type.extreal.extreal',X39))) | s(bool,'T_0') = s(bool,X41)) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f34053,plain,(
  ( ! [X66,X64,X62,X67,X65,X63] : (p(s(bool,X64)) | s(bool,X65) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK22(X63,X62)),s('type.extreal.extreal',X63))) | s(bool,'T_0') = s(bool,X65) | s(bool,X66) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X62))),s('type.extreal.extreal',X63))) | s(bool,X66) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(subsumption_resolution,[],[f34041,f276])).
fof(f34041,plain,(
  ( ! [X66,X64,X62,X67,X65,X63] : (p(s(bool,X64)) | s(bool,X65) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK22(X63,X62)),s('type.extreal.extreal',X63))) | s(bool,'T_0') = s(bool,X65) | s(bool,X66) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X62))),s('type.extreal.extreal',X63))) | s(bool,X66) = s(bool,X67) | s(bool,'T_0') = s(bool,X67) | s(bool,'T_0') = s(bool,X64)) )),
  inference(superposition,[],[f2636,f2054])).
fof(f34048,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (p(s(bool,X40)) | s(bool,X41) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK22(X39,X38)),s('type.extreal.extreal',X39))) | s(bool,'T_0') = s(bool,X41) | s(bool,X42) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X38))),s('type.extreal.extreal',X39))) | s(bool,X42) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(subsumption_resolution,[],[f34037,f276])).
fof(f34037,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (p(s(bool,X40)) | s(bool,X41) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK22(X39,X38)),s('type.extreal.extreal',X39))) | s(bool,'T_0') = s(bool,X41) | s(bool,X42) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X38))),s('type.extreal.extreal',X39))) | s(bool,'T_0') = s(bool,X40) | s(bool,X42) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(superposition,[],[f2636,f2054])).
fof(f33993,plain,(
  ( ! [X37,X41,X38,X42,X40] : (s(bool,X40) = s(bool,happ(s(fun('type.extreal.extreal',bool),X38),s('type.extreal.extreal',X37))) | s(bool,'T_0') = s(bool,X42) | s(bool,X41) = s(bool,X42) | s(bool,X41) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X37),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X38))))) | s(bool,'T_0') = s(bool,X40)) ) | $spl114),
  inference(cnf_transformation,[],[f33993_D])).
fof(f33993_D,plain,(
  ( ! [X37,X41,X38,X42,X40] : (s(bool,X40) = s(bool,happ(s(fun('type.extreal.extreal',bool),X38),s('type.extreal.extreal',X37))) | s(bool,'T_0') = s(bool,X42) | s(bool,X41) = s(bool,X42) | s(bool,X41) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X37),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X38))))) | s(bool,'T_0') = s(bool,X40)) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f33997,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (p(s(bool,X63)) | s(bool,X64) = s(bool,happ(s(fun('type.extreal.extreal',bool),X62),s('type.extreal.extreal',X61))) | s(bool,'T_0') = s(bool,X64) | s(bool,X65) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X61),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X62))))) | s(bool,X65) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f33985,f276])).
fof(f33985,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (p(s(bool,X63)) | s(bool,X64) = s(bool,happ(s(fun('type.extreal.extreal',bool),X62),s('type.extreal.extreal',X61))) | s(bool,'T_0') = s(bool,X64) | s(bool,X65) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X61),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X62))))) | s(bool,X65) = s(bool,X66) | s(bool,'T_0') = s(bool,X66) | s(bool,'T_0') = s(bool,X63)) )),
  inference(superposition,[],[f2626,f2054])).
fof(f33992,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (p(s(bool,X39)) | s(bool,X40) = s(bool,happ(s(fun('type.extreal.extreal',bool),X38),s('type.extreal.extreal',X37))) | s(bool,'T_0') = s(bool,X40) | s(bool,X41) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X37),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X38))))) | s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f33981,f276])).
fof(f33981,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (p(s(bool,X39)) | s(bool,X40) = s(bool,happ(s(fun('type.extreal.extreal',bool),X38),s('type.extreal.extreal',X37))) | s(bool,'T_0') = s(bool,X40) | s(bool,X41) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X37),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X38))))) | s(bool,'T_0') = s(bool,X39) | s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f2626,f2054])).
fof(f33955,plain,(
  ( ! [X19,X17,X20] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X17),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X20) | s(bool,X19) = s(bool,X20) | s(bool,X19) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X17))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X17)))) ) | $spl112),
  inference(cnf_transformation,[],[f33955_D])).
fof(f33955_D,plain,(
  ( ! [X19,X17,X20] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X17),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X20) | s(bool,X19) = s(bool,X20) | s(bool,X19) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X17))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X17)))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f33959,plain,(
  ( ! [X35,X33,X36,X34] : (p(s(bool,X34)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X33),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X33))) | s(bool,X35) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X33))) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(subsumption_resolution,[],[f33947,f276])).
fof(f33947,plain,(
  ( ! [X35,X33,X36,X34] : (p(s(bool,X34)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X33),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X33))) | s(bool,X35) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X33))) | s(bool,X35) = s(bool,X36) | s(bool,'T_0') = s(bool,X36) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f4872,f2054])).
fof(f33954,plain,(
  ( ! [X19,X17,X20,X18] : (p(s(bool,X18)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X17),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X17))) | s(bool,X19) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X17))) | s(bool,X19) = s(bool,X20) | s(bool,'T_0') = s(bool,X20)) )),
  inference(subsumption_resolution,[],[f33943,f276])).
fof(f33943,plain,(
  ( ! [X19,X17,X20,X18] : (p(s(bool,X18)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X17),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X17))) | s(bool,X19) = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X17))) | s(bool,'T_0') = s(bool,X18) | s(bool,X19) = s(bool,X20) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f4872,f2054])).
fof(f32869,plain,(
  ( ! [X45,X43,X41,X46,X44,X40] : (s(bool,X43) = s(bool,happ(s(fun(X40,bool),X41),s(X40,X44))) | s(bool,'T_0') = s(bool,X46) | s(bool,X45) = s(bool,X46) | s(bool,X45) = s(bool,happ(s(fun(X40,bool),X41),s(X40,sK16(X41,X40)))) | s(bool,'T_0') = s(bool,X43)) ) | $spl110),
  inference(cnf_transformation,[],[f32869_D])).
fof(f32869_D,plain,(
  ( ! [X45,X43,X41,X46,X44,X40] : (s(bool,X43) = s(bool,happ(s(fun(X40,bool),X41),s(X40,X44))) | s(bool,'T_0') = s(bool,X46) | s(bool,X45) = s(bool,X46) | s(bool,X45) = s(bool,happ(s(fun(X40,bool),X41),s(X40,sK16(X41,X40)))) | s(bool,'T_0') = s(bool,X43)) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f32873,plain,(
  ( ! [X70,X68,X74,X72,X71,X69,X73] : (p(s(bool,X70)) | s(bool,X71) = s(bool,happ(s(fun(X68,bool),X69),s(X68,X72))) | s(bool,'T_0') = s(bool,X71) | s(bool,X73) = s(bool,happ(s(fun(X68,bool),X69),s(X68,sK16(X69,X68)))) | s(bool,X73) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f32851,f276])).
fof(f32851,plain,(
  ( ! [X70,X68,X74,X72,X71,X69,X73] : (p(s(bool,X70)) | s(bool,X71) = s(bool,happ(s(fun(X68,bool),X69),s(X68,X72))) | s(bool,'T_0') = s(bool,X71) | s(bool,X73) = s(bool,happ(s(fun(X68,bool),X69),s(X68,sK16(X69,X68)))) | s(bool,X73) = s(bool,X74) | s(bool,'T_0') = s(bool,X74) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f2628,f2054])).
fof(f32868,plain,(
  ( ! [X45,X43,X41,X46,X44,X42,X40] : (p(s(bool,X42)) | s(bool,X43) = s(bool,happ(s(fun(X40,bool),X41),s(X40,X44))) | s(bool,'T_0') = s(bool,X43) | s(bool,X45) = s(bool,happ(s(fun(X40,bool),X41),s(X40,sK16(X41,X40)))) | s(bool,X45) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f32847,f276])).
fof(f32847,plain,(
  ( ! [X45,X43,X41,X46,X44,X42,X40] : (p(s(bool,X42)) | s(bool,X43) = s(bool,happ(s(fun(X40,bool),X41),s(X40,X44))) | s(bool,'T_0') = s(bool,X43) | s(bool,X45) = s(bool,happ(s(fun(X40,bool),X41),s(X40,sK16(X41,X40)))) | s(bool,'T_0') = s(bool,X42) | s(bool,X45) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f2628,f2054])).
fof(f32736,plain,(
  ( ! [X45,X43,X41,X46,X44,X40] : (s(bool,X43) = s(bool,happ(s(fun(X40,bool),X41),s(X40,X44))) | s(bool,'T_0') = s(bool,X46) | s(bool,X45) = s(bool,X46) | s(bool,X45) = s(bool,happ(s(fun(X40,bool),X41),s(X40,sK17(X41,X40)))) | s(bool,'T_0') = s(bool,X43)) ) | $spl108),
  inference(cnf_transformation,[],[f32736_D])).
fof(f32736_D,plain,(
  ( ! [X45,X43,X41,X46,X44,X40] : (s(bool,X43) = s(bool,happ(s(fun(X40,bool),X41),s(X40,X44))) | s(bool,'T_0') = s(bool,X46) | s(bool,X45) = s(bool,X46) | s(bool,X45) = s(bool,happ(s(fun(X40,bool),X41),s(X40,sK17(X41,X40)))) | s(bool,'T_0') = s(bool,X43)) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f32740,plain,(
  ( ! [X70,X68,X74,X72,X71,X69,X73] : (p(s(bool,X70)) | s(bool,X71) = s(bool,happ(s(fun(X68,bool),X69),s(X68,X72))) | s(bool,'T_0') = s(bool,X71) | s(bool,X73) = s(bool,happ(s(fun(X68,bool),X69),s(X68,sK17(X69,X68)))) | s(bool,X73) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f32718,f276])).
fof(f32718,plain,(
  ( ! [X70,X68,X74,X72,X71,X69,X73] : (p(s(bool,X70)) | s(bool,X71) = s(bool,happ(s(fun(X68,bool),X69),s(X68,X72))) | s(bool,'T_0') = s(bool,X71) | s(bool,X73) = s(bool,happ(s(fun(X68,bool),X69),s(X68,sK17(X69,X68)))) | s(bool,X73) = s(bool,X74) | s(bool,'T_0') = s(bool,X74) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f2627,f2054])).
fof(f32735,plain,(
  ( ! [X45,X43,X41,X46,X44,X42,X40] : (p(s(bool,X42)) | s(bool,X43) = s(bool,happ(s(fun(X40,bool),X41),s(X40,X44))) | s(bool,'T_0') = s(bool,X43) | s(bool,X45) = s(bool,happ(s(fun(X40,bool),X41),s(X40,sK17(X41,X40)))) | s(bool,X45) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f32714,f276])).
fof(f32714,plain,(
  ( ! [X45,X43,X41,X46,X44,X42,X40] : (p(s(bool,X42)) | s(bool,X43) = s(bool,happ(s(fun(X40,bool),X41),s(X40,X44))) | s(bool,'T_0') = s(bool,X43) | s(bool,X45) = s(bool,happ(s(fun(X40,bool),X41),s(X40,sK17(X41,X40)))) | s(bool,'T_0') = s(bool,X42) | s(bool,X45) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f2627,f2054])).
fof(f31843,plain,(
  ( ! [X716,X713,X715] : (s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),X713))) | s(bool,'T_0') = s(bool,X716) | s(bool,X715) = s(bool,X716) | s(bool,X715) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.min.@_1'(s(fun('type.extreal.extreal',bool),X713))),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X713)))))) ) | $spl104),
  inference(cnf_transformation,[],[f31843_D])).
fof(f31843_D,plain,(
  ( ! [X716,X713,X715] : (s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),X713))) | s(bool,'T_0') = s(bool,X716) | s(bool,X715) = s(bool,X716) | s(bool,X715) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.min.@_1'(s(fun('type.extreal.extreal',bool),X713))),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X713)))))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f32055,plain,(
  ( ! [X716,X714,X713,X715] : (p(s(bool,X714)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),X713))) | s(bool,X715) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.min.@_1'(s(fun('type.extreal.extreal',bool),X713))),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X713))))) | s(bool,X715) = s(bool,X716) | s(bool,'T_0') = s(bool,X716)) )),
  inference(subsumption_resolution,[],[f29580,f276])).
fof(f29580,plain,(
  ( ! [X716,X714,X713,X715] : (p(s(bool,X714)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),X713))) | s(bool,X715) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.min.@_1'(s(fun('type.extreal.extreal',bool),X713))),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X713))))) | s(bool,X715) = s(bool,X716) | s(bool,'T_0') = s(bool,X716) | s(bool,'T_0') = s(bool,X714)) )),
  inference(superposition,[],[f4965,f2054])).
fof(f31839,plain,(
  ( ! [X697,X695,X698,X694] : (s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X694),s('type.extreal.extreal',sK22(X695,X694)))) | s(bool,'T_0') = s(bool,X698) | s(bool,X697) = s(bool,X698) | s(bool,X697) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X694))),s('type.extreal.extreal',X695)))) ) | $spl102),
  inference(cnf_transformation,[],[f31839_D])).
fof(f31839_D,plain,(
  ( ! [X697,X695,X698,X694] : (s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X694),s('type.extreal.extreal',sK22(X695,X694)))) | s(bool,'T_0') = s(bool,X698) | s(bool,X697) = s(bool,X698) | s(bool,X697) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X694))),s('type.extreal.extreal',X695)))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f32052,plain,(
  ( ! [X697,X695,X696,X698,X694] : (p(s(bool,X696)) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X694),s('type.extreal.extreal',sK22(X695,X694)))) | s(bool,X697) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X694))),s('type.extreal.extreal',X695))) | s(bool,X697) = s(bool,X698) | s(bool,'T_0') = s(bool,X698)) )),
  inference(subsumption_resolution,[],[f29576,f276])).
fof(f29576,plain,(
  ( ! [X697,X695,X696,X698,X694] : (p(s(bool,X696)) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X694),s('type.extreal.extreal',sK22(X695,X694)))) | s(bool,X697) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X694))),s('type.extreal.extreal',X695))) | s(bool,X697) = s(bool,X698) | s(bool,'T_0') = s(bool,X698) | s(bool,'T_0') = s(bool,X696)) )),
  inference(superposition,[],[f1867,f2054])).
fof(f31835,plain,(
  ( ! [X666,X670,X667,X669,X671,X672] : (s(bool,X669) = s(bool,happ(s(fun(X666,bool),X667),s(X666,X670))) | s(bool,'T_0') = s(bool,X672) | s(bool,X671) = s(bool,X672) | s(bool,X671) = s(bool,'|dollar|exists_1'(s(fun(X666,bool),X667))) | s(bool,'T_0') = s(bool,X669)) ) | $spl100),
  inference(cnf_transformation,[],[f31835_D])).
fof(f31835_D,plain,(
  ( ! [X666,X670,X667,X669,X671,X672] : (s(bool,X669) = s(bool,happ(s(fun(X666,bool),X667),s(X666,X670))) | s(bool,'T_0') = s(bool,X672) | s(bool,X671) = s(bool,X672) | s(bool,X671) = s(bool,'|dollar|exists_1'(s(fun(X666,bool),X667))) | s(bool,'T_0') = s(bool,X669)) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f32049,plain,(
  ( ! [X666,X668,X670,X667,X669,X671,X672] : (p(s(bool,X668)) | s(bool,X669) = s(bool,happ(s(fun(X666,bool),X667),s(X666,X670))) | s(bool,'T_0') = s(bool,X669) | s(bool,X671) = s(bool,'|dollar|exists_1'(s(fun(X666,bool),X667))) | s(bool,X671) = s(bool,X672) | s(bool,'T_0') = s(bool,X672)) )),
  inference(subsumption_resolution,[],[f29571,f276])).
fof(f29571,plain,(
  ( ! [X666,X668,X670,X667,X669,X671,X672] : (p(s(bool,X668)) | s(bool,X669) = s(bool,happ(s(fun(X666,bool),X667),s(X666,X670))) | s(bool,'T_0') = s(bool,X669) | s(bool,X671) = s(bool,'|dollar|exists_1'(s(fun(X666,bool),X667))) | s(bool,X671) = s(bool,X672) | s(bool,'T_0') = s(bool,X672) | s(bool,'T_0') = s(bool,X668)) )),
  inference(superposition,[],[f2629,f2054])).
fof(f31833,plain,(
  ( ! [X586,X582,X585,X583] : (s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X583),s('type.extreal.extreal',X582))) | s(bool,'T_0') = s(bool,X586) | s(bool,X585) = s(bool,X586) | s(bool,X585) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X582),s('type.extreal.extreal',X583)))) ) | $spl98),
  inference(cnf_transformation,[],[f31833_D])).
fof(f31833_D,plain,(
  ( ! [X586,X582,X585,X583] : (s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X583),s('type.extreal.extreal',X582))) | s(bool,'T_0') = s(bool,X586) | s(bool,X585) = s(bool,X586) | s(bool,X585) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X582),s('type.extreal.extreal',X583)))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f32048,plain,(
  ( ! [X584,X586,X582,X585,X583] : (p(s(bool,X584)) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X583),s('type.extreal.extreal',X582))) | s(bool,X585) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X582),s('type.extreal.extreal',X583))) | s(bool,X585) = s(bool,X586) | s(bool,'T_0') = s(bool,X586)) )),
  inference(subsumption_resolution,[],[f29555,f276])).
fof(f29555,plain,(
  ( ! [X584,X586,X582,X585,X583] : (p(s(bool,X584)) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X583),s('type.extreal.extreal',X582))) | s(bool,X585) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X582),s('type.extreal.extreal',X583))) | s(bool,X585) = s(bool,X586) | s(bool,'T_0') = s(bool,X586) | s(bool,'T_0') = s(bool,X584)) )),
  inference(superposition,[],[f1296,f2054])).
fof(f31831,plain,(
  ( ! [X516,X518,X519,X515] : (s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X515,bool),X516))) | s(bool,'T_0') = s(bool,X519) | s(bool,X518) = s(bool,X519) | s(bool,X518) = s(bool,happ(s(fun(X515,bool),X516),s(X515,sK17(X516,X515))))) ) | $spl96),
  inference(cnf_transformation,[],[f31831_D])).
fof(f31831_D,plain,(
  ( ! [X516,X518,X519,X515] : (s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X515,bool),X516))) | s(bool,'T_0') = s(bool,X519) | s(bool,X518) = s(bool,X519) | s(bool,X518) = s(bool,happ(s(fun(X515,bool),X516),s(X515,sK17(X516,X515))))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f32047,plain,(
  ( ! [X516,X518,X517,X519,X515] : (p(s(bool,X517)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X515,bool),X516))) | s(bool,X518) = s(bool,happ(s(fun(X515,bool),X516),s(X515,sK17(X516,X515)))) | s(bool,X518) = s(bool,X519) | s(bool,'T_0') = s(bool,X519)) )),
  inference(subsumption_resolution,[],[f29544,f276])).
fof(f29544,plain,(
  ( ! [X516,X518,X517,X519,X515] : (p(s(bool,X517)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X515,bool),X516))) | s(bool,X518) = s(bool,happ(s(fun(X515,bool),X516),s(X515,sK17(X516,X515)))) | s(bool,X518) = s(bool,X519) | s(bool,'T_0') = s(bool,X519) | s(bool,'T_0') = s(bool,X517)) )),
  inference(superposition,[],[f4966,f2054])).
fof(f31829,plain,(
  ( ! [X509,X512,X514,X510,X513] : (s(bool,X512) = s(bool,'|dollar|exists_1'(s(fun(X509,bool),X510))) | s(bool,'T_0') = s(bool,X514) | s(bool,X513) = s(bool,X514) | s(bool,X513) = s(bool,happ(s(fun(X509,bool),X510),s(X509,sK17(X510,X509)))) | s(bool,'T_0') = s(bool,X512)) ) | $spl94),
  inference(cnf_transformation,[],[f31829_D])).
fof(f31829_D,plain,(
  ( ! [X509,X512,X514,X510,X513] : (s(bool,X512) = s(bool,'|dollar|exists_1'(s(fun(X509,bool),X510))) | s(bool,'T_0') = s(bool,X514) | s(bool,X513) = s(bool,X514) | s(bool,X513) = s(bool,happ(s(fun(X509,bool),X510),s(X509,sK17(X510,X509)))) | s(bool,'T_0') = s(bool,X512)) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f32046,plain,(
  ( ! [X509,X511,X512,X514,X510,X513] : (p(s(bool,X511)) | s(bool,X512) = s(bool,'|dollar|exists_1'(s(fun(X509,bool),X510))) | s(bool,'T_0') = s(bool,X512) | s(bool,X513) = s(bool,happ(s(fun(X509,bool),X510),s(X509,sK17(X510,X509)))) | s(bool,X513) = s(bool,X514) | s(bool,'T_0') = s(bool,X514)) )),
  inference(subsumption_resolution,[],[f29543,f276])).
fof(f29543,plain,(
  ( ! [X509,X511,X512,X514,X510,X513] : (p(s(bool,X511)) | s(bool,X512) = s(bool,'|dollar|exists_1'(s(fun(X509,bool),X510))) | s(bool,'T_0') = s(bool,X512) | s(bool,X513) = s(bool,happ(s(fun(X509,bool),X510),s(X509,sK17(X510,X509)))) | s(bool,X513) = s(bool,X514) | s(bool,'T_0') = s(bool,X514) | s(bool,'T_0') = s(bool,X511)) )),
  inference(superposition,[],[f4303,f2054])).
fof(f31827,plain,(
  ( ! [X505,X507,X508,X504] : (s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X504,bool),X505))) | s(bool,'T_0') = s(bool,X508) | s(bool,X507) = s(bool,X508) | s(bool,X507) = s(bool,happ(s(fun(X504,bool),X505),s(X504,sK16(X505,X504))))) ) | $spl92),
  inference(cnf_transformation,[],[f31827_D])).
fof(f31827_D,plain,(
  ( ! [X505,X507,X508,X504] : (s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X504,bool),X505))) | s(bool,'T_0') = s(bool,X508) | s(bool,X507) = s(bool,X508) | s(bool,X507) = s(bool,happ(s(fun(X504,bool),X505),s(X504,sK16(X505,X504))))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f32045,plain,(
  ( ! [X505,X507,X508,X504,X506] : (p(s(bool,X506)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X504,bool),X505))) | s(bool,X507) = s(bool,happ(s(fun(X504,bool),X505),s(X504,sK16(X505,X504)))) | s(bool,X507) = s(bool,X508) | s(bool,'T_0') = s(bool,X508)) )),
  inference(subsumption_resolution,[],[f29542,f276])).
fof(f29542,plain,(
  ( ! [X505,X507,X508,X504,X506] : (p(s(bool,X506)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X504,bool),X505))) | s(bool,X507) = s(bool,happ(s(fun(X504,bool),X505),s(X504,sK16(X505,X504)))) | s(bool,X507) = s(bool,X508) | s(bool,'T_0') = s(bool,X508) | s(bool,'T_0') = s(bool,X506)) )),
  inference(superposition,[],[f4967,f2054])).
fof(f31825,plain,(
  ( ! [X499,X502,X498,X501,X503] : (s(bool,X501) = s(bool,'|dollar|exists_1'(s(fun(X498,bool),X499))) | s(bool,'T_0') = s(bool,X503) | s(bool,X502) = s(bool,X503) | s(bool,X502) = s(bool,happ(s(fun(X498,bool),X499),s(X498,sK16(X499,X498)))) | s(bool,'T_0') = s(bool,X501)) ) | $spl90),
  inference(cnf_transformation,[],[f31825_D])).
fof(f31825_D,plain,(
  ( ! [X499,X502,X498,X501,X503] : (s(bool,X501) = s(bool,'|dollar|exists_1'(s(fun(X498,bool),X499))) | s(bool,'T_0') = s(bool,X503) | s(bool,X502) = s(bool,X503) | s(bool,X502) = s(bool,happ(s(fun(X498,bool),X499),s(X498,sK16(X499,X498)))) | s(bool,'T_0') = s(bool,X501)) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f32044,plain,(
  ( ! [X499,X500,X502,X498,X501,X503] : (p(s(bool,X500)) | s(bool,X501) = s(bool,'|dollar|exists_1'(s(fun(X498,bool),X499))) | s(bool,'T_0') = s(bool,X501) | s(bool,X502) = s(bool,happ(s(fun(X498,bool),X499),s(X498,sK16(X499,X498)))) | s(bool,X502) = s(bool,X503) | s(bool,'T_0') = s(bool,X503)) )),
  inference(subsumption_resolution,[],[f29541,f276])).
fof(f29541,plain,(
  ( ! [X499,X500,X502,X498,X501,X503] : (p(s(bool,X500)) | s(bool,X501) = s(bool,'|dollar|exists_1'(s(fun(X498,bool),X499))) | s(bool,'T_0') = s(bool,X501) | s(bool,X502) = s(bool,happ(s(fun(X498,bool),X499),s(X498,sK16(X499,X498)))) | s(bool,X502) = s(bool,X503) | s(bool,'T_0') = s(bool,X503) | s(bool,'T_0') = s(bool,X500)) )),
  inference(superposition,[],[f3897,f2054])).
fof(f31821,plain,(
  ( ! [X455,X459,X456,X458] : (s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X455))),s('type.extreal.extreal',X456))) | s(bool,'T_0') = s(bool,X459) | s(bool,X458) = s(bool,X459) | s(bool,X458) = s(bool,happ(s(fun('type.extreal.extreal',bool),X455),s('type.extreal.extreal',sK22(X456,X455))))) ) | $spl88),
  inference(cnf_transformation,[],[f31821_D])).
fof(f31821_D,plain,(
  ( ! [X455,X459,X456,X458] : (s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X455))),s('type.extreal.extreal',X456))) | s(bool,'T_0') = s(bool,X459) | s(bool,X458) = s(bool,X459) | s(bool,X458) = s(bool,happ(s(fun('type.extreal.extreal',bool),X455),s('type.extreal.extreal',sK22(X456,X455))))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f32041,plain,(
  ( ! [X455,X457,X459,X456,X458] : (p(s(bool,X457)) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X455))),s('type.extreal.extreal',X456))) | s(bool,X458) = s(bool,happ(s(fun('type.extreal.extreal',bool),X455),s('type.extreal.extreal',sK22(X456,X455)))) | s(bool,X458) = s(bool,X459) | s(bool,'T_0') = s(bool,X459)) )),
  inference(subsumption_resolution,[],[f29533,f276])).
fof(f29533,plain,(
  ( ! [X455,X457,X459,X456,X458] : (p(s(bool,X457)) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X455))),s('type.extreal.extreal',X456))) | s(bool,X458) = s(bool,happ(s(fun('type.extreal.extreal',bool),X455),s('type.extreal.extreal',sK22(X456,X455)))) | s(bool,X458) = s(bool,X459) | s(bool,'T_0') = s(bool,X459) | s(bool,'T_0') = s(bool,X457)) )),
  inference(superposition,[],[f1866,f2054])).
fof(f31769,plain,(
  ( ! [X206,X204,X205,X203,X201] : (sP1(X201,X203) | s(bool,'T_0') = s(bool,X206) | s(bool,X205) = s(bool,X206) | s(bool,X201) = s(bool,X205) | s(bool,'T_0') = s(bool,X204) | s(bool,X203) = s(bool,X204)) ) | $spl86),
  inference(cnf_transformation,[],[f31769_D])).
fof(f31769_D,plain,(
  ( ! [X206,X204,X205,X203,X201] : (sP1(X201,X203) | s(bool,'T_0') = s(bool,X206) | s(bool,X205) = s(bool,X206) | s(bool,X201) = s(bool,X205) | s(bool,'T_0') = s(bool,X204) | s(bool,X203) = s(bool,X204)) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f31990,plain,(
  ( ! [X206,X204,X202,X205,X203,X201] : (p(s(bool,X202)) | sP1(X201,X203) | s(bool,X203) = s(bool,X204) | s(bool,'T_0') = s(bool,X204) | s(bool,X201) = s(bool,X205) | s(bool,X205) = s(bool,X206) | s(bool,'T_0') = s(bool,X206)) )),
  inference(subsumption_resolution,[],[f29480,f276])).
fof(f29480,plain,(
  ( ! [X206,X204,X202,X205,X203,X201] : (p(s(bool,X202)) | sP1(X201,X203) | s(bool,X203) = s(bool,X204) | s(bool,'T_0') = s(bool,X204) | s(bool,X201) = s(bool,X205) | s(bool,X205) = s(bool,X206) | s(bool,'T_0') = s(bool,X206) | s(bool,'T_0') = s(bool,X202)) )),
  inference(superposition,[],[f2610,f2054])).
fof(f31767,plain,(
  ( ! [X198,X200,X199,X197,X195] : (sP1(X197,X195) | s(bool,'T_0') = s(bool,X200) | s(bool,X199) = s(bool,X200) | s(bool,X195) = s(bool,X199) | s(bool,'T_0') = s(bool,X198) | s(bool,X197) = s(bool,X198)) ) | $spl84),
  inference(cnf_transformation,[],[f31767_D])).
fof(f31767_D,plain,(
  ( ! [X198,X200,X199,X197,X195] : (sP1(X197,X195) | s(bool,'T_0') = s(bool,X200) | s(bool,X199) = s(bool,X200) | s(bool,X195) = s(bool,X199) | s(bool,'T_0') = s(bool,X198) | s(bool,X197) = s(bool,X198)) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f31989,plain,(
  ( ! [X198,X196,X200,X199,X197,X195] : (p(s(bool,X196)) | sP1(X197,X195) | s(bool,X197) = s(bool,X198) | s(bool,'T_0') = s(bool,X198) | s(bool,X195) = s(bool,X199) | s(bool,X199) = s(bool,X200) | s(bool,'T_0') = s(bool,X200)) )),
  inference(subsumption_resolution,[],[f29479,f276])).
fof(f29479,plain,(
  ( ! [X198,X196,X200,X199,X197,X195] : (p(s(bool,X196)) | sP1(X197,X195) | s(bool,X197) = s(bool,X198) | s(bool,'T_0') = s(bool,X198) | s(bool,X195) = s(bool,X199) | s(bool,X199) = s(bool,X200) | s(bool,'T_0') = s(bool,X200) | s(bool,'T_0') = s(bool,X196)) )),
  inference(superposition,[],[f2609,f2054])).
fof(f31765,plain,(
  ( ! [X182,X180,X183] : (sP1(X180,'T_0') | s(bool,'T_0') = s(bool,X183) | s(bool,X182) = s(bool,X183) | s(bool,X180) = s(bool,X182)) ) | $spl82),
  inference(cnf_transformation,[],[f31765_D])).
fof(f31765_D,plain,(
  ( ! [X182,X180,X183] : (sP1(X180,'T_0') | s(bool,'T_0') = s(bool,X183) | s(bool,X182) = s(bool,X183) | s(bool,X180) = s(bool,X182)) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f31988,plain,(
  ( ! [X182,X180,X183,X181] : (p(s(bool,X181)) | sP1(X180,'T_0') | s(bool,X180) = s(bool,X182) | s(bool,X182) = s(bool,X183) | s(bool,'T_0') = s(bool,X183)) ) | $spl12),
  inference(subsumption_resolution,[],[f29476,f276])).
fof(f29476,plain,(
  ( ! [X182,X180,X183,X181] : (p(s(bool,X181)) | sP1(X180,'T_0') | s(bool,X180) = s(bool,X182) | s(bool,X182) = s(bool,X183) | s(bool,'T_0') = s(bool,X183) | s(bool,'T_0') = s(bool,X181)) ) | $spl12),
  inference(superposition,[],[f2558,f2054])).
fof(f31763,plain,(
  ( ! [X178,X176,X179] : (sP1('T_0',X176) | s(bool,'T_0') = s(bool,X179) | s(bool,X178) = s(bool,X179) | s(bool,X176) = s(bool,X178)) ) | $spl80),
  inference(cnf_transformation,[],[f31763_D])).
fof(f31763_D,plain,(
  ( ! [X178,X176,X179] : (sP1('T_0',X176) | s(bool,'T_0') = s(bool,X179) | s(bool,X178) = s(bool,X179) | s(bool,X176) = s(bool,X178)) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f31987,plain,(
  ( ! [X177,X178,X176,X179] : (p(s(bool,X177)) | sP1('T_0',X176) | s(bool,X176) = s(bool,X178) | s(bool,X178) = s(bool,X179) | s(bool,'T_0') = s(bool,X179)) ) | $spl8),
  inference(subsumption_resolution,[],[f29475,f276])).
fof(f29475,plain,(
  ( ! [X177,X178,X176,X179] : (p(s(bool,X177)) | sP1('T_0',X176) | s(bool,X176) = s(bool,X178) | s(bool,X178) = s(bool,X179) | s(bool,'T_0') = s(bool,X179) | s(bool,'T_0') = s(bool,X177)) ) | $spl8),
  inference(superposition,[],[f2554,f2054])).
fof(f31845,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') != s(bool,X5) | s(bool,'T_0') = s(bool,X7) | s(bool,X6) = s(bool,X7) | s(bool,X5) = s(bool,X6)) ) | $spl106),
  inference(cnf_transformation,[],[f31845_D])).
fof(f31845_D,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') != s(bool,X5) | s(bool,'T_0') = s(bool,X7) | s(bool,X6) = s(bool,X7) | s(bool,X5) = s(bool,X6)) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f31844,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'T_0') != s(bool,X5) | p(s(bool,X4)) | s(bool,X5) = s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f28999,f276])).
fof(f28999,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'T_0') != s(bool,X5) | p(s(bool,X4)) | s(bool,X5) = s(bool,X6) | s(bool,'T_0') = s(bool,X4) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f276,f2054])).
fof(f31842,plain,(
  ( ! [X716,X714,X713,X715] : (p(s(bool,X714)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),X713))) | s(bool,X715) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.min.@_1'(s(fun('type.extreal.extreal',bool),X713))),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X713))))) | s(bool,X715) = s(bool,X716) | s(bool,'T_0') = s(bool,X716)) )),
  inference(subsumption_resolution,[],[f28940,f276])).
fof(f28940,plain,(
  ( ! [X716,X714,X713,X715] : (p(s(bool,X714)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),X713))) | s(bool,X715) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.min.@_1'(s(fun('type.extreal.extreal',bool),X713))),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X713))))) | s(bool,'T_0') = s(bool,X714) | s(bool,X715) = s(bool,X716) | s(bool,'T_0') = s(bool,X716)) )),
  inference(superposition,[],[f4965,f2054])).
fof(f31838,plain,(
  ( ! [X697,X695,X696,X698,X694] : (p(s(bool,X696)) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X694),s('type.extreal.extreal',sK22(X695,X694)))) | s(bool,X697) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X694))),s('type.extreal.extreal',X695))) | s(bool,X697) = s(bool,X698) | s(bool,'T_0') = s(bool,X698)) )),
  inference(subsumption_resolution,[],[f28936,f276])).
fof(f28936,plain,(
  ( ! [X697,X695,X696,X698,X694] : (p(s(bool,X696)) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X694),s('type.extreal.extreal',sK22(X695,X694)))) | s(bool,X697) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X694))),s('type.extreal.extreal',X695))) | s(bool,'T_0') = s(bool,X696) | s(bool,X697) = s(bool,X698) | s(bool,'T_0') = s(bool,X698)) )),
  inference(superposition,[],[f1867,f2054])).
fof(f31834,plain,(
  ( ! [X666,X668,X670,X667,X669,X671,X672] : (p(s(bool,X668)) | s(bool,X669) = s(bool,happ(s(fun(X666,bool),X667),s(X666,X670))) | s(bool,'T_0') = s(bool,X669) | s(bool,X671) = s(bool,'|dollar|exists_1'(s(fun(X666,bool),X667))) | s(bool,X671) = s(bool,X672) | s(bool,'T_0') = s(bool,X672)) )),
  inference(subsumption_resolution,[],[f28931,f276])).
fof(f28931,plain,(
  ( ! [X666,X668,X670,X667,X669,X671,X672] : (p(s(bool,X668)) | s(bool,X669) = s(bool,happ(s(fun(X666,bool),X667),s(X666,X670))) | s(bool,'T_0') = s(bool,X669) | s(bool,X671) = s(bool,'|dollar|exists_1'(s(fun(X666,bool),X667))) | s(bool,'T_0') = s(bool,X668) | s(bool,X671) = s(bool,X672) | s(bool,'T_0') = s(bool,X672)) )),
  inference(superposition,[],[f2629,f2054])).
fof(f31832,plain,(
  ( ! [X584,X586,X582,X585,X583] : (p(s(bool,X584)) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X583),s('type.extreal.extreal',X582))) | s(bool,X585) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X582),s('type.extreal.extreal',X583))) | s(bool,X585) = s(bool,X586) | s(bool,'T_0') = s(bool,X586)) )),
  inference(subsumption_resolution,[],[f28915,f276])).
fof(f28915,plain,(
  ( ! [X584,X586,X582,X585,X583] : (p(s(bool,X584)) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X583),s('type.extreal.extreal',X582))) | s(bool,X585) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X582),s('type.extreal.extreal',X583))) | s(bool,'T_0') = s(bool,X584) | s(bool,X585) = s(bool,X586) | s(bool,'T_0') = s(bool,X586)) )),
  inference(superposition,[],[f1296,f2054])).
fof(f31830,plain,(
  ( ! [X516,X518,X517,X519,X515] : (p(s(bool,X517)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X515,bool),X516))) | s(bool,X518) = s(bool,happ(s(fun(X515,bool),X516),s(X515,sK17(X516,X515)))) | s(bool,X518) = s(bool,X519) | s(bool,'T_0') = s(bool,X519)) )),
  inference(subsumption_resolution,[],[f28904,f276])).
fof(f28904,plain,(
  ( ! [X516,X518,X517,X519,X515] : (p(s(bool,X517)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X515,bool),X516))) | s(bool,X518) = s(bool,happ(s(fun(X515,bool),X516),s(X515,sK17(X516,X515)))) | s(bool,'T_0') = s(bool,X517) | s(bool,X518) = s(bool,X519) | s(bool,'T_0') = s(bool,X519)) )),
  inference(superposition,[],[f4966,f2054])).
fof(f31828,plain,(
  ( ! [X509,X511,X512,X514,X510,X513] : (p(s(bool,X511)) | s(bool,X512) = s(bool,'|dollar|exists_1'(s(fun(X509,bool),X510))) | s(bool,'T_0') = s(bool,X512) | s(bool,X513) = s(bool,happ(s(fun(X509,bool),X510),s(X509,sK17(X510,X509)))) | s(bool,X513) = s(bool,X514) | s(bool,'T_0') = s(bool,X514)) )),
  inference(subsumption_resolution,[],[f28903,f276])).
fof(f28903,plain,(
  ( ! [X509,X511,X512,X514,X510,X513] : (p(s(bool,X511)) | s(bool,X512) = s(bool,'|dollar|exists_1'(s(fun(X509,bool),X510))) | s(bool,'T_0') = s(bool,X512) | s(bool,X513) = s(bool,happ(s(fun(X509,bool),X510),s(X509,sK17(X510,X509)))) | s(bool,'T_0') = s(bool,X511) | s(bool,X513) = s(bool,X514) | s(bool,'T_0') = s(bool,X514)) )),
  inference(superposition,[],[f4303,f2054])).
fof(f31826,plain,(
  ( ! [X505,X507,X508,X504,X506] : (p(s(bool,X506)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X504,bool),X505))) | s(bool,X507) = s(bool,happ(s(fun(X504,bool),X505),s(X504,sK16(X505,X504)))) | s(bool,X507) = s(bool,X508) | s(bool,'T_0') = s(bool,X508)) )),
  inference(subsumption_resolution,[],[f28902,f276])).
fof(f28902,plain,(
  ( ! [X505,X507,X508,X504,X506] : (p(s(bool,X506)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X504,bool),X505))) | s(bool,X507) = s(bool,happ(s(fun(X504,bool),X505),s(X504,sK16(X505,X504)))) | s(bool,'T_0') = s(bool,X506) | s(bool,X507) = s(bool,X508) | s(bool,'T_0') = s(bool,X508)) )),
  inference(superposition,[],[f4967,f2054])).
fof(f31824,plain,(
  ( ! [X499,X500,X502,X498,X501,X503] : (p(s(bool,X500)) | s(bool,X501) = s(bool,'|dollar|exists_1'(s(fun(X498,bool),X499))) | s(bool,'T_0') = s(bool,X501) | s(bool,X502) = s(bool,happ(s(fun(X498,bool),X499),s(X498,sK16(X499,X498)))) | s(bool,X502) = s(bool,X503) | s(bool,'T_0') = s(bool,X503)) )),
  inference(subsumption_resolution,[],[f28901,f276])).
fof(f28901,plain,(
  ( ! [X499,X500,X502,X498,X501,X503] : (p(s(bool,X500)) | s(bool,X501) = s(bool,'|dollar|exists_1'(s(fun(X498,bool),X499))) | s(bool,'T_0') = s(bool,X501) | s(bool,X502) = s(bool,happ(s(fun(X498,bool),X499),s(X498,sK16(X499,X498)))) | s(bool,'T_0') = s(bool,X500) | s(bool,X502) = s(bool,X503) | s(bool,'T_0') = s(bool,X503)) )),
  inference(superposition,[],[f3897,f2054])).
fof(f31820,plain,(
  ( ! [X455,X457,X459,X456,X458] : (p(s(bool,X457)) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X455))),s('type.extreal.extreal',X456))) | s(bool,X458) = s(bool,happ(s(fun('type.extreal.extreal',bool),X455),s('type.extreal.extreal',sK22(X456,X455)))) | s(bool,X458) = s(bool,X459) | s(bool,'T_0') = s(bool,X459)) )),
  inference(subsumption_resolution,[],[f28893,f276])).
fof(f28893,plain,(
  ( ! [X455,X457,X459,X456,X458] : (p(s(bool,X457)) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X455))),s('type.extreal.extreal',X456))) | s(bool,X458) = s(bool,happ(s(fun('type.extreal.extreal',bool),X455),s('type.extreal.extreal',sK22(X456,X455)))) | s(bool,'T_0') = s(bool,X457) | s(bool,X458) = s(bool,X459) | s(bool,'T_0') = s(bool,X459)) )),
  inference(superposition,[],[f1866,f2054])).
fof(f31768,plain,(
  ( ! [X206,X204,X202,X205,X203,X201] : (p(s(bool,X202)) | sP1(X201,X203) | s(bool,X203) = s(bool,X204) | s(bool,'T_0') = s(bool,X204) | s(bool,X201) = s(bool,X205) | s(bool,X205) = s(bool,X206) | s(bool,'T_0') = s(bool,X206)) )),
  inference(subsumption_resolution,[],[f28840,f276])).
fof(f28840,plain,(
  ( ! [X206,X204,X202,X205,X203,X201] : (p(s(bool,X202)) | sP1(X201,X203) | s(bool,X203) = s(bool,X204) | s(bool,'T_0') = s(bool,X204) | s(bool,X201) = s(bool,X205) | s(bool,'T_0') = s(bool,X202) | s(bool,X205) = s(bool,X206) | s(bool,'T_0') = s(bool,X206)) )),
  inference(superposition,[],[f2610,f2054])).
fof(f31766,plain,(
  ( ! [X198,X196,X200,X199,X197,X195] : (p(s(bool,X196)) | sP1(X197,X195) | s(bool,X197) = s(bool,X198) | s(bool,'T_0') = s(bool,X198) | s(bool,X195) = s(bool,X199) | s(bool,X199) = s(bool,X200) | s(bool,'T_0') = s(bool,X200)) )),
  inference(subsumption_resolution,[],[f28839,f276])).
fof(f28839,plain,(
  ( ! [X198,X196,X200,X199,X197,X195] : (p(s(bool,X196)) | sP1(X197,X195) | s(bool,X197) = s(bool,X198) | s(bool,'T_0') = s(bool,X198) | s(bool,X195) = s(bool,X199) | s(bool,'T_0') = s(bool,X196) | s(bool,X199) = s(bool,X200) | s(bool,'T_0') = s(bool,X200)) )),
  inference(superposition,[],[f2609,f2054])).
fof(f31764,plain,(
  ( ! [X182,X180,X183,X181] : (p(s(bool,X181)) | sP1(X180,'T_0') | s(bool,X180) = s(bool,X182) | s(bool,X182) = s(bool,X183) | s(bool,'T_0') = s(bool,X183)) ) | $spl12),
  inference(subsumption_resolution,[],[f28836,f276])).
fof(f28836,plain,(
  ( ! [X182,X180,X183,X181] : (p(s(bool,X181)) | sP1(X180,'T_0') | s(bool,X180) = s(bool,X182) | s(bool,'T_0') = s(bool,X181) | s(bool,X182) = s(bool,X183) | s(bool,'T_0') = s(bool,X183)) ) | $spl12),
  inference(superposition,[],[f2558,f2054])).
fof(f31762,plain,(
  ( ! [X177,X178,X176,X179] : (p(s(bool,X177)) | sP1('T_0',X176) | s(bool,X176) = s(bool,X178) | s(bool,X178) = s(bool,X179) | s(bool,'T_0') = s(bool,X179)) ) | $spl8),
  inference(subsumption_resolution,[],[f28835,f276])).
fof(f28835,plain,(
  ( ! [X177,X178,X176,X179] : (p(s(bool,X177)) | sP1('T_0',X176) | s(bool,X176) = s(bool,X178) | s(bool,'T_0') = s(bool,X177) | s(bool,X178) = s(bool,X179) | s(bool,'T_0') = s(bool,X179)) ) | $spl8),
  inference(superposition,[],[f2554,f2054])).
fof(f31761,plain,(
  ( ! [X175,X173,X174,X172,X170] : (s(bool,X170) != s(bool,X172) | s(bool,'T_0') = s(bool,X175) | s(bool,X174) = s(bool,X175) | s(bool,X170) = s(bool,X174) | s(bool,'T_0') = s(bool,X173) | s(bool,X172) = s(bool,X173)) ) | $spl78),
  inference(cnf_transformation,[],[f31761_D])).
fof(f31761_D,plain,(
  ( ! [X175,X173,X174,X172,X170] : (s(bool,X170) != s(bool,X172) | s(bool,'T_0') = s(bool,X175) | s(bool,X174) = s(bool,X175) | s(bool,X170) = s(bool,X174) | s(bool,'T_0') = s(bool,X173) | s(bool,X172) = s(bool,X173)) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f31760,plain,(
  ( ! [X175,X173,X171,X174,X172,X170] : (p(s(bool,X171)) | s(bool,X170) != s(bool,X172) | s(bool,X172) = s(bool,X173) | s(bool,'T_0') = s(bool,X173) | s(bool,X170) = s(bool,X174) | s(bool,X174) = s(bool,X175) | s(bool,'T_0') = s(bool,X175)) )),
  inference(subsumption_resolution,[],[f28834,f276])).
fof(f28834,plain,(
  ( ! [X175,X173,X171,X174,X172,X170] : (p(s(bool,X171)) | s(bool,X170) != s(bool,X172) | s(bool,X172) = s(bool,X173) | s(bool,'T_0') = s(bool,X173) | s(bool,X170) = s(bool,X174) | s(bool,'T_0') = s(bool,X171) | s(bool,X174) = s(bool,X175) | s(bool,'T_0') = s(bool,X175)) )),
  inference(superposition,[],[f2204,f2054])).
fof(f21561,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | $spl77),
  inference(cnf_transformation,[],[f21561_D])).
fof(f21561_D,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) <=> ~$spl77),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl77])])).
fof(f21558,plain,(
  ( ! [X146,X147] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X146) = s(fun(bool,bool),X147)) ) | ($spl52 | $spl65)),
  inference(forward_demodulation,[],[f21522,f8380])).
fof(f11737,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f11736,f263])).
fof(f11736,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f11228,f10159])).
fof(f10159,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK22(sK14,sK13)))) | $spl2),
  inference(unit_resulting_resolution,[],[f9491,f1867])).
fof(f9491,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13))),s('type.extreal.extreal',sK14)))) | $spl2),
  inference(unit_resulting_resolution,[],[f462,f324])).
fof(f11228,plain,(
  ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK22(sK14,sK13))))) | ($spl0 | $spl2)),
  inference(unit_resulting_resolution,[],[f10679,f461])).
fof(f10679,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK22(sK14,sK13))))) | $spl2),
  inference(unit_resulting_resolution,[],[f10161,f325])).
fof(f10161,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK22(sK14,sK13)),s('type.extreal.extreal',sK14)))) | $spl2),
  inference(unit_resulting_resolution,[],[f9491,f328])).
fof(f10139,plain,(
  $false | ($spl0 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f10138,f263])).
fof(f10138,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f10137,f530])).
fof(f10137,plain,(
  ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK15)))) | ($spl0 | $spl4)),
  inference(subsumption_resolution,[],[f10131,f263])).
fof(f10131,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK15)))) | ($spl0 | $spl4)),
  inference(superposition,[],[f461,f529])).
fof(f529,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15))) | $spl4),
  inference(unit_resulting_resolution,[],[f465,f277])).
fof(f9487,plain,(
  ( ! [X50,X49] : (sP1('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))),X49) | s(bool,'T_0') = s(bool,X50) | s(bool,X49) = s(bool,X50)) ) | $spl72),
  inference(cnf_transformation,[],[f9487_D])).
fof(f9487_D,plain,(
  ( ! [X50,X49] : (sP1('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))),X49) | s(bool,'T_0') = s(bool,X50) | s(bool,X49) = s(bool,X50)) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f9486,plain,(
  ( ! [X50,X48,X49] : (p(s(bool,X48)) | sP1('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))),X49) | s(bool,X49) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) ) | $spl3),
  inference(subsumption_resolution,[],[f9468,f276])).
fof(f9468,plain,(
  ( ! [X50,X48,X49] : (p(s(bool,X48)) | sP1('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))),X49) | s(bool,X49) = s(bool,X50) | s(bool,'T_0') = s(bool,X50) | s(bool,'T_0') = s(bool,X48)) ) | $spl3),
  inference(superposition,[],[f2610,f2014])).
fof(f2014,plain,(
  ( ! [X9] : (s(bool,X9) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13))))) | s(bool,'T_0') = s(bool,X9)) ) | $spl3),
  inference(subsumption_resolution,[],[f2011,f468])).
fof(f2011,plain,(
  ( ! [X9] : (p(s(bool,'F_0')) | s(bool,X9) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13))))) | s(bool,'T_0') = s(bool,X9)) ) | $spl3),
  inference(superposition,[],[f846,f547])).
fof(f547,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13))))) | $spl3),
  inference(unit_resulting_resolution,[],[f463,f281])).
fof(f9485,plain,(
  ( ! [X31,X32] : (sP1('F_0',X31) | s(bool,'T_0') = s(bool,X32) | s(bool,X31) = s(bool,X32)) ) | $spl70),
  inference(cnf_transformation,[],[f9485_D])).
fof(f9485_D,plain,(
  ( ! [X31,X32] : (sP1('F_0',X31) | s(bool,'T_0') = s(bool,X32) | s(bool,X31) = s(bool,X32)) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f9484,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | sP1('F_0',X31) | s(bool,X31) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(subsumption_resolution,[],[f9462,f276])).
fof(f9462,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | sP1('F_0',X31) | s(bool,X31) = s(bool,X32) | s(bool,'T_0') = s(bool,X32) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f2610,f270])).
fof(f9420,plain,(
  ( ! [X50,X49] : (sP1(X49,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13))))) | s(bool,'T_0') = s(bool,X50) | s(bool,X49) = s(bool,X50)) ) | $spl68),
  inference(cnf_transformation,[],[f9420_D])).
fof(f9420_D,plain,(
  ( ! [X50,X49] : (sP1(X49,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13))))) | s(bool,'T_0') = s(bool,X50) | s(bool,X49) = s(bool,X50)) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f9419,plain,(
  ( ! [X50,X48,X49] : (p(s(bool,X48)) | sP1(X49,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13))))) | s(bool,X49) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) ) | $spl3),
  inference(subsumption_resolution,[],[f9401,f276])).
fof(f9401,plain,(
  ( ! [X50,X48,X49] : (p(s(bool,X48)) | sP1(X49,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13))))) | s(bool,X49) = s(bool,X50) | s(bool,'T_0') = s(bool,X50) | s(bool,'T_0') = s(bool,X48)) ) | $spl3),
  inference(superposition,[],[f2609,f2014])).
fof(f9418,plain,(
  ( ! [X31,X32] : (sP1(X31,'F_0') | s(bool,'T_0') = s(bool,X32) | s(bool,X31) = s(bool,X32)) ) | $spl66),
  inference(cnf_transformation,[],[f9418_D])).
fof(f9418_D,plain,(
  ( ! [X31,X32] : (sP1(X31,'F_0') | s(bool,'T_0') = s(bool,X32) | s(bool,X31) = s(bool,X32)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f9417,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | sP1(X31,'F_0') | s(bool,X31) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(subsumption_resolution,[],[f9395,f276])).
fof(f9395,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | sP1(X31,'F_0') | s(bool,X31) = s(bool,X32) | s(bool,'T_0') = s(bool,X32) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f2609,f270])).
fof(f7209,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,g(s(bool,sK23(X3,X4,bool,bool)))) | s(bool,X5) = s(bool,sK23(X3,X4,bool,bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1976,f783])).
fof(f7208,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK23(X0,X1,bool,bool)))) | s(bool,X2) = s(bool,sK23(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1976,f783])).
fof(f7100,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl56),
  inference(cnf_transformation,[],[f7100_D])).
fof(f7100_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f7096,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK23(X0,X1,bool,bool)) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) | $spl52),
  inference(cnf_transformation,[],[f7096_D])).
fof(f7096_D,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK23(X0,X1,bool,bool)) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f7103,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl61),
  inference(cnf_transformation,[],[f7103_D])).
fof(f7103_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f7094,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK23(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1975,f270])).
fof(f7101,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl58),
  inference(cnf_transformation,[],[f7101_D])).
fof(f7101_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f7093,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK23(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1975,f783])).
fof(f7092,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK23(X7,X8,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1975,f783])).
fof(f7095,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl50),
  inference(cnf_transformation,[],[f7095_D])).
fof(f7095_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f7098,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl55),
  inference(cnf_transformation,[],[f7098_D])).
fof(f7098_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl55),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl55])])).
fof(f7091,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK23(X3,X4,bool,bool)) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1975,f783])).
fof(f7090,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK23(X0,X1,bool,bool)) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1975,f783])).
fof(f6888,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5011,f5026])).
fof(f5026,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f3884,f3884,f321])).
fof(f5011,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f3884,f280])).
fof(f6887,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5017,f263])).
fof(f5017,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f525,f3884,f321])).
fof(f6886,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f5019,f263])).
fof(f5019,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f526,f3884,f321])).
fof(f526,plain,(
  s(bool,'T_0') != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13))))) | $spl3),
  inference(unit_resulting_resolution,[],[f463,f276])).
fof(f6885,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f5023,f3884])).
fof(f5023,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f526,f3884,f321])).
fof(f6884,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5027,f263])).
fof(f5027,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f525,f3884,f321])).
fof(f6883,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f5029,f263])).
fof(f5029,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f526,f3884,f321])).
fof(f6882,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f5033,f3884])).
fof(f5033,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f526,f3884,f321])).
fof(f6881,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5040,f263])).
fof(f5040,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f586,f3884,f349])).
fof(f6880,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5041,f3884])).
fof(f5041,plain,(
  ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK15)))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f636,f3884,f349])).
fof(f636,plain,(
  sP1('F_0',happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK15))) | $spl6),
  inference(unit_resulting_resolution,[],[f468,f591])).
fof(f591,plain,(
  ( ! [X2] : (sP1(X2,happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK15))) | p(s(bool,X2))) ) | $spl6),
  inference(subsumption_resolution,[],[f589,f263])).
fof(f589,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | p(s(bool,X2)) | sP1(X2,happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK15)))) ) | $spl6),
  inference(superposition,[],[f353,f530])).
fof(f6879,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5042,f3884])).
fof(f5042,plain,(
  ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK16(sK13,'type.extreal.extreal'))))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f3915,f3884,f349])).
fof(f3915,plain,(
  sP1('F_0',happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK16(sK13,'type.extreal.extreal')))) | $spl6),
  inference(unit_resulting_resolution,[],[f468,f1819,f353])).
fof(f1819,plain,(
  p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK16(sK13,'type.extreal.extreal'))))) | $spl6),
  inference(subsumption_resolution,[],[f1816,f263])).
fof(f1816,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK16(sK13,'type.extreal.extreal'))))) | $spl6),
  inference(superposition,[],[f304,f530])).
fof(f6878,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5043,f3884])).
fof(f5043,plain,(
  ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK17(sK13,'type.extreal.extreal'))))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f4341,f3884,f349])).
fof(f4341,plain,(
  sP1('F_0',happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK17(sK13,'type.extreal.extreal')))) | $spl6),
  inference(unit_resulting_resolution,[],[f468,f1833,f353])).
fof(f1833,plain,(
  p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK17(sK13,'type.extreal.extreal'))))) | $spl6),
  inference(subsumption_resolution,[],[f1830,f263])).
fof(f1830,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK17(sK13,'type.extreal.extreal'))))) | $spl6),
  inference(superposition,[],[f305,f530])).
fof(f6877,plain,(
  $false | ($spl4 | $spl46)),
  inference(subsumption_resolution,[],[f5044,f3884])).
fof(f5044,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15)))) | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f626,f3884,f349])).
fof(f626,plain,(
  sP1('F_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15))) | $spl4),
  inference(unit_resulting_resolution,[],[f468,f590])).
fof(f590,plain,(
  ( ! [X0] : (sP1(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15))) | p(s(bool,X0))) ) | $spl4),
  inference(subsumption_resolution,[],[f587,f263])).
fof(f587,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP1(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15)))) ) | $spl4),
  inference(superposition,[],[f353,f529])).
fof(f6876,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5045,f3884])).
fof(f5045,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1466,f3884,f349])).
fof(f1466,plain,(
  sP1('F_0','|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13))) | $spl6),
  inference(unit_resulting_resolution,[],[f468,f1452,f353])).
fof(f6875,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5055,f263])).
fof(f5055,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f578,f3884,f349])).
fof(f6874,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5062,f3884])).
fof(f5062,plain,(
  ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK15)))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f631,f3884,f349])).
fof(f631,plain,(
  sP1(happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK15)),'F_0') | $spl6),
  inference(unit_resulting_resolution,[],[f468,f583])).
fof(f583,plain,(
  ( ! [X2] : (sP1(happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK15)),X2) | p(s(bool,X2))) ) | $spl6),
  inference(subsumption_resolution,[],[f581,f263])).
fof(f581,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | p(s(bool,X2)) | sP1(happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK15)),X2)) ) | $spl6),
  inference(superposition,[],[f350,f530])).
fof(f6873,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5063,f3884])).
fof(f5063,plain,(
  ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK16(sK13,'type.extreal.extreal'))))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f3914,f3884,f349])).
fof(f3914,plain,(
  sP1(happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK16(sK13,'type.extreal.extreal'))),'F_0') | $spl6),
  inference(unit_resulting_resolution,[],[f468,f1819,f350])).
fof(f6872,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5064,f3884])).
fof(f5064,plain,(
  ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK17(sK13,'type.extreal.extreal'))))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f4340,f3884,f349])).
fof(f4340,plain,(
  sP1(happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK17(sK13,'type.extreal.extreal'))),'F_0') | $spl6),
  inference(unit_resulting_resolution,[],[f468,f1833,f350])).
fof(f6871,plain,(
  $false | ($spl4 | $spl46)),
  inference(subsumption_resolution,[],[f5065,f3884])).
fof(f5065,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15)))) | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f609,f3884,f349])).
fof(f609,plain,(
  sP1('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f468,f582])).
fof(f582,plain,(
  ( ! [X0] : (sP1('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15)),X0) | p(s(bool,X0))) ) | $spl4),
  inference(subsumption_resolution,[],[f579,f263])).
fof(f579,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP1('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15)),X0)) ) | $spl4),
  inference(superposition,[],[f350,f529])).
fof(f6870,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5066,f3884])).
fof(f5066,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1465,f3884,f349])).
fof(f1465,plain,(
  sP1('|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),'F_0') | $spl6),
  inference(unit_resulting_resolution,[],[f468,f1452,f350])).
fof(f6869,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5094,f3884])).
fof(f5094,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f622,f3884,f371])).
fof(f6868,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5109,f3884])).
fof(f5109,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f613,f3884,f372])).
fof(f6867,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5116,f3884])).
fof(f5116,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f622,f3884,f373])).
fof(f6866,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5121,f3884])).
fof(f5121,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f613,f3884,f373])).
fof(f6865,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f5128,f3884])).
fof(f5128,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f888,f468,f3884,f381])).
fof(f888,plain,(
  sP6('T_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))),'F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f263,f547,f398])).
fof(f6864,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5129,f263])).
fof(f5129,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f878,f468,f3884,f381])).
fof(f6863,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5130,f3884])).
fof(f5130,plain,(
  ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK15)))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1427,f468,f3884,f381])).
fof(f1427,plain,(
  sP6('F_0',happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK15)),'F_0') | $spl6),
  inference(unit_resulting_resolution,[],[f496,f487,f480,f468,f384])).
fof(f496,plain,(
  ( ! [X0,X1] : (sP3(X0,happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK15)),X1)) ) | $spl6),
  inference(unit_resulting_resolution,[],[f467,f395])).
fof(f6862,plain,(
  $false | ($spl4 | $spl46)),
  inference(subsumption_resolution,[],[f5131,f3884])).
fof(f5131,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15)))) | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f1428,f468,f3884,f381])).
fof(f1428,plain,(
  sP6('F_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f495,f487,f480,f468,f384])).
fof(f495,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15)),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f465,f395])).
fof(f6861,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5132,f3884])).
fof(f5132,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1927,f468,f3884,f381])).
fof(f1927,plain,(
  sP6('F_0','|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),'F_0') | $spl6),
  inference(unit_resulting_resolution,[],[f468,f482,f487,f1527,f384])).
fof(f1527,plain,(
  ( ! [X0,X1] : (sP3(X0,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),X1)) ) | $spl6),
  inference(unit_resulting_resolution,[],[f1452,f395])).
fof(f6860,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5217,f263])).
fof(f5217,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f1361,f468,f3884,f381])).
fof(f6859,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f5218,f263])).
fof(f5218,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f888,f468,f3884,f381])).
fof(f6858,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5223,f3884])).
fof(f5223,plain,(
  ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK15)))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1363,f468,f3884,f381])).
fof(f1363,plain,(
  sP6(happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK15)),'F_0','F_0') | $spl6),
  inference(unit_resulting_resolution,[],[f493,f486,f480,f468,f383])).
fof(f486,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK15)))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f467,f390])).
fof(f493,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK15)))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f467,f394])).
fof(f6857,plain,(
  $false | ($spl4 | $spl46)),
  inference(subsumption_resolution,[],[f5224,f3884])).
fof(f5224,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15)))) | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f1362,f468,f3884,f381])).
fof(f1362,plain,(
  sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15)),'F_0','F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f492,f485,f480,f468,f383])).
fof(f485,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f465,f390])).
fof(f492,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f465,f394])).
fof(f6856,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5225,f3884])).
fof(f5225,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1915,f468,f3884,f381])).
fof(f1915,plain,(
  sP6('|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),'F_0','F_0') | $spl6),
  inference(unit_resulting_resolution,[],[f468,f480,f489,f1526,f383])).
fof(f1526,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f1452,f394])).
fof(f6855,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5254,f5676])).
fof(f5676,plain,(
  ( ! [X0] : (~sP6(X0,'T_0','F_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f525,f3884,f399])).
fof(f5254,plain,(
  ( ! [X0] : (sP6(X0,'T_0','F_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f494,f487,f477,f3884,f382])).
fof(f6854,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f5259,f5678])).
fof(f5678,plain,(
  ( ! [X0] : (~sP6(X0,'T_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f526,f3884,f399])).
fof(f5259,plain,(
  ( ! [X0] : (sP6(X0,'T_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f494,f488,f477,f3884,f382])).
fof(f488,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))),X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f463,f391])).
fof(f6853,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5370,f5676])).
fof(f5370,plain,(
  ( ! [X0] : (sP6(X0,'T_0','F_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f494,f487,f480,f3884,f382])).
fof(f6852,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5375,f904])).
fof(f5375,plain,(
  sP6('T_0','T_0','F_0') | $spl46),
  inference(unit_resulting_resolution,[],[f494,f484,f480,f3884,f382])).
fof(f6851,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5378,f1485])).
fof(f1485,plain,(
  ~sP6('T_0','|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),'F_0') | $spl6),
  inference(unit_resulting_resolution,[],[f263,f468,f1452,f381])).
fof(f5378,plain,(
  sP6('T_0','|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),'F_0') | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1527,f484,f480,f3884,f382])).
fof(f6850,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5380,f5676])).
fof(f5380,plain,(
  sP6(happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK15)),'T_0','F_0') | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f494,f486,f480,f3884,f382])).
fof(f6849,plain,(
  $false | ($spl4 | $spl46)),
  inference(subsumption_resolution,[],[f5385,f5676])).
fof(f5385,plain,(
  sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15)),'T_0','F_0') | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f494,f485,f480,f3884,f382])).
fof(f6848,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5390,f1535])).
fof(f1535,plain,(
  ~sP6('|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),'T_0','F_0') | $spl6),
  inference(unit_resulting_resolution,[],[f525,f1452,f399])).
fof(f5390,plain,(
  sP6('|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),'T_0','F_0') | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f494,f1525,f480,f3884,f382])).
fof(f1525,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f1452,f390])).
fof(f6847,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5393,f1486])).
fof(f1486,plain,(
  ~sP6('|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),'F_0') | $spl6),
  inference(unit_resulting_resolution,[],[f1452,f468,f1452,f381])).
fof(f5393,plain,(
  sP6('|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),'F_0') | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1527,f1525,f480,f3884,f382])).
fof(f6846,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f5406,f5678])).
fof(f5406,plain,(
  ( ! [X0] : (sP6(X0,'T_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f494,f488,f481,f3884,f382])).
fof(f481,plain,(
  ( ! [X0,X1] : (sP5(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))),X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f463,f387])).
fof(f6845,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f5411,f905])).
fof(f905,plain,(
  ~sP6('T_0','T_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13))))) | $spl3),
  inference(unit_resulting_resolution,[],[f263,f526,f399])).
fof(f5411,plain,(
  sP6('T_0','T_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f494,f484,f481,f3884,f382])).
fof(f6844,plain,(
  $false | ($spl3 | $spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5416,f5678])).
fof(f5416,plain,(
  sP6(happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK15)),'T_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13))))) | ($spl3 | $spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f494,f486,f481,f3884,f382])).
fof(f6843,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(subsumption_resolution,[],[f5421,f5678])).
fof(f5421,plain,(
  sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15)),'T_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13))))) | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f494,f485,f481,f3884,f382])).
fof(f6842,plain,(
  $false | ($spl3 | $spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5426,f1536])).
fof(f1536,plain,(
  ~sP6('|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),'T_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13))))) | ($spl3 | $spl6)),
  inference(unit_resulting_resolution,[],[f526,f1452,f399])).
fof(f5426,plain,(
  sP6('|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),'T_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13))))) | ($spl3 | $spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f494,f1525,f481,f3884,f382])).
fof(f6841,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f5454,f5676])).
fof(f5454,plain,(
  sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))),'T_0','F_0') | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f494,f487,f483,f3884,f382])).
fof(f483,plain,(
  ( ! [X0,X1] : (sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))),X0,X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f463,f388])).
fof(f6840,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f5458,f5678])).
fof(f5458,plain,(
  sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))),'T_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f494,f488,f483,f3884,f382])).
fof(f6839,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5470,f5552])).
fof(f5552,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f3884,f386])).
fof(f5470,plain,(
  ~sP5('T_0','F_0','T_0') | $spl46),
  inference(unit_resulting_resolution,[],[f494,f487,f904,f3884,f382])).
fof(f6838,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5471,f5598])).
fof(f5598,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f3884,f390])).
fof(f5471,plain,(
  ~sP4('T_0','F_0','T_0') | $spl46),
  inference(unit_resulting_resolution,[],[f494,f480,f904,f3884,f382])).
fof(f6837,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5472,f5604])).
fof(f5604,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f3884,f396])).
fof(f5472,plain,(
  ~sP3('F_0','T_0','T_0') | $spl46),
  inference(unit_resulting_resolution,[],[f487,f480,f904,f3884,f382])).
fof(f6836,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f5474,f5552])).
fof(f5474,plain,(
  ~sP5('T_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))),'T_0') | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f494,f488,f905,f3884,f382])).
fof(f6835,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f5475,f5598])).
fof(f5475,plain,(
  ~sP4('T_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))),'T_0') | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f494,f481,f905,f3884,f382])).
fof(f6834,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f5476,f5604])).
fof(f5476,plain,(
  ~sP3('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))),'T_0','T_0') | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f488,f481,f905,f3884,f382])).
fof(f6833,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5482,f5552])).
fof(f5482,plain,(
  ~sP5('T_0','F_0','|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1527,f487,f1485,f3884,f382])).
fof(f6832,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5483,f5598])).
fof(f5483,plain,(
  ~sP4('|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),'F_0','T_0') | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1527,f480,f1485,f3884,f382])).
fof(f6831,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5484,f5604])).
fof(f5484,plain,(
  ~sP3('F_0','|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),'T_0') | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f487,f480,f1485,f3884,f382])).
fof(f6830,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5499,f5552])).
fof(f5499,plain,(
  ~sP5('|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),'F_0','T_0') | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f494,f487,f1535,f3884,f382])).
fof(f6829,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5500,f5598])).
fof(f5500,plain,(
  ~sP4('T_0','F_0','|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f494,f480,f1535,f3884,f382])).
fof(f6828,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5501,f5604])).
fof(f5501,plain,(
  ~sP3('F_0','T_0','|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f487,f480,f1535,f3884,f382])).
fof(f6827,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5504,f5552])).
fof(f5504,plain,(
  ~sP5('|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),'F_0','|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1527,f487,f1486,f3884,f382])).
fof(f6826,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5505,f5598])).
fof(f5505,plain,(
  ~sP4('|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),'F_0','|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1527,f480,f1486,f3884,f382])).
fof(f6825,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5506,f5604])).
fof(f5506,plain,(
  ~sP3('F_0','|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f487,f480,f1486,f3884,f382])).
fof(f6824,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f5509,f3884])).
fof(f5509,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f481,f468,f3884,f385])).
fof(f6823,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5510,f3884])).
fof(f5510,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f482,f468,f3884,f385])).
fof(f6822,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f5511,f3884])).
fof(f5511,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f468,f3884,f385])).
fof(f6821,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5518,f5552])).
fof(f5518,plain,(
  ( ! [X0] : (~sP5(X0,'T_0','F_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f468,f263,f3884,f385])).
fof(f6820,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5528,f5552])).
fof(f5528,plain,(
  ( ! [X0,X1] : (~sP5(X0,X1,'F_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f468,f3884,f3884,f385])).
fof(f6819,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5530,f3884])).
fof(f5530,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f480,f468,f3884,f385])).
fof(f6818,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f5531,f3884])).
fof(f5531,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f481,f468,f3884,f385])).
fof(f6817,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f5533,f3884])).
fof(f5533,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f483,f468,f3884,f385])).
fof(f6816,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5540,f5552])).
fof(f5540,plain,(
  ( ! [X0] : (~sP5('T_0',X0,'F_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f468,f263,f3884,f385])).
fof(f6815,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5550,f5552])).
fof(f5550,plain,(
  ( ! [X0,X1] : (~sP5(X0,X1,'F_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f468,f3884,f3884,f385])).
fof(f6814,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f5555,f3884])).
fof(f5555,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f488,f468,f3884,f389])).
fof(f6813,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5556,f3884])).
fof(f5556,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f489,f468,f3884,f389])).
fof(f6812,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f5557,f3884])).
fof(f5557,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f490,f468,f3884,f389])).
fof(f490,plain,(
  ( ! [X0,X1] : (sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))),X0,X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f463,f392])).
fof(f6811,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5564,f5598])).
fof(f5564,plain,(
  ( ! [X0] : (~sP4(X0,'T_0','F_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f468,f263,f3884,f389])).
fof(f6810,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5574,f5598])).
fof(f5574,plain,(
  ( ! [X0,X1] : (~sP4(X0,X1,'F_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f468,f3884,f3884,f389])).
fof(f6809,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5576,f3884])).
fof(f5576,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f487,f468,f3884,f389])).
fof(f6808,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f5577,f3884])).
fof(f5577,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f488,f468,f3884,f389])).
fof(f6807,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f5579,f3884])).
fof(f5579,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f490,f468,f3884,f389])).
fof(f6806,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5586,f5598])).
fof(f5586,plain,(
  ( ! [X0] : (~sP4('T_0',X0,'F_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f468,f263,f3884,f389])).
fof(f6805,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5596,f5598])).
fof(f5596,plain,(
  ( ! [X0,X1] : (~sP4(X0,X1,'F_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f468,f3884,f3884,f389])).
fof(f6801,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5614,f530])).
fof(f5614,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK15))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f918,f3884,f398])).
fof(f918,plain,(
  ~sP6('F_0','T_0',happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK15))) | $spl6),
  inference(unit_resulting_resolution,[],[f468,f530,f400])).
fof(f6800,plain,(
  $false | ($spl4 | $spl46)),
  inference(subsumption_resolution,[],[f5615,f529])).
fof(f5615,plain,(
  s(bool,'T_0') != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15))) | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f916,f3884,f398])).
fof(f916,plain,(
  ~sP6('F_0','T_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15))) | $spl4),
  inference(unit_resulting_resolution,[],[f468,f529,f400])).
fof(f6799,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5616,f1457])).
fof(f5616,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1896,f3884,f398])).
fof(f1896,plain,(
  ~sP6('F_0','T_0','|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13))) | $spl6),
  inference(unit_resulting_resolution,[],[f1518,f379])).
fof(f1518,plain,(
  ~sP4('T_0','|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),'F_0') | $spl6),
  inference(unit_resulting_resolution,[],[f263,f468,f1452,f389])).
fof(f6798,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f5618,f547])).
fof(f5618,plain,(
  s(bool,'F_0') != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f919,f3884,f398])).
fof(f919,plain,(
  ~sP6('F_0','F_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13))))) | $spl3),
  inference(unit_resulting_resolution,[],[f468,f547,f400])).
fof(f6797,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5619,f530])).
fof(f5619,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK15))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f914,f3884,f398])).
fof(f914,plain,(
  ~sP6('F_0',happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK15)),'T_0') | $spl6),
  inference(unit_resulting_resolution,[],[f468,f530,f400])).
fof(f6796,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5620,f303])).
fof(f5620,plain,(
  ( ! [X0,X1] : (s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) != s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f1764,f3884,f398])).
fof(f6795,plain,(
  $false | ($spl4 | $spl46)),
  inference(subsumption_resolution,[],[f5621,f529])).
fof(f5621,plain,(
  s(bool,'T_0') != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15))) | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f912,f3884,f398])).
fof(f912,plain,(
  ~sP6('F_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15)),'T_0') | $spl4),
  inference(unit_resulting_resolution,[],[f468,f529,f400])).
fof(f6794,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f5622,f547])).
fof(f5622,plain,(
  s(bool,'F_0') != s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f915,f3884,f398])).
fof(f915,plain,(
  ~sP6('F_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))),'F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f468,f547,f400])).
fof(f6793,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5623,f1457])).
fof(f5623,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1891,f3884,f398])).
fof(f1891,plain,(
  ~sP6('F_0','|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),'T_0') | $spl6),
  inference(unit_resulting_resolution,[],[f1511,f379])).
fof(f1511,plain,(
  ~sP4('|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),'T_0','F_0') | $spl6),
  inference(unit_resulting_resolution,[],[f263,f468,f1452,f389])).
fof(f6792,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5624,f303])).
fof(f5624,plain,(
  ( ! [X0,X1] : (s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) != s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f1766,f3884,f398])).
fof(f6790,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5660,f525])).
fof(f5660,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl46),
  inference(unit_resulting_resolution,[],[f878,f3884,f399])).
fof(f6789,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f5661,f526])).
fof(f5661,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f879,f3884,f399])).
fof(f879,plain,(
  sP6('F_0','T_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13))))) | $spl3),
  inference(unit_resulting_resolution,[],[f468,f526,f397])).
fof(f6788,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5662,f525])).
fof(f5662,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl46),
  inference(unit_resulting_resolution,[],[f876,f3884,f399])).
fof(f6787,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f6786,f525])).
fof(f6786,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl6 | $spl46)),
  inference(forward_demodulation,[],[f5663,f530])).
fof(f5663,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK15))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1376,f3884,f399])).
fof(f1376,plain,(
  sP6('F_0','F_0',happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK15))) | $spl6),
  inference(unit_resulting_resolution,[],[f499,f489,f482,f468,f383])).
fof(f499,plain,(
  ( ! [X0,X1] : (sP3(happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK15)),X0,X1)) ) | $spl6),
  inference(unit_resulting_resolution,[],[f467,f396])).
fof(f6785,plain,(
  $false | ($spl4 | $spl46)),
  inference(subsumption_resolution,[],[f6784,f525])).
fof(f6784,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl4 | $spl46)),
  inference(forward_demodulation,[],[f5664,f529])).
fof(f5664,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15))) | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f1377,f3884,f399])).
fof(f1377,plain,(
  sP6('F_0','F_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15))) | $spl4),
  inference(unit_resulting_resolution,[],[f498,f489,f482,f468,f383])).
fof(f498,plain,(
  ( ! [X0,X1] : (sP3('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15)),X0,X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f465,f396])).
fof(f6783,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f6782,f525])).
fof(f6782,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl6 | $spl46)),
  inference(forward_demodulation,[],[f5665,f1457])).
fof(f5665,plain,(
  s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1943,f3884,f399])).
fof(f1943,plain,(
  sP6('F_0','F_0','|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13))) | $spl6),
  inference(unit_resulting_resolution,[],[f468,f482,f489,f1528,f383])).
fof(f1528,plain,(
  ( ! [X0,X1] : (sP3('|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),X0,X1)) ) | $spl6),
  inference(unit_resulting_resolution,[],[f1452,f396])).
fof(f6781,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f6780,f525])).
fof(f6780,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl6 | $spl46)),
  inference(forward_demodulation,[],[f5666,f530])).
fof(f5666,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK15))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1427,f3884,f399])).
fof(f6779,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f5667,f526])).
fof(f5667,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f877,f3884,f399])).
fof(f877,plain,(
  sP6('F_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))),'T_0') | $spl3),
  inference(unit_resulting_resolution,[],[f468,f526,f397])).
fof(f6778,plain,(
  $false | ($spl4 | $spl46)),
  inference(subsumption_resolution,[],[f6777,f525])).
fof(f6777,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl4 | $spl46)),
  inference(forward_demodulation,[],[f5668,f529])).
fof(f5668,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15))) | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f1428,f3884,f399])).
fof(f6776,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f6775,f525])).
fof(f6775,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl6 | $spl46)),
  inference(forward_demodulation,[],[f5669,f1457])).
fof(f5669,plain,(
  s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1927,f3884,f399])).
fof(f6774,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5697,f5945])).
fof(f5945,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f3884,f3884,f3884,f414])).
fof(f5697,plain,(
  ( ! [X0,X1] : (~sP8(X0,X1,'F_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f468,f3884,f402])).
fof(f6773,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5712,f5945])).
fof(f5712,plain,(
  ( ! [X0,X1] : (~sP8(X0,'F_0',X1)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f468,f3884,f403])).
fof(f6772,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5715,f3884])).
fof(f5715,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f505,f645,f3884,f404])).
fof(f6771,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5716,f3884])).
fof(f5716,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f500,f645,f3884,f404])).
fof(f6770,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f5718,f3884])).
fof(f5718,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f504,f645,f3884,f404])).
fof(f504,plain,(
  ( ! [X0,X1] : (sP7(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))),X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f463,f410])).
fof(f6769,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5720,f3884])).
fof(f5720,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f505,f1540,f3884,f404])).
fof(f1540,plain,(
  ( ! [X0] : (~sP8('|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),'F_0',X0)) ) | $spl6),
  inference(unit_resulting_resolution,[],[f468,f1452,f403])).
fof(f6768,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5721,f3884])).
fof(f5721,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1569,f1540,f3884,f404])).
fof(f1569,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f1452,f409])).
fof(f6767,plain,(
  $false | ($spl3 | $spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5723,f3884])).
fof(f5723,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))))) | ($spl3 | $spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f504,f1540,f3884,f404])).
fof(f6766,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5730,f5845])).
fof(f5845,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f3884,f409])).
fof(f5730,plain,(
  ~sP7('F_0','T_0','T_0') | $spl46),
  inference(unit_resulting_resolution,[],[f645,f263,f3884,f404])).
fof(f6765,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5734,f5845])).
fof(f5734,plain,(
  ~sP7('F_0','T_0','|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1540,f263,f3884,f404])).
fof(f6764,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5744,f5845])).
fof(f5744,plain,(
  ( ! [X0] : (~sP7(X0,'F_0','T_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f644,f3884,f3884,f404])).
fof(f6763,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5746,f5845])).
fof(f5746,plain,(
  ( ! [X0] : (~sP7('F_0',X0,'T_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f645,f3884,f3884,f404])).
fof(f6762,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5753,f5845])).
fof(f5753,plain,(
  ( ! [X0] : (~sP7(X0,'F_0','|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)))) ) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1539,f3884,f3884,f404])).
fof(f1539,plain,(
  ( ! [X0] : (~sP8('|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),X0,'F_0')) ) | $spl6),
  inference(unit_resulting_resolution,[],[f468,f1452,f402])).
fof(f6761,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5755,f5845])).
fof(f5755,plain,(
  ( ! [X0] : (~sP7('F_0',X0,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)))) ) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1540,f3884,f3884,f404])).
fof(f6760,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5757,f3884])).
fof(f5757,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f503,f644,f3884,f404])).
fof(f6759,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5758,f3884])).
fof(f5758,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f500,f644,f3884,f404])).
fof(f6758,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f5760,f3884])).
fof(f5760,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f506,f644,f3884,f404])).
fof(f506,plain,(
  ( ! [X0,X1] : (sP7('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))),X0,X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f463,f411])).
fof(f6757,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5762,f3884])).
fof(f5762,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f503,f1539,f3884,f404])).
fof(f6756,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5763,f3884])).
fof(f5763,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1569,f1539,f3884,f404])).
fof(f6755,plain,(
  $false | ($spl3 | $spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5765,f3884])).
fof(f5765,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))))) | ($spl3 | $spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f506,f1539,f3884,f404])).
fof(f6754,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5773,f5845])).
fof(f5773,plain,(
  ~sP7('T_0','F_0','T_0') | $spl46),
  inference(unit_resulting_resolution,[],[f644,f263,f3884,f404])).
fof(f6753,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5778,f5845])).
fof(f5778,plain,(
  ~sP7('T_0','F_0','|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1539,f263,f3884,f404])).
fof(f6752,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5788,f5845])).
fof(f5788,plain,(
  ( ! [X0] : (~sP7(X0,'F_0','T_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f644,f3884,f3884,f404])).
fof(f6751,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5790,f5845])).
fof(f5790,plain,(
  ( ! [X0] : (~sP7('F_0',X0,'T_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f645,f3884,f3884,f404])).
fof(f6750,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5797,f5845])).
fof(f5797,plain,(
  ( ! [X0] : (~sP7(X0,'F_0','|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)))) ) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1539,f3884,f3884,f404])).
fof(f6749,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5799,f5845])).
fof(f5799,plain,(
  ( ! [X0] : (~sP7('F_0',X0,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)))) ) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1540,f3884,f3884,f404])).
fof(f6748,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f5802,f3884])).
fof(f5802,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f504,f468,f3884,f408])).
fof(f6747,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5803,f3884])).
fof(f5803,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f505,f468,f3884,f408])).
fof(f6746,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f5804,f3884])).
fof(f5804,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f506,f468,f3884,f408])).
fof(f6745,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5811,f5845])).
fof(f5811,plain,(
  ( ! [X0] : (~sP7(X0,'T_0','F_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f468,f263,f3884,f408])).
fof(f6744,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5821,f5845])).
fof(f5821,plain,(
  ( ! [X0,X1] : (~sP7(X0,X1,'F_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f468,f3884,f3884,f408])).
fof(f6743,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5823,f3884])).
fof(f5823,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f503,f468,f3884,f408])).
fof(f6742,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f5824,f3884])).
fof(f5824,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f504,f468,f3884,f408])).
fof(f6741,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f5826,f3884])).
fof(f5826,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f506,f468,f3884,f408])).
fof(f6740,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5833,f5845])).
fof(f5833,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'F_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f468,f263,f3884,f408])).
fof(f6739,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5843,f5845])).
fof(f5843,plain,(
  ( ! [X0,X1] : (~sP7(X0,X1,'F_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f468,f3884,f3884,f408])).
fof(f6738,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5848,f3884])).
fof(f5848,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f645,f263,f3884,f414])).
fof(f6737,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5850,f3884])).
fof(f5850,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f645,f3884,f3884,f414])).
fof(f6736,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5851,f263])).
fof(f5851,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f1041,f3884,f3884,f414])).
fof(f6735,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5852,f263])).
fof(f5852,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1584,f3884,f3884,f414])).
fof(f1584,plain,(
  ~sP8('F_0','|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),'T_0') | $spl6),
  inference(unit_resulting_resolution,[],[f263,f468,f1452,f417])).
fof(f6734,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5853,f3884])).
fof(f5853,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1588,f3884,f3884,f414])).
fof(f1588,plain,(
  ~sP8('F_0','T_0','|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13))) | $spl6),
  inference(unit_resulting_resolution,[],[f263,f468,f1452,f417])).
fof(f6733,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5854,f3884])).
fof(f5854,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1589,f3884,f3884,f414])).
fof(f1589,plain,(
  ~sP8('F_0','|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13))) | $spl6),
  inference(unit_resulting_resolution,[],[f1452,f468,f1452,f417])).
fof(f6732,plain,(
  $false | ($spl4 | $spl46)),
  inference(subsumption_resolution,[],[f5855,f3884])).
fof(f5855,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15)))) | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f1309,f3884,f3884,f414])).
fof(f1309,plain,(
  ~sP8('F_0','T_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15))) | $spl4),
  inference(unit_resulting_resolution,[],[f1297,f401])).
fof(f1297,plain,(
  ~sP7('T_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15)),'F_0') | $spl4),
  inference(unit_resulting_resolution,[],[f263,f468,f1037])).
fof(f1037,plain,(
  ( ! [X8,X9] : (~sP7(X8,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15)),X9) | p(s(bool,X9)) | ~p(s(bool,X8))) ) | $spl4),
  inference(subsumption_resolution,[],[f1030,f263])).
fof(f1030,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X8)) | p(s(bool,X9)) | ~sP7(X8,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15)),X9)) ) | $spl4),
  inference(superposition,[],[f408,f529])).
fof(f6731,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5857,f3884])).
fof(f5857,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1540,f3884,f3884,f414])).
fof(f6730,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5858,f263])).
fof(f5858,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f645,f263,f3884,f414])).
fof(f6729,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5861,f3884])).
fof(f5861,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1540,f263,f3884,f414])).
fof(f6728,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5866,f263])).
fof(f5866,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f644,f3884,f3884,f414])).
fof(f6727,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5867,f263])).
fof(f5867,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f645,f3884,f3884,f414])).
fof(f6726,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5873,f3884])).
fof(f5873,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1539,f3884,f3884,f414])).
fof(f6725,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5874,f3884])).
fof(f5874,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1540,f3884,f3884,f414])).
fof(f6724,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5879,f3884])).
fof(f5879,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f644,f263,f3884,f414])).
fof(f6723,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5881,f3884])).
fof(f5881,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f644,f3884,f3884,f414])).
fof(f6722,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5883,f263])).
fof(f5883,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f1041,f3884,f3884,f414])).
fof(f6721,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5884,f3884])).
fof(f5884,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1584,f3884,f3884,f414])).
fof(f6720,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5885,f263])).
fof(f5885,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1588,f3884,f3884,f414])).
fof(f6719,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5886,f3884])).
fof(f5886,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1589,f3884,f3884,f414])).
fof(f6718,plain,(
  $false | ($spl4 | $spl46)),
  inference(subsumption_resolution,[],[f5887,f263])).
fof(f5887,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f1309,f3884,f3884,f414])).
fof(f6717,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5888,f3884])).
fof(f5888,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1539,f3884,f3884,f414])).
fof(f6716,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5890,f263])).
fof(f5890,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f644,f263,f3884,f414])).
fof(f6715,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5894,f3884])).
fof(f5894,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1539,f263,f3884,f414])).
fof(f6714,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5899,f263])).
fof(f5899,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f644,f3884,f3884,f414])).
fof(f6713,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5900,f263])).
fof(f5900,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f645,f3884,f3884,f414])).
fof(f6712,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5906,f3884])).
fof(f5906,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1539,f3884,f3884,f414])).
fof(f6711,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5907,f3884])).
fof(f5907,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1540,f3884,f3884,f414])).
fof(f6710,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5913,f263])).
fof(f5913,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f1041,f263,f3884,f414])).
fof(f6709,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5914,f3884])).
fof(f5914,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1584,f263,f3884,f414])).
fof(f6708,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5916,f3884])).
fof(f5916,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f644,f3884,f3884,f414])).
fof(f6707,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5918,f263])).
fof(f5918,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f1041,f3884,f3884,f414])).
fof(f6706,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5919,f3884])).
fof(f5919,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1584,f3884,f3884,f414])).
fof(f6705,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5920,f263])).
fof(f5920,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1588,f3884,f3884,f414])).
fof(f6704,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5921,f3884])).
fof(f5921,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1589,f3884,f3884,f414])).
fof(f6703,plain,(
  $false | ($spl4 | $spl46)),
  inference(subsumption_resolution,[],[f5922,f263])).
fof(f5922,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f1309,f3884,f3884,f414])).
fof(f6702,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5923,f3884])).
fof(f5923,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1539,f3884,f3884,f414])).
fof(f6701,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5926,f263])).
fof(f5926,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f1041,f263,f3884,f414])).
fof(f6700,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5927,f3884])).
fof(f5927,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1588,f263,f3884,f414])).
fof(f6699,plain,(
  $false | ($spl4 | $spl46)),
  inference(subsumption_resolution,[],[f5928,f3884])).
fof(f5928,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15)))) | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f1309,f263,f3884,f414])).
fof(f6698,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5935,f3884])).
fof(f5935,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f645,f3884,f3884,f414])).
fof(f6697,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5936,f263])).
fof(f5936,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f1041,f3884,f3884,f414])).
fof(f6696,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5937,f263])).
fof(f5937,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1584,f3884,f3884,f414])).
fof(f6695,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5938,f3884])).
fof(f5938,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1588,f3884,f3884,f414])).
fof(f6694,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5939,f3884])).
fof(f5939,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1589,f3884,f3884,f414])).
fof(f6693,plain,(
  $false | ($spl4 | $spl46)),
  inference(subsumption_resolution,[],[f5940,f3884])).
fof(f5940,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15)))) | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f1309,f3884,f3884,f414])).
fof(f6692,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f5942,f3884])).
fof(f5942,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1540,f3884,f3884,f414])).
fof(f6691,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f5948,f3884])).
fof(f5948,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f773,f468,f3884,f417])).
fof(f773,plain,(
  ( ! [X0] : (sP8('F_0',X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f504,f468,f454])).
fof(f6690,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5949,f3884])).
fof(f5949,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f654,f468,f3884,f417])).
fof(f6689,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f5950,f3884])).
fof(f5950,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f775,f468,f3884,f417])).
fof(f775,plain,(
  ( ! [X0] : (sP8('F_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))),X0)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f506,f468,f454])).
fof(f6688,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f5976,f3884])).
fof(f5976,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f647,f468,f3884,f417])).
fof(f6687,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f5977,f3884])).
fof(f5977,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f773,f468,f3884,f417])).
fof(f6686,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f5979,f3884])).
fof(f5979,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f775,f468,f3884,f417])).
fof(f6685,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f6022,f6514])).
fof(f6514,plain,(
  ( ! [X0,X1] : (sP10('F_0',X0,X1)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f513,f3884,f453])).
fof(f6022,plain,(
  ( ! [X0,X1] : (~sP10('F_0',X0,X1)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f468,f3884,f418])).
fof(f6684,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f6041,f6514])).
fof(f6041,plain,(
  ( ! [X0,X1] : (~sP10('F_0',X0,X1)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f468,f3884,f419])).
fof(f6683,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f6063,f6116])).
fof(f6116,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f3884,f427])).
fof(f6063,plain,(
  ~sP9('F_0','T_0','F_0') | $spl46),
  inference(unit_resulting_resolution,[],[f659,f468,f3884,f422])).
fof(f6682,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f6065,f6116])).
fof(f6065,plain,(
  ~sP9('F_0','|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),'F_0') | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1598,f468,f3884,f422])).
fof(f1598,plain,(
  ( ! [X0] : (~sP10('F_0','|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),X0)) ) | $spl6),
  inference(unit_resulting_resolution,[],[f468,f1452,f419])).
fof(f6681,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f6087,f6116])).
fof(f6087,plain,(
  ~sP9('F_0','F_0','T_0') | $spl46),
  inference(unit_resulting_resolution,[],[f658,f468,f3884,f423])).
fof(f6680,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f6088,f6116])).
fof(f6088,plain,(
  ~sP9('F_0','F_0','|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1592,f468,f3884,f423])).
fof(f1592,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f468,f1452,f418])).
fof(f6679,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f6112,f6116])).
fof(f6112,plain,(
  ( ! [X0] : (~sP9(X0,'F_0','F_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f468,f468,f3884,f425])).
fof(f6678,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f6122,f263])).
fof(f6122,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f1142,f3884,f430])).
fof(f6677,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f6123,f3884])).
fof(f6123,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1629,f3884,f430])).
fof(f1629,plain,(
  ~sP10('|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),'F_0','F_0') | $spl6),
  inference(unit_resulting_resolution,[],[f468,f468,f1452,f432])).
fof(f6676,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f6128,f263])).
fof(f6128,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f658,f3884,f430])).
fof(f6675,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f6129,f3884])).
fof(f6129,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1592,f3884,f430])).
fof(f6674,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f6130,f3884])).
fof(f6130,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f659,f3884,f430])).
fof(f6673,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f6131,f3884])).
fof(f6131,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1598,f3884,f430])).
fof(f6672,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f6142,f263])).
fof(f6142,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f1142,f3884,f431])).
fof(f6671,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f6143,f3884])).
fof(f6143,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1629,f3884,f431])).
fof(f6670,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f6148,f3884])).
fof(f6148,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f658,f3884,f431])).
fof(f6669,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f6149,f3884])).
fof(f6149,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1592,f3884,f431])).
fof(f6668,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f6150,f263])).
fof(f6150,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f659,f3884,f431])).
fof(f6667,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f6151,f3884])).
fof(f6151,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1598,f3884,f431])).
fof(f6666,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f6194,f6489])).
fof(f6489,plain,(
  ( ! [X0,X1] : (sP12('F_0',X0,X1)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f520,f3884,f452])).
fof(f6194,plain,(
  ( ! [X0,X1] : (~sP12('F_0',X0,X1)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f468,f3884,f436])).
fof(f6665,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f6197,f3884])).
fof(f6197,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f520,f690,f3884,f439])).
fof(f6664,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f6198,f3884])).
fof(f6198,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f517,f690,f3884,f439])).
fof(f6663,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f6200,f3884])).
fof(f6200,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f516,f690,f3884,f439])).
fof(f516,plain,(
  ( ! [X0,X1] : (sP11(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f463,f443])).
fof(f6662,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f6201,f3884])).
fof(f6201,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f520,f1630,f3884,f439])).
fof(f1630,plain,(
  ( ! [X0] : (~sP12('F_0','|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),X0)) ) | $spl6),
  inference(unit_resulting_resolution,[],[f468,f1452,f436])).
fof(f6661,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f6202,f3884])).
fof(f6202,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1669,f1630,f3884,f439])).
fof(f1669,plain,(
  ( ! [X0,X1] : (sP11(X0,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),X1)) ) | $spl6),
  inference(unit_resulting_resolution,[],[f1452,f444])).
fof(f6660,plain,(
  $false | ($spl3 | $spl6 | $spl46)),
  inference(subsumption_resolution,[],[f6204,f3884])).
fof(f6204,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))))) | ($spl3 | $spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f516,f1630,f3884,f439])).
fof(f6659,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f6211,f6363])).
fof(f6363,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f3884,f444])).
fof(f6211,plain,(
  ~sP11('F_0','T_0','T_0') | $spl46),
  inference(unit_resulting_resolution,[],[f690,f263,f3884,f439])).
fof(f6658,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f6213,f6363])).
fof(f6213,plain,(
  ~sP11('F_0','|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),'T_0') | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1630,f263,f3884,f439])).
fof(f6657,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f6225,f6363])).
fof(f6225,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'F_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f688,f3884,f3884,f439])).
fof(f6656,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f6227,f6363])).
fof(f6227,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f690,f3884,f3884,f439])).
fof(f6655,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f6229,f6363])).
fof(f6229,plain,(
  ( ! [X0] : (~sP11('F_0','|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),X0)) ) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1630,f3884,f3884,f439])).
fof(f6654,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f6254,f6363])).
fof(f6254,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'F_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f688,f3884,f3884,f439])).
fof(f6653,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f6255,f6363])).
fof(f6255,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f690,f3884,f3884,f439])).
fof(f6652,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f6256,f6363])).
fof(f6256,plain,(
  ( ! [X0] : (~sP11('F_0','|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),X0)) ) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1630,f3884,f3884,f439])).
fof(f6651,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f6287,f6363])).
fof(f6287,plain,(
  ~sP11('F_0','F_0','F_0') | $spl46),
  inference(unit_resulting_resolution,[],[f688,f468,f3884,f440])).
fof(f6650,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f6309,f6363])).
fof(f6309,plain,(
  ~sP11('F_0','F_0','F_0') | $spl46),
  inference(unit_resulting_resolution,[],[f688,f468,f3884,f441])).
fof(f6649,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f6334,f3884])).
fof(f6334,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f516,f468,f3884,f442])).
fof(f6648,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f6335,f3884])).
fof(f6335,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f520,f468,f3884,f442])).
fof(f6647,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f6336,f3884])).
fof(f6336,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f521,f468,f3884,f442])).
fof(f521,plain,(
  ( ! [X0,X1] : (sP11('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))),X0,X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f463,f445])).
fof(f6646,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f6337,f6363])).
fof(f6337,plain,(
  ( ! [X0] : (~sP11(X0,'F_0','T_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f263,f468,f3884,f442])).
fof(f6645,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f6339,f6363])).
fof(f6339,plain,(
  ( ! [X0,X1] : (~sP11(X0,'F_0',X1)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f3884,f468,f3884,f442])).
fof(f6644,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f6355,f3884])).
fof(f6355,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f515,f468,f3884,f442])).
fof(f6643,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f6356,f3884])).
fof(f6356,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f516,f468,f3884,f442])).
fof(f6642,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f6358,f3884])).
fof(f6358,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f521,f468,f3884,f442])).
fof(f6641,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f6359,f6363])).
fof(f6359,plain,(
  ( ! [X0] : (~sP11('T_0','F_0',X0)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f263,f468,f3884,f442])).
fof(f6640,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f6361,f6363])).
fof(f6361,plain,(
  ( ! [X0,X1] : (~sP11(X0,'F_0',X1)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f3884,f468,f3884,f442])).
fof(f6639,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f6385,f263])).
fof(f6385,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f1233,f3884,f448])).
fof(f6638,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f6386,f263])).
fof(f6386,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1678,f3884,f448])).
fof(f1678,plain,(
  ~sP12('T_0','F_0','|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13))) | $spl6),
  inference(unit_resulting_resolution,[],[f263,f468,f1452,f449])).
fof(f6637,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f6387,f3884])).
fof(f6387,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1683,f3884,f448])).
fof(f1683,plain,(
  ~sP12('|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),'F_0','T_0') | $spl6),
  inference(unit_resulting_resolution,[],[f263,f468,f1452,f449])).
fof(f6636,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f6392,f3884])).
fof(f6392,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f688,f3884,f448])).
fof(f6635,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f6393,f263])).
fof(f6393,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f690,f3884,f448])).
fof(f6634,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f6394,f3884])).
fof(f6394,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1630,f3884,f448])).
fof(f6633,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f6402,f263])).
fof(f6402,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f698,f468,f3884,f449])).
fof(f6632,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f6403,f263])).
fof(f6403,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f733,f468,f3884,f449])).
fof(f733,plain,(
  ( ! [X0] : (sP12('T_0',X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f516,f263,f452])).
fof(f6631,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f6406,f3884])).
fof(f6406,plain,(
  ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK15)))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1189,f468,f3884,f449])).
fof(f1189,plain,(
  ( ! [X0] : (sP12(happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK15)),X0,'F_0')) ) | $spl6),
  inference(unit_resulting_resolution,[],[f468,f703])).
fof(f703,plain,(
  ( ! [X4,X5] : (sP12(happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK15)),X5,X4) | p(s(bool,X4))) ) | $spl6),
  inference(subsumption_resolution,[],[f701,f263])).
fof(f701,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | p(s(bool,X4)) | sP12(happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK15)),X5,X4)) ) | $spl6),
  inference(superposition,[],[f447,f530])).
fof(f6630,plain,(
  $false | ($spl4 | $spl46)),
  inference(subsumption_resolution,[],[f6407,f3884])).
fof(f6407,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15)))) | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f707,f468,f3884,f449])).
fof(f707,plain,(
  ( ! [X0] : (sP12('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15)),X0,'F_0')) ) | $spl4),
  inference(unit_resulting_resolution,[],[f468,f702])).
fof(f702,plain,(
  ( ! [X0,X1] : (sP12('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15)),X1,X0) | p(s(bool,X0))) ) | $spl4),
  inference(subsumption_resolution,[],[f699,f263])).
fof(f699,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP12('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15)),X1,X0)) ) | $spl4),
  inference(superposition,[],[f447,f529])).
fof(f6629,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f6408,f3884])).
fof(f6408,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f1202,f468,f3884,f449])).
fof(f1202,plain,(
  sP12('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))),'F_0','T_0') | $spl3),
  inference(unit_resulting_resolution,[],[f521,f263,f468,f441])).
fof(f6628,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f6409,f3884])).
fof(f6409,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1685,f468,f3884,f449])).
fof(f1685,plain,(
  ( ! [X0] : (sP12('|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),X0,'F_0')) ) | $spl6),
  inference(unit_resulting_resolution,[],[f515,f1452,f452])).
fof(f6627,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f6436,f3884])).
fof(f6436,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))))) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f733,f468,f3884,f449])).
fof(f6626,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f6437,f263])).
fof(f6437,plain,(
  ~p(s(bool,'T_0')) | $spl46),
  inference(unit_resulting_resolution,[],[f1201,f468,f3884,f449])).
fof(f6625,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f6438,f3884])).
fof(f6438,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)))) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1670,f468,f3884,f449])).
fof(f1670,plain,(
  sP12('F_0','F_0','|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13))) | $spl6),
  inference(unit_resulting_resolution,[],[f468,f468,f1452,f446])).
fof(f6624,plain,(
  $false | ($spl3 | $spl46)),
  inference(subsumption_resolution,[],[f6441,f263])).
fof(f6441,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f1202,f468,f3884,f449])).
fof(f6623,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f6471,f515])).
fof(f6471,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'F_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f688,f3884,f452])).
fof(f6622,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f6472,f517])).
fof(f6472,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f690,f3884,f452])).
fof(f6621,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f6473,f1669])).
fof(f6473,plain,(
  ( ! [X0] : (~sP11('F_0','|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),X0)) ) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1630,f3884,f452])).
fof(f6620,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f6492,f507])).
fof(f6492,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) ) | $spl46),
  inference(unit_resulting_resolution,[],[f658,f3884,f453])).
fof(f6619,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f6493,f1617])).
fof(f1617,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f1452,f426])).
fof(f6493,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)))) ) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1592,f3884,f453])).
fof(f6618,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f6494,f510])).
fof(f6494,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) ) | $spl46),
  inference(unit_resulting_resolution,[],[f659,f3884,f453])).
fof(f6617,plain,(
  $false | ($spl6 | $spl46)),
  inference(subsumption_resolution,[],[f6495,f1618])).
fof(f1618,plain,(
  ( ! [X0,X1] : (sP9(X0,'|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),X1)) ) | $spl6),
  inference(unit_resulting_resolution,[],[f1452,f427])).
fof(f6495,plain,(
  ( ! [X0] : (~sP9('F_0','|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),X0)) ) | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1598,f3884,f453])).
fof(f6616,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f6573,f3884])).
fof(f6573,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f3884,f324])).
fof(f6615,plain,(
  $false | $spl46),
  inference(subsumption_resolution,[],[f6590,f3884])).
fof(f6590,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f3884,f324])).
fof(f6591,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f3884,f3884,f324])).
fof(f6574,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f3884,f3884,f324])).
fof(f6560,plain,(
  $false | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1683,f3884,f3839])).
fof(f3839,plain,(
  ( ! [X30,X29] : (sP12('|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),X29,X30) | ~p(s(bool,X29))) ) | $spl6),
  inference(subsumption_resolution,[],[f3800,f263])).
fof(f3800,plain,(
  ( ! [X30,X29] : (~p(s(bool,'T_0')) | ~p(s(bool,X29)) | sP12('|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),X29,X30)) ) | $spl6),
  inference(superposition,[],[f448,f1457])).
fof(f6558,plain,(
  $false | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1629,f3884,f3835])).
fof(f3835,plain,(
  ( ! [X21,X22] : (sP10('|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),X21,X22) | ~p(s(bool,X21))) ) | $spl6),
  inference(subsumption_resolution,[],[f3796,f263])).
fof(f3796,plain,(
  ( ! [X21,X22] : (~p(s(bool,'T_0')) | ~p(s(bool,X21)) | sP10('|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),X21,X22)) ) | $spl6),
  inference(superposition,[],[f431,f1457])).
fof(f6556,plain,(
  $false | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1629,f3884,f3834])).
fof(f3834,plain,(
  ( ! [X19,X20] : (sP10('|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),X20,X19) | ~p(s(bool,X19))) ) | $spl6),
  inference(subsumption_resolution,[],[f3795,f263])).
fof(f3795,plain,(
  ( ! [X19,X20] : (~p(s(bool,'T_0')) | ~p(s(bool,X19)) | sP10('|dollar|exists_1'(s(fun('type.extreal.extreal',bool),sK13)),X20,X19)) ) | $spl6),
  inference(superposition,[],[f430,f1457])).
fof(f6551,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f468,f506,f3884,f1037])).
fof(f6549,plain,(
  $false | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f468,f505,f3884,f1037])).
fof(f6542,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f468,f490,f3884,f983])).
fof(f983,plain,(
  ( ! [X8,X9] : (~sP4(X8,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15)),X9) | p(s(bool,X9)) | ~p(s(bool,X8))) ) | $spl4),
  inference(subsumption_resolution,[],[f976,f263])).
fof(f976,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X8)) | p(s(bool,X9)) | ~sP4(X8,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15)),X9)) ) | $spl4),
  inference(superposition,[],[f389,f529])).
fof(f6540,plain,(
  $false | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f468,f489,f3884,f983])).
fof(f6533,plain,(
  $false | ($spl3 | $spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f468,f483,f3884,f958])).
fof(f958,plain,(
  ( ! [X8,X9] : (~sP5(X8,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15)),X9) | p(s(bool,X9)) | ~p(s(bool,X8))) ) | $spl4),
  inference(subsumption_resolution,[],[f951,f263])).
fof(f951,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X8)) | p(s(bool,X9)) | ~sP5(X8,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15)),X9)) ) | $spl4),
  inference(superposition,[],[f385,f529])).
fof(f6531,plain,(
  $false | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f468,f482,f3884,f958])).
fof(f6518,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f613,f3884,f455])).
fof(f6515,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f658,f513,f3884,f453])).
fof(f6513,plain,(
  $false | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1598,f1618,f3884,f453])).
fof(f6511,plain,(
  $false | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f658,f511,f3884,f453])).
fof(f511,plain,(
  ( ! [X0,X1] : (sP9(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15)),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f465,f427])).
fof(f6509,plain,(
  $false | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f658,f512,f3884,f453])).
fof(f512,plain,(
  ( ! [X0,X1] : (sP9(X0,happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK15)),X1)) ) | $spl6),
  inference(unit_resulting_resolution,[],[f467,f427])).
fof(f6507,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f659,f510,f3884,f453])).
fof(f6505,plain,(
  $false | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1592,f1617,f3884,f453])).
fof(f6503,plain,(
  $false | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f659,f508,f3884,f453])).
fof(f508,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15)))) ) | $spl4),
  inference(unit_resulting_resolution,[],[f465,f426])).
fof(f6501,plain,(
  $false | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f659,f509,f3884,f453])).
fof(f509,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK15)))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f467,f426])).
fof(f6499,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f658,f507,f3884,f453])).
fof(f6490,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f688,f520,f3884,f452])).
fof(f6488,plain,(
  $false | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1630,f1669,f3884,f452])).
fof(f6486,plain,(
  $false | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f688,f518,f3884,f452])).
fof(f518,plain,(
  ( ! [X0,X1] : (sP11(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15)),X1)) ) | $spl4),
  inference(unit_resulting_resolution,[],[f465,f444])).
fof(f6484,plain,(
  $false | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f688,f519,f3884,f452])).
fof(f519,plain,(
  ( ! [X0,X1] : (sP11(X0,happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK15)),X1)) ) | $spl6),
  inference(unit_resulting_resolution,[],[f467,f444])).
fof(f6482,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f690,f517,f3884,f452])).
fof(f6480,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f690,f516,f3884,f452])).
fof(f6478,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f688,f515,f3884,f452])).
fof(f6470,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f698,f468,f3884,f3884,f449])).
fof(f6453,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f1201,f468,f263,f3884,f449])).
fof(f6434,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f698,f468,f3884,f3884,f449])).
fof(f6417,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f698,f468,f263,f3884,f449])).
fof(f6401,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f688,f3884,f3884,f448])).
fof(f6399,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f690,f263,f3884,f448])).
fof(f6391,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f688,f3884,f3884,f448])).
fof(f6389,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f1233,f263,f3884,f448])).
fof(f6381,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f688,f468,f3884,f447])).
fof(f6373,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f688,f468,f468,f3884,f446])).
fof(f6364,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f1218,f3884,f444])).
fof(f6362,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f515,f3884,f468,f3884,f442])).
fof(f6360,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f515,f263,f468,f3884,f442])).
fof(f6340,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f515,f3884,f468,f3884,f442])).
fof(f6338,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f520,f263,f468,f3884,f442])).
fof(f6317,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f688,f520,f468,f3884,f441])).
fof(f6314,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f688,f515,f468,f3884,f441])).
fof(f6295,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f688,f520,f468,f3884,f440])).
fof(f6292,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f688,f515,f468,f3884,f440])).
fof(f6273,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f688,f520,f3884,f3884,f439])).
fof(f6271,plain,(
  $false | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1630,f1669,f3884,f3884,f439])).
fof(f6269,plain,(
  $false | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f688,f518,f3884,f3884,f439])).
fof(f6267,plain,(
  $false | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f688,f519,f3884,f3884,f439])).
fof(f6265,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f690,f517,f3884,f3884,f439])).
fof(f6263,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f690,f516,f3884,f3884,f439])).
fof(f6261,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f688,f515,f3884,f3884,f439])).
fof(f6230,plain,(
  $false | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f520,f1630,f3884,f3884,f439])).
fof(f6228,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f520,f690,f3884,f3884,f439])).
fof(f6226,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f520,f688,f3884,f3884,f439])).
fof(f6214,plain,(
  $false | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f520,f1630,f263,f3884,f439])).
fof(f6212,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f520,f690,f263,f3884,f439])).
fof(f6195,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f1201,f468,f3884,f436])).
fof(f6178,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f1111,f468,f468,f3884,f432])).
fof(f6157,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f658,f3884,f3884,f431])).
fof(f6155,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f659,f263,f3884,f431])).
fof(f6147,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f658,f3884,f3884,f431])).
fof(f6145,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f1142,f263,f3884,f431])).
fof(f6137,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f658,f3884,f3884,f430])).
fof(f6135,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f658,f263,f3884,f430])).
fof(f6127,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f658,f3884,f3884,f430])).
fof(f6125,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f1142,f263,f3884,f430])).
fof(f6117,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f1126,f3884,f427])).
fof(f6115,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f1126,f3884,f426])).
fof(f6113,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f513,f468,f468,f3884,f425])).
fof(f6098,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f658,f513,f468,f3884,f423])).
fof(f6096,plain,(
  $false | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1592,f1617,f468,f3884,f423])).
fof(f6092,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f658,f507,f468,f3884,f423])).
fof(f6066,plain,(
  $false | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f513,f1598,f468,f3884,f422])).
fof(f6064,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f513,f659,f468,f3884,f422])).
fof(f6042,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f1111,f468,f3884,f419])).
fof(f6023,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f1111,f468,f3884,f418])).
fof(f6004,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f647,f468,f3884,f3884,f417])).
fof(f5989,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f647,f468,f263,f3884,f417])).
fof(f5975,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f647,f468,f3884,f3884,f417])).
fof(f5960,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f654,f468,f263,f3884,f417])).
fof(f5946,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f644,f3884,f3884,f3884,f414])).
fof(f5944,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f645,f263,f3884,f3884,f414])).
fof(f5933,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f644,f3884,f263,f3884,f414])).
fof(f5931,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f1041,f263,f263,f3884,f414])).
fof(f5911,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f644,f3884,f3884,f3884,f414])).
fof(f5909,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f644,f263,f3884,f3884,f414])).
fof(f5898,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f644,f3884,f263,f3884,f414])).
fof(f5896,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f644,f263,f263,f3884,f414])).
fof(f5878,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f644,f3884,f3884,f3884,f414])).
fof(f5876,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f644,f263,f3884,f3884,f414])).
fof(f5865,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f645,f3884,f263,f3884,f414])).
fof(f5863,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f645,f263,f263,f3884,f414])).
fof(f5846,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f1026,f3884,f409])).
fof(f5844,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f503,f468,f3884,f3884,f408])).
fof(f5834,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f503,f468,f263,f3884,f408])).
fof(f5822,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f503,f468,f3884,f3884,f408])).
fof(f5812,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f505,f468,f263,f3884,f408])).
fof(f5800,plain,(
  $false | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f505,f1540,f3884,f3884,f404])).
fof(f5798,plain,(
  $false | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f503,f1539,f3884,f3884,f404])).
fof(f5791,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f505,f645,f3884,f3884,f404])).
fof(f5789,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f503,f644,f3884,f3884,f404])).
fof(f5779,plain,(
  $false | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f503,f1539,f263,f3884,f404])).
fof(f5774,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f503,f644,f263,f3884,f404])).
fof(f5756,plain,(
  $false | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f505,f1540,f3884,f3884,f404])).
fof(f5754,plain,(
  $false | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f503,f1539,f3884,f3884,f404])).
fof(f5747,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f505,f645,f3884,f3884,f404])).
fof(f5745,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f503,f644,f3884,f3884,f404])).
fof(f5735,plain,(
  $false | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f505,f1540,f263,f3884,f404])).
fof(f5731,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f505,f645,f263,f3884,f404])).
fof(f5713,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f654,f468,f3884,f403])).
fof(f5698,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f647,f468,f3884,f402])).
fof(f5683,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f877,f526,f3884,f399])).
fof(f5681,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f876,f525,f3884,f399])).
fof(f5679,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f879,f526,f3884,f399])).
fof(f5677,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f878,f525,f3884,f399])).
fof(f5647,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f919,f547,f3884,f398])).
fof(f5645,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f1766,f303,f3884,f398])).
fof(f5643,plain,(
  $false | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1896,f1457,f3884,f398])).
fof(f5641,plain,(
  $false | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f918,f530,f3884,f398])).
fof(f5639,plain,(
  $false | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f916,f529,f3884,f398])).
fof(f5637,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f915,f547,f3884,f398])).
fof(f5635,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f1764,f303,f3884,f398])).
fof(f5633,plain,(
  $false | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1891,f1457,f3884,f398])).
fof(f5631,plain,(
  $false | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f914,f530,f3884,f398])).
fof(f5629,plain,(
  $false | ($spl4 | $spl46)),
  inference(unit_resulting_resolution,[],[f912,f529,f3884,f398])).
fof(f6614,plain,(
  $false | $spl46),
  inference(trivial_inequality_removal,[],[f5617])).
fof(f5617,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl46),
  inference(unit_resulting_resolution,[],[f1008,f3884,f398])).
fof(f6613,plain,(
  $false | $spl46),
  inference(trivial_inequality_removal,[],[f5613])).
fof(f5613,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl46),
  inference(unit_resulting_resolution,[],[f987,f3884,f398])).
fof(f5605,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f995,f3884,f396])).
fof(f5603,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f995,f3884,f395])).
fof(f5601,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f995,f3884,f394])).
fof(f5599,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f972,f3884,f390])).
fof(f5597,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f487,f468,f3884,f3884,f389])).
fof(f5587,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f487,f468,f263,f3884,f389])).
fof(f5575,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f487,f468,f3884,f3884,f389])).
fof(f5565,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f489,f468,f263,f3884,f389])).
fof(f5553,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f947,f3884,f386])).
fof(f5551,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f480,f468,f3884,f3884,f385])).
fof(f5541,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f480,f468,f263,f3884,f385])).
fof(f5529,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f480,f468,f3884,f3884,f385])).
fof(f5519,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f482,f468,f263,f3884,f385])).
fof(f5507,plain,(
  $false | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1527,f487,f480,f1486,f3884,f382])).
fof(f5502,plain,(
  $false | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f494,f487,f480,f1535,f3884,f382])).
fof(f5485,plain,(
  $false | ($spl6 | $spl46)),
  inference(unit_resulting_resolution,[],[f1527,f487,f480,f1485,f3884,f382])).
fof(f5477,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f494,f488,f481,f905,f3884,f382])).
fof(f5473,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f494,f487,f480,f904,f3884,f382])).
fof(f5229,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f1361,f3884,f468,f3884,f381])).
fof(f5227,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f1361,f263,f468,f3884,f381])).
fof(f5176,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f1361,f468,f3884,f3884,f381])).
fof(f5146,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f878,f468,f263,f3884,f381])).
fof(f5126,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f613,f3884,f3884,f373])).
fof(f5124,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f613,f263,f3884,f373])).
fof(f5120,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f613,f3884,f3884,f373])).
fof(f5118,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f622,f263,f3884,f373])).
fof(f5114,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f613,f3884,f3884,f372])).
fof(f5112,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f613,f263,f3884,f372])).
fof(f5108,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f613,f3884,f3884,f372])).
fof(f5103,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f613,f3884,f3884,f371])).
fof(f5098,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f613,f3884,f3884,f371])).
fof(f5096,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f622,f263,f3884,f371])).
fof(f5092,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f565,f468,f3884,f353])).
fof(f5081,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f565,f468,f3884,f350])).
fof(f5070,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f578,f3884,f3884,f349])).
fof(f5068,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f578,f263,f3884,f349])).
fof(f5054,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f578,f3884,f3884,f349])).
fof(f5052,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f586,f263,f3884,f349])).
fof(f5038,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f602,f3884,f335])).
fof(f5036,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f602,f3884,f334])).
fof(f5034,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f3884,f526,f3884,f321])).
fof(f5032,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f3884,f525,f3884,f321])).
fof(f5030,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f263,f526,f3884,f321])).
fof(f5028,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f263,f525,f3884,f321])).
fof(f5024,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f3884,f526,f3884,f321])).
fof(f5022,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f3884,f525,f3884,f321])).
fof(f5020,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f263,f526,f3884,f321])).
fof(f5018,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f263,f525,f3884,f321])).
fof(f5012,plain,(
  $false | ($spl3 | $spl46)),
  inference(unit_resulting_resolution,[],[f547,f3884,f280])).
fof(f5010,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f525,f3884,f277])).
fof(f5008,plain,(
  $false | $spl46),
  inference(unit_resulting_resolution,[],[f468,f3884])).
fof(f5005,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X6))),s('type.extreal.extreal',X7))) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X6),s('type.extreal.extreal',sK22(X7,X6))))) ) | $spl48),
  inference(cnf_transformation,[],[f5005_D])).
fof(f5005_D,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X6))),s('type.extreal.extreal',X7))) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X6),s('type.extreal.extreal',sK22(X7,X6))))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f5006,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X11)) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X9))),s('type.extreal.extreal',X10))) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X9),s('type.extreal.extreal',sK22(X10,X9))))) )),
  inference(subsumption_resolution,[],[f5002,f276])).
fof(f5002,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X11)) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X9))),s('type.extreal.extreal',X10))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X9),s('type.extreal.extreal',sK22(X10,X9))))) )),
  inference(superposition,[],[f1866,f783])).
fof(f5004,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X8)) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X6))),s('type.extreal.extreal',X7))) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X6),s('type.extreal.extreal',sK22(X7,X6))))) )),
  inference(subsumption_resolution,[],[f5001,f276])).
fof(f5001,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X8)) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X6))),s('type.extreal.extreal',X7))) | s(bool,'T_0') = s(bool,happ(s(fun('type.extreal.extreal',bool),X6),s('type.extreal.extreal',sK22(X7,X6)))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f1866,f783])).
fof(f3883,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X6))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7)))) ) | $spl44),
  inference(cnf_transformation,[],[f3883_D])).
fof(f3883_D,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X6))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7)))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f3885,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X11)) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X10),s('type.extreal.extreal',X9))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X9),s('type.extreal.extreal',X10)))) )),
  inference(subsumption_resolution,[],[f3879,f276])).
fof(f3879,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X11)) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X10),s('type.extreal.extreal',X9))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X9),s('type.extreal.extreal',X10)))) )),
  inference(superposition,[],[f1296,f783])).
fof(f3882,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X8)) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X6))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7)))) )),
  inference(subsumption_resolution,[],[f3878,f276])).
fof(f3878,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X8)) | s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X6))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f1296,f783])).
fof(f2823,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f262,f261,f260,f263,f468,f270,f277,f276,f281,f280,f303,f305,f304,f469,f312,f310,f316,f314,f317,f321,f318,f325,f324,f328,f327,f326,f335,f334,f470,f353,f350,f349,f348,f366,f367,f373,f372,f371,f455,f369,f368,f384,f383,f382,f381,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f396,f395,f394,f393,f400,f399,f398,f397,f454,f404,f403,f402,f401,f411,f410,f409,f408,f417,f414,f413,f412,f424,f423,f422,f453,f420,f419,f418,f428,f427,f426,f425,f432,f431,f430,f429,f441,f440,f439,f452,f437,f436,f435,f445,f444,f443,f442,f449,f448,f447,f446,f471,f474,f477,f480,f482,f484,f487,f489,f491,f494,f497,f500,f503,f505,f507,f510,f513,f515,f517,f520,f522,f525,f565,f575,f578,f586,f602,f613,f622,f639,f644,f645,f647,f654,f658,f659,f662,f674,f688,f690,f698,f710,f783,f846,f878,f876,f904,f901,f947,f972,f987,f995,f1008,f1012,f1026,f1041,f1111,f1126,f1142,f1201,f1218,f1233,f1296,f1313,f1361,f1767,f1766,f1765,f1764,f1768,f1769,f1779,f1780,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1817,f1831,f1840,f1866,f1867,f1898,f1923,f1947,f1959,f1975,f1976,f2039,f2054,f2064,f2134,f2135,f2136,f2137,f2138,f2139,f2140,f2204,f2553,f2557,f2561,f2564,f2567,f2570,f2573,f2576,f2579,f2582,f2585,f2588,f2591,f2594,f2597,f2600,f2603,f2605,f2606,f2607,f2608,f2609,f2610,f2611,f2612,f2613,f2614,f2615,f2616,f2617,f2618,f2619,f2620,f2621,f2622,f2623,f2624,f2625,f2626,f2627,f2628,f2629,f2630,f2631,f2632,f2633,f2634,f2635,f2636,f2637,f2290,f2291,f2292,f2293,f2294,f2295,f2296,f2643,f2645,f2648,f2650,f2652,f2654,f2656,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2555,f2768,f2786,f2788,f2790,f2792,f2794,f2796,f2798,f2800,f2802,f2804,f2806,f2808,f2810,f2812,f2814,f2816,f2818,f2820,f2822])).
fof(f2822,plain,(
  ( ! [X92,X90,X91,X89] : (~sP9(happ(s(fun(X89,bool),X90),s(X89,'const.min.@_1'(s(fun(X89,bool),X90)))),X91,X92) | sP10(happ(s(fun(X89,bool),X90),s(X89,'const.min.@_1'(s(fun(X89,bool),X90)))),X91,X92)) ) | $spl10),
  inference(subsumption_resolution,[],[f2750,f263])).
fof(f2750,plain,(
  ( ! [X92,X90,X91,X89] : (~p(s(bool,'T_0')) | ~sP9(happ(s(fun(X89,bool),X90),s(X89,'const.min.@_1'(s(fun(X89,bool),X90)))),X91,X92) | sP10(happ(s(fun(X89,bool),X90),s(X89,'const.min.@_1'(s(fun(X89,bool),X90)))),X91,X92)) ) | $spl10),
  inference(backward_demodulation,[],[f2555,f1797])).
fof(f2820,plain,(
  ( ! [X88,X87,X85,X86] : (~sP11(happ(s(fun(X85,bool),X86),s(X85,'const.min.@_1'(s(fun(X85,bool),X86)))),X87,X88) | sP12(happ(s(fun(X85,bool),X86),s(X85,'const.min.@_1'(s(fun(X85,bool),X86)))),X87,X88)) ) | $spl10),
  inference(subsumption_resolution,[],[f2749,f263])).
fof(f2749,plain,(
  ( ! [X88,X87,X85,X86] : (~p(s(bool,'T_0')) | ~sP11(happ(s(fun(X85,bool),X86),s(X85,'const.min.@_1'(s(fun(X85,bool),X86)))),X87,X88) | sP12(happ(s(fun(X85,bool),X86),s(X85,'const.min.@_1'(s(fun(X85,bool),X86)))),X87,X88)) ) | $spl10),
  inference(backward_demodulation,[],[f2555,f1796])).
fof(f2818,plain,(
  ( ! [X83,X81,X84,X82] : (~p(s(bool,X83)) | p(s(bool,X84)) | ~sP12(happ(s(fun(X81,bool),X82),s(X81,'const.min.@_1'(s(fun(X81,bool),X82)))),X84,X83)) ) | $spl10),
  inference(subsumption_resolution,[],[f2748,f263])).
fof(f2748,plain,(
  ( ! [X83,X81,X84,X82] : (~p(s(bool,'T_0')) | ~p(s(bool,X83)) | p(s(bool,X84)) | ~sP12(happ(s(fun(X81,bool),X82),s(X81,'const.min.@_1'(s(fun(X81,bool),X82)))),X84,X83)) ) | $spl10),
  inference(backward_demodulation,[],[f2555,f1795])).
fof(f2816,plain,(
  ( ! [X80,X78,X79,X77] : (~p(s(bool,X79)) | sP12(happ(s(fun(X77,bool),X78),s(X77,'const.min.@_1'(s(fun(X77,bool),X78)))),X79,X80)) ) | $spl10),
  inference(subsumption_resolution,[],[f2747,f263])).
fof(f2747,plain,(
  ( ! [X80,X78,X79,X77] : (~p(s(bool,'T_0')) | ~p(s(bool,X79)) | sP12(happ(s(fun(X77,bool),X78),s(X77,'const.min.@_1'(s(fun(X77,bool),X78)))),X79,X80)) ) | $spl10),
  inference(backward_demodulation,[],[f2555,f1794])).
fof(f2814,plain,(
  ( ! [X76,X74,X75,X73] : (p(s(bool,X75)) | sP12(happ(s(fun(X73,bool),X74),s(X73,'const.min.@_1'(s(fun(X73,bool),X74)))),X76,X75)) ) | $spl10),
  inference(subsumption_resolution,[],[f2746,f263])).
fof(f2746,plain,(
  ( ! [X76,X74,X75,X73] : (~p(s(bool,'T_0')) | p(s(bool,X75)) | sP12(happ(s(fun(X73,bool),X74),s(X73,'const.min.@_1'(s(fun(X73,bool),X74)))),X76,X75)) ) | $spl10),
  inference(backward_demodulation,[],[f2555,f1793])).
fof(f2812,plain,(
  ( ! [X70,X72,X71,X69] : (p(s(bool,X71)) | ~p(s(bool,X72)) | ~sP11(X72,X71,happ(s(fun(X69,bool),X70),s(X69,'const.min.@_1'(s(fun(X69,bool),X70)))))) ) | $spl10),
  inference(subsumption_resolution,[],[f2745,f263])).
fof(f2745,plain,(
  ( ! [X70,X72,X71,X69] : (~p(s(bool,'T_0')) | p(s(bool,X71)) | ~p(s(bool,X72)) | ~sP11(X72,X71,happ(s(fun(X69,bool),X70),s(X69,'const.min.@_1'(s(fun(X69,bool),X70)))))) ) | $spl10),
  inference(backward_demodulation,[],[f2555,f1792])).
fof(f2810,plain,(
  ( ! [X68,X66,X67,X65] : (p(s(bool,X67)) | p(s(bool,X68)) | ~sP10(happ(s(fun(X65,bool),X66),s(X65,'const.min.@_1'(s(fun(X65,bool),X66)))),X68,X67)) ) | $spl10),
  inference(subsumption_resolution,[],[f2744,f263])).
fof(f2744,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,'T_0')) | p(s(bool,X67)) | p(s(bool,X68)) | ~sP10(happ(s(fun(X65,bool),X66),s(X65,'const.min.@_1'(s(fun(X65,bool),X66)))),X68,X67)) ) | $spl10),
  inference(backward_demodulation,[],[f2555,f1791])).
fof(f2808,plain,(
  ( ! [X61,X64,X62,X63] : (~p(s(bool,X63)) | sP10(happ(s(fun(X61,bool),X62),s(X61,'const.min.@_1'(s(fun(X61,bool),X62)))),X63,X64)) ) | $spl10),
  inference(subsumption_resolution,[],[f2743,f263])).
fof(f2743,plain,(
  ( ! [X61,X64,X62,X63] : (~p(s(bool,'T_0')) | ~p(s(bool,X63)) | sP10(happ(s(fun(X61,bool),X62),s(X61,'const.min.@_1'(s(fun(X61,bool),X62)))),X63,X64)) ) | $spl10),
  inference(backward_demodulation,[],[f2555,f1790])).
fof(f2806,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,X59)) | sP10(happ(s(fun(X57,bool),X58),s(X57,'const.min.@_1'(s(fun(X57,bool),X58)))),X60,X59)) ) | $spl10),
  inference(subsumption_resolution,[],[f2742,f263])).
fof(f2742,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,'T_0')) | ~p(s(bool,X59)) | sP10(happ(s(fun(X57,bool),X58),s(X57,'const.min.@_1'(s(fun(X57,bool),X58)))),X60,X59)) ) | $spl10),
  inference(backward_demodulation,[],[f2555,f1789])).
fof(f2804,plain,(
  ( ! [X54,X56,X55,X53] : (p(s(bool,X55)) | p(s(bool,X56)) | ~sP9(happ(s(fun(X53,bool),X54),s(X53,'const.min.@_1'(s(fun(X53,bool),X54)))),X55,X56)) ) | $spl10),
  inference(subsumption_resolution,[],[f2741,f263])).
fof(f2741,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,'T_0')) | p(s(bool,X55)) | p(s(bool,X56)) | ~sP9(happ(s(fun(X53,bool),X54),s(X53,'const.min.@_1'(s(fun(X53,bool),X54)))),X55,X56)) ) | $spl10),
  inference(backward_demodulation,[],[f2555,f1788])).
fof(f2802,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X51)) | p(s(bool,X52)) | ~sP8(X52,X51,happ(s(fun(X49,bool),X50),s(X49,'const.min.@_1'(s(fun(X49,bool),X50)))))) ) | $spl10),
  inference(subsumption_resolution,[],[f2740,f263])).
fof(f2740,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X51)) | p(s(bool,X52)) | ~sP8(X52,X51,happ(s(fun(X49,bool),X50),s(X49,'const.min.@_1'(s(fun(X49,bool),X50)))))) ) | $spl10),
  inference(backward_demodulation,[],[f2555,f1787])).
fof(f2800,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,X47)) | p(s(bool,X48)) | ~sP7(X47,happ(s(fun(X45,bool),X46),s(X45,'const.min.@_1'(s(fun(X45,bool),X46)))),X48)) ) | $spl10),
  inference(subsumption_resolution,[],[f2739,f263])).
fof(f2739,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,'T_0')) | ~p(s(bool,X47)) | p(s(bool,X48)) | ~sP7(X47,happ(s(fun(X45,bool),X46),s(X45,'const.min.@_1'(s(fun(X45,bool),X46)))),X48)) ) | $spl10),
  inference(backward_demodulation,[],[f2555,f1786])).
fof(f2798,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,X43)) | p(s(bool,X44)) | ~sP4(X43,happ(s(fun(X41,bool),X42),s(X41,'const.min.@_1'(s(fun(X41,bool),X42)))),X44)) ) | $spl10),
  inference(subsumption_resolution,[],[f2738,f263])).
fof(f2738,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,'T_0')) | ~p(s(bool,X43)) | p(s(bool,X44)) | ~sP4(X43,happ(s(fun(X41,bool),X42),s(X41,'const.min.@_1'(s(fun(X41,bool),X42)))),X44)) ) | $spl10),
  inference(backward_demodulation,[],[f2555,f1785])).
fof(f2796,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,X39)) | p(s(bool,X40)) | ~sP5(X39,happ(s(fun(X37,bool),X38),s(X37,'const.min.@_1'(s(fun(X37,bool),X38)))),X40)) ) | $spl10),
  inference(subsumption_resolution,[],[f2737,f263])).
fof(f2737,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X39)) | p(s(bool,X40)) | ~sP5(X39,happ(s(fun(X37,bool),X38),s(X37,'const.min.@_1'(s(fun(X37,bool),X38)))),X40)) ) | $spl10),
  inference(backward_demodulation,[],[f2555,f1784])).
fof(f2794,plain,(
  ( ! [X35,X33,X36,X34] : (sP6(X35,X36,happ(s(fun(X33,bool),X34),s(X33,'const.min.@_1'(s(fun(X33,bool),X34))))) | ~sP5(X35,happ(s(fun(X33,bool),X34),s(X33,'const.min.@_1'(s(fun(X33,bool),X34)))),X36) | ~sP4(X36,happ(s(fun(X33,bool),X34),s(X33,'const.min.@_1'(s(fun(X33,bool),X34)))),X35) | ~sP3(happ(s(fun(X33,bool),X34),s(X33,'const.min.@_1'(s(fun(X33,bool),X34)))),X36,X35)) ) | $spl10),
  inference(subsumption_resolution,[],[f2736,f263])).
fof(f2736,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,'T_0')) | sP6(X35,X36,happ(s(fun(X33,bool),X34),s(X33,'const.min.@_1'(s(fun(X33,bool),X34))))) | ~sP5(X35,happ(s(fun(X33,bool),X34),s(X33,'const.min.@_1'(s(fun(X33,bool),X34)))),X36) | ~sP4(X36,happ(s(fun(X33,bool),X34),s(X33,'const.min.@_1'(s(fun(X33,bool),X34)))),X35) | ~sP3(happ(s(fun(X33,bool),X34),s(X33,'const.min.@_1'(s(fun(X33,bool),X34)))),X36,X35)) ) | $spl10),
  inference(backward_demodulation,[],[f2555,f1783])).
fof(f2792,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X32)) | sP1(X32,happ(s(fun(X30,bool),X31),s(X30,'const.min.@_1'(s(fun(X30,bool),X31)))))) ) | $spl10),
  inference(subsumption_resolution,[],[f2735,f263])).
fof(f2735,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,'T_0')) | p(s(bool,X32)) | sP1(X32,happ(s(fun(X30,bool),X31),s(X30,'const.min.@_1'(s(fun(X30,bool),X31)))))) ) | $spl10),
  inference(backward_demodulation,[],[f2555,f1782])).
fof(f2790,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X29)) | sP1(happ(s(fun(X27,bool),X28),s(X27,'const.min.@_1'(s(fun(X27,bool),X28)))),X29)) ) | $spl10),
  inference(subsumption_resolution,[],[f2734,f263])).
fof(f2734,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,'T_0')) | p(s(bool,X29)) | sP1(happ(s(fun(X27,bool),X28),s(X27,'const.min.@_1'(s(fun(X27,bool),X28)))),X29)) ) | $spl10),
  inference(backward_demodulation,[],[f2555,f1781])).
fof(f2788,plain,(
  ( ! [X26,X24,X25] : (sP0(happ(s(fun(X24,bool),X25),s(X24,'const.min.@_1'(s(fun(X24,bool),X25)))),X26)) ) | $spl10),
  inference(subsumption_resolution,[],[f2733,f263])).
fof(f2733,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,'T_0')) | sP0(happ(s(fun(X24,bool),X25),s(X24,'const.min.@_1'(s(fun(X24,bool),X25)))),X26)) ) | $spl10),
  inference(backward_demodulation,[],[f2555,f1780])).
fof(f2786,plain,(
  ( ! [X23,X21,X22] : (sP0(X23,happ(s(fun(X21,bool),X22),s(X21,'const.min.@_1'(s(fun(X21,bool),X22)))))) ) | $spl10),
  inference(subsumption_resolution,[],[f2732,f263])).
fof(f2732,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,'T_0')) | sP0(X23,happ(s(fun(X21,bool),X22),s(X21,'const.min.@_1'(s(fun(X21,bool),X22)))))) ) | $spl10),
  inference(backward_demodulation,[],[f2555,f1779])).
fof(f2768,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl10),
  inference(subsumption_resolution,[],[f2677,f2555])).
fof(f2677,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl10),
  inference(backward_demodulation,[],[f2555,f280])).
fof(f2821,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f262,f261,f260,f263,f468,f270,f277,f276,f281,f280,f303,f305,f304,f469,f312,f310,f316,f314,f317,f321,f318,f325,f324,f328,f327,f326,f335,f334,f470,f353,f350,f349,f348,f366,f367,f373,f372,f371,f455,f369,f368,f384,f383,f382,f381,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f396,f395,f394,f393,f400,f399,f398,f397,f454,f404,f403,f402,f401,f411,f410,f409,f408,f417,f414,f413,f412,f424,f423,f422,f453,f420,f419,f418,f428,f427,f426,f425,f432,f431,f430,f429,f441,f440,f439,f452,f437,f436,f435,f445,f444,f443,f442,f449,f448,f447,f446,f471,f474,f477,f480,f482,f484,f487,f489,f491,f494,f497,f500,f503,f505,f507,f510,f513,f515,f517,f520,f522,f525,f565,f575,f578,f586,f602,f613,f622,f639,f644,f645,f647,f654,f658,f659,f662,f674,f688,f690,f698,f710,f783,f846,f878,f876,f904,f901,f947,f972,f987,f995,f1008,f1012,f1026,f1041,f1111,f1126,f1142,f1201,f1218,f1233,f1296,f1313,f1361,f1767,f1766,f1765,f1764,f1768,f1769,f1779,f1780,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1817,f1831,f1840,f1866,f1867,f1898,f1923,f1947,f1959,f1975,f1976,f2039,f2054,f2064,f2134,f2135,f2136,f2137,f2138,f2139,f2140,f2204,f2553,f2557,f2561,f2564,f2567,f2570,f2573,f2576,f2579,f2582,f2585,f2588,f2591,f2594,f2597,f2600,f2603,f2605,f2606,f2607,f2608,f2609,f2610,f2611,f2612,f2613,f2614,f2615,f2616,f2617,f2618,f2619,f2620,f2621,f2622,f2623,f2624,f2625,f2626,f2627,f2628,f2629,f2630,f2631,f2632,f2633,f2634,f2635,f2636,f2637,f2290,f2291,f2292,f2293,f2294,f2295,f2296,f2643,f2645,f2648,f2650,f2652,f2654,f2656,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2555,f2768,f2786,f2788,f2790,f2792,f2794,f2796,f2798,f2800,f2802,f2804,f2806,f2808,f2810,f2812,f2814,f2816,f2818,f2820])).
fof(f2819,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f262,f261,f260,f263,f468,f270,f277,f276,f281,f280,f303,f305,f304,f469,f312,f310,f316,f314,f317,f321,f318,f325,f324,f328,f327,f326,f335,f334,f470,f353,f350,f349,f348,f366,f367,f373,f372,f371,f455,f369,f368,f384,f383,f382,f381,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f396,f395,f394,f393,f400,f399,f398,f397,f454,f404,f403,f402,f401,f411,f410,f409,f408,f417,f414,f413,f412,f424,f423,f422,f453,f420,f419,f418,f428,f427,f426,f425,f432,f431,f430,f429,f441,f440,f439,f452,f437,f436,f435,f445,f444,f443,f442,f449,f448,f447,f446,f471,f474,f477,f480,f482,f484,f487,f489,f491,f494,f497,f500,f503,f505,f507,f510,f513,f515,f517,f520,f522,f525,f565,f575,f578,f586,f602,f613,f622,f639,f644,f645,f647,f654,f658,f659,f662,f674,f688,f690,f698,f710,f783,f846,f878,f876,f904,f901,f947,f972,f987,f995,f1008,f1012,f1026,f1041,f1111,f1126,f1142,f1201,f1218,f1233,f1296,f1313,f1361,f1767,f1766,f1765,f1764,f1768,f1769,f1779,f1780,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1817,f1831,f1840,f1866,f1867,f1898,f1923,f1947,f1959,f1975,f1976,f2039,f2054,f2064,f2134,f2135,f2136,f2137,f2138,f2139,f2140,f2204,f2553,f2557,f2561,f2564,f2567,f2570,f2573,f2576,f2579,f2582,f2585,f2588,f2591,f2594,f2597,f2600,f2603,f2605,f2606,f2607,f2608,f2609,f2610,f2611,f2612,f2613,f2614,f2615,f2616,f2617,f2618,f2619,f2620,f2621,f2622,f2623,f2624,f2625,f2626,f2627,f2628,f2629,f2630,f2631,f2632,f2633,f2634,f2635,f2636,f2637,f2290,f2291,f2292,f2293,f2294,f2295,f2296,f2643,f2645,f2648,f2650,f2652,f2654,f2656,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2555,f2768,f2786,f2788,f2790,f2792,f2794,f2796,f2798,f2800,f2802,f2804,f2806,f2808,f2810,f2812,f2814,f2816,f2818])).
fof(f2817,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f262,f261,f260,f263,f468,f270,f277,f276,f281,f280,f303,f305,f304,f469,f312,f310,f316,f314,f317,f321,f318,f325,f324,f328,f327,f326,f335,f334,f470,f353,f350,f349,f348,f366,f367,f373,f372,f371,f455,f369,f368,f384,f383,f382,f381,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f396,f395,f394,f393,f400,f399,f398,f397,f454,f404,f403,f402,f401,f411,f410,f409,f408,f417,f414,f413,f412,f424,f423,f422,f453,f420,f419,f418,f428,f427,f426,f425,f432,f431,f430,f429,f441,f440,f439,f452,f437,f436,f435,f445,f444,f443,f442,f449,f448,f447,f446,f471,f474,f477,f480,f482,f484,f487,f489,f491,f494,f497,f500,f503,f505,f507,f510,f513,f515,f517,f520,f522,f525,f565,f575,f578,f586,f602,f613,f622,f639,f644,f645,f647,f654,f658,f659,f662,f674,f688,f690,f698,f710,f783,f846,f878,f876,f904,f901,f947,f972,f987,f995,f1008,f1012,f1026,f1041,f1111,f1126,f1142,f1201,f1218,f1233,f1296,f1313,f1361,f1767,f1766,f1765,f1764,f1768,f1769,f1779,f1780,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1817,f1831,f1840,f1866,f1867,f1898,f1923,f1947,f1959,f1975,f1976,f2039,f2054,f2064,f2134,f2135,f2136,f2137,f2138,f2139,f2140,f2204,f2553,f2557,f2561,f2564,f2567,f2570,f2573,f2576,f2579,f2582,f2585,f2588,f2591,f2594,f2597,f2600,f2603,f2605,f2606,f2607,f2608,f2609,f2610,f2611,f2612,f2613,f2614,f2615,f2616,f2617,f2618,f2619,f2620,f2621,f2622,f2623,f2624,f2625,f2626,f2627,f2628,f2629,f2630,f2631,f2632,f2633,f2634,f2635,f2636,f2637,f2290,f2291,f2292,f2293,f2294,f2295,f2296,f2643,f2645,f2648,f2650,f2652,f2654,f2656,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2555,f2768,f2786,f2788,f2790,f2792,f2794,f2796,f2798,f2800,f2802,f2804,f2806,f2808,f2810,f2812,f2814,f2816])).
fof(f2815,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f262,f261,f260,f263,f468,f270,f277,f276,f281,f280,f303,f305,f304,f469,f312,f310,f316,f314,f317,f321,f318,f325,f324,f328,f327,f326,f335,f334,f470,f353,f350,f349,f348,f366,f367,f373,f372,f371,f455,f369,f368,f384,f383,f382,f381,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f396,f395,f394,f393,f400,f399,f398,f397,f454,f404,f403,f402,f401,f411,f410,f409,f408,f417,f414,f413,f412,f424,f423,f422,f453,f420,f419,f418,f428,f427,f426,f425,f432,f431,f430,f429,f441,f440,f439,f452,f437,f436,f435,f445,f444,f443,f442,f449,f448,f447,f446,f471,f474,f477,f480,f482,f484,f487,f489,f491,f494,f497,f500,f503,f505,f507,f510,f513,f515,f517,f520,f522,f525,f565,f575,f578,f586,f602,f613,f622,f639,f644,f645,f647,f654,f658,f659,f662,f674,f688,f690,f698,f710,f783,f846,f878,f876,f904,f901,f947,f972,f987,f995,f1008,f1012,f1026,f1041,f1111,f1126,f1142,f1201,f1218,f1233,f1296,f1313,f1361,f1767,f1766,f1765,f1764,f1768,f1769,f1779,f1780,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1817,f1831,f1840,f1866,f1867,f1898,f1923,f1947,f1959,f1975,f1976,f2039,f2054,f2064,f2134,f2135,f2136,f2137,f2138,f2139,f2140,f2204,f2553,f2557,f2561,f2564,f2567,f2570,f2573,f2576,f2579,f2582,f2585,f2588,f2591,f2594,f2597,f2600,f2603,f2605,f2606,f2607,f2608,f2609,f2610,f2611,f2612,f2613,f2614,f2615,f2616,f2617,f2618,f2619,f2620,f2621,f2622,f2623,f2624,f2625,f2626,f2627,f2628,f2629,f2630,f2631,f2632,f2633,f2634,f2635,f2636,f2637,f2290,f2291,f2292,f2293,f2294,f2295,f2296,f2643,f2645,f2648,f2650,f2652,f2654,f2656,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2555,f2768,f2786,f2788,f2790,f2792,f2794,f2796,f2798,f2800,f2802,f2804,f2806,f2808,f2810,f2812,f2814])).
fof(f2813,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f262,f261,f260,f263,f468,f270,f277,f276,f281,f280,f303,f305,f304,f469,f312,f310,f316,f314,f317,f321,f318,f325,f324,f328,f327,f326,f335,f334,f470,f353,f350,f349,f348,f366,f367,f373,f372,f371,f455,f369,f368,f384,f383,f382,f381,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f396,f395,f394,f393,f400,f399,f398,f397,f454,f404,f403,f402,f401,f411,f410,f409,f408,f417,f414,f413,f412,f424,f423,f422,f453,f420,f419,f418,f428,f427,f426,f425,f432,f431,f430,f429,f441,f440,f439,f452,f437,f436,f435,f445,f444,f443,f442,f449,f448,f447,f446,f471,f474,f477,f480,f482,f484,f487,f489,f491,f494,f497,f500,f503,f505,f507,f510,f513,f515,f517,f520,f522,f525,f565,f575,f578,f586,f602,f613,f622,f639,f644,f645,f647,f654,f658,f659,f662,f674,f688,f690,f698,f710,f783,f846,f878,f876,f904,f901,f947,f972,f987,f995,f1008,f1012,f1026,f1041,f1111,f1126,f1142,f1201,f1218,f1233,f1296,f1313,f1361,f1767,f1766,f1765,f1764,f1768,f1769,f1779,f1780,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1817,f1831,f1840,f1866,f1867,f1898,f1923,f1947,f1959,f1975,f1976,f2039,f2054,f2064,f2134,f2135,f2136,f2137,f2138,f2139,f2140,f2204,f2553,f2557,f2561,f2564,f2567,f2570,f2573,f2576,f2579,f2582,f2585,f2588,f2591,f2594,f2597,f2600,f2603,f2605,f2606,f2607,f2608,f2609,f2610,f2611,f2612,f2613,f2614,f2615,f2616,f2617,f2618,f2619,f2620,f2621,f2622,f2623,f2624,f2625,f2626,f2627,f2628,f2629,f2630,f2631,f2632,f2633,f2634,f2635,f2636,f2637,f2290,f2291,f2292,f2293,f2294,f2295,f2296,f2643,f2645,f2648,f2650,f2652,f2654,f2656,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2555,f2768,f2786,f2788,f2790,f2792,f2794,f2796,f2798,f2800,f2802,f2804,f2806,f2808,f2810,f2812])).
fof(f2811,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f262,f261,f260,f263,f468,f270,f277,f276,f281,f280,f303,f305,f304,f469,f312,f310,f316,f314,f317,f321,f318,f325,f324,f328,f327,f326,f335,f334,f470,f353,f350,f349,f348,f366,f367,f373,f372,f371,f455,f369,f368,f384,f383,f382,f381,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f396,f395,f394,f393,f400,f399,f398,f397,f454,f404,f403,f402,f401,f411,f410,f409,f408,f417,f414,f413,f412,f424,f423,f422,f453,f420,f419,f418,f428,f427,f426,f425,f432,f431,f430,f429,f441,f440,f439,f452,f437,f436,f435,f445,f444,f443,f442,f449,f448,f447,f446,f471,f474,f477,f480,f482,f484,f487,f489,f491,f494,f497,f500,f503,f505,f507,f510,f513,f515,f517,f520,f522,f525,f565,f575,f578,f586,f602,f613,f622,f639,f644,f645,f647,f654,f658,f659,f662,f674,f688,f690,f698,f710,f783,f846,f878,f876,f904,f901,f947,f972,f987,f995,f1008,f1012,f1026,f1041,f1111,f1126,f1142,f1201,f1218,f1233,f1296,f1313,f1361,f1767,f1766,f1765,f1764,f1768,f1769,f1779,f1780,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1817,f1831,f1840,f1866,f1867,f1898,f1923,f1947,f1959,f1975,f1976,f2039,f2054,f2064,f2134,f2135,f2136,f2137,f2138,f2139,f2140,f2204,f2553,f2557,f2561,f2564,f2567,f2570,f2573,f2576,f2579,f2582,f2585,f2588,f2591,f2594,f2597,f2600,f2603,f2605,f2606,f2607,f2608,f2609,f2610,f2611,f2612,f2613,f2614,f2615,f2616,f2617,f2618,f2619,f2620,f2621,f2622,f2623,f2624,f2625,f2626,f2627,f2628,f2629,f2630,f2631,f2632,f2633,f2634,f2635,f2636,f2637,f2290,f2291,f2292,f2293,f2294,f2295,f2296,f2643,f2645,f2648,f2650,f2652,f2654,f2656,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2555,f2768,f2786,f2788,f2790,f2792,f2794,f2796,f2798,f2800,f2802,f2804,f2806,f2808,f2810])).
fof(f2809,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f262,f261,f260,f263,f468,f270,f277,f276,f281,f280,f303,f305,f304,f469,f312,f310,f316,f314,f317,f321,f318,f325,f324,f328,f327,f326,f335,f334,f470,f353,f350,f349,f348,f366,f367,f373,f372,f371,f455,f369,f368,f384,f383,f382,f381,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f396,f395,f394,f393,f400,f399,f398,f397,f454,f404,f403,f402,f401,f411,f410,f409,f408,f417,f414,f413,f412,f424,f423,f422,f453,f420,f419,f418,f428,f427,f426,f425,f432,f431,f430,f429,f441,f440,f439,f452,f437,f436,f435,f445,f444,f443,f442,f449,f448,f447,f446,f471,f474,f477,f480,f482,f484,f487,f489,f491,f494,f497,f500,f503,f505,f507,f510,f513,f515,f517,f520,f522,f525,f565,f575,f578,f586,f602,f613,f622,f639,f644,f645,f647,f654,f658,f659,f662,f674,f688,f690,f698,f710,f783,f846,f878,f876,f904,f901,f947,f972,f987,f995,f1008,f1012,f1026,f1041,f1111,f1126,f1142,f1201,f1218,f1233,f1296,f1313,f1361,f1767,f1766,f1765,f1764,f1768,f1769,f1779,f1780,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1817,f1831,f1840,f1866,f1867,f1898,f1923,f1947,f1959,f1975,f1976,f2039,f2054,f2064,f2134,f2135,f2136,f2137,f2138,f2139,f2140,f2204,f2553,f2557,f2561,f2564,f2567,f2570,f2573,f2576,f2579,f2582,f2585,f2588,f2591,f2594,f2597,f2600,f2603,f2605,f2606,f2607,f2608,f2609,f2610,f2611,f2612,f2613,f2614,f2615,f2616,f2617,f2618,f2619,f2620,f2621,f2622,f2623,f2624,f2625,f2626,f2627,f2628,f2629,f2630,f2631,f2632,f2633,f2634,f2635,f2636,f2637,f2290,f2291,f2292,f2293,f2294,f2295,f2296,f2643,f2645,f2648,f2650,f2652,f2654,f2656,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2555,f2768,f2786,f2788,f2790,f2792,f2794,f2796,f2798,f2800,f2802,f2804,f2806,f2808])).
fof(f2807,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f262,f261,f260,f263,f468,f270,f277,f276,f281,f280,f303,f305,f304,f469,f312,f310,f316,f314,f317,f321,f318,f325,f324,f328,f327,f326,f335,f334,f470,f353,f350,f349,f348,f366,f367,f373,f372,f371,f455,f369,f368,f384,f383,f382,f381,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f396,f395,f394,f393,f400,f399,f398,f397,f454,f404,f403,f402,f401,f411,f410,f409,f408,f417,f414,f413,f412,f424,f423,f422,f453,f420,f419,f418,f428,f427,f426,f425,f432,f431,f430,f429,f441,f440,f439,f452,f437,f436,f435,f445,f444,f443,f442,f449,f448,f447,f446,f471,f474,f477,f480,f482,f484,f487,f489,f491,f494,f497,f500,f503,f505,f507,f510,f513,f515,f517,f520,f522,f525,f565,f575,f578,f586,f602,f613,f622,f639,f644,f645,f647,f654,f658,f659,f662,f674,f688,f690,f698,f710,f783,f846,f878,f876,f904,f901,f947,f972,f987,f995,f1008,f1012,f1026,f1041,f1111,f1126,f1142,f1201,f1218,f1233,f1296,f1313,f1361,f1767,f1766,f1765,f1764,f1768,f1769,f1779,f1780,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1817,f1831,f1840,f1866,f1867,f1898,f1923,f1947,f1959,f1975,f1976,f2039,f2054,f2064,f2134,f2135,f2136,f2137,f2138,f2139,f2140,f2204,f2553,f2557,f2561,f2564,f2567,f2570,f2573,f2576,f2579,f2582,f2585,f2588,f2591,f2594,f2597,f2600,f2603,f2605,f2606,f2607,f2608,f2609,f2610,f2611,f2612,f2613,f2614,f2615,f2616,f2617,f2618,f2619,f2620,f2621,f2622,f2623,f2624,f2625,f2626,f2627,f2628,f2629,f2630,f2631,f2632,f2633,f2634,f2635,f2636,f2637,f2290,f2291,f2292,f2293,f2294,f2295,f2296,f2643,f2645,f2648,f2650,f2652,f2654,f2656,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2555,f2768,f2786,f2788,f2790,f2792,f2794,f2796,f2798,f2800,f2802,f2804,f2806])).
fof(f2805,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f262,f261,f260,f263,f468,f270,f277,f276,f281,f280,f303,f305,f304,f469,f312,f310,f316,f314,f317,f321,f318,f325,f324,f328,f327,f326,f335,f334,f470,f353,f350,f349,f348,f366,f367,f373,f372,f371,f455,f369,f368,f384,f383,f382,f381,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f396,f395,f394,f393,f400,f399,f398,f397,f454,f404,f403,f402,f401,f411,f410,f409,f408,f417,f414,f413,f412,f424,f423,f422,f453,f420,f419,f418,f428,f427,f426,f425,f432,f431,f430,f429,f441,f440,f439,f452,f437,f436,f435,f445,f444,f443,f442,f449,f448,f447,f446,f471,f474,f477,f480,f482,f484,f487,f489,f491,f494,f497,f500,f503,f505,f507,f510,f513,f515,f517,f520,f522,f525,f565,f575,f578,f586,f602,f613,f622,f639,f644,f645,f647,f654,f658,f659,f662,f674,f688,f690,f698,f710,f783,f846,f878,f876,f904,f901,f947,f972,f987,f995,f1008,f1012,f1026,f1041,f1111,f1126,f1142,f1201,f1218,f1233,f1296,f1313,f1361,f1767,f1766,f1765,f1764,f1768,f1769,f1779,f1780,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1817,f1831,f1840,f1866,f1867,f1898,f1923,f1947,f1959,f1975,f1976,f2039,f2054,f2064,f2134,f2135,f2136,f2137,f2138,f2139,f2140,f2204,f2553,f2557,f2561,f2564,f2567,f2570,f2573,f2576,f2579,f2582,f2585,f2588,f2591,f2594,f2597,f2600,f2603,f2605,f2606,f2607,f2608,f2609,f2610,f2611,f2612,f2613,f2614,f2615,f2616,f2617,f2618,f2619,f2620,f2621,f2622,f2623,f2624,f2625,f2626,f2627,f2628,f2629,f2630,f2631,f2632,f2633,f2634,f2635,f2636,f2637,f2290,f2291,f2292,f2293,f2294,f2295,f2296,f2643,f2645,f2648,f2650,f2652,f2654,f2656,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2555,f2768,f2786,f2788,f2790,f2792,f2794,f2796,f2798,f2800,f2802,f2804])).
fof(f2803,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f262,f261,f260,f263,f468,f270,f277,f276,f281,f280,f303,f305,f304,f469,f312,f310,f316,f314,f317,f321,f318,f325,f324,f328,f327,f326,f335,f334,f470,f353,f350,f349,f348,f366,f367,f373,f372,f371,f455,f369,f368,f384,f383,f382,f381,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f396,f395,f394,f393,f400,f399,f398,f397,f454,f404,f403,f402,f401,f411,f410,f409,f408,f417,f414,f413,f412,f424,f423,f422,f453,f420,f419,f418,f428,f427,f426,f425,f432,f431,f430,f429,f441,f440,f439,f452,f437,f436,f435,f445,f444,f443,f442,f449,f448,f447,f446,f471,f474,f477,f480,f482,f484,f487,f489,f491,f494,f497,f500,f503,f505,f507,f510,f513,f515,f517,f520,f522,f525,f565,f575,f578,f586,f602,f613,f622,f639,f644,f645,f647,f654,f658,f659,f662,f674,f688,f690,f698,f710,f783,f846,f878,f876,f904,f901,f947,f972,f987,f995,f1008,f1012,f1026,f1041,f1111,f1126,f1142,f1201,f1218,f1233,f1296,f1313,f1361,f1767,f1766,f1765,f1764,f1768,f1769,f1779,f1780,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1817,f1831,f1840,f1866,f1867,f1898,f1923,f1947,f1959,f1975,f1976,f2039,f2054,f2064,f2134,f2135,f2136,f2137,f2138,f2139,f2140,f2204,f2553,f2557,f2561,f2564,f2567,f2570,f2573,f2576,f2579,f2582,f2585,f2588,f2591,f2594,f2597,f2600,f2603,f2605,f2606,f2607,f2608,f2609,f2610,f2611,f2612,f2613,f2614,f2615,f2616,f2617,f2618,f2619,f2620,f2621,f2622,f2623,f2624,f2625,f2626,f2627,f2628,f2629,f2630,f2631,f2632,f2633,f2634,f2635,f2636,f2637,f2290,f2291,f2292,f2293,f2294,f2295,f2296,f2643,f2645,f2648,f2650,f2652,f2654,f2656,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2555,f2768,f2786,f2788,f2790,f2792,f2794,f2796,f2798,f2800,f2802])).
fof(f2801,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f262,f261,f260,f263,f468,f270,f277,f276,f281,f280,f303,f305,f304,f469,f312,f310,f316,f314,f317,f321,f318,f325,f324,f328,f327,f326,f335,f334,f470,f353,f350,f349,f348,f366,f367,f373,f372,f371,f455,f369,f368,f384,f383,f382,f381,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f396,f395,f394,f393,f400,f399,f398,f397,f454,f404,f403,f402,f401,f411,f410,f409,f408,f417,f414,f413,f412,f424,f423,f422,f453,f420,f419,f418,f428,f427,f426,f425,f432,f431,f430,f429,f441,f440,f439,f452,f437,f436,f435,f445,f444,f443,f442,f449,f448,f447,f446,f471,f474,f477,f480,f482,f484,f487,f489,f491,f494,f497,f500,f503,f505,f507,f510,f513,f515,f517,f520,f522,f525,f565,f575,f578,f586,f602,f613,f622,f639,f644,f645,f647,f654,f658,f659,f662,f674,f688,f690,f698,f710,f783,f846,f878,f876,f904,f901,f947,f972,f987,f995,f1008,f1012,f1026,f1041,f1111,f1126,f1142,f1201,f1218,f1233,f1296,f1313,f1361,f1767,f1766,f1765,f1764,f1768,f1769,f1779,f1780,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1817,f1831,f1840,f1866,f1867,f1898,f1923,f1947,f1959,f1975,f1976,f2039,f2054,f2064,f2134,f2135,f2136,f2137,f2138,f2139,f2140,f2204,f2553,f2557,f2561,f2564,f2567,f2570,f2573,f2576,f2579,f2582,f2585,f2588,f2591,f2594,f2597,f2600,f2603,f2605,f2606,f2607,f2608,f2609,f2610,f2611,f2612,f2613,f2614,f2615,f2616,f2617,f2618,f2619,f2620,f2621,f2622,f2623,f2624,f2625,f2626,f2627,f2628,f2629,f2630,f2631,f2632,f2633,f2634,f2635,f2636,f2637,f2290,f2291,f2292,f2293,f2294,f2295,f2296,f2643,f2645,f2648,f2650,f2652,f2654,f2656,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2555,f2768,f2786,f2788,f2790,f2792,f2794,f2796,f2798,f2800])).
fof(f2799,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f262,f261,f260,f263,f468,f270,f277,f276,f281,f280,f303,f305,f304,f469,f312,f310,f316,f314,f317,f321,f318,f325,f324,f328,f327,f326,f335,f334,f470,f353,f350,f349,f348,f366,f367,f373,f372,f371,f455,f369,f368,f384,f383,f382,f381,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f396,f395,f394,f393,f400,f399,f398,f397,f454,f404,f403,f402,f401,f411,f410,f409,f408,f417,f414,f413,f412,f424,f423,f422,f453,f420,f419,f418,f428,f427,f426,f425,f432,f431,f430,f429,f441,f440,f439,f452,f437,f436,f435,f445,f444,f443,f442,f449,f448,f447,f446,f471,f474,f477,f480,f482,f484,f487,f489,f491,f494,f497,f500,f503,f505,f507,f510,f513,f515,f517,f520,f522,f525,f565,f575,f578,f586,f602,f613,f622,f639,f644,f645,f647,f654,f658,f659,f662,f674,f688,f690,f698,f710,f783,f846,f878,f876,f904,f901,f947,f972,f987,f995,f1008,f1012,f1026,f1041,f1111,f1126,f1142,f1201,f1218,f1233,f1296,f1313,f1361,f1767,f1766,f1765,f1764,f1768,f1769,f1779,f1780,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1817,f1831,f1840,f1866,f1867,f1898,f1923,f1947,f1959,f1975,f1976,f2039,f2054,f2064,f2134,f2135,f2136,f2137,f2138,f2139,f2140,f2204,f2553,f2557,f2561,f2564,f2567,f2570,f2573,f2576,f2579,f2582,f2585,f2588,f2591,f2594,f2597,f2600,f2603,f2605,f2606,f2607,f2608,f2609,f2610,f2611,f2612,f2613,f2614,f2615,f2616,f2617,f2618,f2619,f2620,f2621,f2622,f2623,f2624,f2625,f2626,f2627,f2628,f2629,f2630,f2631,f2632,f2633,f2634,f2635,f2636,f2637,f2290,f2291,f2292,f2293,f2294,f2295,f2296,f2643,f2645,f2648,f2650,f2652,f2654,f2656,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2555,f2768,f2786,f2788,f2790,f2792,f2794,f2796,f2798])).
fof(f2797,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f262,f261,f260,f263,f468,f270,f277,f276,f281,f280,f303,f305,f304,f469,f312,f310,f316,f314,f317,f321,f318,f325,f324,f328,f327,f326,f335,f334,f470,f353,f350,f349,f348,f366,f367,f373,f372,f371,f455,f369,f368,f384,f383,f382,f381,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f396,f395,f394,f393,f400,f399,f398,f397,f454,f404,f403,f402,f401,f411,f410,f409,f408,f417,f414,f413,f412,f424,f423,f422,f453,f420,f419,f418,f428,f427,f426,f425,f432,f431,f430,f429,f441,f440,f439,f452,f437,f436,f435,f445,f444,f443,f442,f449,f448,f447,f446,f471,f474,f477,f480,f482,f484,f487,f489,f491,f494,f497,f500,f503,f505,f507,f510,f513,f515,f517,f520,f522,f525,f565,f575,f578,f586,f602,f613,f622,f639,f644,f645,f647,f654,f658,f659,f662,f674,f688,f690,f698,f710,f783,f846,f878,f876,f904,f901,f947,f972,f987,f995,f1008,f1012,f1026,f1041,f1111,f1126,f1142,f1201,f1218,f1233,f1296,f1313,f1361,f1767,f1766,f1765,f1764,f1768,f1769,f1779,f1780,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1817,f1831,f1840,f1866,f1867,f1898,f1923,f1947,f1959,f1975,f1976,f2039,f2054,f2064,f2134,f2135,f2136,f2137,f2138,f2139,f2140,f2204,f2553,f2557,f2561,f2564,f2567,f2570,f2573,f2576,f2579,f2582,f2585,f2588,f2591,f2594,f2597,f2600,f2603,f2605,f2606,f2607,f2608,f2609,f2610,f2611,f2612,f2613,f2614,f2615,f2616,f2617,f2618,f2619,f2620,f2621,f2622,f2623,f2624,f2625,f2626,f2627,f2628,f2629,f2630,f2631,f2632,f2633,f2634,f2635,f2636,f2637,f2290,f2291,f2292,f2293,f2294,f2295,f2296,f2643,f2645,f2648,f2650,f2652,f2654,f2656,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2555,f2768,f2786,f2788,f2790,f2792,f2794,f2796])).
fof(f2795,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f262,f261,f260,f263,f468,f270,f277,f276,f281,f280,f303,f305,f304,f469,f312,f310,f316,f314,f317,f321,f318,f325,f324,f328,f327,f326,f335,f334,f470,f353,f350,f349,f348,f366,f367,f373,f372,f371,f455,f369,f368,f384,f383,f382,f381,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f396,f395,f394,f393,f400,f399,f398,f397,f454,f404,f403,f402,f401,f411,f410,f409,f408,f417,f414,f413,f412,f424,f423,f422,f453,f420,f419,f418,f428,f427,f426,f425,f432,f431,f430,f429,f441,f440,f439,f452,f437,f436,f435,f445,f444,f443,f442,f449,f448,f447,f446,f471,f474,f477,f480,f482,f484,f487,f489,f491,f494,f497,f500,f503,f505,f507,f510,f513,f515,f517,f520,f522,f525,f565,f575,f578,f586,f602,f613,f622,f639,f644,f645,f647,f654,f658,f659,f662,f674,f688,f690,f698,f710,f783,f846,f878,f876,f904,f901,f947,f972,f987,f995,f1008,f1012,f1026,f1041,f1111,f1126,f1142,f1201,f1218,f1233,f1296,f1313,f1361,f1767,f1766,f1765,f1764,f1768,f1769,f1779,f1780,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1817,f1831,f1840,f1866,f1867,f1898,f1923,f1947,f1959,f1975,f1976,f2039,f2054,f2064,f2134,f2135,f2136,f2137,f2138,f2139,f2140,f2204,f2553,f2557,f2561,f2564,f2567,f2570,f2573,f2576,f2579,f2582,f2585,f2588,f2591,f2594,f2597,f2600,f2603,f2605,f2606,f2607,f2608,f2609,f2610,f2611,f2612,f2613,f2614,f2615,f2616,f2617,f2618,f2619,f2620,f2621,f2622,f2623,f2624,f2625,f2626,f2627,f2628,f2629,f2630,f2631,f2632,f2633,f2634,f2635,f2636,f2637,f2290,f2291,f2292,f2293,f2294,f2295,f2296,f2643,f2645,f2648,f2650,f2652,f2654,f2656,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2555,f2768,f2786,f2788,f2790,f2792,f2794])).
fof(f2793,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f262,f261,f260,f263,f468,f270,f277,f276,f281,f280,f303,f305,f304,f469,f312,f310,f316,f314,f317,f321,f318,f325,f324,f328,f327,f326,f335,f334,f470,f353,f350,f349,f348,f366,f367,f373,f372,f371,f455,f369,f368,f384,f383,f382,f381,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f396,f395,f394,f393,f400,f399,f398,f397,f454,f404,f403,f402,f401,f411,f410,f409,f408,f417,f414,f413,f412,f424,f423,f422,f453,f420,f419,f418,f428,f427,f426,f425,f432,f431,f430,f429,f441,f440,f439,f452,f437,f436,f435,f445,f444,f443,f442,f449,f448,f447,f446,f471,f474,f477,f480,f482,f484,f487,f489,f491,f494,f497,f500,f503,f505,f507,f510,f513,f515,f517,f520,f522,f525,f565,f575,f578,f586,f602,f613,f622,f639,f644,f645,f647,f654,f658,f659,f662,f674,f688,f690,f698,f710,f783,f846,f878,f876,f904,f901,f947,f972,f987,f995,f1008,f1012,f1026,f1041,f1111,f1126,f1142,f1201,f1218,f1233,f1296,f1313,f1361,f1767,f1766,f1765,f1764,f1768,f1769,f1779,f1780,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1817,f1831,f1840,f1866,f1867,f1898,f1923,f1947,f1959,f1975,f1976,f2039,f2054,f2064,f2134,f2135,f2136,f2137,f2138,f2139,f2140,f2204,f2553,f2557,f2561,f2564,f2567,f2570,f2573,f2576,f2579,f2582,f2585,f2588,f2591,f2594,f2597,f2600,f2603,f2605,f2606,f2607,f2608,f2609,f2610,f2611,f2612,f2613,f2614,f2615,f2616,f2617,f2618,f2619,f2620,f2621,f2622,f2623,f2624,f2625,f2626,f2627,f2628,f2629,f2630,f2631,f2632,f2633,f2634,f2635,f2636,f2637,f2290,f2291,f2292,f2293,f2294,f2295,f2296,f2643,f2645,f2648,f2650,f2652,f2654,f2656,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2555,f2768,f2786,f2788,f2790,f2792])).
fof(f2791,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f262,f261,f260,f263,f468,f270,f277,f276,f281,f280,f303,f305,f304,f469,f312,f310,f316,f314,f317,f321,f318,f325,f324,f328,f327,f326,f335,f334,f470,f353,f350,f349,f348,f366,f367,f373,f372,f371,f455,f369,f368,f384,f383,f382,f381,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f396,f395,f394,f393,f400,f399,f398,f397,f454,f404,f403,f402,f401,f411,f410,f409,f408,f417,f414,f413,f412,f424,f423,f422,f453,f420,f419,f418,f428,f427,f426,f425,f432,f431,f430,f429,f441,f440,f439,f452,f437,f436,f435,f445,f444,f443,f442,f449,f448,f447,f446,f471,f474,f477,f480,f482,f484,f487,f489,f491,f494,f497,f500,f503,f505,f507,f510,f513,f515,f517,f520,f522,f525,f565,f575,f578,f586,f602,f613,f622,f639,f644,f645,f647,f654,f658,f659,f662,f674,f688,f690,f698,f710,f783,f846,f878,f876,f904,f901,f947,f972,f987,f995,f1008,f1012,f1026,f1041,f1111,f1126,f1142,f1201,f1218,f1233,f1296,f1313,f1361,f1767,f1766,f1765,f1764,f1768,f1769,f1779,f1780,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1817,f1831,f1840,f1866,f1867,f1898,f1923,f1947,f1959,f1975,f1976,f2039,f2054,f2064,f2134,f2135,f2136,f2137,f2138,f2139,f2140,f2204,f2553,f2557,f2561,f2564,f2567,f2570,f2573,f2576,f2579,f2582,f2585,f2588,f2591,f2594,f2597,f2600,f2603,f2605,f2606,f2607,f2608,f2609,f2610,f2611,f2612,f2613,f2614,f2615,f2616,f2617,f2618,f2619,f2620,f2621,f2622,f2623,f2624,f2625,f2626,f2627,f2628,f2629,f2630,f2631,f2632,f2633,f2634,f2635,f2636,f2637,f2290,f2291,f2292,f2293,f2294,f2295,f2296,f2643,f2645,f2648,f2650,f2652,f2654,f2656,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2555,f2768,f2786,f2788,f2790])).
fof(f2789,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f262,f261,f260,f263,f468,f270,f277,f276,f281,f280,f303,f305,f304,f469,f312,f310,f316,f314,f317,f321,f318,f325,f324,f328,f327,f326,f335,f334,f470,f353,f350,f349,f348,f366,f367,f373,f372,f371,f455,f369,f368,f384,f383,f382,f381,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f396,f395,f394,f393,f400,f399,f398,f397,f454,f404,f403,f402,f401,f411,f410,f409,f408,f417,f414,f413,f412,f424,f423,f422,f453,f420,f419,f418,f428,f427,f426,f425,f432,f431,f430,f429,f441,f440,f439,f452,f437,f436,f435,f445,f444,f443,f442,f449,f448,f447,f446,f471,f474,f477,f480,f482,f484,f487,f489,f491,f494,f497,f500,f503,f505,f507,f510,f513,f515,f517,f520,f522,f525,f565,f575,f578,f586,f602,f613,f622,f639,f644,f645,f647,f654,f658,f659,f662,f674,f688,f690,f698,f710,f783,f846,f878,f876,f904,f901,f947,f972,f987,f995,f1008,f1012,f1026,f1041,f1111,f1126,f1142,f1201,f1218,f1233,f1296,f1313,f1361,f1767,f1766,f1765,f1764,f1768,f1769,f1779,f1780,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1817,f1831,f1840,f1866,f1867,f1898,f1923,f1947,f1959,f1975,f1976,f2039,f2054,f2064,f2134,f2135,f2136,f2137,f2138,f2139,f2140,f2204,f2553,f2557,f2561,f2564,f2567,f2570,f2573,f2576,f2579,f2582,f2585,f2588,f2591,f2594,f2597,f2600,f2603,f2605,f2606,f2607,f2608,f2609,f2610,f2611,f2612,f2613,f2614,f2615,f2616,f2617,f2618,f2619,f2620,f2621,f2622,f2623,f2624,f2625,f2626,f2627,f2628,f2629,f2630,f2631,f2632,f2633,f2634,f2635,f2636,f2637,f2290,f2291,f2292,f2293,f2294,f2295,f2296,f2643,f2645,f2648,f2650,f2652,f2654,f2656,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2555,f2768,f2786,f2788])).
fof(f2787,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f262,f261,f260,f263,f468,f270,f277,f276,f281,f280,f303,f305,f304,f469,f312,f310,f316,f314,f317,f321,f318,f325,f324,f328,f327,f326,f335,f334,f470,f353,f350,f349,f348,f366,f367,f373,f372,f371,f455,f369,f368,f384,f383,f382,f381,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f396,f395,f394,f393,f400,f399,f398,f397,f454,f404,f403,f402,f401,f411,f410,f409,f408,f417,f414,f413,f412,f424,f423,f422,f453,f420,f419,f418,f428,f427,f426,f425,f432,f431,f430,f429,f441,f440,f439,f452,f437,f436,f435,f445,f444,f443,f442,f449,f448,f447,f446,f471,f474,f477,f480,f482,f484,f487,f489,f491,f494,f497,f500,f503,f505,f507,f510,f513,f515,f517,f520,f522,f525,f565,f575,f578,f586,f602,f613,f622,f639,f644,f645,f647,f654,f658,f659,f662,f674,f688,f690,f698,f710,f783,f846,f878,f876,f904,f901,f947,f972,f987,f995,f1008,f1012,f1026,f1041,f1111,f1126,f1142,f1201,f1218,f1233,f1296,f1313,f1361,f1767,f1766,f1765,f1764,f1768,f1769,f1779,f1780,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1817,f1831,f1840,f1866,f1867,f1898,f1923,f1947,f1959,f1975,f1976,f2039,f2054,f2064,f2134,f2135,f2136,f2137,f2138,f2139,f2140,f2204,f2553,f2557,f2561,f2564,f2567,f2570,f2573,f2576,f2579,f2582,f2585,f2588,f2591,f2594,f2597,f2600,f2603,f2605,f2606,f2607,f2608,f2609,f2610,f2611,f2612,f2613,f2614,f2615,f2616,f2617,f2618,f2619,f2620,f2621,f2622,f2623,f2624,f2625,f2626,f2627,f2628,f2629,f2630,f2631,f2632,f2633,f2634,f2635,f2636,f2637,f2290,f2291,f2292,f2293,f2294,f2295,f2296,f2643,f2645,f2648,f2650,f2652,f2654,f2656,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2555,f2768,f2786])).
fof(f2785,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f2784,f263])).
fof(f2784,plain,(
  ~p(s(bool,'T_0')) | $spl10),
  inference(forward_demodulation,[],[f2783,f2555])).
fof(f2783,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))))) ) | $spl10),
  inference(subsumption_resolution,[],[f2726,f263])).
fof(f2726,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))))) ) | $spl10),
  inference(backward_demodulation,[],[f2555,f324])).
fof(f2770,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f2679,f263])).
fof(f2679,plain,(
  ~p(s(bool,'T_0')) | $spl10),
  inference(backward_demodulation,[],[f2555,f468])).
fof(f2769,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f262,f261,f260,f263,f468,f270,f277,f276,f281,f280,f303,f305,f304,f469,f312,f310,f316,f314,f317,f321,f318,f325,f324,f328,f327,f326,f335,f334,f470,f353,f350,f349,f348,f366,f367,f373,f372,f371,f455,f369,f368,f384,f383,f382,f381,f380,f379,f378,f388,f387,f386,f385,f392,f391,f390,f389,f396,f395,f394,f393,f400,f399,f398,f397,f454,f404,f403,f402,f401,f411,f410,f409,f408,f417,f414,f413,f412,f424,f423,f422,f453,f420,f419,f418,f428,f427,f426,f425,f432,f431,f430,f429,f441,f440,f439,f452,f437,f436,f435,f445,f444,f443,f442,f449,f448,f447,f446,f471,f474,f477,f480,f482,f484,f487,f489,f491,f494,f497,f500,f503,f505,f507,f510,f513,f515,f517,f520,f522,f525,f565,f575,f578,f586,f602,f613,f622,f639,f644,f645,f647,f654,f658,f659,f662,f674,f688,f690,f698,f710,f783,f846,f878,f876,f904,f901,f947,f972,f987,f995,f1008,f1012,f1026,f1041,f1111,f1126,f1142,f1201,f1218,f1233,f1296,f1313,f1361,f1767,f1766,f1765,f1764,f1768,f1769,f1779,f1780,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1817,f1831,f1840,f1866,f1867,f1898,f1923,f1947,f1959,f1975,f1976,f2039,f2054,f2064,f2134,f2135,f2136,f2137,f2138,f2139,f2140,f2204,f2553,f2557,f2561,f2564,f2567,f2570,f2573,f2576,f2579,f2582,f2585,f2588,f2591,f2594,f2597,f2600,f2603,f2605,f2606,f2607,f2608,f2609,f2610,f2611,f2612,f2613,f2614,f2615,f2616,f2617,f2618,f2619,f2620,f2621,f2622,f2623,f2624,f2625,f2626,f2627,f2628,f2629,f2630,f2631,f2632,f2633,f2634,f2635,f2636,f2637,f2290,f2291,f2292,f2293,f2294,f2295,f2296,f2643,f2645,f2648,f2650,f2652,f2654,f2656,f2658,f2660,f2662,f2664,f2666,f2668,f2670,f2672,f2674,f2676,f2555,f2768])).
fof(f2604,plain,(
  ( ! [X2,X1] : (sP10('T_0',X1,X2) | ~sP9('T_0',X1,X2)) ) | $spl42),
  inference(cnf_transformation,[],[f2604_D])).
fof(f2604_D,plain,(
  ( ! [X2,X1] : (sP10('T_0',X1,X2) | ~sP9('T_0',X1,X2)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f2676,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP9('T_0',X1,X2) | sP10('T_0',X1,X2)) )),
  inference(condensation,[],[f2675])).
fof(f2675,plain,(
  ( ! [X94,X92,X95,X93] : (~sP9('T_0',X93,X94) | sP10('T_0',X93,X94) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(subsumption_resolution,[],[f2324,f846])).
fof(f2324,plain,(
  ( ! [X94,X92,X95,X93] : (~p(s(bool,X92)) | ~sP9('T_0',X93,X94) | sP10('T_0',X93,X94) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(superposition,[],[f453,f783])).
fof(f2601,plain,(
  ( ! [X2,X1] : (sP12('T_0',X1,X2) | ~sP11('T_0',X1,X2)) ) | $spl40),
  inference(cnf_transformation,[],[f2601_D])).
fof(f2601_D,plain,(
  ( ! [X2,X1] : (sP12('T_0',X1,X2) | ~sP11('T_0',X1,X2)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f2674,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP11('T_0',X1,X2) | sP12('T_0',X1,X2)) )),
  inference(condensation,[],[f2673])).
fof(f2673,plain,(
  ( ! [X90,X88,X91,X89] : (~sP11('T_0',X89,X90) | sP12('T_0',X89,X90) | s(bool,X88) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(subsumption_resolution,[],[f2323,f846])).
fof(f2323,plain,(
  ( ! [X90,X88,X91,X89] : (~p(s(bool,X88)) | ~sP11('T_0',X89,X90) | sP12('T_0',X89,X90) | s(bool,X88) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(superposition,[],[f452,f783])).
fof(f2598,plain,(
  ( ! [X2,X1] : (~sP12('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl38),
  inference(cnf_transformation,[],[f2598_D])).
fof(f2598_D,plain,(
  ( ! [X2,X1] : (~sP12('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f2672,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP12('T_0',X2,X1)) )),
  inference(condensation,[],[f2671])).
fof(f2671,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,X85)) | p(s(bool,X86)) | ~sP12('T_0',X86,X85) | s(bool,X84) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(subsumption_resolution,[],[f2322,f846])).
fof(f2322,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,X84)) | ~p(s(bool,X85)) | p(s(bool,X86)) | ~sP12('T_0',X86,X85) | s(bool,X84) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(superposition,[],[f449,f783])).
fof(f2595,plain,(
  ( ! [X2,X1] : (sP12('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl36),
  inference(cnf_transformation,[],[f2595_D])).
fof(f2595_D,plain,(
  ( ! [X2,X1] : (sP12('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f2670,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP12('T_0',X1,X2)) )),
  inference(condensation,[],[f2669])).
fof(f2669,plain,(
  ( ! [X80,X83,X81,X82] : (~p(s(bool,X81)) | sP12('T_0',X81,X82) | s(bool,X80) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(subsumption_resolution,[],[f2321,f846])).
fof(f2321,plain,(
  ( ! [X80,X83,X81,X82] : (~p(s(bool,X80)) | ~p(s(bool,X81)) | sP12('T_0',X81,X82) | s(bool,X80) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(superposition,[],[f448,f783])).
fof(f2592,plain,(
  ( ! [X2,X1] : (sP12('T_0',X2,X1) | p(s(bool,X1))) ) | $spl34),
  inference(cnf_transformation,[],[f2592_D])).
fof(f2592_D,plain,(
  ( ! [X2,X1] : (sP12('T_0',X2,X1) | p(s(bool,X1))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f2668,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP12('T_0',X2,X1)) )),
  inference(condensation,[],[f2667])).
fof(f2667,plain,(
  ( ! [X78,X76,X79,X77] : (p(s(bool,X77)) | sP12('T_0',X78,X77) | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(subsumption_resolution,[],[f2320,f846])).
fof(f2320,plain,(
  ( ! [X78,X76,X79,X77] : (~p(s(bool,X76)) | p(s(bool,X77)) | sP12('T_0',X78,X77) | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(superposition,[],[f447,f783])).
fof(f2589,plain,(
  ( ! [X2,X1] : (~sP11(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) | $spl32),
  inference(cnf_transformation,[],[f2589_D])).
fof(f2589_D,plain,(
  ( ! [X2,X1] : (~sP11(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f2666,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP11(X2,X1,'T_0')) )),
  inference(condensation,[],[f2665])).
fof(f2665,plain,(
  ( ! [X74,X72,X75,X73] : (p(s(bool,X73)) | ~p(s(bool,X74)) | ~sP11(X74,X73,'T_0') | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(subsumption_resolution,[],[f2319,f846])).
fof(f2319,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,X72)) | p(s(bool,X73)) | ~p(s(bool,X74)) | ~sP11(X74,X73,'T_0') | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f442,f783])).
fof(f2586,plain,(
  ( ! [X2,X1] : (~sP10('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl30),
  inference(cnf_transformation,[],[f2586_D])).
fof(f2586_D,plain,(
  ( ! [X2,X1] : (~sP10('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f2664,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP10('T_0',X2,X1)) )),
  inference(condensation,[],[f2663])).
fof(f2663,plain,(
  ( ! [X70,X68,X71,X69] : (p(s(bool,X69)) | p(s(bool,X70)) | ~sP10('T_0',X70,X69) | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(subsumption_resolution,[],[f2318,f846])).
fof(f2318,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,X68)) | p(s(bool,X69)) | p(s(bool,X70)) | ~sP10('T_0',X70,X69) | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(superposition,[],[f432,f783])).
fof(f2583,plain,(
  ( ! [X2,X1] : (sP10('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl28),
  inference(cnf_transformation,[],[f2583_D])).
fof(f2583_D,plain,(
  ( ! [X2,X1] : (sP10('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f2662,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP10('T_0',X1,X2)) )),
  inference(condensation,[],[f2661])).
fof(f2661,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,X65)) | sP10('T_0',X65,X66) | s(bool,X64) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(subsumption_resolution,[],[f2317,f846])).
fof(f2317,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,X64)) | ~p(s(bool,X65)) | sP10('T_0',X65,X66) | s(bool,X64) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(superposition,[],[f431,f783])).
fof(f2580,plain,(
  ( ! [X2,X1] : (sP10('T_0',X2,X1) | ~p(s(bool,X1))) ) | $spl26),
  inference(cnf_transformation,[],[f2580_D])).
fof(f2580_D,plain,(
  ( ! [X2,X1] : (sP10('T_0',X2,X1) | ~p(s(bool,X1))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f2660,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP10('T_0',X2,X1)) )),
  inference(condensation,[],[f2659])).
fof(f2659,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,X61)) | sP10('T_0',X62,X61) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(subsumption_resolution,[],[f2316,f846])).
fof(f2316,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,X60)) | ~p(s(bool,X61)) | sP10('T_0',X62,X61) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(superposition,[],[f430,f783])).
fof(f2577,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) | $spl24),
  inference(cnf_transformation,[],[f2577_D])).
fof(f2577_D,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f2658,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f2657])).
fof(f2657,plain,(
  ( ! [X59,X57,X58,X56] : (p(s(bool,X57)) | p(s(bool,X58)) | ~sP9('T_0',X57,X58) | s(bool,X56) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(subsumption_resolution,[],[f2315,f846])).
fof(f2315,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,X56)) | p(s(bool,X57)) | p(s(bool,X58)) | ~sP9('T_0',X57,X58) | s(bool,X56) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f425,f783])).
fof(f2574,plain,(
  ( ! [X2,X1] : (~sP8(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl22),
  inference(cnf_transformation,[],[f2574_D])).
fof(f2574_D,plain,(
  ( ! [X2,X1] : (~sP8(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f2656,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X2,X1,'T_0')) )),
  inference(condensation,[],[f2655])).
fof(f2655,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X53)) | p(s(bool,X54)) | ~sP8(X54,X53,'T_0') | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(subsumption_resolution,[],[f2314,f846])).
fof(f2314,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X52)) | ~p(s(bool,X53)) | p(s(bool,X54)) | ~sP8(X54,X53,'T_0') | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f417,f783])).
fof(f2571,plain,(
  ( ! [X2,X1] : (~sP7(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl20),
  inference(cnf_transformation,[],[f2571_D])).
fof(f2571_D,plain,(
  ( ! [X2,X1] : (~sP7(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f2654,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X1,'T_0',X2)) )),
  inference(condensation,[],[f2653])).
fof(f2653,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X49)) | p(s(bool,X50)) | ~sP7(X49,'T_0',X50) | s(bool,X48) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(subsumption_resolution,[],[f2313,f846])).
fof(f2313,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X48)) | ~p(s(bool,X49)) | p(s(bool,X50)) | ~sP7(X49,'T_0',X50) | s(bool,X48) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f408,f783])).
fof(f2568,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl18),
  inference(cnf_transformation,[],[f2568_D])).
fof(f2568_D,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f2652,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f2651])).
fof(f2651,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X45)) | p(s(bool,X46)) | ~sP4(X45,'T_0',X46) | s(bool,X44) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f2312,f846])).
fof(f2312,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X44)) | ~p(s(bool,X45)) | p(s(bool,X46)) | ~sP4(X45,'T_0',X46) | s(bool,X44) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f389,f783])).
fof(f2565,plain,(
  ( ! [X2,X1] : (~sP5(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl16),
  inference(cnf_transformation,[],[f2565_D])).
fof(f2565_D,plain,(
  ( ! [X2,X1] : (~sP5(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f2650,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f2649])).
fof(f2649,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,X41)) | p(s(bool,X42)) | ~sP5(X41,'T_0',X42) | s(bool,X40) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(subsumption_resolution,[],[f2311,f846])).
fof(f2311,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,X40)) | ~p(s(bool,X41)) | p(s(bool,X42)) | ~sP5(X41,'T_0',X42) | s(bool,X40) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(superposition,[],[f385,f783])).
fof(f2562,plain,(
  ( ! [X2,X1] : (~sP5(X1,'T_0',X2) | ~sP4(X2,'T_0',X1) | sP6(X1,X2,'T_0')) ) | $spl14),
  inference(cnf_transformation,[],[f2562_D])).
fof(f2562_D,plain,(
  ( ! [X2,X1] : (~sP5(X1,'T_0',X2) | ~sP4(X2,'T_0',X1) | sP6(X1,X2,'T_0')) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f2648,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP6(X1,X2,'T_0') | ~sP5(X1,'T_0',X2) | ~sP4(X2,'T_0',X1)) )),
  inference(subsumption_resolution,[],[f2647,f497])).
fof(f2647,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP6(X1,X2,'T_0') | ~sP5(X1,'T_0',X2) | ~sP4(X2,'T_0',X1) | ~sP3('T_0',X2,X1)) )),
  inference(condensation,[],[f2646])).
fof(f2646,plain,(
  ( ! [X39,X37,X38,X36] : (sP6(X37,X38,'T_0') | ~sP5(X37,'T_0',X38) | ~sP4(X38,'T_0',X37) | ~sP3('T_0',X38,X37) | s(bool,X36) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(subsumption_resolution,[],[f2310,f846])).
fof(f2310,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,X36)) | sP6(X37,X38,'T_0') | ~sP5(X37,'T_0',X38) | ~sP4(X38,'T_0',X37) | ~sP3('T_0',X38,X37) | s(bool,X36) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f382,f783])).
fof(f2558,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) | $spl12),
  inference(cnf_transformation,[],[f2558_D])).
fof(f2558_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f2645,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f2644])).
fof(f2644,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X34)) | sP1(X34,'T_0') | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(subsumption_resolution,[],[f2309,f846])).
fof(f2309,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | p(s(bool,X34)) | sP1(X34,'T_0') | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f353,f783])).
fof(f2554,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) | $spl8),
  inference(cnf_transformation,[],[f2554_D])).
fof(f2554_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f2643,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f2642])).
fof(f2642,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X31)) | sP1('T_0',X31) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(subsumption_resolution,[],[f2308,f846])).
fof(f2308,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | sP1('T_0',X31) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f350,f783])).
fof(f2603,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP9('T_0',X1,X2) | sP10('T_0',X1,X2)) )),
  inference(condensation,[],[f2602])).
fof(f2602,plain,(
  ( ! [X94,X92,X95,X93] : (~sP9('T_0',X93,X94) | sP10('T_0',X93,X94) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(subsumption_resolution,[],[f2231,f846])).
fof(f2231,plain,(
  ( ! [X94,X92,X95,X93] : (~p(s(bool,X92)) | ~sP9('T_0',X93,X94) | sP10('T_0',X93,X94) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(superposition,[],[f453,f783])).
fof(f2600,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP11('T_0',X1,X2) | sP12('T_0',X1,X2)) )),
  inference(condensation,[],[f2599])).
fof(f2599,plain,(
  ( ! [X90,X88,X91,X89] : (~sP11('T_0',X89,X90) | sP12('T_0',X89,X90) | s(bool,X88) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(subsumption_resolution,[],[f2230,f846])).
fof(f2230,plain,(
  ( ! [X90,X88,X91,X89] : (~p(s(bool,X88)) | ~sP11('T_0',X89,X90) | sP12('T_0',X89,X90) | s(bool,X88) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(superposition,[],[f452,f783])).
fof(f2597,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP12('T_0',X2,X1)) )),
  inference(condensation,[],[f2596])).
fof(f2596,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,X85)) | p(s(bool,X86)) | ~sP12('T_0',X86,X85) | s(bool,X84) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(subsumption_resolution,[],[f2229,f846])).
fof(f2229,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,X84)) | ~p(s(bool,X85)) | p(s(bool,X86)) | ~sP12('T_0',X86,X85) | s(bool,X84) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(superposition,[],[f449,f783])).
fof(f2594,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP12('T_0',X1,X2)) )),
  inference(condensation,[],[f2593])).
fof(f2593,plain,(
  ( ! [X80,X83,X81,X82] : (~p(s(bool,X81)) | sP12('T_0',X81,X82) | s(bool,X80) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(subsumption_resolution,[],[f2228,f846])).
fof(f2228,plain,(
  ( ! [X80,X83,X81,X82] : (~p(s(bool,X80)) | ~p(s(bool,X81)) | sP12('T_0',X81,X82) | s(bool,X80) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(superposition,[],[f448,f783])).
fof(f2591,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP12('T_0',X2,X1)) )),
  inference(condensation,[],[f2590])).
fof(f2590,plain,(
  ( ! [X78,X76,X79,X77] : (p(s(bool,X77)) | sP12('T_0',X78,X77) | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(subsumption_resolution,[],[f2227,f846])).
fof(f2227,plain,(
  ( ! [X78,X76,X79,X77] : (~p(s(bool,X76)) | p(s(bool,X77)) | sP12('T_0',X78,X77) | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(superposition,[],[f447,f783])).
fof(f2588,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP11(X2,X1,'T_0')) )),
  inference(condensation,[],[f2587])).
fof(f2587,plain,(
  ( ! [X74,X72,X75,X73] : (p(s(bool,X73)) | ~p(s(bool,X74)) | ~sP11(X74,X73,'T_0') | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(subsumption_resolution,[],[f2226,f846])).
fof(f2226,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,X72)) | p(s(bool,X73)) | ~p(s(bool,X74)) | ~sP11(X74,X73,'T_0') | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f442,f783])).
fof(f2585,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP10('T_0',X2,X1)) )),
  inference(condensation,[],[f2584])).
fof(f2584,plain,(
  ( ! [X70,X68,X71,X69] : (p(s(bool,X69)) | p(s(bool,X70)) | ~sP10('T_0',X70,X69) | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(subsumption_resolution,[],[f2225,f846])).
fof(f2225,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,X68)) | p(s(bool,X69)) | p(s(bool,X70)) | ~sP10('T_0',X70,X69) | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(superposition,[],[f432,f783])).
fof(f2582,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP10('T_0',X1,X2)) )),
  inference(condensation,[],[f2581])).
fof(f2581,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,X65)) | sP10('T_0',X65,X66) | s(bool,X64) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(subsumption_resolution,[],[f2224,f846])).
fof(f2224,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,X64)) | ~p(s(bool,X65)) | sP10('T_0',X65,X66) | s(bool,X64) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(superposition,[],[f431,f783])).
fof(f2579,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP10('T_0',X2,X1)) )),
  inference(condensation,[],[f2578])).
fof(f2578,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,X61)) | sP10('T_0',X62,X61) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(subsumption_resolution,[],[f2223,f846])).
fof(f2223,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,X60)) | ~p(s(bool,X61)) | sP10('T_0',X62,X61) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(superposition,[],[f430,f783])).
fof(f2576,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f2575])).
fof(f2575,plain,(
  ( ! [X59,X57,X58,X56] : (p(s(bool,X57)) | p(s(bool,X58)) | ~sP9('T_0',X57,X58) | s(bool,X56) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(subsumption_resolution,[],[f2222,f846])).
fof(f2222,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,X56)) | p(s(bool,X57)) | p(s(bool,X58)) | ~sP9('T_0',X57,X58) | s(bool,X56) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f425,f783])).
fof(f2573,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X2,X1,'T_0')) )),
  inference(condensation,[],[f2572])).
fof(f2572,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X53)) | p(s(bool,X54)) | ~sP8(X54,X53,'T_0') | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(subsumption_resolution,[],[f2221,f846])).
fof(f2221,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X52)) | ~p(s(bool,X53)) | p(s(bool,X54)) | ~sP8(X54,X53,'T_0') | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f417,f783])).
fof(f2570,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X1,'T_0',X2)) )),
  inference(condensation,[],[f2569])).
fof(f2569,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X49)) | p(s(bool,X50)) | ~sP7(X49,'T_0',X50) | s(bool,X48) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(subsumption_resolution,[],[f2220,f846])).
fof(f2220,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X48)) | ~p(s(bool,X49)) | p(s(bool,X50)) | ~sP7(X49,'T_0',X50) | s(bool,X48) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f408,f783])).
fof(f2567,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f2566])).
fof(f2566,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X45)) | p(s(bool,X46)) | ~sP4(X45,'T_0',X46) | s(bool,X44) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f2219,f846])).
fof(f2219,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X44)) | ~p(s(bool,X45)) | p(s(bool,X46)) | ~sP4(X45,'T_0',X46) | s(bool,X44) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f389,f783])).
fof(f2564,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f2563])).
fof(f2563,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,X41)) | p(s(bool,X42)) | ~sP5(X41,'T_0',X42) | s(bool,X40) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(subsumption_resolution,[],[f2218,f846])).
fof(f2218,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,X40)) | ~p(s(bool,X41)) | p(s(bool,X42)) | ~sP5(X41,'T_0',X42) | s(bool,X40) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(superposition,[],[f385,f783])).
fof(f2561,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP6(X1,X2,'T_0') | ~sP5(X1,'T_0',X2) | ~sP4(X2,'T_0',X1)) )),
  inference(subsumption_resolution,[],[f2560,f497])).
fof(f2560,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP6(X1,X2,'T_0') | ~sP5(X1,'T_0',X2) | ~sP4(X2,'T_0',X1) | ~sP3('T_0',X2,X1)) )),
  inference(condensation,[],[f2559])).
fof(f2559,plain,(
  ( ! [X39,X37,X38,X36] : (sP6(X37,X38,'T_0') | ~sP5(X37,'T_0',X38) | ~sP4(X38,'T_0',X37) | ~sP3('T_0',X38,X37) | s(bool,X36) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(subsumption_resolution,[],[f2217,f846])).
fof(f2217,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,X36)) | sP6(X37,X38,'T_0') | ~sP5(X37,'T_0',X38) | ~sP4(X38,'T_0',X37) | ~sP3('T_0',X38,X37) | s(bool,X36) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f382,f783])).
fof(f2557,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f2556])).
fof(f2556,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X34)) | sP1(X34,'T_0') | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(subsumption_resolution,[],[f2216,f846])).
fof(f2216,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | p(s(bool,X34)) | sP1(X34,'T_0') | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f353,f783])).
fof(f2553,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f2552])).
fof(f2552,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X31)) | sP1('T_0',X31) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(subsumption_resolution,[],[f2215,f846])).
fof(f2215,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | sP1('T_0',X31) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f350,f783])).
fof(f467,plain,(
  p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK15)))) | $spl6),
  inference(cnf_transformation,[],[f467_D])).
fof(f467_D,plain,(
  p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK15)))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f462,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))))) | $spl2),
  inference(cnf_transformation,[],[f462_D])).
fof(f462_D,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f260,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))))) | p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK15))))),
  inference(cnf_transformation,[],[f168])).
fof(f168,plain,(
  ((p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',sK15)))) & p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15))))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13))))))) & (! [X2] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',X2)))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',X2))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13,sK14,sK15])],[f167])).
fof(f167,plain,(
  ? [X0,X1] : ((? [X3] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X3)))) & p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X3))))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X0))))))) & (! [X2] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X2)))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X0))))))))),
  inference(rectify,[],[f166])).
fof(f166,plain,(
  ? [X0,X1] : ((? [X2] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X2)))) & p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X0))))))) & (! [X2] : (~p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X2)))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X0))))))))),
  inference(nnf_transformation,[],[f115])).
fof(f115,plain,(
  ? [X0,X1] : (? [X2] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X2)))) & p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))))) <~> p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X0)))))))),
  inference(ennf_transformation,[],[f47])).
fof(f47,plain,(
  ~! [X0,X1] : (? [X2] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),X0),s('type.extreal.extreal',X2)))) & p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))))) <=> p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X0)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (? [X8] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),X6),s('type.extreal.extreal',X8)))) & p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X8))))) <=> p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X6)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (? [X8] : (p(s(bool,happ(s(fun('type.extreal.extreal',bool),X6),s('type.extreal.extreal',X8)))) & p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X8))))) <=> p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),X6)))))))),
  file('thm.extreal.sup_lt','thm.extreal.sup_lt')).
fof(f465,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15)))) | $spl4),
  inference(cnf_transformation,[],[f465_D])).
fof(f465_D,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f261,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15))))),
  inference(cnf_transformation,[],[f168])).
fof(f461,plain,(
  ( ! [X2] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',X2)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',X2))))) ) | $spl0),
  inference(cnf_transformation,[],[f461_D])).
fof(f461_D,plain,(
  ( ! [X2] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',X2)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',X2))))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f463,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))))) | $spl3),
  inference(cnf_transformation,[],[f463_D])).
fof(f463_D,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f262,plain,(
  ( ! [X2] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_sup_1'(s(fun('type.extreal.extreal',bool),sK13)))))) | ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',X2)))) | ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),sK13),s('type.extreal.extreal',X2))))) )),
  inference(cnf_transformation,[],[f168])).
% SZS output end Proof for thm.extreal
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1799
% Passive clauses: 6362
% Generated clauses: 57996
% Final active clauses: 1630
% Final passive clauses: 3616
% Input formulas: 45
% Initial clauses: 192
% 
% Duplicate literals: 2565
% Trivial inequalities: 91
% Fw subsumption resolutions: 2271
% Fw demodulations: 8985
% Bw demodulations: 105
% Condensations: 3471
% Global subsumptions: 38
% 
% Simple tautologies: 155
% Equational tautologies: 1335
% Forward subsumptions: 32114
% Bw demodulations to eq. taut.: 2
% 
% Unit resulting resolution: 26551
% Factoring: 6
% Forward superposition: 6711
% Backward superposition: 6585
% Self superposition: 156
% Equality factoring: 137
% Equality resolution: 3
% 
% Split clauses: 211
% Split components: 435
% SAT solver clauses: 7294
% SAT solver unit clauses: 333
% SAT solver binary clauses: 2678
% 
% Sat splits: 211
% Sat splitting refutations: 470
% 
% Memory used [KB]: 30063
% Time elapsed: 3.790 s
% ------------------------------
% ------------------------------
% Success in time 5.105 s
