Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.relation
% (73428)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 194
% Passive clauses: 3319
% Generated clauses: 17482
% Final active clauses: 172
% Final passive clauses: 2911
% Input formulas: 28
% Initial clauses: 76
% 
% Duplicate literals: 664
% Trivial inequalities: 305
% Fw subsumption resolutions: 2509
% Fw demodulations: 2639
% Bw demodulations: 18
% 
% Simple tautologies: 380
% Equational tautologies: 19
% Forward subsumptions: 6521
% 
% Binary resolution: 1016
% Forward superposition: 6014
% Backward superposition: 3823
% Self superposition: 79
% Equality factoring: 45
% Equality resolution: 56
% 
% Split clauses: 739
% Split components: 1490
% SAT solver clauses: 1339
% SAT solver unit clauses: 15
% SAT solver binary clauses: 148
% 
% Sat splits: 739
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1338
% TWLsolver calls for satisfiability: 250
% 
% Memory used [KB]: 10618
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.relation
% (73443)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 108
% Passive clauses: 1397
% Generated clauses: 9342
% Final active clauses: 96
% Final passive clauses: 1226
% Input formulas: 28
% Initial clauses: 76
% 
% Duplicate literals: 342
% Trivial inequalities: 51
% Fw subsumption resolutions: 390
% Fw demodulations: 1080
% Bw demodulations: 72
% Condensations: 544
% Global subsumptions: 8
% 
% Simple tautologies: 97
% Equational tautologies: 17
% Forward subsumptions: 3397
% 
% Binary resolution: 340
% Factoring: 1
% Forward superposition: 2931
% Backward superposition: 2005
% Self superposition: 56
% Equality factoring: 36
% Equality resolution: 6
% 
% Split clauses: 42
% Split components: 89
% SAT solver clauses: 2043
% SAT solver unit clauses: 79
% SAT solver binary clauses: 39
% SAT solver learnt clauses: 16
% 
% Sat splits: 42
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 2011
% TWLsolver calls for satisfiability: 367
% 
% Memory used [KB]: 13688
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.relation
% (73465)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 495
% Passive clauses: 6358
% Generated clauses: 145749
% Final active clauses: 65
% Final passive clauses: 282
% Input formulas: 28
% Initial clauses: 76
% Discarded non-redundant clauses: 157
% 
% Duplicate literals: 6278
% Trivial inequalities: 1080
% Fw subsumption resolutions: 15449
% Fw demodulations: 22644
% Bw demodulations: 4856
% Condensations: 15834
% Global subsumptions: 323
% 
% Simple tautologies: 1133
% Equational tautologies: 1028
% Forward subsumptions: 74413
% Fw demodulations to eq. taut.: 17
% Bw demodulations to eq. taut.: 348
% 
% Unit resulting resolution: 3945
% Forward superposition: 38912
% Backward superposition: 27737
% Self superposition: 593
% Equality factoring: 300
% Equality resolution: 69
% 
% Split clauses: 163
% Split components: 333
% SAT solver clauses: 7366
% SAT solver unit clauses: 45
% SAT solver binary clauses: 316
% 
% Sat splits: 163
% Sat splitting refutations: 356
% 
% Memory used [KB]: 66907
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.relation
% (73735)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 749
% Passive clauses: 21200
% Generated clauses: 102323
% Final active clauses: 266
% Final passive clauses: 15045
% Input formulas: 28
% Initial clauses: 75
% 
% Selected by SInE selection: 27
% SInE iterations: 1
% 
% Duplicate literals: 2240
% Trivial inequalities: 2180
% Fw subsumption resolutions: 18052
% Fw demodulations: 5804
% Bw demodulations: 343
% 
% Simple tautologies: 1983
% Equational tautologies: 91
% Forward subsumptions: 40284
% Bw demodulations to eq. taut.: 32
% 
% Binary resolution: 15183
% Factoring: 17
% Forward superposition: 25357
% Backward superposition: 30256
% Self superposition: 355
% Equality factoring: 288
% Equality resolution: 128
% 
% Split clauses: 7201
% Split components: 15429
% SAT solver clauses: 7293
% SAT solver unit clauses: 37
% SAT solver binary clauses: 3480
% 
% Sat splits: 7201
% Sat splitting refutations: 43
% 
% Memory used [KB]: 41960
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 142 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.relation
% (73861)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 71
% Passive clauses: 1811
% Generated clauses: 11534
% Final active clauses: 34
% Final passive clauses: 561
% Input formulas: 28
% Initial clauses: 76
% 
% Duplicate literals: 337
% Trivial inequalities: 259
% Fw subsumption resolutions: 1686
% Fw demodulations: 843
% Bw demodulations: 1150
% Global subsumptions: 73
% 
% Simple tautologies: 399
% Equational tautologies: 65
% Forward subsumptions: 5438
% Bw demodulations to eq. taut.: 38
% 
% Unit resulting resolution: 392
% Factoring: 4
% Forward superposition: 2977
% Backward superposition: 2404
% Self superposition: 74
% Equality factoring: 50
% Equality resolution: 10
% 
% Split clauses: 17
% Split components: 35
% SAT solver clauses: 2070
% SAT solver unit clauses: 135
% SAT solver binary clauses: 112
% 
% Sat splits: 17
% Sat splitting refutations: 167
% 
% Memory used [KB]: 8187
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 135 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.relation
% (73879)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 353
% Passive clauses: 3380
% Generated clauses: 18155
% Final active clauses: 334
% Final passive clauses: 2954
% Input formulas: 28
% Initial clauses: 76
% 
% Duplicate literals: 406
% Trivial inequalities: 500
% Fw subsumption resolutions: 991
% Fw demodulations: 742
% Bw demodulations: 10
% Condensations: 1520
% 
% Simple tautologies: 496
% Equational tautologies: 45
% Forward subsumptions: 10053
% Backward subsumptions: 123
% 
% Binary resolution: 5174
% Factoring: 23
% Forward superposition: 6410
% Backward superposition: 2351
% Equality factoring: 2
% Equality resolution: 40
% 
% SAT solver clauses: 3382
% SAT solver unit clauses: 31
% SAT solver binary clauses: 155
% 
% Memory used [KB]: 16630
% Time elapsed: 1.600 s
% ------------------------------
% ------------------------------
% remaining time: 119 next slice time: 9
lrs+10_8:1_bd=preordered:bs=on:ccuc=first:fsr=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:nicw=on:sas=minisat:stl=120:sos=on:sscc=on:sser=off:sfr=on:ssfp=4000:ssfq=1.0:ssnc=none:sp=reverse_arity:urr=on_7 on thm.relation
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.relation
% SZS output start Proof for thm.relation
fof(f5132,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f124,f201,f123,f203,f122,f205,f221,f220,f219,f218,f232,f217,f216,f215,f209,f208,f230,f207,f206,f226,f227,f235,f234,f261,f260,f257,f256,f255,f249,f248,f270,f247,f246,f266,f267,f273,f272,f306,f305,f302,f301,f300,f299,f296,f295,f292,f291,f290,f289,f282,f281,f280,f279,f314,f317,f319,f320,f322,f323,f325,f310,f311,f316,f703,f702,f698,f697,f694,f693,f692,f691,f690,f689,f688,f687,f682,f681,f678,f677,f672,f671,f670,f669,f664,f662,f658,f657,f656,f655,f652,f651,f650,f649,f648,f647,f646,f645,f640,f638,f634,f633,f632,f631,f624,f623,f622,f621,f720,f708,f747,f712,f713,f745,f809,f805,f804,f803,f802,f801,f796,f789,f783,f782,f781,f779,f775,f774,f773,f772,f771,f766,f762,f756,f755,f822,f754,f753,f752,f751,f811,f851,f812,f849,f817,f818,f852,f847,f907,f906,f903,f902,f901,f900,f899,f898,f897,f896,f895,f894,f890,f889,f886,f885,f882,f881,f880,f879,f878,f877,f876,f875,f874,f873,f872,f871,f870,f869,f868,f867,f860,f859,f858,f857,f921,f914,f915,f926,f985,f981,f980,f979,f978,f977,f972,f965,f959,f958,f957,f955,f951,f950,f949,f948,f947,f942,f938,f932,f931,f998,f930,f929,f928,f927,f987,f1027,f988,f1025,f993,f994,f1028,f1023,f1063,f1062,f1061,f1060,f1059,f1055,f1052,f1051,f1050,f1049,f1048,f1047,f1046,f1045,f1044,f1038,f1037,f925,f1036,f1035,f1034,f1033,f1032,f1031,f1082,f1083,f1084,f1071,f1072,f1087,f920,f1141,f1140,f1137,f1136,f1135,f1134,f1133,f1132,f1131,f1130,f1129,f1128,f1124,f1123,f1120,f1119,f1116,f1115,f1114,f1113,f1112,f1111,f1110,f1109,f1108,f1107,f1106,f1105,f1104,f1103,f1102,f1101,f1094,f1093,f1092,f1091,f1155,f1148,f1149,f1194,f1193,f1192,f1191,f1190,f1186,f1183,f1182,f1181,f1180,f1179,f1178,f1177,f1176,f1175,f1169,f1168,f1206,f1167,f1166,f1165,f1213,f1164,f1215,f1163,f1162,f1217,f1218,f1219,f1202,f1203,f1220,f1211,f1221,f1256,f1255,f1254,f1253,f1252,f1248,f1245,f1244,f1243,f1242,f1241,f1240,f1239,f1238,f1237,f1231,f1230,f1268,f1229,f1228,f1227,f1275,f1226,f1277,f1225,f1224,f1279,f1280,f1281,f1264,f1265,f1282,f1273,f1317,f1316,f1315,f1314,f1313,f1309,f1306,f1305,f1304,f1303,f1302,f1301,f1300,f1299,f1298,f1292,f1291,f1329,f1290,f1289,f1288,f1336,f1287,f1338,f1286,f1285,f1340,f1341,f1342,f1325,f1326,f1343,f1334,f1378,f1377,f1376,f1375,f1374,f1370,f1367,f1366,f1365,f1364,f1363,f1362,f1361,f1360,f1359,f1353,f1352,f1390,f1351,f1350,f1349,f1078,f1348,f1080,f1347,f1346,f1397,f1398,f1399,f1386,f1387,f1400,f1395,f2226,f2220,f2219,f2218,f2217,f2216,f2215,f2214,f2213,f2212,f2205,f2196,f2190,f2189,f2188,f2186,f2180,f2179,f2178,f2177,f2176,f2175,f2174,f2173,f2172,f2165,f2161,f2155,f2154,f2239,f2153,f2152,f2234,f2235,f2279,f2274,f2358,f2352,f2351,f2350,f2349,f2348,f2347,f2346,f2345,f2344,f2337,f2328,f2322,f2321,f2320,f2318,f2312,f2311,f2310,f2309,f2308,f2307,f2306,f2305,f2304,f2297,f2293,f2287,f2286,f2371,f2285,f2284,f2366,f2367,f2411,f2406,f2490,f2484,f2483,f2482,f2481,f2480,f2479,f2478,f2477,f2476,f2469,f2460,f2454,f2453,f2452,f2450,f2444,f2443,f2442,f2441,f2440,f2439,f2438,f2437,f2436,f2429,f2425,f2419,f2418,f2503,f2417,f2416,f2498,f2499,f2543,f2538,f2622,f2616,f2615,f2614,f2613,f2612,f2611,f2610,f2609,f2608,f2601,f2592,f2586,f2585,f2584,f2582,f2576,f2575,f2574,f2573,f2572,f2571,f2570,f2569,f2568,f2561,f2557,f2551,f2550,f2635,f2549,f2548,f2630,f2631,f2675,f2670,f3664,f3663,f3662,f3661,f3660,f3659,f3658,f3657,f3656,f3655,f3654,f3653,f3652,f3651,f3650,f3649,f3648,f3636,f3633,f3632,f3631,f3630,f3629,f3628,f3627,f3626,f3625,f3624,f3623,f3622,f3621,f3620,f3619,f3618,f3617,f3616,f3615,f3614,f3613,f3612,f3611,f3610,f3609,f3603,f3602,f3676,f3601,f3600,f3597,f3596,f3691,f3692,f3693,f3672,f3673,f3694,f3689,f3939,f3938,f3937,f3936,f3935,f3934,f3933,f3932,f3931,f3930,f3929,f3928,f3927,f3926,f3925,f3924,f3923,f3910,f3907,f3906,f3905,f3904,f3903,f3902,f3901,f3900,f3899,f3898,f3897,f3896,f3895,f3894,f3893,f3892,f3891,f3890,f3889,f3888,f3887,f3886,f3885,f3884,f3883,f3882,f3881,f3875,f3874,f3951,f3873,f3872,f3869,f3868,f3968,f3969,f3970,f3947,f3948,f3971,f3965,f4045,f4044,f4043,f4042,f4041,f4040,f4039,f4038,f4037,f4036,f4035,f4034,f4033,f4032,f4031,f4030,f4029,f4016,f4013,f4012,f4011,f4010,f4009,f4008,f4007,f4006,f4005,f4004,f4003,f4002,f4001,f4000,f3999,f3998,f3997,f3996,f3995,f3994,f3993,f3992,f3991,f3990,f3989,f3988,f3987,f3981,f3980,f4057,f3979,f3978,f3975,f3974,f4073,f4074,f4075,f4053,f4054,f4076,f4071,f4150,f4149,f4148,f4147,f4146,f4145,f4144,f4143,f4142,f4141,f4140,f4139,f4138,f4137,f4136,f4135,f4134,f4121,f4118,f4117,f4116,f4115,f4114,f4113,f4112,f4111,f4110,f4109,f4108,f4107,f4106,f4105,f4104,f4103,f4102,f4101,f4100,f4099,f4098,f4097,f4096,f4095,f4094,f4093,f4092,f4086,f4085,f4162,f4084,f4083,f4080,f4079,f4179,f4180,f4181,f4158,f4159,f4182,f4176,f4256,f4255,f4254,f4253,f4252,f4251,f4250,f4249,f4248,f4247,f4246,f4245,f4244,f4243,f4242,f4241,f4240,f4227,f4224,f4223,f4222,f4221,f4220,f4219,f4218,f4217,f4216,f4215,f4214,f4213,f4212,f4211,f4210,f4209,f4208,f4207,f4206,f4205,f4204,f4203,f4202,f4201,f4200,f4199,f4198,f4192,f4191,f4268,f4190,f4189,f4186,f4185,f4284,f4285,f4286,f4264,f4265,f4287,f4282,f4361,f4360,f4359,f4358,f4357,f4356,f4355,f4354,f4353,f4352,f4351,f4350,f4349,f4348,f4347,f4346,f4345,f4332,f4329,f4328,f4327,f4326,f4325,f4324,f4323,f4322,f4321,f4320,f4319,f4318,f4317,f4316,f4315,f4314,f4313,f4312,f4311,f4310,f4309,f4308,f4307,f4306,f4305,f4304,f4303,f4297,f4296,f4373,f4295,f4294,f4291,f4290,f4390,f4391,f4392,f4369,f4370,f4393,f4387,f4467,f4466,f4465,f4464,f4463,f4462,f4461,f4460,f4459,f4458,f4457,f4456,f4455,f4454,f4453,f4452,f4451,f4438,f4435,f4434,f4433,f4432,f4431,f4430,f4429,f4428,f4427,f4426,f4425,f4424,f4423,f4422,f4421,f4420,f4419,f4418,f4417,f4416,f4415,f4414,f4413,f4412,f4411,f4410,f4409,f4403,f4402,f4479,f4401,f4400,f4397,f4396,f4495,f4496,f4497,f4475,f4476,f4498,f4493,f4572,f4571,f4570,f4569,f4568,f4567,f4566,f4565,f4564,f4563,f4562,f4561,f4560,f4559,f4558,f4557,f4556,f4543,f4540,f4539,f4538,f4537,f4536,f4535,f4534,f4533,f4532,f4531,f4530,f4529,f4528,f4527,f4526,f4525,f4524,f4523,f4522,f4521,f4520,f4519,f4518,f4517,f4516,f4515,f4514,f4508,f4507,f4584,f4506,f4505,f4502,f4501,f4601,f4602,f4603,f4580,f4581,f4604,f4598,f4619,f4620,f4640,f4641,f4654,f4655,f4668,f4669,f4682,f4683,f4696,f4697,f4710,f4711,f4724,f4725,f4733,f4731,f748,f2807,f2808,f749,f2809,f1160,f4926,f4927,f5130,f4735,f1086,f5131])).
fof(f5131,plain,(
  s(bool,'F_0') != s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1))))) | p(s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)))))) | ~p(s(bool,'F_0'))),
  introduced(theory_axiom,[])).
fof(f1086,plain,(
  ~p(s(bool,'T_0')) | $spl63),
  inference(cnf_transformation,[],[f1086_D])).
fof(f1086_D,plain,(
  ~p(s(bool,'T_0')) <=> ~$spl63),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl63])])).
fof(f4735,plain,(
  p(s(bool,'F_0')) | $spl208),
  inference(cnf_transformation,[],[f4735_D])).
fof(f4735_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f5130,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'T_0')) | ($spl44 | $spl90 | $spl106)),
  inference(forward_demodulation,[],[f5128,f926])).
fof(f5128,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ($spl90 | $spl106)),
  inference(superposition,[],[f1281,f1395])).
fof(f4927,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) | $spl210),
  inference(cnf_transformation,[],[f4927_D])).
fof(f4927_D,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f4926,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) | ($spl10 | $spl64)),
  inference(duplicate_literal_removal,[],[f4925])).
fof(f4925,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) | ($spl10 | $spl64)),
  inference(forward_demodulation,[],[f4826,f234])).
fof(f4826,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK1))))) ) | $spl64),
  inference(superposition,[],[f167,f1154])).
fof(f1154,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl64),
  inference(cnf_transformation,[],[f1154_D])).
fof(f1154_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f167,plain,(
  ( ! [X6,X0,X7,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X7)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X6)))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X5,X6,X7] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X7)))))) & ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK3(X1,X0)))),s(X0,sK4(X1,X0))))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK4(X1,X0)))),s(X0,sK5(X1,X0))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK3(X1,X0)))),s(X0,sK5(X1,X0)))))) | p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3,sK4,sK5])],[f106])).
fof(f106,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X5,X6,X7] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X7)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X7)))))) & (? [X2,X3,X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4))))) | p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(rectify,[],[f105])).
fof(f105,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X2,X3,X4] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))))) & (? [X2,X3,X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4))))) | p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(nnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3,X4] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4))))))),
  inference(flattening,[],[f72])).
fof(f72,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3,X4] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4))))))),
  inference(ennf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3,X4] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4))))))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3,X4] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X4))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4))))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6] : ! [X19] : (p(s(bool,'const.relation.transitive_1'(s(fun(X6,fun(X6,bool)),X19)))) <=> ! [X10,X11,X20] : ((p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X19),s(X6,X10))),s(X6,X11)))) & p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X19),s(X6,X11))),s(X6,X20))))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X19),s(X6,X10))),s(X6,X20))))))),
  file('thm.relation.transitive_RINTER','thm.relation.transitive_def')).
fof(f1160,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl66),
  inference(cnf_transformation,[],[f1160_D])).
fof(f1160_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f2809,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl65),
  inference(unit_resulting_resolution,[],[f1155,f128])).
fof(f128,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.relation.transitive_RINTER',aHL_BOOL_CASES)).
fof(f749,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl33),
  inference(cnf_transformation,[],[f749_D])).
fof(f749_D,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f2808,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl65),
  inference(unit_resulting_resolution,[],[f1155,f130])).
fof(f130,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X9] : (s(bool,'T_0') = s(bool,X9) <=> p(s(bool,X9)))),
  file('thm.relation.transitive_RINTER','thm.bool.EQ_CLAUSES|split|0')).
fof(f2807,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl65),
  inference(unit_resulting_resolution,[],[f1155,f132])).
fof(f132,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X9] : (s(bool,'T_0') = s(bool,X9) <=> p(s(bool,X9)))),
  file('thm.relation.transitive_RINTER','thm.bool.EQ_CLAUSES|split|1')).
fof(f748,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl32),
  inference(cnf_transformation,[],[f748_D])).
fof(f748_D,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f4731,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ($spl78 | $spl94)),
  inference(superposition,[],[f1219,f1334])).
fof(f4733,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'T_0')) | ($spl66 | $spl78 | $spl94)),
  inference(forward_demodulation,[],[f4731,f1160])).
fof(f4725,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X1))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X1))),s(sK0,sK7(X0,happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))))) ) | $spl206),
  inference(cnf_transformation,[],[f4725_D])).
fof(f4725_D,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X1))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X1))),s(sK0,sK7(X0,happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))))) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f4724,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X1))),s(sK0,sK7(X0,happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X1))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) | ($spl14 | $spl52)),
  inference(forward_demodulation,[],[f4713,f272])).
fof(f4713,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X1))),s(sK0,sK7(X0,happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X1))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK2))))) ) | $spl52),
  inference(resolution,[],[f1381,f167])).
fof(f1381,plain,(
  ( ! [X2] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7(X2,happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | p(s(bool,X2))) ) | $spl52),
  inference(resolution,[],[f1027,f177])).
fof(f177,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2)) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0,X1,X2] : (((p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X2,X1,X0))))) & ~p(s(bool,X2))) | ! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | p(s(bool,X2))) & ((p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK7(X2,X1,X0))))) & ~p(s(bool,X2))) | ! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6,sK7])],[f113])).
fof(f113,plain,(
  ! [X0,X1,X2] : ((? [X6] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) & ~p(s(bool,X2))) | ! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | p(s(bool,X2))) & ((? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | ! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2)))))),
  inference(rectify,[],[f112])).
fof(f112,plain,(
  ! [X0,X1,X2] : ((? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,X2))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2))) & ((? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | ! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2)))))),
  inference(flattening,[],[f111])).
fof(f111,plain,(
  ! [X0,X1,X2] : ((? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,X2))) | (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2)))) & ((? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | ! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2)))))),
  inference(nnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1,X2] : (! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))) <=> (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2))))),
  inference(ennf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) => p(s(bool,X2))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) => p(s(bool,X2))))),
  inference(flattening,[],[f55])).
fof(f55,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) => p(s(bool,X2))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) => p(s(bool,X2))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : ! [X12,X13] : (! [X10] : (p(s(bool,happ(s(fun(X6,bool),X12),s(X6,X10)))) => p(s(bool,X13))) <=> (? [X10] : p(s(bool,happ(s(fun(X6,bool),X12),s(X6,X10)))) => p(s(bool,X13))))),
  file('thm.relation.transitive_RINTER','thm.bool.LEFT_FORALL_IMP_THM')).
fof(f4711,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X1))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X1))),s(sK0,sK6(X0,happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))))) ) | $spl204),
  inference(cnf_transformation,[],[f4711_D])).
fof(f4711_D,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X1))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X1))),s(sK0,sK6(X0,happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))))) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f4710,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X1))),s(sK0,sK6(X0,happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X1))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) | ($spl14 | $spl52)),
  inference(forward_demodulation,[],[f4699,f272])).
fof(f4699,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X1))),s(sK0,sK6(X0,happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X1))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK2))))) ) | $spl52),
  inference(resolution,[],[f1380,f167])).
fof(f1380,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6(X1,happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | p(s(bool,X1))) ) | $spl52),
  inference(resolution,[],[f1027,f175])).
fof(f175,plain,(
  ( ! [X2,X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | p(s(bool,X2)) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK6(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f114])).
fof(f4697,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,sK7(X0,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))))) ) | $spl202),
  inference(cnf_transformation,[],[f4697_D])).
fof(f4697_D,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,sK7(X0,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))))) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f4696,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,sK7(X0,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) | ($spl10 | $spl50)),
  inference(forward_demodulation,[],[f4685,f234])).
fof(f4685,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,sK7(X0,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK1))))) ) | $spl50),
  inference(resolution,[],[f1320,f167])).
fof(f1320,plain,(
  ( ! [X2] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7(X2,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | p(s(bool,X2))) ) | $spl50),
  inference(resolution,[],[f1025,f177])).
fof(f4683,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,sK6(X0,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))))) ) | $spl200),
  inference(cnf_transformation,[],[f4683_D])).
fof(f4683_D,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,sK6(X0,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))))) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f4682,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,sK6(X0,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) | ($spl10 | $spl50)),
  inference(forward_demodulation,[],[f4671,f234])).
fof(f4671,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,sK6(X0,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK1))))) ) | $spl50),
  inference(resolution,[],[f1319,f167])).
fof(f1319,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6(X1,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | p(s(bool,X1))) ) | $spl50),
  inference(resolution,[],[f1025,f175])).
fof(f4669,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X1))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X1))),s(sK0,sK7(X0,happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))))) ) | $spl198),
  inference(cnf_transformation,[],[f4669_D])).
fof(f4669_D,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X1))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X1))),s(sK0,sK7(X0,happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))))) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f4668,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X1))),s(sK0,sK7(X0,happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X1))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) | ($spl14 | $spl40)),
  inference(forward_demodulation,[],[f4657,f272])).
fof(f4657,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X1))),s(sK0,sK7(X0,happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X1))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK2))))) ) | $spl40),
  inference(resolution,[],[f1259,f167])).
fof(f1259,plain,(
  ( ! [X2] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7(X2,happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | p(s(bool,X2))) ) | $spl40),
  inference(resolution,[],[f851,f177])).
fof(f4655,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X1))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X1))),s(sK0,sK6(X0,happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))))) ) | $spl196),
  inference(cnf_transformation,[],[f4655_D])).
fof(f4655_D,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X1))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X1))),s(sK0,sK6(X0,happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))))) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f4654,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X1))),s(sK0,sK6(X0,happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X1))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) | ($spl14 | $spl40)),
  inference(forward_demodulation,[],[f4643,f272])).
fof(f4643,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X1))),s(sK0,sK6(X0,happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X1))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK2))))) ) | $spl40),
  inference(resolution,[],[f1258,f167])).
fof(f1258,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6(X1,happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | p(s(bool,X1))) ) | $spl40),
  inference(resolution,[],[f851,f175])).
fof(f4641,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,sK7(X0,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))))) ) | $spl194),
  inference(cnf_transformation,[],[f4641_D])).
fof(f4641_D,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,sK7(X0,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))))) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f4640,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,sK7(X0,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) | ($spl10 | $spl38)),
  inference(forward_demodulation,[],[f4629,f234])).
fof(f4629,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,sK7(X0,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK1))))) ) | $spl38),
  inference(resolution,[],[f1197,f167])).
fof(f1197,plain,(
  ( ! [X2] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7(X2,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | p(s(bool,X2))) ) | $spl38),
  inference(resolution,[],[f849,f177])).
fof(f4620,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,sK6(X0,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))))) ) | $spl192),
  inference(cnf_transformation,[],[f4620_D])).
fof(f4620_D,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,sK6(X0,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))))) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f4619,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,sK6(X0,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) | ($spl10 | $spl38)),
  inference(forward_demodulation,[],[f4608,f234])).
fof(f4608,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,sK6(X0,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X1))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK1))))) ) | $spl38),
  inference(resolution,[],[f1196,f167])).
fof(f1196,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6(X1,happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | p(s(bool,X1))) ) | $spl38),
  inference(resolution,[],[f849,f175])).
fof(f4598,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl186),
  inference(cnf_transformation,[],[f4598_D])).
fof(f4598_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f4604,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl98),
  inference(global_subsumption,[],[f124,f123,f122,f1338,f4572,f4543,f4540,f4539,f4538,f4508,f4507,f4506,f4505,f4502,f4501,f4580,f4581,f4582])).
fof(f4582,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl98),
  inference(superposition,[],[f1338,f128])).
fof(f4581,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl98),
  inference(resolution,[],[f1338,f130])).
fof(f4580,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl98),
  inference(resolution,[],[f1338,f132])).
fof(f4603,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) | $spl190),
  inference(cnf_transformation,[],[f4603_D])).
fof(f4603_D,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f4602,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) | ($spl10 | $spl98)),
  inference(forward_demodulation,[],[f4573,f234])).
fof(f4573,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK1))))) ) | $spl98),
  inference(resolution,[],[f1338,f167])).
fof(f4601,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | $spl188),
  inference(cnf_transformation,[],[f4601_D])).
fof(f4601_D,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f4501,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | $spl98),
  inference(unit_resulting_resolution,[],[f1338,f1338,f183])).
fof(f183,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.relation.RINTER_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5))))) )),
  inference(cnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,'const.relation.RINTER_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5)))) | (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5)))))) & (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5)))) | p(s(bool,'const.relation.RINTER_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5))))))),
  inference(flattening,[],[f116])).
fof(f116,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,'const.relation.RINTER_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5)))) | (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5)))))) & ((~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5))))) | p(s(bool,'const.relation.RINTER_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5))))))),
  inference(nnf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (p(s(bool,'const.relation.RINTER_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5)))) <=> (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5))))))),
  inference(flattening,[],[f61])).
fof(f61,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : (p(s(bool,'const.relation.RINTER_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5)))) <=> (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5))))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6,X21] : ! [X8,X7,X10,X11] : (p(s(bool,'const.relation.RINTER_4'(s(fun(X6,fun(X21,bool)),X8),s(fun(X6,fun(X21,bool)),X7),s(X6,X10),s(X21,X11)))) <=> (p(s(bool,happ(s(fun(X21,bool),happ(s(fun(X6,fun(X21,bool)),X8),s(X6,X10))),s(X21,X11)))) & p(s(bool,happ(s(fun(X21,bool),happ(s(fun(X6,fun(X21,bool)),X7),s(X6,X10))),s(X21,X11))))))),
  file('thm.relation.transitive_RINTER','thm.relation.RINTER')).
fof(f4502,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | $spl98),
  inference(unit_resulting_resolution,[],[f1338,f1338,f183])).
fof(f4505,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl98),
  inference(unit_resulting_resolution,[],[f1338,f130])).
fof(f4506,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl98),
  inference(unit_resulting_resolution,[],[f1338,f132])).
fof(f4584,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl185),
  inference(cnf_transformation,[],[f4584_D])).
fof(f4584_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) <=> ~$spl185),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl185])])).
fof(f4507,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl98),
  inference(unit_resulting_resolution,[],[f1338,f133])).
fof(f133,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f35])).
fof(f35,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X9] : (s(bool,'F_0') = s(bool,X9) <=> ~p(s(bool,X9)))),
  file('thm.relation.transitive_RINTER','thm.bool.EQ_CLAUSES|split|2')).
fof(f4508,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl98),
  inference(unit_resulting_resolution,[],[f1338,f135])).
fof(f135,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X9] : (s(bool,'F_0') = s(bool,X9) <=> ~p(s(bool,X9)))),
  file('thm.relation.transitive_RINTER','thm.bool.EQ_CLAUSES|split|3')).
fof(f4514,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl18 | $spl98)),
  inference(unit_resulting_resolution,[],[f127,f127,f316,f1338,f191])).
fof(f191,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,X3)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f78])).
fof(f78,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X10,X17,X11,X18] : ((s(bool,X10) = s(bool,X17) & (p(s(bool,X17)) => s(bool,X11) = s(bool,X18))) => ((p(s(bool,X10)) => p(s(bool,X11))) <=> (p(s(bool,X17)) => p(s(bool,X18)))))),
  file('thm.relation.transitive_RINTER','thm.bool.IMP_CONG')).
fof(f127,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f32])).
fof(f32,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.relation.transitive_RINTER',aHL_FALSITY)).
fof(f4515,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl44 | $spl98)),
  inference(unit_resulting_resolution,[],[f127,f127,f926,f1338,f191])).
fof(f4516,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl66 | $spl98)),
  inference(unit_resulting_resolution,[],[f127,f127,f1160,f1338,f191])).
fof(f4517,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl33 | $spl98)),
  inference(unit_resulting_resolution,[],[f127,f127,f1410,f1338,f191])).
fof(f1410,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),X0),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) ) | $spl33),
  inference(unit_resulting_resolution,[],[f1089,f136])).
fof(f136,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f84])).
fof(f1089,plain,(
  ( ! [X0] : (~p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),X0),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) | $spl33),
  inference(unit_resulting_resolution,[],[f749,f181])).
fof(f181,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.relation.RINTER_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X4))),s(X1,X5))))) )),
  inference(cnf_transformation,[],[f117])).
fof(f4518,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl10 | $spl98)),
  inference(unit_resulting_resolution,[],[f127,f126,f234,f1338,f193])).
fof(f193,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f121])).
fof(f126,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.relation.transitive_RINTER','thm.bool.TRUTH')).
fof(f4519,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl10 | $spl98)),
  inference(unit_resulting_resolution,[],[f127,f125,f234,f1338,f193])).
fof(f125,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.relation.transitive_RINTER',aHL_TRUTH)).
fof(f4520,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl14 | $spl98)),
  inference(unit_resulting_resolution,[],[f127,f126,f272,f1338,f193])).
fof(f4521,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl14 | $spl98)),
  inference(unit_resulting_resolution,[],[f127,f125,f272,f1338,f193])).
fof(f4522,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl36 | $spl98)),
  inference(unit_resulting_resolution,[],[f127,f126,f847,f1338,f193])).
fof(f4523,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl36 | $spl98)),
  inference(unit_resulting_resolution,[],[f127,f125,f847,f1338,f193])).
fof(f4524,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl48 | $spl98)),
  inference(unit_resulting_resolution,[],[f127,f126,f1023,f1338,f193])).
fof(f4525,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl48 | $spl98)),
  inference(unit_resulting_resolution,[],[f127,f125,f1023,f1338,f193])).
fof(f4526,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl70 | $spl98)),
  inference(unit_resulting_resolution,[],[f127,f126,f1211,f1338,f193])).
fof(f4527,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl70 | $spl98)),
  inference(unit_resulting_resolution,[],[f127,f125,f1211,f1338,f193])).
fof(f4528,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl82 | $spl98)),
  inference(unit_resulting_resolution,[],[f127,f126,f1273,f1338,f193])).
fof(f4529,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl82 | $spl98)),
  inference(unit_resulting_resolution,[],[f127,f125,f1273,f1338,f193])).
fof(f4530,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl94 | $spl98)),
  inference(unit_resulting_resolution,[],[f127,f126,f1334,f1338,f193])).
fof(f4531,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl94 | $spl98)),
  inference(unit_resulting_resolution,[],[f127,f125,f1334,f1338,f193])).
fof(f4532,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl98 | $spl106)),
  inference(unit_resulting_resolution,[],[f127,f126,f1395,f1338,f193])).
fof(f4533,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl98 | $spl106)),
  inference(unit_resulting_resolution,[],[f127,f125,f1395,f1338,f193])).
fof(f4534,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl18 | $spl98)),
  inference(unit_resulting_resolution,[],[f127,f127,f316,f1338,f195])).
fof(f195,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f121])).
fof(f4535,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl44 | $spl98)),
  inference(unit_resulting_resolution,[],[f127,f127,f926,f1338,f195])).
fof(f4536,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl66 | $spl98)),
  inference(unit_resulting_resolution,[],[f127,f127,f1160,f1338,f195])).
fof(f4537,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl33 | $spl98)),
  inference(unit_resulting_resolution,[],[f127,f127,f1410,f1338,f195])).
fof(f4538,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl98),
  inference(unit_resulting_resolution,[],[f127,f127,f126,f1338,f196])).
fof(f196,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f121])).
fof(f4539,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl98),
  inference(unit_resulting_resolution,[],[f127,f127,f125,f1338,f196])).
fof(f4540,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl98),
  inference(unit_resulting_resolution,[],[f127,f127,f1338,f1338,f196])).
fof(f4543,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl98),
  inference(unit_resulting_resolution,[],[f127,f127,f1338,f1338,f196])).
fof(f4556,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl10 | $spl98)),
  inference(unit_resulting_resolution,[],[f127,f126,f234,f1338,f197])).
fof(f197,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X2)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f121])).
fof(f4557,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl10 | $spl98)),
  inference(unit_resulting_resolution,[],[f127,f125,f234,f1338,f197])).
fof(f4558,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl14 | $spl98)),
  inference(unit_resulting_resolution,[],[f127,f126,f272,f1338,f197])).
fof(f4559,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl14 | $spl98)),
  inference(unit_resulting_resolution,[],[f127,f125,f272,f1338,f197])).
fof(f4560,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl36 | $spl98)),
  inference(unit_resulting_resolution,[],[f127,f126,f847,f1338,f197])).
fof(f4561,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl36 | $spl98)),
  inference(unit_resulting_resolution,[],[f127,f125,f847,f1338,f197])).
fof(f4562,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl48 | $spl98)),
  inference(unit_resulting_resolution,[],[f127,f126,f1023,f1338,f197])).
fof(f4563,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl48 | $spl98)),
  inference(unit_resulting_resolution,[],[f127,f125,f1023,f1338,f197])).
fof(f4564,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl70 | $spl98)),
  inference(unit_resulting_resolution,[],[f127,f126,f1211,f1338,f197])).
fof(f4565,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl70 | $spl98)),
  inference(unit_resulting_resolution,[],[f127,f125,f1211,f1338,f197])).
fof(f4566,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl82 | $spl98)),
  inference(unit_resulting_resolution,[],[f127,f126,f1273,f1338,f197])).
fof(f4567,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl82 | $spl98)),
  inference(unit_resulting_resolution,[],[f127,f125,f1273,f1338,f197])).
fof(f4568,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl94 | $spl98)),
  inference(unit_resulting_resolution,[],[f127,f126,f1334,f1338,f197])).
fof(f4569,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl94 | $spl98)),
  inference(unit_resulting_resolution,[],[f127,f125,f1334,f1338,f197])).
fof(f4570,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl98 | $spl106)),
  inference(unit_resulting_resolution,[],[f127,f126,f1395,f1338,f197])).
fof(f4571,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl98 | $spl106)),
  inference(unit_resulting_resolution,[],[f127,f125,f1395,f1338,f197])).
fof(f4572,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl98),
  inference(unit_resulting_resolution,[],[f127,f127,f1338,f199])).
fof(f199,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(duplicate_literal_removal,[],[f194])).
fof(f194,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f121])).
fof(f4493,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl178),
  inference(cnf_transformation,[],[f4493_D])).
fof(f4493_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f4498,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl96),
  inference(global_subsumption,[],[f124,f123,f122,f1336,f4467,f4438,f4435,f4434,f4433,f4403,f4402,f4401,f4400,f4397,f4396,f4475,f4476,f4477])).
fof(f4477,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl96),
  inference(superposition,[],[f1336,f128])).
fof(f4476,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl96),
  inference(resolution,[],[f1336,f130])).
fof(f4475,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl96),
  inference(resolution,[],[f1336,f132])).
fof(f4497,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) | $spl182),
  inference(cnf_transformation,[],[f4497_D])).
fof(f4497_D,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f4496,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) | ($spl10 | $spl96)),
  inference(forward_demodulation,[],[f4468,f234])).
fof(f4468,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK1))))) ) | $spl96),
  inference(resolution,[],[f1336,f167])).
fof(f4495,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | $spl180),
  inference(cnf_transformation,[],[f4495_D])).
fof(f4495_D,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f4396,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | $spl96),
  inference(unit_resulting_resolution,[],[f1336,f1336,f183])).
fof(f4397,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | $spl96),
  inference(unit_resulting_resolution,[],[f1336,f1336,f183])).
fof(f4400,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl96),
  inference(unit_resulting_resolution,[],[f1336,f130])).
fof(f4401,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl96),
  inference(unit_resulting_resolution,[],[f1336,f132])).
fof(f4479,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl177),
  inference(cnf_transformation,[],[f4479_D])).
fof(f4479_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) <=> ~$spl177),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl177])])).
fof(f4402,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl96),
  inference(unit_resulting_resolution,[],[f1336,f133])).
fof(f4403,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl96),
  inference(unit_resulting_resolution,[],[f1336,f135])).
fof(f4409,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl18 | $spl96)),
  inference(unit_resulting_resolution,[],[f127,f127,f316,f1336,f191])).
fof(f4410,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl44 | $spl96)),
  inference(unit_resulting_resolution,[],[f127,f127,f926,f1336,f191])).
fof(f4411,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl66 | $spl96)),
  inference(unit_resulting_resolution,[],[f127,f127,f1160,f1336,f191])).
fof(f4412,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl33 | $spl96)),
  inference(unit_resulting_resolution,[],[f127,f127,f1410,f1336,f191])).
fof(f4413,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl10 | $spl96)),
  inference(unit_resulting_resolution,[],[f127,f126,f234,f1336,f193])).
fof(f4414,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl10 | $spl96)),
  inference(unit_resulting_resolution,[],[f127,f125,f234,f1336,f193])).
fof(f4415,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl14 | $spl96)),
  inference(unit_resulting_resolution,[],[f127,f126,f272,f1336,f193])).
fof(f4416,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl14 | $spl96)),
  inference(unit_resulting_resolution,[],[f127,f125,f272,f1336,f193])).
fof(f4417,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl36 | $spl96)),
  inference(unit_resulting_resolution,[],[f127,f126,f847,f1336,f193])).
fof(f4418,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl36 | $spl96)),
  inference(unit_resulting_resolution,[],[f127,f125,f847,f1336,f193])).
fof(f4419,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl48 | $spl96)),
  inference(unit_resulting_resolution,[],[f127,f126,f1023,f1336,f193])).
fof(f4420,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl48 | $spl96)),
  inference(unit_resulting_resolution,[],[f127,f125,f1023,f1336,f193])).
fof(f4421,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl70 | $spl96)),
  inference(unit_resulting_resolution,[],[f127,f126,f1211,f1336,f193])).
fof(f4422,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl70 | $spl96)),
  inference(unit_resulting_resolution,[],[f127,f125,f1211,f1336,f193])).
fof(f4423,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl82 | $spl96)),
  inference(unit_resulting_resolution,[],[f127,f126,f1273,f1336,f193])).
fof(f4424,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl82 | $spl96)),
  inference(unit_resulting_resolution,[],[f127,f125,f1273,f1336,f193])).
fof(f4425,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl94 | $spl96)),
  inference(unit_resulting_resolution,[],[f127,f126,f1334,f1336,f193])).
fof(f4426,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl94 | $spl96)),
  inference(unit_resulting_resolution,[],[f127,f125,f1334,f1336,f193])).
fof(f4427,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl96 | $spl106)),
  inference(unit_resulting_resolution,[],[f127,f126,f1395,f1336,f193])).
fof(f4428,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl96 | $spl106)),
  inference(unit_resulting_resolution,[],[f127,f125,f1395,f1336,f193])).
fof(f4429,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl18 | $spl96)),
  inference(unit_resulting_resolution,[],[f127,f127,f316,f1336,f195])).
fof(f4430,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl44 | $spl96)),
  inference(unit_resulting_resolution,[],[f127,f127,f926,f1336,f195])).
fof(f4431,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl66 | $spl96)),
  inference(unit_resulting_resolution,[],[f127,f127,f1160,f1336,f195])).
fof(f4432,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl33 | $spl96)),
  inference(unit_resulting_resolution,[],[f127,f127,f1410,f1336,f195])).
fof(f4433,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl96),
  inference(unit_resulting_resolution,[],[f127,f127,f126,f1336,f196])).
fof(f4434,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl96),
  inference(unit_resulting_resolution,[],[f127,f127,f125,f1336,f196])).
fof(f4435,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl96),
  inference(unit_resulting_resolution,[],[f127,f127,f1336,f1336,f196])).
fof(f4438,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl96),
  inference(unit_resulting_resolution,[],[f127,f127,f1336,f1336,f196])).
fof(f4451,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl10 | $spl96)),
  inference(unit_resulting_resolution,[],[f127,f126,f234,f1336,f197])).
fof(f4452,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl10 | $spl96)),
  inference(unit_resulting_resolution,[],[f127,f125,f234,f1336,f197])).
fof(f4453,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl14 | $spl96)),
  inference(unit_resulting_resolution,[],[f127,f126,f272,f1336,f197])).
fof(f4454,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl14 | $spl96)),
  inference(unit_resulting_resolution,[],[f127,f125,f272,f1336,f197])).
fof(f4455,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl36 | $spl96)),
  inference(unit_resulting_resolution,[],[f127,f126,f847,f1336,f197])).
fof(f4456,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl36 | $spl96)),
  inference(unit_resulting_resolution,[],[f127,f125,f847,f1336,f197])).
fof(f4457,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl48 | $spl96)),
  inference(unit_resulting_resolution,[],[f127,f126,f1023,f1336,f197])).
fof(f4458,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl48 | $spl96)),
  inference(unit_resulting_resolution,[],[f127,f125,f1023,f1336,f197])).
fof(f4459,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl70 | $spl96)),
  inference(unit_resulting_resolution,[],[f127,f126,f1211,f1336,f197])).
fof(f4460,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl70 | $spl96)),
  inference(unit_resulting_resolution,[],[f127,f125,f1211,f1336,f197])).
fof(f4461,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl82 | $spl96)),
  inference(unit_resulting_resolution,[],[f127,f126,f1273,f1336,f197])).
fof(f4462,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl82 | $spl96)),
  inference(unit_resulting_resolution,[],[f127,f125,f1273,f1336,f197])).
fof(f4463,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl94 | $spl96)),
  inference(unit_resulting_resolution,[],[f127,f126,f1334,f1336,f197])).
fof(f4464,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl94 | $spl96)),
  inference(unit_resulting_resolution,[],[f127,f125,f1334,f1336,f197])).
fof(f4465,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl96 | $spl106)),
  inference(unit_resulting_resolution,[],[f127,f126,f1395,f1336,f197])).
fof(f4466,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl96 | $spl106)),
  inference(unit_resulting_resolution,[],[f127,f125,f1395,f1336,f197])).
fof(f4467,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl96),
  inference(unit_resulting_resolution,[],[f127,f127,f1336,f199])).
fof(f4387,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl170),
  inference(cnf_transformation,[],[f4387_D])).
fof(f4387_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f4393,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl86),
  inference(global_subsumption,[],[f124,f123,f122,f1277,f4361,f4332,f4329,f4328,f4327,f4297,f4296,f4295,f4294,f4291,f4290,f4369,f4370,f4371])).
fof(f4371,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl86),
  inference(superposition,[],[f1277,f128])).
fof(f4370,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl86),
  inference(resolution,[],[f1277,f130])).
fof(f4369,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl86),
  inference(resolution,[],[f1277,f132])).
fof(f4392,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) | $spl174),
  inference(cnf_transformation,[],[f4392_D])).
fof(f4392_D,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f4391,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) | ($spl14 | $spl86)),
  inference(forward_demodulation,[],[f4362,f272])).
fof(f4362,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK2))))) ) | $spl86),
  inference(resolution,[],[f1277,f167])).
fof(f4390,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | $spl172),
  inference(cnf_transformation,[],[f4390_D])).
fof(f4390_D,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f4290,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | $spl86),
  inference(unit_resulting_resolution,[],[f1277,f1277,f183])).
fof(f4291,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | $spl86),
  inference(unit_resulting_resolution,[],[f1277,f1277,f183])).
fof(f4294,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl86),
  inference(unit_resulting_resolution,[],[f1277,f130])).
fof(f4295,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl86),
  inference(unit_resulting_resolution,[],[f1277,f132])).
fof(f4373,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl169),
  inference(cnf_transformation,[],[f4373_D])).
fof(f4373_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) <=> ~$spl169),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl169])])).
fof(f4296,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl86),
  inference(unit_resulting_resolution,[],[f1277,f133])).
fof(f4297,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl86),
  inference(unit_resulting_resolution,[],[f1277,f135])).
fof(f4303,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl18 | $spl86)),
  inference(unit_resulting_resolution,[],[f127,f127,f316,f1277,f191])).
fof(f4304,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl44 | $spl86)),
  inference(unit_resulting_resolution,[],[f127,f127,f926,f1277,f191])).
fof(f4305,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl66 | $spl86)),
  inference(unit_resulting_resolution,[],[f127,f127,f1160,f1277,f191])).
fof(f4306,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl33 | $spl86)),
  inference(unit_resulting_resolution,[],[f127,f127,f1410,f1277,f191])).
fof(f4307,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl10 | $spl86)),
  inference(unit_resulting_resolution,[],[f127,f126,f234,f1277,f193])).
fof(f4308,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl10 | $spl86)),
  inference(unit_resulting_resolution,[],[f127,f125,f234,f1277,f193])).
fof(f4309,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl14 | $spl86)),
  inference(unit_resulting_resolution,[],[f127,f126,f272,f1277,f193])).
fof(f4310,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl14 | $spl86)),
  inference(unit_resulting_resolution,[],[f127,f125,f272,f1277,f193])).
fof(f4311,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl36 | $spl86)),
  inference(unit_resulting_resolution,[],[f127,f126,f847,f1277,f193])).
fof(f4312,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl36 | $spl86)),
  inference(unit_resulting_resolution,[],[f127,f125,f847,f1277,f193])).
fof(f4313,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl48 | $spl86)),
  inference(unit_resulting_resolution,[],[f127,f126,f1023,f1277,f193])).
fof(f4314,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl48 | $spl86)),
  inference(unit_resulting_resolution,[],[f127,f125,f1023,f1277,f193])).
fof(f4315,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl70 | $spl86)),
  inference(unit_resulting_resolution,[],[f127,f126,f1211,f1277,f193])).
fof(f4316,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl70 | $spl86)),
  inference(unit_resulting_resolution,[],[f127,f125,f1211,f1277,f193])).
fof(f4317,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl82 | $spl86)),
  inference(unit_resulting_resolution,[],[f127,f126,f1273,f1277,f193])).
fof(f4318,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl82 | $spl86)),
  inference(unit_resulting_resolution,[],[f127,f125,f1273,f1277,f193])).
fof(f4319,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl86 | $spl94)),
  inference(unit_resulting_resolution,[],[f127,f126,f1334,f1277,f193])).
fof(f4320,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl86 | $spl94)),
  inference(unit_resulting_resolution,[],[f127,f125,f1334,f1277,f193])).
fof(f4321,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl86 | $spl106)),
  inference(unit_resulting_resolution,[],[f127,f126,f1395,f1277,f193])).
fof(f4322,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl86 | $spl106)),
  inference(unit_resulting_resolution,[],[f127,f125,f1395,f1277,f193])).
fof(f4323,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl18 | $spl86)),
  inference(unit_resulting_resolution,[],[f127,f127,f316,f1277,f195])).
fof(f4324,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl44 | $spl86)),
  inference(unit_resulting_resolution,[],[f127,f127,f926,f1277,f195])).
fof(f4325,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl66 | $spl86)),
  inference(unit_resulting_resolution,[],[f127,f127,f1160,f1277,f195])).
fof(f4326,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl33 | $spl86)),
  inference(unit_resulting_resolution,[],[f127,f127,f1410,f1277,f195])).
fof(f4327,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl86),
  inference(unit_resulting_resolution,[],[f127,f127,f126,f1277,f196])).
fof(f4328,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl86),
  inference(unit_resulting_resolution,[],[f127,f127,f125,f1277,f196])).
fof(f4329,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl86),
  inference(unit_resulting_resolution,[],[f127,f127,f1277,f1277,f196])).
fof(f4332,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl86),
  inference(unit_resulting_resolution,[],[f127,f127,f1277,f1277,f196])).
fof(f4345,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl10 | $spl86)),
  inference(unit_resulting_resolution,[],[f127,f126,f234,f1277,f197])).
fof(f4346,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl10 | $spl86)),
  inference(unit_resulting_resolution,[],[f127,f125,f234,f1277,f197])).
fof(f4347,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl14 | $spl86)),
  inference(unit_resulting_resolution,[],[f127,f126,f272,f1277,f197])).
fof(f4348,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl14 | $spl86)),
  inference(unit_resulting_resolution,[],[f127,f125,f272,f1277,f197])).
fof(f4349,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl36 | $spl86)),
  inference(unit_resulting_resolution,[],[f127,f126,f847,f1277,f197])).
fof(f4350,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl36 | $spl86)),
  inference(unit_resulting_resolution,[],[f127,f125,f847,f1277,f197])).
fof(f4351,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl48 | $spl86)),
  inference(unit_resulting_resolution,[],[f127,f126,f1023,f1277,f197])).
fof(f4352,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl48 | $spl86)),
  inference(unit_resulting_resolution,[],[f127,f125,f1023,f1277,f197])).
fof(f4353,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl70 | $spl86)),
  inference(unit_resulting_resolution,[],[f127,f126,f1211,f1277,f197])).
fof(f4354,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl70 | $spl86)),
  inference(unit_resulting_resolution,[],[f127,f125,f1211,f1277,f197])).
fof(f4355,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl82 | $spl86)),
  inference(unit_resulting_resolution,[],[f127,f126,f1273,f1277,f197])).
fof(f4356,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl82 | $spl86)),
  inference(unit_resulting_resolution,[],[f127,f125,f1273,f1277,f197])).
fof(f4357,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl86 | $spl94)),
  inference(unit_resulting_resolution,[],[f127,f126,f1334,f1277,f197])).
fof(f4358,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl86 | $spl94)),
  inference(unit_resulting_resolution,[],[f127,f125,f1334,f1277,f197])).
fof(f4359,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl86 | $spl106)),
  inference(unit_resulting_resolution,[],[f127,f126,f1395,f1277,f197])).
fof(f4360,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl86 | $spl106)),
  inference(unit_resulting_resolution,[],[f127,f125,f1395,f1277,f197])).
fof(f4361,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl86),
  inference(unit_resulting_resolution,[],[f127,f127,f1277,f199])).
fof(f4282,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl162),
  inference(cnf_transformation,[],[f4282_D])).
fof(f4282_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f4287,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl84),
  inference(global_subsumption,[],[f124,f123,f122,f1275,f4256,f4227,f4224,f4223,f4222,f4192,f4191,f4190,f4189,f4186,f4185,f4264,f4265,f4266])).
fof(f4266,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl84),
  inference(superposition,[],[f1275,f128])).
fof(f4265,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl84),
  inference(resolution,[],[f1275,f130])).
fof(f4264,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl84),
  inference(resolution,[],[f1275,f132])).
fof(f4286,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) | $spl166),
  inference(cnf_transformation,[],[f4286_D])).
fof(f4286_D,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f4285,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) | ($spl14 | $spl84)),
  inference(forward_demodulation,[],[f4257,f272])).
fof(f4257,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK2))))) ) | $spl84),
  inference(resolution,[],[f1275,f167])).
fof(f4284,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | $spl164),
  inference(cnf_transformation,[],[f4284_D])).
fof(f4284_D,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f4185,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | $spl84),
  inference(unit_resulting_resolution,[],[f1275,f1275,f183])).
fof(f4186,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | $spl84),
  inference(unit_resulting_resolution,[],[f1275,f1275,f183])).
fof(f4189,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl84),
  inference(unit_resulting_resolution,[],[f1275,f130])).
fof(f4190,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl84),
  inference(unit_resulting_resolution,[],[f1275,f132])).
fof(f4268,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl161),
  inference(cnf_transformation,[],[f4268_D])).
fof(f4268_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) <=> ~$spl161),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl161])])).
fof(f4191,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl84),
  inference(unit_resulting_resolution,[],[f1275,f133])).
fof(f4192,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl84),
  inference(unit_resulting_resolution,[],[f1275,f135])).
fof(f4198,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl18 | $spl84)),
  inference(unit_resulting_resolution,[],[f127,f127,f316,f1275,f191])).
fof(f4199,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl44 | $spl84)),
  inference(unit_resulting_resolution,[],[f127,f127,f926,f1275,f191])).
fof(f4200,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl66 | $spl84)),
  inference(unit_resulting_resolution,[],[f127,f127,f1160,f1275,f191])).
fof(f4201,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl33 | $spl84)),
  inference(unit_resulting_resolution,[],[f127,f127,f1410,f1275,f191])).
fof(f4202,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl10 | $spl84)),
  inference(unit_resulting_resolution,[],[f127,f126,f234,f1275,f193])).
fof(f4203,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl10 | $spl84)),
  inference(unit_resulting_resolution,[],[f127,f125,f234,f1275,f193])).
fof(f4204,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl14 | $spl84)),
  inference(unit_resulting_resolution,[],[f127,f126,f272,f1275,f193])).
fof(f4205,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl14 | $spl84)),
  inference(unit_resulting_resolution,[],[f127,f125,f272,f1275,f193])).
fof(f4206,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl36 | $spl84)),
  inference(unit_resulting_resolution,[],[f127,f126,f847,f1275,f193])).
fof(f4207,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl36 | $spl84)),
  inference(unit_resulting_resolution,[],[f127,f125,f847,f1275,f193])).
fof(f4208,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl48 | $spl84)),
  inference(unit_resulting_resolution,[],[f127,f126,f1023,f1275,f193])).
fof(f4209,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl48 | $spl84)),
  inference(unit_resulting_resolution,[],[f127,f125,f1023,f1275,f193])).
fof(f4210,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl70 | $spl84)),
  inference(unit_resulting_resolution,[],[f127,f126,f1211,f1275,f193])).
fof(f4211,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl70 | $spl84)),
  inference(unit_resulting_resolution,[],[f127,f125,f1211,f1275,f193])).
fof(f4212,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl82 | $spl84)),
  inference(unit_resulting_resolution,[],[f127,f126,f1273,f1275,f193])).
fof(f4213,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl82 | $spl84)),
  inference(unit_resulting_resolution,[],[f127,f125,f1273,f1275,f193])).
fof(f4214,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl84 | $spl94)),
  inference(unit_resulting_resolution,[],[f127,f126,f1334,f1275,f193])).
fof(f4215,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl84 | $spl94)),
  inference(unit_resulting_resolution,[],[f127,f125,f1334,f1275,f193])).
fof(f4216,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl84 | $spl106)),
  inference(unit_resulting_resolution,[],[f127,f126,f1395,f1275,f193])).
fof(f4217,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl84 | $spl106)),
  inference(unit_resulting_resolution,[],[f127,f125,f1395,f1275,f193])).
fof(f4218,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl18 | $spl84)),
  inference(unit_resulting_resolution,[],[f127,f127,f316,f1275,f195])).
fof(f4219,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl44 | $spl84)),
  inference(unit_resulting_resolution,[],[f127,f127,f926,f1275,f195])).
fof(f4220,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl66 | $spl84)),
  inference(unit_resulting_resolution,[],[f127,f127,f1160,f1275,f195])).
fof(f4221,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl33 | $spl84)),
  inference(unit_resulting_resolution,[],[f127,f127,f1410,f1275,f195])).
fof(f4222,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl84),
  inference(unit_resulting_resolution,[],[f127,f127,f126,f1275,f196])).
fof(f4223,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl84),
  inference(unit_resulting_resolution,[],[f127,f127,f125,f1275,f196])).
fof(f4224,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl84),
  inference(unit_resulting_resolution,[],[f127,f127,f1275,f1275,f196])).
fof(f4227,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl84),
  inference(unit_resulting_resolution,[],[f127,f127,f1275,f1275,f196])).
fof(f4240,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl10 | $spl84)),
  inference(unit_resulting_resolution,[],[f127,f126,f234,f1275,f197])).
fof(f4241,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl10 | $spl84)),
  inference(unit_resulting_resolution,[],[f127,f125,f234,f1275,f197])).
fof(f4242,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl14 | $spl84)),
  inference(unit_resulting_resolution,[],[f127,f126,f272,f1275,f197])).
fof(f4243,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl14 | $spl84)),
  inference(unit_resulting_resolution,[],[f127,f125,f272,f1275,f197])).
fof(f4244,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl36 | $spl84)),
  inference(unit_resulting_resolution,[],[f127,f126,f847,f1275,f197])).
fof(f4245,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl36 | $spl84)),
  inference(unit_resulting_resolution,[],[f127,f125,f847,f1275,f197])).
fof(f4246,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl48 | $spl84)),
  inference(unit_resulting_resolution,[],[f127,f126,f1023,f1275,f197])).
fof(f4247,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl48 | $spl84)),
  inference(unit_resulting_resolution,[],[f127,f125,f1023,f1275,f197])).
fof(f4248,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl70 | $spl84)),
  inference(unit_resulting_resolution,[],[f127,f126,f1211,f1275,f197])).
fof(f4249,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl70 | $spl84)),
  inference(unit_resulting_resolution,[],[f127,f125,f1211,f1275,f197])).
fof(f4250,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl82 | $spl84)),
  inference(unit_resulting_resolution,[],[f127,f126,f1273,f1275,f197])).
fof(f4251,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl82 | $spl84)),
  inference(unit_resulting_resolution,[],[f127,f125,f1273,f1275,f197])).
fof(f4252,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl84 | $spl94)),
  inference(unit_resulting_resolution,[],[f127,f126,f1334,f1275,f197])).
fof(f4253,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl84 | $spl94)),
  inference(unit_resulting_resolution,[],[f127,f125,f1334,f1275,f197])).
fof(f4254,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl84 | $spl106)),
  inference(unit_resulting_resolution,[],[f127,f126,f1395,f1275,f197])).
fof(f4255,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl84 | $spl106)),
  inference(unit_resulting_resolution,[],[f127,f125,f1395,f1275,f197])).
fof(f4256,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl84),
  inference(unit_resulting_resolution,[],[f127,f127,f1275,f199])).
fof(f4176,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl154),
  inference(cnf_transformation,[],[f4176_D])).
fof(f4176_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f4182,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl74),
  inference(global_subsumption,[],[f124,f123,f122,f1215,f4150,f4121,f4118,f4117,f4116,f4086,f4085,f4084,f4083,f4080,f4079,f4158,f4159,f4160])).
fof(f4160,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl74),
  inference(superposition,[],[f1215,f128])).
fof(f4159,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl74),
  inference(resolution,[],[f1215,f130])).
fof(f4158,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl74),
  inference(resolution,[],[f1215,f132])).
fof(f4181,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) | $spl158),
  inference(cnf_transformation,[],[f4181_D])).
fof(f4181_D,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f4180,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) | ($spl10 | $spl74)),
  inference(forward_demodulation,[],[f4151,f234])).
fof(f4151,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK1))))) ) | $spl74),
  inference(resolution,[],[f1215,f167])).
fof(f4179,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | $spl156),
  inference(cnf_transformation,[],[f4179_D])).
fof(f4179_D,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f4079,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | $spl74),
  inference(unit_resulting_resolution,[],[f1215,f1215,f183])).
fof(f4080,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | $spl74),
  inference(unit_resulting_resolution,[],[f1215,f1215,f183])).
fof(f4083,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl74),
  inference(unit_resulting_resolution,[],[f1215,f130])).
fof(f4084,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl74),
  inference(unit_resulting_resolution,[],[f1215,f132])).
fof(f4162,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl153),
  inference(cnf_transformation,[],[f4162_D])).
fof(f4162_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) <=> ~$spl153),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl153])])).
fof(f4085,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl74),
  inference(unit_resulting_resolution,[],[f1215,f133])).
fof(f4086,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl74),
  inference(unit_resulting_resolution,[],[f1215,f135])).
fof(f4092,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl18 | $spl74)),
  inference(unit_resulting_resolution,[],[f127,f127,f316,f1215,f191])).
fof(f4093,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl44 | $spl74)),
  inference(unit_resulting_resolution,[],[f127,f127,f926,f1215,f191])).
fof(f4094,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl66 | $spl74)),
  inference(unit_resulting_resolution,[],[f127,f127,f1160,f1215,f191])).
fof(f4095,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl33 | $spl74)),
  inference(unit_resulting_resolution,[],[f127,f127,f1410,f1215,f191])).
fof(f4096,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl10 | $spl74)),
  inference(unit_resulting_resolution,[],[f127,f126,f234,f1215,f193])).
fof(f4097,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl10 | $spl74)),
  inference(unit_resulting_resolution,[],[f127,f125,f234,f1215,f193])).
fof(f4098,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl14 | $spl74)),
  inference(unit_resulting_resolution,[],[f127,f126,f272,f1215,f193])).
fof(f4099,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl14 | $spl74)),
  inference(unit_resulting_resolution,[],[f127,f125,f272,f1215,f193])).
fof(f4100,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl36 | $spl74)),
  inference(unit_resulting_resolution,[],[f127,f126,f847,f1215,f193])).
fof(f4101,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl36 | $spl74)),
  inference(unit_resulting_resolution,[],[f127,f125,f847,f1215,f193])).
fof(f4102,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl48 | $spl74)),
  inference(unit_resulting_resolution,[],[f127,f126,f1023,f1215,f193])).
fof(f4103,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl48 | $spl74)),
  inference(unit_resulting_resolution,[],[f127,f125,f1023,f1215,f193])).
fof(f4104,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl70 | $spl74)),
  inference(unit_resulting_resolution,[],[f127,f126,f1211,f1215,f193])).
fof(f4105,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl70 | $spl74)),
  inference(unit_resulting_resolution,[],[f127,f125,f1211,f1215,f193])).
fof(f4106,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl74 | $spl82)),
  inference(unit_resulting_resolution,[],[f127,f126,f1273,f1215,f193])).
fof(f4107,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl74 | $spl82)),
  inference(unit_resulting_resolution,[],[f127,f125,f1273,f1215,f193])).
fof(f4108,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl74 | $spl94)),
  inference(unit_resulting_resolution,[],[f127,f126,f1334,f1215,f193])).
fof(f4109,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl74 | $spl94)),
  inference(unit_resulting_resolution,[],[f127,f125,f1334,f1215,f193])).
fof(f4110,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl74 | $spl106)),
  inference(unit_resulting_resolution,[],[f127,f126,f1395,f1215,f193])).
fof(f4111,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl74 | $spl106)),
  inference(unit_resulting_resolution,[],[f127,f125,f1395,f1215,f193])).
fof(f4112,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl18 | $spl74)),
  inference(unit_resulting_resolution,[],[f127,f127,f316,f1215,f195])).
fof(f4113,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl44 | $spl74)),
  inference(unit_resulting_resolution,[],[f127,f127,f926,f1215,f195])).
fof(f4114,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl66 | $spl74)),
  inference(unit_resulting_resolution,[],[f127,f127,f1160,f1215,f195])).
fof(f4115,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl33 | $spl74)),
  inference(unit_resulting_resolution,[],[f127,f127,f1410,f1215,f195])).
fof(f4116,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl74),
  inference(unit_resulting_resolution,[],[f127,f127,f126,f1215,f196])).
fof(f4117,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl74),
  inference(unit_resulting_resolution,[],[f127,f127,f125,f1215,f196])).
fof(f4118,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl74),
  inference(unit_resulting_resolution,[],[f127,f127,f1215,f1215,f196])).
fof(f4121,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl74),
  inference(unit_resulting_resolution,[],[f127,f127,f1215,f1215,f196])).
fof(f4134,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl10 | $spl74)),
  inference(unit_resulting_resolution,[],[f127,f126,f234,f1215,f197])).
fof(f4135,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl10 | $spl74)),
  inference(unit_resulting_resolution,[],[f127,f125,f234,f1215,f197])).
fof(f4136,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl14 | $spl74)),
  inference(unit_resulting_resolution,[],[f127,f126,f272,f1215,f197])).
fof(f4137,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl14 | $spl74)),
  inference(unit_resulting_resolution,[],[f127,f125,f272,f1215,f197])).
fof(f4138,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl36 | $spl74)),
  inference(unit_resulting_resolution,[],[f127,f126,f847,f1215,f197])).
fof(f4139,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl36 | $spl74)),
  inference(unit_resulting_resolution,[],[f127,f125,f847,f1215,f197])).
fof(f4140,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl48 | $spl74)),
  inference(unit_resulting_resolution,[],[f127,f126,f1023,f1215,f197])).
fof(f4141,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl48 | $spl74)),
  inference(unit_resulting_resolution,[],[f127,f125,f1023,f1215,f197])).
fof(f4142,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl70 | $spl74)),
  inference(unit_resulting_resolution,[],[f127,f126,f1211,f1215,f197])).
fof(f4143,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl70 | $spl74)),
  inference(unit_resulting_resolution,[],[f127,f125,f1211,f1215,f197])).
fof(f4144,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl74 | $spl82)),
  inference(unit_resulting_resolution,[],[f127,f126,f1273,f1215,f197])).
fof(f4145,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl74 | $spl82)),
  inference(unit_resulting_resolution,[],[f127,f125,f1273,f1215,f197])).
fof(f4146,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl74 | $spl94)),
  inference(unit_resulting_resolution,[],[f127,f126,f1334,f1215,f197])).
fof(f4147,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl74 | $spl94)),
  inference(unit_resulting_resolution,[],[f127,f125,f1334,f1215,f197])).
fof(f4148,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl74 | $spl106)),
  inference(unit_resulting_resolution,[],[f127,f126,f1395,f1215,f197])).
fof(f4149,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl74 | $spl106)),
  inference(unit_resulting_resolution,[],[f127,f125,f1395,f1215,f197])).
fof(f4150,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl74),
  inference(unit_resulting_resolution,[],[f127,f127,f1215,f199])).
fof(f4071,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl146),
  inference(cnf_transformation,[],[f4071_D])).
fof(f4071_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f4076,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl72),
  inference(global_subsumption,[],[f124,f123,f122,f1213,f4045,f4016,f4013,f4012,f4011,f3981,f3980,f3979,f3978,f3975,f3974,f4053,f4054,f4055])).
fof(f4055,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl72),
  inference(superposition,[],[f1213,f128])).
fof(f4054,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl72),
  inference(resolution,[],[f1213,f130])).
fof(f4053,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl72),
  inference(resolution,[],[f1213,f132])).
fof(f4075,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) | $spl150),
  inference(cnf_transformation,[],[f4075_D])).
fof(f4075_D,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f4074,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) | ($spl10 | $spl72)),
  inference(forward_demodulation,[],[f4046,f234])).
fof(f4046,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK1))))) ) | $spl72),
  inference(resolution,[],[f1213,f167])).
fof(f4073,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | $spl148),
  inference(cnf_transformation,[],[f4073_D])).
fof(f4073_D,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f3974,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | $spl72),
  inference(unit_resulting_resolution,[],[f1213,f1213,f183])).
fof(f3975,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | $spl72),
  inference(unit_resulting_resolution,[],[f1213,f1213,f183])).
fof(f3978,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl72),
  inference(unit_resulting_resolution,[],[f1213,f130])).
fof(f3979,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl72),
  inference(unit_resulting_resolution,[],[f1213,f132])).
fof(f4057,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl145),
  inference(cnf_transformation,[],[f4057_D])).
fof(f4057_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) <=> ~$spl145),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl145])])).
fof(f3980,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl72),
  inference(unit_resulting_resolution,[],[f1213,f133])).
fof(f3981,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl72),
  inference(unit_resulting_resolution,[],[f1213,f135])).
fof(f3987,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl18 | $spl72)),
  inference(unit_resulting_resolution,[],[f127,f127,f316,f1213,f191])).
fof(f3988,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl44 | $spl72)),
  inference(unit_resulting_resolution,[],[f127,f127,f926,f1213,f191])).
fof(f3989,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl66 | $spl72)),
  inference(unit_resulting_resolution,[],[f127,f127,f1160,f1213,f191])).
fof(f3990,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl33 | $spl72)),
  inference(unit_resulting_resolution,[],[f127,f127,f1410,f1213,f191])).
fof(f3991,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl10 | $spl72)),
  inference(unit_resulting_resolution,[],[f127,f126,f234,f1213,f193])).
fof(f3992,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl10 | $spl72)),
  inference(unit_resulting_resolution,[],[f127,f125,f234,f1213,f193])).
fof(f3993,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl14 | $spl72)),
  inference(unit_resulting_resolution,[],[f127,f126,f272,f1213,f193])).
fof(f3994,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl14 | $spl72)),
  inference(unit_resulting_resolution,[],[f127,f125,f272,f1213,f193])).
fof(f3995,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl36 | $spl72)),
  inference(unit_resulting_resolution,[],[f127,f126,f847,f1213,f193])).
fof(f3996,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl36 | $spl72)),
  inference(unit_resulting_resolution,[],[f127,f125,f847,f1213,f193])).
fof(f3997,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl48 | $spl72)),
  inference(unit_resulting_resolution,[],[f127,f126,f1023,f1213,f193])).
fof(f3998,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl48 | $spl72)),
  inference(unit_resulting_resolution,[],[f127,f125,f1023,f1213,f193])).
fof(f3999,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl70 | $spl72)),
  inference(unit_resulting_resolution,[],[f127,f126,f1211,f1213,f193])).
fof(f4000,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl70 | $spl72)),
  inference(unit_resulting_resolution,[],[f127,f125,f1211,f1213,f193])).
fof(f4001,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl72 | $spl82)),
  inference(unit_resulting_resolution,[],[f127,f126,f1273,f1213,f193])).
fof(f4002,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl72 | $spl82)),
  inference(unit_resulting_resolution,[],[f127,f125,f1273,f1213,f193])).
fof(f4003,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl72 | $spl94)),
  inference(unit_resulting_resolution,[],[f127,f126,f1334,f1213,f193])).
fof(f4004,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl72 | $spl94)),
  inference(unit_resulting_resolution,[],[f127,f125,f1334,f1213,f193])).
fof(f4005,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl72 | $spl106)),
  inference(unit_resulting_resolution,[],[f127,f126,f1395,f1213,f193])).
fof(f4006,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl72 | $spl106)),
  inference(unit_resulting_resolution,[],[f127,f125,f1395,f1213,f193])).
fof(f4007,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl18 | $spl72)),
  inference(unit_resulting_resolution,[],[f127,f127,f316,f1213,f195])).
fof(f4008,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl44 | $spl72)),
  inference(unit_resulting_resolution,[],[f127,f127,f926,f1213,f195])).
fof(f4009,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl66 | $spl72)),
  inference(unit_resulting_resolution,[],[f127,f127,f1160,f1213,f195])).
fof(f4010,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl33 | $spl72)),
  inference(unit_resulting_resolution,[],[f127,f127,f1410,f1213,f195])).
fof(f4011,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl72),
  inference(unit_resulting_resolution,[],[f127,f127,f126,f1213,f196])).
fof(f4012,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl72),
  inference(unit_resulting_resolution,[],[f127,f127,f125,f1213,f196])).
fof(f4013,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl72),
  inference(unit_resulting_resolution,[],[f127,f127,f1213,f1213,f196])).
fof(f4016,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl72),
  inference(unit_resulting_resolution,[],[f127,f127,f1213,f1213,f196])).
fof(f4029,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl10 | $spl72)),
  inference(unit_resulting_resolution,[],[f127,f126,f234,f1213,f197])).
fof(f4030,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl10 | $spl72)),
  inference(unit_resulting_resolution,[],[f127,f125,f234,f1213,f197])).
fof(f4031,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl14 | $spl72)),
  inference(unit_resulting_resolution,[],[f127,f126,f272,f1213,f197])).
fof(f4032,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl14 | $spl72)),
  inference(unit_resulting_resolution,[],[f127,f125,f272,f1213,f197])).
fof(f4033,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl36 | $spl72)),
  inference(unit_resulting_resolution,[],[f127,f126,f847,f1213,f197])).
fof(f4034,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl36 | $spl72)),
  inference(unit_resulting_resolution,[],[f127,f125,f847,f1213,f197])).
fof(f4035,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl48 | $spl72)),
  inference(unit_resulting_resolution,[],[f127,f126,f1023,f1213,f197])).
fof(f4036,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl48 | $spl72)),
  inference(unit_resulting_resolution,[],[f127,f125,f1023,f1213,f197])).
fof(f4037,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl70 | $spl72)),
  inference(unit_resulting_resolution,[],[f127,f126,f1211,f1213,f197])).
fof(f4038,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl70 | $spl72)),
  inference(unit_resulting_resolution,[],[f127,f125,f1211,f1213,f197])).
fof(f4039,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl72 | $spl82)),
  inference(unit_resulting_resolution,[],[f127,f126,f1273,f1213,f197])).
fof(f4040,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl72 | $spl82)),
  inference(unit_resulting_resolution,[],[f127,f125,f1273,f1213,f197])).
fof(f4041,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl72 | $spl94)),
  inference(unit_resulting_resolution,[],[f127,f126,f1334,f1213,f197])).
fof(f4042,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl72 | $spl94)),
  inference(unit_resulting_resolution,[],[f127,f125,f1334,f1213,f197])).
fof(f4043,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl72 | $spl106)),
  inference(unit_resulting_resolution,[],[f127,f126,f1395,f1213,f197])).
fof(f4044,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl72 | $spl106)),
  inference(unit_resulting_resolution,[],[f127,f125,f1395,f1213,f197])).
fof(f4045,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl72),
  inference(unit_resulting_resolution,[],[f127,f127,f1213,f199])).
fof(f3965,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl138),
  inference(cnf_transformation,[],[f3965_D])).
fof(f3965_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f3971,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl56),
  inference(global_subsumption,[],[f124,f123,f122,f1080,f3939,f3910,f3907,f3906,f3905,f3875,f3874,f3873,f3872,f3869,f3868,f3947,f3948,f3949])).
fof(f3949,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl56),
  inference(superposition,[],[f1080,f128])).
fof(f3948,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl56),
  inference(resolution,[],[f1080,f130])).
fof(f3947,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl56),
  inference(resolution,[],[f1080,f132])).
fof(f3970,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) | $spl142),
  inference(cnf_transformation,[],[f3970_D])).
fof(f3970_D,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f3969,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) | ($spl14 | $spl56)),
  inference(forward_demodulation,[],[f3940,f272])).
fof(f3940,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK2))))) ) | $spl56),
  inference(resolution,[],[f1080,f167])).
fof(f3968,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | $spl140),
  inference(cnf_transformation,[],[f3968_D])).
fof(f3968_D,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f3868,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | $spl56),
  inference(unit_resulting_resolution,[],[f1080,f1080,f183])).
fof(f3869,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | $spl56),
  inference(unit_resulting_resolution,[],[f1080,f1080,f183])).
fof(f3872,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl56),
  inference(unit_resulting_resolution,[],[f1080,f130])).
fof(f3873,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl56),
  inference(unit_resulting_resolution,[],[f1080,f132])).
fof(f3951,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl137),
  inference(cnf_transformation,[],[f3951_D])).
fof(f3951_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) <=> ~$spl137),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl137])])).
fof(f3874,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl56),
  inference(unit_resulting_resolution,[],[f1080,f133])).
fof(f3875,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl56),
  inference(unit_resulting_resolution,[],[f1080,f135])).
fof(f3881,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl18 | $spl56)),
  inference(unit_resulting_resolution,[],[f127,f127,f316,f1080,f191])).
fof(f3882,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl44 | $spl56)),
  inference(unit_resulting_resolution,[],[f127,f127,f926,f1080,f191])).
fof(f3883,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl56 | $spl66)),
  inference(unit_resulting_resolution,[],[f127,f127,f1160,f1080,f191])).
fof(f3884,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl33 | $spl56)),
  inference(unit_resulting_resolution,[],[f127,f127,f1410,f1080,f191])).
fof(f3885,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl10 | $spl56)),
  inference(unit_resulting_resolution,[],[f127,f126,f234,f1080,f193])).
fof(f3886,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl10 | $spl56)),
  inference(unit_resulting_resolution,[],[f127,f125,f234,f1080,f193])).
fof(f3887,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl14 | $spl56)),
  inference(unit_resulting_resolution,[],[f127,f126,f272,f1080,f193])).
fof(f3888,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl14 | $spl56)),
  inference(unit_resulting_resolution,[],[f127,f125,f272,f1080,f193])).
fof(f3889,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl36 | $spl56)),
  inference(unit_resulting_resolution,[],[f127,f126,f847,f1080,f193])).
fof(f3890,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl36 | $spl56)),
  inference(unit_resulting_resolution,[],[f127,f125,f847,f1080,f193])).
fof(f3891,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl48 | $spl56)),
  inference(unit_resulting_resolution,[],[f127,f126,f1023,f1080,f193])).
fof(f3892,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl48 | $spl56)),
  inference(unit_resulting_resolution,[],[f127,f125,f1023,f1080,f193])).
fof(f3893,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl56 | $spl70)),
  inference(unit_resulting_resolution,[],[f127,f126,f1211,f1080,f193])).
fof(f3894,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl56 | $spl70)),
  inference(unit_resulting_resolution,[],[f127,f125,f1211,f1080,f193])).
fof(f3895,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl56 | $spl82)),
  inference(unit_resulting_resolution,[],[f127,f126,f1273,f1080,f193])).
fof(f3896,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl56 | $spl82)),
  inference(unit_resulting_resolution,[],[f127,f125,f1273,f1080,f193])).
fof(f3897,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl56 | $spl94)),
  inference(unit_resulting_resolution,[],[f127,f126,f1334,f1080,f193])).
fof(f3898,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl56 | $spl94)),
  inference(unit_resulting_resolution,[],[f127,f125,f1334,f1080,f193])).
fof(f3899,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl56 | $spl106)),
  inference(unit_resulting_resolution,[],[f127,f126,f1395,f1080,f193])).
fof(f3900,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl56 | $spl106)),
  inference(unit_resulting_resolution,[],[f127,f125,f1395,f1080,f193])).
fof(f3901,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl18 | $spl56)),
  inference(unit_resulting_resolution,[],[f127,f127,f316,f1080,f195])).
fof(f3902,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl44 | $spl56)),
  inference(unit_resulting_resolution,[],[f127,f127,f926,f1080,f195])).
fof(f3903,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl56 | $spl66)),
  inference(unit_resulting_resolution,[],[f127,f127,f1160,f1080,f195])).
fof(f3904,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl33 | $spl56)),
  inference(unit_resulting_resolution,[],[f127,f127,f1410,f1080,f195])).
fof(f3905,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl56),
  inference(unit_resulting_resolution,[],[f127,f127,f126,f1080,f196])).
fof(f3906,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl56),
  inference(unit_resulting_resolution,[],[f127,f127,f125,f1080,f196])).
fof(f3907,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl56),
  inference(unit_resulting_resolution,[],[f127,f127,f1080,f1080,f196])).
fof(f3910,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl56),
  inference(unit_resulting_resolution,[],[f127,f127,f1080,f1080,f196])).
fof(f3923,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl10 | $spl56)),
  inference(unit_resulting_resolution,[],[f127,f126,f234,f1080,f197])).
fof(f3924,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl10 | $spl56)),
  inference(unit_resulting_resolution,[],[f127,f125,f234,f1080,f197])).
fof(f3925,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl14 | $spl56)),
  inference(unit_resulting_resolution,[],[f127,f126,f272,f1080,f197])).
fof(f3926,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl14 | $spl56)),
  inference(unit_resulting_resolution,[],[f127,f125,f272,f1080,f197])).
fof(f3927,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl36 | $spl56)),
  inference(unit_resulting_resolution,[],[f127,f126,f847,f1080,f197])).
fof(f3928,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl36 | $spl56)),
  inference(unit_resulting_resolution,[],[f127,f125,f847,f1080,f197])).
fof(f3929,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl48 | $spl56)),
  inference(unit_resulting_resolution,[],[f127,f126,f1023,f1080,f197])).
fof(f3930,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl48 | $spl56)),
  inference(unit_resulting_resolution,[],[f127,f125,f1023,f1080,f197])).
fof(f3931,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl56 | $spl70)),
  inference(unit_resulting_resolution,[],[f127,f126,f1211,f1080,f197])).
fof(f3932,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl56 | $spl70)),
  inference(unit_resulting_resolution,[],[f127,f125,f1211,f1080,f197])).
fof(f3933,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl56 | $spl82)),
  inference(unit_resulting_resolution,[],[f127,f126,f1273,f1080,f197])).
fof(f3934,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl56 | $spl82)),
  inference(unit_resulting_resolution,[],[f127,f125,f1273,f1080,f197])).
fof(f3935,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl56 | $spl94)),
  inference(unit_resulting_resolution,[],[f127,f126,f1334,f1080,f197])).
fof(f3936,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl56 | $spl94)),
  inference(unit_resulting_resolution,[],[f127,f125,f1334,f1080,f197])).
fof(f3937,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl56 | $spl106)),
  inference(unit_resulting_resolution,[],[f127,f126,f1395,f1080,f197])).
fof(f3938,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl56 | $spl106)),
  inference(unit_resulting_resolution,[],[f127,f125,f1395,f1080,f197])).
fof(f3939,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl56),
  inference(unit_resulting_resolution,[],[f127,f127,f1080,f199])).
fof(f3689,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl130),
  inference(cnf_transformation,[],[f3689_D])).
fof(f3689_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f3694,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl54),
  inference(global_subsumption,[],[f124,f123,f122,f1078,f3664,f3636,f3633,f3632,f3631,f3603,f3602,f3601,f3600,f3597,f3596,f3672,f3673,f3674])).
fof(f3674,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl54),
  inference(superposition,[],[f1078,f128])).
fof(f3673,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl54),
  inference(resolution,[],[f1078,f130])).
fof(f3672,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl54),
  inference(resolution,[],[f1078,f132])).
fof(f3693,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) | $spl134),
  inference(cnf_transformation,[],[f3693_D])).
fof(f3693_D,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f3692,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) | ($spl14 | $spl54)),
  inference(forward_demodulation,[],[f3665,f272])).
fof(f3665,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK2))))) ) | $spl54),
  inference(resolution,[],[f1078,f167])).
fof(f3691,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | $spl132),
  inference(cnf_transformation,[],[f3691_D])).
fof(f3691_D,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f3596,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | $spl54),
  inference(unit_resulting_resolution,[],[f1078,f1078,f183])).
fof(f3597,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | $spl54),
  inference(unit_resulting_resolution,[],[f1078,f1078,f183])).
fof(f3600,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl54),
  inference(unit_resulting_resolution,[],[f1078,f130])).
fof(f3601,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl54),
  inference(unit_resulting_resolution,[],[f1078,f132])).
fof(f3676,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl129),
  inference(cnf_transformation,[],[f3676_D])).
fof(f3676_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) <=> ~$spl129),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl129])])).
fof(f3602,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl54),
  inference(unit_resulting_resolution,[],[f1078,f133])).
fof(f3603,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl54),
  inference(unit_resulting_resolution,[],[f1078,f135])).
fof(f3609,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl18 | $spl54)),
  inference(unit_resulting_resolution,[],[f127,f127,f316,f1078,f191])).
fof(f3610,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl44 | $spl54)),
  inference(unit_resulting_resolution,[],[f127,f127,f926,f1078,f191])).
fof(f3611,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl54 | $spl66)),
  inference(unit_resulting_resolution,[],[f127,f127,f1160,f1078,f191])).
fof(f3612,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl10 | $spl54)),
  inference(unit_resulting_resolution,[],[f127,f126,f234,f1078,f193])).
fof(f3613,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl10 | $spl54)),
  inference(unit_resulting_resolution,[],[f127,f125,f234,f1078,f193])).
fof(f3614,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl14 | $spl54)),
  inference(unit_resulting_resolution,[],[f127,f126,f272,f1078,f193])).
fof(f3615,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl14 | $spl54)),
  inference(unit_resulting_resolution,[],[f127,f125,f272,f1078,f193])).
fof(f3616,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl36 | $spl54)),
  inference(unit_resulting_resolution,[],[f127,f126,f847,f1078,f193])).
fof(f3617,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl36 | $spl54)),
  inference(unit_resulting_resolution,[],[f127,f125,f847,f1078,f193])).
fof(f3618,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl48 | $spl54)),
  inference(unit_resulting_resolution,[],[f127,f126,f1023,f1078,f193])).
fof(f3619,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl48 | $spl54)),
  inference(unit_resulting_resolution,[],[f127,f125,f1023,f1078,f193])).
fof(f3620,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl54 | $spl70)),
  inference(unit_resulting_resolution,[],[f127,f126,f1211,f1078,f193])).
fof(f3621,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl54 | $spl70)),
  inference(unit_resulting_resolution,[],[f127,f125,f1211,f1078,f193])).
fof(f3622,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl54 | $spl82)),
  inference(unit_resulting_resolution,[],[f127,f126,f1273,f1078,f193])).
fof(f3623,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl54 | $spl82)),
  inference(unit_resulting_resolution,[],[f127,f125,f1273,f1078,f193])).
fof(f3624,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl54 | $spl94)),
  inference(unit_resulting_resolution,[],[f127,f126,f1334,f1078,f193])).
fof(f3625,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl54 | $spl94)),
  inference(unit_resulting_resolution,[],[f127,f125,f1334,f1078,f193])).
fof(f3626,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl54 | $spl106)),
  inference(unit_resulting_resolution,[],[f127,f126,f1395,f1078,f193])).
fof(f3627,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl54 | $spl106)),
  inference(unit_resulting_resolution,[],[f127,f125,f1395,f1078,f193])).
fof(f3628,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl18 | $spl54)),
  inference(unit_resulting_resolution,[],[f127,f127,f316,f1078,f195])).
fof(f3629,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl44 | $spl54)),
  inference(unit_resulting_resolution,[],[f127,f127,f926,f1078,f195])).
fof(f3630,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl54 | $spl66)),
  inference(unit_resulting_resolution,[],[f127,f127,f1160,f1078,f195])).
fof(f3631,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl54),
  inference(unit_resulting_resolution,[],[f127,f127,f126,f1078,f196])).
fof(f3632,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl54),
  inference(unit_resulting_resolution,[],[f127,f127,f125,f1078,f196])).
fof(f3633,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl54),
  inference(unit_resulting_resolution,[],[f127,f127,f1078,f1078,f196])).
fof(f3636,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl54),
  inference(unit_resulting_resolution,[],[f127,f127,f1078,f1078,f196])).
fof(f3648,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl10 | $spl54)),
  inference(unit_resulting_resolution,[],[f127,f126,f234,f1078,f197])).
fof(f3649,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl10 | $spl54)),
  inference(unit_resulting_resolution,[],[f127,f125,f234,f1078,f197])).
fof(f3650,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl14 | $spl54)),
  inference(unit_resulting_resolution,[],[f127,f126,f272,f1078,f197])).
fof(f3651,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl14 | $spl54)),
  inference(unit_resulting_resolution,[],[f127,f125,f272,f1078,f197])).
fof(f3652,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl36 | $spl54)),
  inference(unit_resulting_resolution,[],[f127,f126,f847,f1078,f197])).
fof(f3653,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl36 | $spl54)),
  inference(unit_resulting_resolution,[],[f127,f125,f847,f1078,f197])).
fof(f3654,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl48 | $spl54)),
  inference(unit_resulting_resolution,[],[f127,f126,f1023,f1078,f197])).
fof(f3655,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl48 | $spl54)),
  inference(unit_resulting_resolution,[],[f127,f125,f1023,f1078,f197])).
fof(f3656,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl54 | $spl70)),
  inference(unit_resulting_resolution,[],[f127,f126,f1211,f1078,f197])).
fof(f3657,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl54 | $spl70)),
  inference(unit_resulting_resolution,[],[f127,f125,f1211,f1078,f197])).
fof(f3658,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl54 | $spl82)),
  inference(unit_resulting_resolution,[],[f127,f126,f1273,f1078,f197])).
fof(f3659,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl54 | $spl82)),
  inference(unit_resulting_resolution,[],[f127,f125,f1273,f1078,f197])).
fof(f3660,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl54 | $spl94)),
  inference(unit_resulting_resolution,[],[f127,f126,f1334,f1078,f197])).
fof(f3661,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl54 | $spl94)),
  inference(unit_resulting_resolution,[],[f127,f125,f1334,f1078,f197])).
fof(f3662,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl54 | $spl106)),
  inference(unit_resulting_resolution,[],[f127,f126,f1395,f1078,f197])).
fof(f3663,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | ($spl54 | $spl106)),
  inference(unit_resulting_resolution,[],[f127,f125,f1395,f1078,f197])).
fof(f3664,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))) | $spl54),
  inference(unit_resulting_resolution,[],[f127,f127,f1078,f199])).
fof(f2670,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl126),
  inference(cnf_transformation,[],[f2670_D])).
fof(f2670_D,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f2675,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl108),
  inference(global_subsumption,[],[f124,f123,f122,f1397,f2622,f2608,f2601,f2592,f2586,f2585,f2584,f2568,f2561,f2551,f2550,f2549,f2548,f2630,f2631,f2632])).
fof(f2632,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl108),
  inference(superposition,[],[f1397,f128])).
fof(f2631,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl108),
  inference(resolution,[],[f1397,f130])).
fof(f2630,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl108),
  inference(resolution,[],[f1397,f132])).
fof(f2548,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl108),
  inference(unit_resulting_resolution,[],[f1397,f130])).
fof(f2549,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl108),
  inference(unit_resulting_resolution,[],[f1397,f132])).
fof(f2635,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl125),
  inference(cnf_transformation,[],[f2635_D])).
fof(f2635_D,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) <=> ~$spl125),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl125])])).
fof(f2550,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl108),
  inference(unit_resulting_resolution,[],[f1397,f133])).
fof(f2551,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl108),
  inference(unit_resulting_resolution,[],[f1397,f135])).
fof(f2557,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl18 | $spl108)),
  inference(unit_resulting_resolution,[],[f127,f127,f316,f1397,f191])).
fof(f2561,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl108),
  inference(unit_resulting_resolution,[],[f1397,f127,f180,f1397,f193])).
fof(f180,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'const.relation.RINTER_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5))) = s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),'const.relation.RINTER_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3))),s(X0,X4))),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(bool,'const.relation.RINTER_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5))) = s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),'const.relation.RINTER_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3))),s(X0,X4))),s(X1,X5)))),
  inference(flattening,[],[f59])).
fof(f59,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : s(bool,'const.relation.RINTER_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5))) = s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),'const.relation.RINTER_2'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3))),s(X0,X4))),s(X1,X5)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X6,X21] : ! [X22,X23,X24,X25] : s(bool,'const.relation.RINTER_4'(s(fun(X6,fun(X21,bool)),X22),s(fun(X6,fun(X21,bool)),X23),s(X6,X24),s(X21,X25))) = s(bool,happ(s(fun(X21,bool),happ(s(fun(X6,fun(X21,bool)),'const.relation.RINTER_2'(s(fun(X6,fun(X21,bool)),X22),s(fun(X6,fun(X21,bool)),X23))),s(X6,X24))),s(X21,X25)))),
  file('thm.relation.transitive_RINTER','ARITY_2_4_const.relation.RINTER')).
fof(f2568,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl108),
  inference(unit_resulting_resolution,[],[f127,f1397,f180,f1397,f193])).
fof(f2569,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl108)),
  inference(unit_resulting_resolution,[],[f127,f126,f234,f1397,f193])).
fof(f2570,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl108)),
  inference(unit_resulting_resolution,[],[f127,f125,f234,f1397,f193])).
fof(f2571,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl108)),
  inference(unit_resulting_resolution,[],[f127,f126,f272,f1397,f193])).
fof(f2572,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl108)),
  inference(unit_resulting_resolution,[],[f127,f125,f272,f1397,f193])).
fof(f2573,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl36 | $spl108)),
  inference(unit_resulting_resolution,[],[f127,f126,f847,f1397,f193])).
fof(f2574,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl36 | $spl108)),
  inference(unit_resulting_resolution,[],[f127,f125,f847,f1397,f193])).
fof(f2575,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl48 | $spl108)),
  inference(unit_resulting_resolution,[],[f127,f126,f1023,f1397,f193])).
fof(f2576,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl48 | $spl108)),
  inference(unit_resulting_resolution,[],[f127,f125,f1023,f1397,f193])).
fof(f2582,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl18 | $spl108)),
  inference(unit_resulting_resolution,[],[f127,f127,f316,f1397,f195])).
fof(f2584,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl108),
  inference(unit_resulting_resolution,[],[f127,f127,f126,f1397,f196])).
fof(f2585,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl108),
  inference(unit_resulting_resolution,[],[f127,f127,f125,f1397,f196])).
fof(f2586,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl108),
  inference(unit_resulting_resolution,[],[f127,f127,f1397,f1397,f196])).
fof(f2592,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl108),
  inference(unit_resulting_resolution,[],[f127,f127,f1397,f1397,f196])).
fof(f2601,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl108),
  inference(unit_resulting_resolution,[],[f1397,f127,f180,f1397,f197])).
fof(f2608,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl108),
  inference(unit_resulting_resolution,[],[f127,f1397,f180,f1397,f197])).
fof(f2609,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl108)),
  inference(unit_resulting_resolution,[],[f127,f126,f234,f1397,f197])).
fof(f2610,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl108)),
  inference(unit_resulting_resolution,[],[f127,f125,f234,f1397,f197])).
fof(f2611,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl108)),
  inference(unit_resulting_resolution,[],[f127,f126,f272,f1397,f197])).
fof(f2612,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl108)),
  inference(unit_resulting_resolution,[],[f127,f125,f272,f1397,f197])).
fof(f2613,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl36 | $spl108)),
  inference(unit_resulting_resolution,[],[f127,f126,f847,f1397,f197])).
fof(f2614,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl36 | $spl108)),
  inference(unit_resulting_resolution,[],[f127,f125,f847,f1397,f197])).
fof(f2615,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl48 | $spl108)),
  inference(unit_resulting_resolution,[],[f127,f126,f1023,f1397,f197])).
fof(f2616,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl48 | $spl108)),
  inference(unit_resulting_resolution,[],[f127,f125,f1023,f1397,f197])).
fof(f2622,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl108),
  inference(unit_resulting_resolution,[],[f127,f127,f1397,f199])).
fof(f2538,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl122),
  inference(cnf_transformation,[],[f2538_D])).
fof(f2538_D,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f2543,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl100),
  inference(global_subsumption,[],[f124,f123,f122,f1340,f2490,f2476,f2469,f2460,f2454,f2453,f2452,f2436,f2429,f2419,f2418,f2417,f2416,f2498,f2499,f2500])).
fof(f2500,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl100),
  inference(superposition,[],[f1340,f128])).
fof(f2499,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl100),
  inference(resolution,[],[f1340,f130])).
fof(f2498,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl100),
  inference(resolution,[],[f1340,f132])).
fof(f2416,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl100),
  inference(unit_resulting_resolution,[],[f1340,f130])).
fof(f2417,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl100),
  inference(unit_resulting_resolution,[],[f1340,f132])).
fof(f2503,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl121),
  inference(cnf_transformation,[],[f2503_D])).
fof(f2503_D,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) <=> ~$spl121),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl121])])).
fof(f2418,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl100),
  inference(unit_resulting_resolution,[],[f1340,f133])).
fof(f2419,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl100),
  inference(unit_resulting_resolution,[],[f1340,f135])).
fof(f2425,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl18 | $spl100)),
  inference(unit_resulting_resolution,[],[f127,f127,f316,f1340,f191])).
fof(f2429,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl100),
  inference(unit_resulting_resolution,[],[f1340,f127,f180,f1340,f193])).
fof(f2436,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl100),
  inference(unit_resulting_resolution,[],[f127,f1340,f180,f1340,f193])).
fof(f2437,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl100)),
  inference(unit_resulting_resolution,[],[f127,f126,f234,f1340,f193])).
fof(f2438,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl100)),
  inference(unit_resulting_resolution,[],[f127,f125,f234,f1340,f193])).
fof(f2439,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl100)),
  inference(unit_resulting_resolution,[],[f127,f126,f272,f1340,f193])).
fof(f2440,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl100)),
  inference(unit_resulting_resolution,[],[f127,f125,f272,f1340,f193])).
fof(f2441,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl36 | $spl100)),
  inference(unit_resulting_resolution,[],[f127,f126,f847,f1340,f193])).
fof(f2442,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl36 | $spl100)),
  inference(unit_resulting_resolution,[],[f127,f125,f847,f1340,f193])).
fof(f2443,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl48 | $spl100)),
  inference(unit_resulting_resolution,[],[f127,f126,f1023,f1340,f193])).
fof(f2444,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl48 | $spl100)),
  inference(unit_resulting_resolution,[],[f127,f125,f1023,f1340,f193])).
fof(f2450,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl18 | $spl100)),
  inference(unit_resulting_resolution,[],[f127,f127,f316,f1340,f195])).
fof(f2452,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl100),
  inference(unit_resulting_resolution,[],[f127,f127,f126,f1340,f196])).
fof(f2453,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl100),
  inference(unit_resulting_resolution,[],[f127,f127,f125,f1340,f196])).
fof(f2454,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl100),
  inference(unit_resulting_resolution,[],[f127,f127,f1340,f1340,f196])).
fof(f2460,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl100),
  inference(unit_resulting_resolution,[],[f127,f127,f1340,f1340,f196])).
fof(f2469,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl100),
  inference(unit_resulting_resolution,[],[f1340,f127,f180,f1340,f197])).
fof(f2476,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl100),
  inference(unit_resulting_resolution,[],[f127,f1340,f180,f1340,f197])).
fof(f2477,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl100)),
  inference(unit_resulting_resolution,[],[f127,f126,f234,f1340,f197])).
fof(f2478,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl100)),
  inference(unit_resulting_resolution,[],[f127,f125,f234,f1340,f197])).
fof(f2479,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl100)),
  inference(unit_resulting_resolution,[],[f127,f126,f272,f1340,f197])).
fof(f2480,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl100)),
  inference(unit_resulting_resolution,[],[f127,f125,f272,f1340,f197])).
fof(f2481,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl36 | $spl100)),
  inference(unit_resulting_resolution,[],[f127,f126,f847,f1340,f197])).
fof(f2482,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl36 | $spl100)),
  inference(unit_resulting_resolution,[],[f127,f125,f847,f1340,f197])).
fof(f2483,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl48 | $spl100)),
  inference(unit_resulting_resolution,[],[f127,f126,f1023,f1340,f197])).
fof(f2484,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl48 | $spl100)),
  inference(unit_resulting_resolution,[],[f127,f125,f1023,f1340,f197])).
fof(f2490,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl100),
  inference(unit_resulting_resolution,[],[f127,f127,f1340,f199])).
fof(f2406,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl118),
  inference(cnf_transformation,[],[f2406_D])).
fof(f2406_D,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f2411,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl88),
  inference(global_subsumption,[],[f124,f123,f122,f1279,f2358,f2344,f2337,f2328,f2322,f2321,f2320,f2304,f2297,f2287,f2286,f2285,f2284,f2366,f2367,f2368])).
fof(f2368,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl88),
  inference(superposition,[],[f1279,f128])).
fof(f2367,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl88),
  inference(resolution,[],[f1279,f130])).
fof(f2366,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl88),
  inference(resolution,[],[f1279,f132])).
fof(f2284,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl88),
  inference(unit_resulting_resolution,[],[f1279,f130])).
fof(f2285,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl88),
  inference(unit_resulting_resolution,[],[f1279,f132])).
fof(f2371,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl117),
  inference(cnf_transformation,[],[f2371_D])).
fof(f2371_D,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) <=> ~$spl117),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl117])])).
fof(f2286,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl88),
  inference(unit_resulting_resolution,[],[f1279,f133])).
fof(f2287,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl88),
  inference(unit_resulting_resolution,[],[f1279,f135])).
fof(f2293,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl18 | $spl88)),
  inference(unit_resulting_resolution,[],[f127,f127,f316,f1279,f191])).
fof(f2297,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl88),
  inference(unit_resulting_resolution,[],[f1279,f127,f180,f1279,f193])).
fof(f2304,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl88),
  inference(unit_resulting_resolution,[],[f127,f1279,f180,f1279,f193])).
fof(f2305,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl88)),
  inference(unit_resulting_resolution,[],[f127,f126,f234,f1279,f193])).
fof(f2306,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl88)),
  inference(unit_resulting_resolution,[],[f127,f125,f234,f1279,f193])).
fof(f2307,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl88)),
  inference(unit_resulting_resolution,[],[f127,f126,f272,f1279,f193])).
fof(f2308,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl88)),
  inference(unit_resulting_resolution,[],[f127,f125,f272,f1279,f193])).
fof(f2309,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl36 | $spl88)),
  inference(unit_resulting_resolution,[],[f127,f126,f847,f1279,f193])).
fof(f2310,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl36 | $spl88)),
  inference(unit_resulting_resolution,[],[f127,f125,f847,f1279,f193])).
fof(f2311,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl48 | $spl88)),
  inference(unit_resulting_resolution,[],[f127,f126,f1023,f1279,f193])).
fof(f2312,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl48 | $spl88)),
  inference(unit_resulting_resolution,[],[f127,f125,f1023,f1279,f193])).
fof(f2318,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl18 | $spl88)),
  inference(unit_resulting_resolution,[],[f127,f127,f316,f1279,f195])).
fof(f2320,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl88),
  inference(unit_resulting_resolution,[],[f127,f127,f126,f1279,f196])).
fof(f2321,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl88),
  inference(unit_resulting_resolution,[],[f127,f127,f125,f1279,f196])).
fof(f2322,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl88),
  inference(unit_resulting_resolution,[],[f127,f127,f1279,f1279,f196])).
fof(f2328,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl88),
  inference(unit_resulting_resolution,[],[f127,f127,f1279,f1279,f196])).
fof(f2337,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl88),
  inference(unit_resulting_resolution,[],[f1279,f127,f180,f1279,f197])).
fof(f2344,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl88),
  inference(unit_resulting_resolution,[],[f127,f1279,f180,f1279,f197])).
fof(f2345,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl88)),
  inference(unit_resulting_resolution,[],[f127,f126,f234,f1279,f197])).
fof(f2346,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl88)),
  inference(unit_resulting_resolution,[],[f127,f125,f234,f1279,f197])).
fof(f2347,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl88)),
  inference(unit_resulting_resolution,[],[f127,f126,f272,f1279,f197])).
fof(f2348,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl88)),
  inference(unit_resulting_resolution,[],[f127,f125,f272,f1279,f197])).
fof(f2349,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl36 | $spl88)),
  inference(unit_resulting_resolution,[],[f127,f126,f847,f1279,f197])).
fof(f2350,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl36 | $spl88)),
  inference(unit_resulting_resolution,[],[f127,f125,f847,f1279,f197])).
fof(f2351,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl48 | $spl88)),
  inference(unit_resulting_resolution,[],[f127,f126,f1023,f1279,f197])).
fof(f2352,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl48 | $spl88)),
  inference(unit_resulting_resolution,[],[f127,f125,f1023,f1279,f197])).
fof(f2358,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl88),
  inference(unit_resulting_resolution,[],[f127,f127,f1279,f199])).
fof(f2274,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl114),
  inference(cnf_transformation,[],[f2274_D])).
fof(f2274_D,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f2279,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl76),
  inference(global_subsumption,[],[f124,f123,f122,f1217,f2226,f2212,f2205,f2196,f2190,f2189,f2188,f2172,f2165,f2155,f2154,f2153,f2152,f2234,f2235,f2236])).
fof(f2236,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl76),
  inference(superposition,[],[f1217,f128])).
fof(f2235,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl76),
  inference(resolution,[],[f1217,f130])).
fof(f2234,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl76),
  inference(resolution,[],[f1217,f132])).
fof(f2152,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl76),
  inference(unit_resulting_resolution,[],[f1217,f130])).
fof(f2153,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl76),
  inference(unit_resulting_resolution,[],[f1217,f132])).
fof(f2239,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl113),
  inference(cnf_transformation,[],[f2239_D])).
fof(f2239_D,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) <=> ~$spl113),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl113])])).
fof(f2154,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl76),
  inference(unit_resulting_resolution,[],[f1217,f133])).
fof(f2155,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl76),
  inference(unit_resulting_resolution,[],[f1217,f135])).
fof(f2161,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl18 | $spl76)),
  inference(unit_resulting_resolution,[],[f127,f127,f316,f1217,f191])).
fof(f2165,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl76),
  inference(unit_resulting_resolution,[],[f1217,f127,f180,f1217,f193])).
fof(f2172,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl76),
  inference(unit_resulting_resolution,[],[f127,f1217,f180,f1217,f193])).
fof(f2173,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl76)),
  inference(unit_resulting_resolution,[],[f127,f126,f234,f1217,f193])).
fof(f2174,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl76)),
  inference(unit_resulting_resolution,[],[f127,f125,f234,f1217,f193])).
fof(f2175,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl76)),
  inference(unit_resulting_resolution,[],[f127,f126,f272,f1217,f193])).
fof(f2176,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl76)),
  inference(unit_resulting_resolution,[],[f127,f125,f272,f1217,f193])).
fof(f2177,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl36 | $spl76)),
  inference(unit_resulting_resolution,[],[f127,f126,f847,f1217,f193])).
fof(f2178,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl36 | $spl76)),
  inference(unit_resulting_resolution,[],[f127,f125,f847,f1217,f193])).
fof(f2179,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl48 | $spl76)),
  inference(unit_resulting_resolution,[],[f127,f126,f1023,f1217,f193])).
fof(f2180,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl48 | $spl76)),
  inference(unit_resulting_resolution,[],[f127,f125,f1023,f1217,f193])).
fof(f2186,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl18 | $spl76)),
  inference(unit_resulting_resolution,[],[f127,f127,f316,f1217,f195])).
fof(f2188,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl76),
  inference(unit_resulting_resolution,[],[f127,f127,f126,f1217,f196])).
fof(f2189,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl76),
  inference(unit_resulting_resolution,[],[f127,f127,f125,f1217,f196])).
fof(f2190,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl76),
  inference(unit_resulting_resolution,[],[f127,f127,f1217,f1217,f196])).
fof(f2196,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl76),
  inference(unit_resulting_resolution,[],[f127,f127,f1217,f1217,f196])).
fof(f2205,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl76),
  inference(unit_resulting_resolution,[],[f1217,f127,f180,f1217,f197])).
fof(f2212,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl76),
  inference(unit_resulting_resolution,[],[f127,f1217,f180,f1217,f197])).
fof(f2213,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl76)),
  inference(unit_resulting_resolution,[],[f127,f126,f234,f1217,f197])).
fof(f2214,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl76)),
  inference(unit_resulting_resolution,[],[f127,f125,f234,f1217,f197])).
fof(f2215,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl76)),
  inference(unit_resulting_resolution,[],[f127,f126,f272,f1217,f197])).
fof(f2216,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl76)),
  inference(unit_resulting_resolution,[],[f127,f125,f272,f1217,f197])).
fof(f2217,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl36 | $spl76)),
  inference(unit_resulting_resolution,[],[f127,f126,f847,f1217,f197])).
fof(f2218,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl36 | $spl76)),
  inference(unit_resulting_resolution,[],[f127,f125,f847,f1217,f197])).
fof(f2219,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl48 | $spl76)),
  inference(unit_resulting_resolution,[],[f127,f126,f1023,f1217,f197])).
fof(f2220,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl48 | $spl76)),
  inference(unit_resulting_resolution,[],[f127,f125,f1023,f1217,f197])).
fof(f2226,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl76),
  inference(unit_resulting_resolution,[],[f127,f127,f1217,f199])).
fof(f1395,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl106),
  inference(cnf_transformation,[],[f1395_D])).
fof(f1395_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f1400,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl52),
  inference(global_subsumption,[],[f124,f123,f122,f1027,f1378,f1370,f1367,f1366,f1365,f1353,f1352,f1351,f1350,f1349,f1348,f1347,f1346,f1380,f1381,f1386,f1387,f1388])).
fof(f1388,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl52),
  inference(superposition,[],[f1027,f128])).
fof(f1387,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl52),
  inference(resolution,[],[f1027,f130])).
fof(f1386,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl52),
  inference(resolution,[],[f1027,f132])).
fof(f1399,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) | $spl110),
  inference(cnf_transformation,[],[f1399_D])).
fof(f1399_D,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f1398,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) | ($spl14 | $spl52)),
  inference(forward_demodulation,[],[f1379,f272])).
fof(f1379,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK2))))) ) | $spl52),
  inference(resolution,[],[f1027,f167])).
fof(f1397,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl108),
  inference(cnf_transformation,[],[f1397_D])).
fof(f1397_D,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f1346,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl52),
  inference(unit_resulting_resolution,[],[f1027,f1027,f183])).
fof(f1347,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl52),
  inference(unit_resulting_resolution,[],[f1027,f1027,f183])).
fof(f1080,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | $spl56),
  inference(cnf_transformation,[],[f1080_D])).
fof(f1080_D,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f1348,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | $spl52),
  inference(unit_resulting_resolution,[],[f127,f1027,f175])).
fof(f1078,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | $spl54),
  inference(cnf_transformation,[],[f1078_D])).
fof(f1078_D,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f1349,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | $spl52),
  inference(unit_resulting_resolution,[],[f127,f1027,f177])).
fof(f1350,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl52),
  inference(unit_resulting_resolution,[],[f1027,f130])).
fof(f1351,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl52),
  inference(unit_resulting_resolution,[],[f1027,f132])).
fof(f1390,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl105),
  inference(cnf_transformation,[],[f1390_D])).
fof(f1390_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) <=> ~$spl105),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl105])])).
fof(f1352,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl52),
  inference(unit_resulting_resolution,[],[f1027,f133])).
fof(f1353,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl52),
  inference(unit_resulting_resolution,[],[f1027,f135])).
fof(f1359,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl18 | $spl52)),
  inference(unit_resulting_resolution,[],[f127,f127,f316,f1027,f191])).
fof(f1360,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl52)),
  inference(unit_resulting_resolution,[],[f127,f126,f234,f1027,f193])).
fof(f1361,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl52)),
  inference(unit_resulting_resolution,[],[f127,f125,f234,f1027,f193])).
fof(f1362,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl52)),
  inference(unit_resulting_resolution,[],[f127,f126,f272,f1027,f193])).
fof(f1363,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl52)),
  inference(unit_resulting_resolution,[],[f127,f125,f272,f1027,f193])).
fof(f1364,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl18 | $spl52)),
  inference(unit_resulting_resolution,[],[f127,f127,f316,f1027,f195])).
fof(f1365,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl52),
  inference(unit_resulting_resolution,[],[f127,f127,f126,f1027,f196])).
fof(f1366,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl52),
  inference(unit_resulting_resolution,[],[f127,f127,f125,f1027,f196])).
fof(f1367,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl52),
  inference(unit_resulting_resolution,[],[f127,f127,f1027,f1027,f196])).
fof(f1370,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl52),
  inference(unit_resulting_resolution,[],[f127,f127,f1027,f1027,f196])).
fof(f1374,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl52)),
  inference(unit_resulting_resolution,[],[f127,f126,f234,f1027,f197])).
fof(f1375,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl52)),
  inference(unit_resulting_resolution,[],[f127,f125,f234,f1027,f197])).
fof(f1376,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl52)),
  inference(unit_resulting_resolution,[],[f127,f126,f272,f1027,f197])).
fof(f1377,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl52)),
  inference(unit_resulting_resolution,[],[f127,f125,f272,f1027,f197])).
fof(f1378,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl52),
  inference(unit_resulting_resolution,[],[f127,f127,f1027,f199])).
fof(f1334,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl94),
  inference(cnf_transformation,[],[f1334_D])).
fof(f1334_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f1343,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl50),
  inference(global_subsumption,[],[f124,f123,f122,f1025,f1317,f1309,f1306,f1305,f1304,f1292,f1291,f1290,f1289,f1288,f1287,f1286,f1285,f1319,f1320,f1325,f1326,f1327])).
fof(f1327,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl50),
  inference(superposition,[],[f1025,f128])).
fof(f1326,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl50),
  inference(resolution,[],[f1025,f130])).
fof(f1325,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl50),
  inference(resolution,[],[f1025,f132])).
fof(f1342,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) | $spl102),
  inference(cnf_transformation,[],[f1342_D])).
fof(f1342_D,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f1341,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) | ($spl10 | $spl50)),
  inference(forward_demodulation,[],[f1318,f234])).
fof(f1318,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK1))))) ) | $spl50),
  inference(resolution,[],[f1025,f167])).
fof(f1340,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl100),
  inference(cnf_transformation,[],[f1340_D])).
fof(f1340_D,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f1285,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl50),
  inference(unit_resulting_resolution,[],[f1025,f1025,f183])).
fof(f1286,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl50),
  inference(unit_resulting_resolution,[],[f1025,f1025,f183])).
fof(f1338,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | $spl98),
  inference(cnf_transformation,[],[f1338_D])).
fof(f1338_D,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f1287,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | $spl50),
  inference(unit_resulting_resolution,[],[f127,f1025,f175])).
fof(f1336,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | $spl96),
  inference(cnf_transformation,[],[f1336_D])).
fof(f1336_D,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f1288,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | $spl50),
  inference(unit_resulting_resolution,[],[f127,f1025,f177])).
fof(f1289,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl50),
  inference(unit_resulting_resolution,[],[f1025,f130])).
fof(f1290,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl50),
  inference(unit_resulting_resolution,[],[f1025,f132])).
fof(f1329,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl93),
  inference(cnf_transformation,[],[f1329_D])).
fof(f1329_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) <=> ~$spl93),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl93])])).
fof(f1291,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl50),
  inference(unit_resulting_resolution,[],[f1025,f133])).
fof(f1292,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl50),
  inference(unit_resulting_resolution,[],[f1025,f135])).
fof(f1298,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl18 | $spl50)),
  inference(unit_resulting_resolution,[],[f127,f127,f316,f1025,f191])).
fof(f1299,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl50)),
  inference(unit_resulting_resolution,[],[f127,f126,f234,f1025,f193])).
fof(f1300,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl50)),
  inference(unit_resulting_resolution,[],[f127,f125,f234,f1025,f193])).
fof(f1301,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl50)),
  inference(unit_resulting_resolution,[],[f127,f126,f272,f1025,f193])).
fof(f1302,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl50)),
  inference(unit_resulting_resolution,[],[f127,f125,f272,f1025,f193])).
fof(f1303,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl18 | $spl50)),
  inference(unit_resulting_resolution,[],[f127,f127,f316,f1025,f195])).
fof(f1304,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl50),
  inference(unit_resulting_resolution,[],[f127,f127,f126,f1025,f196])).
fof(f1305,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl50),
  inference(unit_resulting_resolution,[],[f127,f127,f125,f1025,f196])).
fof(f1306,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl50),
  inference(unit_resulting_resolution,[],[f127,f127,f1025,f1025,f196])).
fof(f1309,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl50),
  inference(unit_resulting_resolution,[],[f127,f127,f1025,f1025,f196])).
fof(f1313,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl50)),
  inference(unit_resulting_resolution,[],[f127,f126,f234,f1025,f197])).
fof(f1314,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl50)),
  inference(unit_resulting_resolution,[],[f127,f125,f234,f1025,f197])).
fof(f1315,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl50)),
  inference(unit_resulting_resolution,[],[f127,f126,f272,f1025,f197])).
fof(f1316,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl50)),
  inference(unit_resulting_resolution,[],[f127,f125,f272,f1025,f197])).
fof(f1317,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl50),
  inference(unit_resulting_resolution,[],[f127,f127,f1025,f199])).
fof(f1273,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl82),
  inference(cnf_transformation,[],[f1273_D])).
fof(f1273_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f1282,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl40),
  inference(global_subsumption,[],[f124,f123,f122,f851,f1256,f1248,f1245,f1244,f1243,f1231,f1230,f1229,f1228,f1227,f1226,f1225,f1224,f1258,f1259,f1264,f1265,f1266])).
fof(f1266,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl40),
  inference(superposition,[],[f851,f128])).
fof(f1265,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl40),
  inference(resolution,[],[f851,f130])).
fof(f1264,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl40),
  inference(resolution,[],[f851,f132])).
fof(f1281,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) | $spl90),
  inference(cnf_transformation,[],[f1281_D])).
fof(f1281_D,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f1280,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) | ($spl14 | $spl40)),
  inference(forward_demodulation,[],[f1257,f272])).
fof(f1257,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK2))))) ) | $spl40),
  inference(resolution,[],[f851,f167])).
fof(f1279,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl88),
  inference(cnf_transformation,[],[f1279_D])).
fof(f1279_D,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f1224,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl40),
  inference(unit_resulting_resolution,[],[f851,f851,f183])).
fof(f1225,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl40),
  inference(unit_resulting_resolution,[],[f851,f851,f183])).
fof(f1277,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | $spl86),
  inference(cnf_transformation,[],[f1277_D])).
fof(f1277_D,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f1226,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | $spl40),
  inference(unit_resulting_resolution,[],[f127,f851,f175])).
fof(f1275,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | $spl84),
  inference(cnf_transformation,[],[f1275_D])).
fof(f1275_D,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f1227,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | $spl40),
  inference(unit_resulting_resolution,[],[f127,f851,f177])).
fof(f1228,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl40),
  inference(unit_resulting_resolution,[],[f851,f130])).
fof(f1229,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl40),
  inference(unit_resulting_resolution,[],[f851,f132])).
fof(f1268,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl81),
  inference(cnf_transformation,[],[f1268_D])).
fof(f1268_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) <=> ~$spl81),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl81])])).
fof(f1230,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl40),
  inference(unit_resulting_resolution,[],[f851,f133])).
fof(f1231,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl40),
  inference(unit_resulting_resolution,[],[f851,f135])).
fof(f1237,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl18 | $spl40)),
  inference(unit_resulting_resolution,[],[f127,f127,f316,f851,f191])).
fof(f1238,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl40)),
  inference(unit_resulting_resolution,[],[f127,f126,f234,f851,f193])).
fof(f1239,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl40)),
  inference(unit_resulting_resolution,[],[f127,f125,f234,f851,f193])).
fof(f1240,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl40)),
  inference(unit_resulting_resolution,[],[f127,f126,f272,f851,f193])).
fof(f1241,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl40)),
  inference(unit_resulting_resolution,[],[f127,f125,f272,f851,f193])).
fof(f1242,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl18 | $spl40)),
  inference(unit_resulting_resolution,[],[f127,f127,f316,f851,f195])).
fof(f1243,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl40),
  inference(unit_resulting_resolution,[],[f127,f127,f126,f851,f196])).
fof(f1244,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl40),
  inference(unit_resulting_resolution,[],[f127,f127,f125,f851,f196])).
fof(f1245,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl40),
  inference(unit_resulting_resolution,[],[f127,f127,f851,f851,f196])).
fof(f1248,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl40),
  inference(unit_resulting_resolution,[],[f127,f127,f851,f851,f196])).
fof(f1252,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl40)),
  inference(unit_resulting_resolution,[],[f127,f126,f234,f851,f197])).
fof(f1253,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl40)),
  inference(unit_resulting_resolution,[],[f127,f125,f234,f851,f197])).
fof(f1254,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl40)),
  inference(unit_resulting_resolution,[],[f127,f126,f272,f851,f197])).
fof(f1255,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl40)),
  inference(unit_resulting_resolution,[],[f127,f125,f272,f851,f197])).
fof(f1256,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl40),
  inference(unit_resulting_resolution,[],[f127,f127,f851,f199])).
fof(f1221,plain,(
  s(bool,'T_0') != s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK1))) | p(s(bool,'T_0')) | ~p(s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK1))))),
  introduced(theory_axiom,[])).
fof(f1211,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl70),
  inference(cnf_transformation,[],[f1211_D])).
fof(f1211_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f1220,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl38),
  inference(global_subsumption,[],[f124,f123,f122,f849,f1194,f1186,f1183,f1182,f1181,f1169,f1168,f1167,f1166,f1165,f1164,f1163,f1162,f1196,f1197,f1202,f1203,f1204])).
fof(f1204,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl38),
  inference(superposition,[],[f849,f128])).
fof(f1203,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl38),
  inference(resolution,[],[f849,f130])).
fof(f1202,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl38),
  inference(resolution,[],[f849,f132])).
fof(f1219,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) | $spl78),
  inference(cnf_transformation,[],[f1219_D])).
fof(f1219_D,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f1218,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) | ($spl10 | $spl38)),
  inference(forward_demodulation,[],[f1195,f234])).
fof(f1195,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,X0))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK1))))) ) | $spl38),
  inference(resolution,[],[f849,f167])).
fof(f1217,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl76),
  inference(cnf_transformation,[],[f1217_D])).
fof(f1217_D,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f1162,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl38),
  inference(unit_resulting_resolution,[],[f849,f849,f183])).
fof(f1163,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl38),
  inference(unit_resulting_resolution,[],[f849,f849,f183])).
fof(f1215,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | $spl74),
  inference(cnf_transformation,[],[f1215_D])).
fof(f1215_D,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f1164,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | $spl38),
  inference(unit_resulting_resolution,[],[f127,f849,f175])).
fof(f1213,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | $spl72),
  inference(cnf_transformation,[],[f1213_D])).
fof(f1213_D,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f1165,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | $spl38),
  inference(unit_resulting_resolution,[],[f127,f849,f177])).
fof(f1166,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl38),
  inference(unit_resulting_resolution,[],[f849,f130])).
fof(f1167,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl38),
  inference(unit_resulting_resolution,[],[f849,f132])).
fof(f1206,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl69),
  inference(cnf_transformation,[],[f1206_D])).
fof(f1206_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) <=> ~$spl69),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl69])])).
fof(f1168,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl38),
  inference(unit_resulting_resolution,[],[f849,f133])).
fof(f1169,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl38),
  inference(unit_resulting_resolution,[],[f849,f135])).
fof(f1175,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl18 | $spl38)),
  inference(unit_resulting_resolution,[],[f127,f127,f316,f849,f191])).
fof(f1176,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl38)),
  inference(unit_resulting_resolution,[],[f127,f126,f234,f849,f193])).
fof(f1177,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl38)),
  inference(unit_resulting_resolution,[],[f127,f125,f234,f849,f193])).
fof(f1178,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl38)),
  inference(unit_resulting_resolution,[],[f127,f126,f272,f849,f193])).
fof(f1179,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl38)),
  inference(unit_resulting_resolution,[],[f127,f125,f272,f849,f193])).
fof(f1180,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl18 | $spl38)),
  inference(unit_resulting_resolution,[],[f127,f127,f316,f849,f195])).
fof(f1181,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl38),
  inference(unit_resulting_resolution,[],[f127,f127,f126,f849,f196])).
fof(f1182,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl38),
  inference(unit_resulting_resolution,[],[f127,f127,f125,f849,f196])).
fof(f1183,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl38),
  inference(unit_resulting_resolution,[],[f127,f127,f849,f849,f196])).
fof(f1186,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl38),
  inference(unit_resulting_resolution,[],[f127,f127,f849,f849,f196])).
fof(f1190,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl38)),
  inference(unit_resulting_resolution,[],[f127,f126,f234,f849,f197])).
fof(f1191,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl38)),
  inference(unit_resulting_resolution,[],[f127,f125,f234,f849,f197])).
fof(f1192,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl38)),
  inference(unit_resulting_resolution,[],[f127,f126,f272,f849,f197])).
fof(f1193,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl38)),
  inference(unit_resulting_resolution,[],[f127,f125,f272,f849,f197])).
fof(f1194,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl38),
  inference(unit_resulting_resolution,[],[f127,f127,f849,f199])).
fof(f1149,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl33),
  inference(resolution,[],[f749,f134])).
fof(f134,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f83])).
fof(f1148,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl33),
  inference(resolution,[],[f749,f136])).
fof(f1155,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl65),
  inference(cnf_transformation,[],[f1155_D])).
fof(f1155_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) <=> ~$spl65),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl65])])).
fof(f1091,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl33),
  inference(unit_resulting_resolution,[],[f749,f129])).
fof(f129,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f81])).
fof(f1092,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl33),
  inference(unit_resulting_resolution,[],[f749,f131])).
fof(f131,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f82])).
fof(f1093,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl33),
  inference(unit_resulting_resolution,[],[f749,f134])).
fof(f1094,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl33),
  inference(unit_resulting_resolution,[],[f749,f136])).
fof(f1101,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl18 | $spl33)),
  inference(unit_resulting_resolution,[],[f126,f127,f316,f749,f191])).
fof(f1102,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl18 | $spl33)),
  inference(unit_resulting_resolution,[],[f125,f127,f316,f749,f191])).
fof(f1103,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl33)),
  inference(unit_resulting_resolution,[],[f126,f126,f234,f749,f193])).
fof(f1104,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl33)),
  inference(unit_resulting_resolution,[],[f125,f126,f234,f749,f193])).
fof(f1105,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl33)),
  inference(unit_resulting_resolution,[],[f126,f125,f234,f749,f193])).
fof(f1106,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl33)),
  inference(unit_resulting_resolution,[],[f125,f125,f234,f749,f193])).
fof(f1107,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl33)),
  inference(unit_resulting_resolution,[],[f126,f126,f272,f749,f193])).
fof(f1108,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl33)),
  inference(unit_resulting_resolution,[],[f125,f126,f272,f749,f193])).
fof(f1109,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl33)),
  inference(unit_resulting_resolution,[],[f126,f125,f272,f749,f193])).
fof(f1110,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl33)),
  inference(unit_resulting_resolution,[],[f125,f125,f272,f749,f193])).
fof(f1111,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl18 | $spl33)),
  inference(unit_resulting_resolution,[],[f126,f127,f316,f749,f195])).
fof(f1112,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl18 | $spl33)),
  inference(unit_resulting_resolution,[],[f125,f127,f316,f749,f195])).
fof(f1113,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl33),
  inference(unit_resulting_resolution,[],[f127,f126,f126,f749,f196])).
fof(f1114,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl33),
  inference(unit_resulting_resolution,[],[f127,f125,f126,f749,f196])).
fof(f1115,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl33),
  inference(unit_resulting_resolution,[],[f127,f126,f125,f749,f196])).
fof(f1116,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl33),
  inference(unit_resulting_resolution,[],[f127,f125,f125,f749,f196])).
fof(f1119,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl33),
  inference(unit_resulting_resolution,[],[f126,f749,f126,f749,f196])).
fof(f1120,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl33),
  inference(unit_resulting_resolution,[],[f125,f749,f126,f749,f196])).
fof(f1123,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl33),
  inference(unit_resulting_resolution,[],[f126,f749,f125,f749,f196])).
fof(f1124,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl33),
  inference(unit_resulting_resolution,[],[f125,f749,f125,f749,f196])).
fof(f1128,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl33)),
  inference(unit_resulting_resolution,[],[f126,f126,f234,f749,f197])).
fof(f1129,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl33)),
  inference(unit_resulting_resolution,[],[f125,f126,f234,f749,f197])).
fof(f1130,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl33)),
  inference(unit_resulting_resolution,[],[f126,f125,f234,f749,f197])).
fof(f1131,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl33)),
  inference(unit_resulting_resolution,[],[f125,f125,f234,f749,f197])).
fof(f1132,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl33)),
  inference(unit_resulting_resolution,[],[f126,f126,f272,f749,f197])).
fof(f1133,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl33)),
  inference(unit_resulting_resolution,[],[f125,f126,f272,f749,f197])).
fof(f1134,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl33)),
  inference(unit_resulting_resolution,[],[f126,f125,f272,f749,f197])).
fof(f1135,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl33)),
  inference(unit_resulting_resolution,[],[f125,f125,f272,f749,f197])).
fof(f1136,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl33),
  inference(unit_resulting_resolution,[],[f127,f126,f749,f199])).
fof(f1137,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl33),
  inference(unit_resulting_resolution,[],[f127,f125,f749,f199])).
fof(f1140,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl33),
  inference(unit_resulting_resolution,[],[f126,f749,f749,f199])).
fof(f1141,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl33),
  inference(unit_resulting_resolution,[],[f125,f749,f749,f199])).
fof(f920,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl42),
  inference(cnf_transformation,[],[f920_D])).
fof(f920_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f1087,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl30),
  inference(global_subsumption,[],[f124,f123,f122,f746,f1063,f1055,f1052,f1051,f1050,f1038,f1037,f1036,f1035,f1034,f1033,f1032,f1031,f1065,f1066,f1071,f1072,f1073])).
fof(f1073,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl30),
  inference(superposition,[],[f746,f128])).
fof(f1066,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7(X2,happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))))) ) | $spl30),
  inference(resolution,[],[f746,f177])).
fof(f1065,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6(X1,happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0)))))) ) | $spl30),
  inference(resolution,[],[f746,f175])).
fof(f746,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl30),
  inference(cnf_transformation,[],[f746_D])).
fof(f746_D,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f1072,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl30),
  inference(resolution,[],[f746,f130])).
fof(f1071,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl30),
  inference(resolution,[],[f746,f132])).
fof(f1084,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) | $spl60),
  inference(cnf_transformation,[],[f1084_D])).
fof(f1084_D,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f1083,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))))) ) | ($spl14 | $spl30)),
  inference(forward_demodulation,[],[f1064,f272])).
fof(f1064,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,X0))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK2))))) ) | $spl30),
  inference(resolution,[],[f746,f167])).
fof(f1082,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl58),
  inference(cnf_transformation,[],[f1082_D])).
fof(f1082_D,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f1031,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl30),
  inference(unit_resulting_resolution,[],[f746,f746,f183])).
fof(f1032,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl30),
  inference(unit_resulting_resolution,[],[f746,f746,f183])).
fof(f1033,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK6('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | $spl30),
  inference(unit_resulting_resolution,[],[f127,f746,f175])).
fof(f1034,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK7('F_0',happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))),sK0))))) | $spl30),
  inference(unit_resulting_resolution,[],[f127,f746,f177])).
fof(f1035,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl30),
  inference(unit_resulting_resolution,[],[f746,f130])).
fof(f1036,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl30),
  inference(unit_resulting_resolution,[],[f746,f132])).
fof(f925,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl45),
  inference(cnf_transformation,[],[f925_D])).
fof(f925_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f1037,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl30),
  inference(unit_resulting_resolution,[],[f746,f133])).
fof(f1038,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl30),
  inference(unit_resulting_resolution,[],[f746,f135])).
fof(f1044,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl18 | $spl30)),
  inference(unit_resulting_resolution,[],[f127,f127,f316,f746,f191])).
fof(f1045,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl30)),
  inference(unit_resulting_resolution,[],[f127,f126,f234,f746,f193])).
fof(f1046,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl30)),
  inference(unit_resulting_resolution,[],[f127,f125,f234,f746,f193])).
fof(f1047,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl30)),
  inference(unit_resulting_resolution,[],[f127,f126,f272,f746,f193])).
fof(f1048,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl30)),
  inference(unit_resulting_resolution,[],[f127,f125,f272,f746,f193])).
fof(f1049,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl18 | $spl30)),
  inference(unit_resulting_resolution,[],[f127,f127,f316,f746,f195])).
fof(f1050,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl30),
  inference(unit_resulting_resolution,[],[f127,f127,f126,f746,f196])).
fof(f1051,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl30),
  inference(unit_resulting_resolution,[],[f127,f127,f125,f746,f196])).
fof(f1052,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl30),
  inference(unit_resulting_resolution,[],[f127,f127,f746,f746,f196])).
fof(f1055,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl30),
  inference(unit_resulting_resolution,[],[f127,f127,f746,f746,f196])).
fof(f1059,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl30)),
  inference(unit_resulting_resolution,[],[f127,f126,f234,f746,f197])).
fof(f1060,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl30)),
  inference(unit_resulting_resolution,[],[f127,f125,f234,f746,f197])).
fof(f1061,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl30)),
  inference(unit_resulting_resolution,[],[f127,f126,f272,f746,f197])).
fof(f1062,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl30)),
  inference(unit_resulting_resolution,[],[f127,f125,f272,f746,f197])).
fof(f1063,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl30),
  inference(unit_resulting_resolution,[],[f127,f127,f746,f199])).
fof(f1023,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl48),
  inference(cnf_transformation,[],[f1023_D])).
fof(f1023_D,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f1028,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl24),
  inference(global_subsumption,[],[f124,f123,f122,f325,f985,f977,f972,f965,f959,f958,f957,f947,f942,f932,f931,f930,f929,f928,f927,f987,f988,f993,f994,f995])).
fof(f995,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl24),
  inference(superposition,[],[f325,f128])).
fof(f994,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl24),
  inference(resolution,[],[f325,f130])).
fof(f993,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl24),
  inference(resolution,[],[f325,f132])).
fof(f1025,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl50),
  inference(cnf_transformation,[],[f1025_D])).
fof(f1025_D,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f988,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl24),
  inference(resolution,[],[f325,f182])).
fof(f182,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.relation.RINTER_4'(s(fun(X0,fun(X1,bool)),X2),s(fun(X0,fun(X1,bool)),X3),s(X0,X4),s(X1,X5)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X3),s(X0,X4))),s(X1,X5))))) )),
  inference(cnf_transformation,[],[f117])).
fof(f1027,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl52),
  inference(cnf_transformation,[],[f1027_D])).
fof(f1027_D,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f987,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl24),
  inference(resolution,[],[f325,f181])).
fof(f927,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl24),
  inference(unit_resulting_resolution,[],[f325,f181])).
fof(f928,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl24),
  inference(unit_resulting_resolution,[],[f325,f182])).
fof(f929,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl24),
  inference(unit_resulting_resolution,[],[f325,f130])).
fof(f930,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl24),
  inference(unit_resulting_resolution,[],[f325,f132])).
fof(f998,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl47),
  inference(cnf_transformation,[],[f998_D])).
fof(f998_D,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f931,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl24),
  inference(unit_resulting_resolution,[],[f325,f133])).
fof(f932,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl24),
  inference(unit_resulting_resolution,[],[f325,f135])).
fof(f938,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl18 | $spl24)),
  inference(unit_resulting_resolution,[],[f127,f127,f316,f325,f191])).
fof(f942,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl24),
  inference(unit_resulting_resolution,[],[f325,f127,f180,f325,f193])).
fof(f947,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl24),
  inference(unit_resulting_resolution,[],[f127,f325,f180,f325,f193])).
fof(f948,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl24)),
  inference(unit_resulting_resolution,[],[f127,f126,f234,f325,f193])).
fof(f949,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl24)),
  inference(unit_resulting_resolution,[],[f127,f125,f234,f325,f193])).
fof(f950,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl24)),
  inference(unit_resulting_resolution,[],[f127,f126,f272,f325,f193])).
fof(f951,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl24)),
  inference(unit_resulting_resolution,[],[f127,f125,f272,f325,f193])).
fof(f955,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl18 | $spl24)),
  inference(unit_resulting_resolution,[],[f127,f127,f316,f325,f195])).
fof(f957,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl24),
  inference(unit_resulting_resolution,[],[f127,f127,f126,f325,f196])).
fof(f958,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl24),
  inference(unit_resulting_resolution,[],[f127,f127,f125,f325,f196])).
fof(f959,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl24),
  inference(unit_resulting_resolution,[],[f127,f127,f325,f325,f196])).
fof(f965,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl24),
  inference(unit_resulting_resolution,[],[f127,f127,f325,f325,f196])).
fof(f972,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl24),
  inference(unit_resulting_resolution,[],[f325,f127,f180,f325,f197])).
fof(f977,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl24),
  inference(unit_resulting_resolution,[],[f127,f325,f180,f325,f197])).
fof(f978,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl24)),
  inference(unit_resulting_resolution,[],[f127,f126,f234,f325,f197])).
fof(f979,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl24)),
  inference(unit_resulting_resolution,[],[f127,f125,f234,f325,f197])).
fof(f980,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl24)),
  inference(unit_resulting_resolution,[],[f127,f126,f272,f325,f197])).
fof(f981,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl24)),
  inference(unit_resulting_resolution,[],[f127,f125,f272,f325,f197])).
fof(f985,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl24),
  inference(unit_resulting_resolution,[],[f127,f127,f325,f199])).
fof(f926,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl44),
  inference(cnf_transformation,[],[f926_D])).
fof(f926_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f915,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl31),
  inference(resolution,[],[f747,f134])).
fof(f914,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl31),
  inference(resolution,[],[f747,f136])).
fof(f921,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl43),
  inference(cnf_transformation,[],[f921_D])).
fof(f921_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f857,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl31),
  inference(unit_resulting_resolution,[],[f747,f129])).
fof(f858,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl31),
  inference(unit_resulting_resolution,[],[f747,f131])).
fof(f859,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl31),
  inference(unit_resulting_resolution,[],[f747,f134])).
fof(f860,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl31),
  inference(unit_resulting_resolution,[],[f747,f136])).
fof(f867,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl18 | $spl31)),
  inference(unit_resulting_resolution,[],[f126,f127,f316,f747,f191])).
fof(f868,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl18 | $spl31)),
  inference(unit_resulting_resolution,[],[f125,f127,f316,f747,f191])).
fof(f869,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl31)),
  inference(unit_resulting_resolution,[],[f126,f126,f234,f747,f193])).
fof(f870,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl31)),
  inference(unit_resulting_resolution,[],[f125,f126,f234,f747,f193])).
fof(f871,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl31)),
  inference(unit_resulting_resolution,[],[f126,f125,f234,f747,f193])).
fof(f872,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl31)),
  inference(unit_resulting_resolution,[],[f125,f125,f234,f747,f193])).
fof(f873,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl31)),
  inference(unit_resulting_resolution,[],[f126,f126,f272,f747,f193])).
fof(f874,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl31)),
  inference(unit_resulting_resolution,[],[f125,f126,f272,f747,f193])).
fof(f875,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl31)),
  inference(unit_resulting_resolution,[],[f126,f125,f272,f747,f193])).
fof(f876,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl31)),
  inference(unit_resulting_resolution,[],[f125,f125,f272,f747,f193])).
fof(f877,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl18 | $spl31)),
  inference(unit_resulting_resolution,[],[f126,f127,f316,f747,f195])).
fof(f878,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl18 | $spl31)),
  inference(unit_resulting_resolution,[],[f125,f127,f316,f747,f195])).
fof(f879,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl31),
  inference(unit_resulting_resolution,[],[f127,f126,f126,f747,f196])).
fof(f880,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl31),
  inference(unit_resulting_resolution,[],[f127,f125,f126,f747,f196])).
fof(f881,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl31),
  inference(unit_resulting_resolution,[],[f127,f126,f125,f747,f196])).
fof(f882,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl31),
  inference(unit_resulting_resolution,[],[f127,f125,f125,f747,f196])).
fof(f885,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl31),
  inference(unit_resulting_resolution,[],[f126,f747,f126,f747,f196])).
fof(f886,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl31),
  inference(unit_resulting_resolution,[],[f125,f747,f126,f747,f196])).
fof(f889,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl31),
  inference(unit_resulting_resolution,[],[f126,f747,f125,f747,f196])).
fof(f890,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl31),
  inference(unit_resulting_resolution,[],[f125,f747,f125,f747,f196])).
fof(f894,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl31)),
  inference(unit_resulting_resolution,[],[f126,f126,f234,f747,f197])).
fof(f895,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl31)),
  inference(unit_resulting_resolution,[],[f125,f126,f234,f747,f197])).
fof(f896,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl31)),
  inference(unit_resulting_resolution,[],[f126,f125,f234,f747,f197])).
fof(f897,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl31)),
  inference(unit_resulting_resolution,[],[f125,f125,f234,f747,f197])).
fof(f898,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl31)),
  inference(unit_resulting_resolution,[],[f126,f126,f272,f747,f197])).
fof(f899,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl31)),
  inference(unit_resulting_resolution,[],[f125,f126,f272,f747,f197])).
fof(f900,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl31)),
  inference(unit_resulting_resolution,[],[f126,f125,f272,f747,f197])).
fof(f901,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl31)),
  inference(unit_resulting_resolution,[],[f125,f125,f272,f747,f197])).
fof(f902,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl31),
  inference(unit_resulting_resolution,[],[f127,f126,f747,f199])).
fof(f903,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl31),
  inference(unit_resulting_resolution,[],[f127,f125,f747,f199])).
fof(f906,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl31),
  inference(unit_resulting_resolution,[],[f126,f747,f747,f199])).
fof(f907,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl31),
  inference(unit_resulting_resolution,[],[f125,f747,f747,f199])).
fof(f847,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl36),
  inference(cnf_transformation,[],[f847_D])).
fof(f847_D,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f852,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl22),
  inference(global_subsumption,[],[f124,f123,f122,f322,f809,f801,f796,f789,f783,f782,f781,f771,f766,f756,f755,f754,f753,f752,f751,f811,f812,f817,f818,f819])).
fof(f819,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl22),
  inference(superposition,[],[f322,f128])).
fof(f818,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl22),
  inference(resolution,[],[f322,f130])).
fof(f817,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl22),
  inference(resolution,[],[f322,f132])).
fof(f849,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl38),
  inference(cnf_transformation,[],[f849_D])).
fof(f849_D,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f812,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl22),
  inference(resolution,[],[f322,f182])).
fof(f851,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl40),
  inference(cnf_transformation,[],[f851_D])).
fof(f851_D,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f811,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl22),
  inference(resolution,[],[f322,f181])).
fof(f751,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl22),
  inference(unit_resulting_resolution,[],[f322,f181])).
fof(f752,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl22),
  inference(unit_resulting_resolution,[],[f322,f182])).
fof(f753,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl22),
  inference(unit_resulting_resolution,[],[f322,f130])).
fof(f754,plain,(
  s(bool,'T_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl22),
  inference(unit_resulting_resolution,[],[f322,f132])).
fof(f822,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl35),
  inference(cnf_transformation,[],[f822_D])).
fof(f822_D,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f755,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl22),
  inference(unit_resulting_resolution,[],[f322,f133])).
fof(f756,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl22),
  inference(unit_resulting_resolution,[],[f322,f135])).
fof(f762,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl18 | $spl22)),
  inference(unit_resulting_resolution,[],[f127,f127,f316,f322,f191])).
fof(f766,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl22),
  inference(unit_resulting_resolution,[],[f322,f127,f180,f322,f193])).
fof(f771,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl22),
  inference(unit_resulting_resolution,[],[f127,f322,f180,f322,f193])).
fof(f772,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl22)),
  inference(unit_resulting_resolution,[],[f127,f126,f234,f322,f193])).
fof(f773,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl22)),
  inference(unit_resulting_resolution,[],[f127,f125,f234,f322,f193])).
fof(f774,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl22)),
  inference(unit_resulting_resolution,[],[f127,f126,f272,f322,f193])).
fof(f775,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl22)),
  inference(unit_resulting_resolution,[],[f127,f125,f272,f322,f193])).
fof(f779,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl18 | $spl22)),
  inference(unit_resulting_resolution,[],[f127,f127,f316,f322,f195])).
fof(f781,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl22),
  inference(unit_resulting_resolution,[],[f127,f127,f126,f322,f196])).
fof(f782,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl22),
  inference(unit_resulting_resolution,[],[f127,f127,f125,f322,f196])).
fof(f783,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl22),
  inference(unit_resulting_resolution,[],[f127,f127,f322,f322,f196])).
fof(f789,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl22),
  inference(unit_resulting_resolution,[],[f127,f127,f322,f322,f196])).
fof(f796,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl22),
  inference(unit_resulting_resolution,[],[f322,f127,f180,f322,f197])).
fof(f801,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl22),
  inference(unit_resulting_resolution,[],[f127,f322,f180,f322,f197])).
fof(f802,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl22)),
  inference(unit_resulting_resolution,[],[f127,f126,f234,f322,f197])).
fof(f803,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl22)),
  inference(unit_resulting_resolution,[],[f127,f125,f234,f322,f197])).
fof(f804,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl22)),
  inference(unit_resulting_resolution,[],[f127,f126,f272,f322,f197])).
fof(f805,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl22)),
  inference(unit_resulting_resolution,[],[f127,f125,f272,f322,f197])).
fof(f809,plain,(
  s(bool,'F_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl22),
  inference(unit_resulting_resolution,[],[f127,f127,f322,f199])).
fof(f745,plain,(
  s(bool,'F_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl28),
  inference(cnf_transformation,[],[f745_D])).
fof(f745_D,plain,(
  s(bool,'F_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f713,plain,(
  s(bool,'F_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl21),
  inference(resolution,[],[f319,f134])).
fof(f712,plain,(
  s(bool,'F_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl21),
  inference(resolution,[],[f319,f136])).
fof(f747,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl31),
  inference(cnf_transformation,[],[f747_D])).
fof(f747_D,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f708,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK2),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl21),
  inference(resolution,[],[f319,f183])).
fof(f720,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl27),
  inference(cnf_transformation,[],[f720_D])).
fof(f720_D,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f621,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl21),
  inference(unit_resulting_resolution,[],[f319,f129])).
fof(f622,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl21),
  inference(unit_resulting_resolution,[],[f319,f131])).
fof(f623,plain,(
  s(bool,'F_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl21),
  inference(unit_resulting_resolution,[],[f319,f134])).
fof(f624,plain,(
  s(bool,'F_0') = s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl21),
  inference(unit_resulting_resolution,[],[f319,f136])).
fof(f631,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl21),
  inference(unit_resulting_resolution,[],[f126,f319,f180,f319,f191])).
fof(f632,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl21),
  inference(unit_resulting_resolution,[],[f125,f319,f180,f319,f191])).
fof(f633,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl18 | $spl21)),
  inference(unit_resulting_resolution,[],[f126,f127,f316,f319,f191])).
fof(f634,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl18 | $spl21)),
  inference(unit_resulting_resolution,[],[f125,f127,f316,f319,f191])).
fof(f638,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl21),
  inference(unit_resulting_resolution,[],[f319,f126,f180,f319,f191])).
fof(f640,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl21),
  inference(unit_resulting_resolution,[],[f319,f125,f180,f319,f191])).
fof(f645,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl21)),
  inference(unit_resulting_resolution,[],[f126,f126,f234,f319,f193])).
fof(f646,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl21)),
  inference(unit_resulting_resolution,[],[f125,f126,f234,f319,f193])).
fof(f647,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl21)),
  inference(unit_resulting_resolution,[],[f126,f125,f234,f319,f193])).
fof(f648,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl21)),
  inference(unit_resulting_resolution,[],[f125,f125,f234,f319,f193])).
fof(f649,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl21)),
  inference(unit_resulting_resolution,[],[f126,f126,f272,f319,f193])).
fof(f650,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl21)),
  inference(unit_resulting_resolution,[],[f125,f126,f272,f319,f193])).
fof(f651,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl21)),
  inference(unit_resulting_resolution,[],[f126,f125,f272,f319,f193])).
fof(f652,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl21)),
  inference(unit_resulting_resolution,[],[f125,f125,f272,f319,f193])).
fof(f655,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl21),
  inference(unit_resulting_resolution,[],[f126,f319,f180,f319,f195])).
fof(f656,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl21),
  inference(unit_resulting_resolution,[],[f125,f319,f180,f319,f195])).
fof(f657,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl18 | $spl21)),
  inference(unit_resulting_resolution,[],[f126,f127,f316,f319,f195])).
fof(f658,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl18 | $spl21)),
  inference(unit_resulting_resolution,[],[f125,f127,f316,f319,f195])).
fof(f662,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl21),
  inference(unit_resulting_resolution,[],[f319,f126,f180,f319,f195])).
fof(f664,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl21),
  inference(unit_resulting_resolution,[],[f319,f125,f180,f319,f195])).
fof(f669,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl21),
  inference(unit_resulting_resolution,[],[f127,f126,f126,f319,f196])).
fof(f670,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl21),
  inference(unit_resulting_resolution,[],[f127,f125,f126,f319,f196])).
fof(f671,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl21),
  inference(unit_resulting_resolution,[],[f127,f126,f125,f319,f196])).
fof(f672,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl21),
  inference(unit_resulting_resolution,[],[f127,f125,f125,f319,f196])).
fof(f677,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl21),
  inference(unit_resulting_resolution,[],[f126,f319,f126,f319,f196])).
fof(f678,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl21),
  inference(unit_resulting_resolution,[],[f125,f319,f126,f319,f196])).
fof(f681,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl21),
  inference(unit_resulting_resolution,[],[f126,f319,f125,f319,f196])).
fof(f682,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl21),
  inference(unit_resulting_resolution,[],[f125,f319,f125,f319,f196])).
fof(f687,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl21)),
  inference(unit_resulting_resolution,[],[f126,f126,f234,f319,f197])).
fof(f688,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl21)),
  inference(unit_resulting_resolution,[],[f125,f126,f234,f319,f197])).
fof(f689,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl21)),
  inference(unit_resulting_resolution,[],[f126,f125,f234,f319,f197])).
fof(f690,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl10 | $spl21)),
  inference(unit_resulting_resolution,[],[f125,f125,f234,f319,f197])).
fof(f691,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl21)),
  inference(unit_resulting_resolution,[],[f126,f126,f272,f319,f197])).
fof(f692,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl21)),
  inference(unit_resulting_resolution,[],[f125,f126,f272,f319,f197])).
fof(f693,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl21)),
  inference(unit_resulting_resolution,[],[f126,f125,f272,f319,f197])).
fof(f694,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | ($spl14 | $spl21)),
  inference(unit_resulting_resolution,[],[f125,f125,f272,f319,f197])).
fof(f697,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl21),
  inference(unit_resulting_resolution,[],[f127,f126,f319,f199])).
fof(f698,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl21),
  inference(unit_resulting_resolution,[],[f127,f125,f319,f199])).
fof(f702,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl21),
  inference(unit_resulting_resolution,[],[f126,f319,f319,f199])).
fof(f703,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))) | $spl21),
  inference(unit_resulting_resolution,[],[f125,f319,f319,f199])).
fof(f316,plain,(
  s(bool,'F_0') = s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1))))) | $spl18),
  inference(cnf_transformation,[],[f316_D])).
fof(f316_D,plain,(
  s(bool,'F_0') = s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1))))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f311,plain,(
  s(bool,'F_0') = s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1))))) | $spl1),
  inference(resolution,[],[f201,f134])).
fof(f310,plain,(
  s(bool,'F_0') = s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1))))) | $spl1),
  inference(resolution,[],[f201,f136])).
fof(f325,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl24),
  inference(cnf_transformation,[],[f325_D])).
fof(f325_D,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f323,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl1),
  inference(forward_demodulation,[],[f276,f180])).
fof(f276,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl1),
  inference(unit_resulting_resolution,[],[f201,f168])).
fof(f168,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK3(X1,X0)))),s(X0,sK4(X1,X0))))) | p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f107])).
fof(f322,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl22),
  inference(cnf_transformation,[],[f322_D])).
fof(f322_D,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f320,plain,(
  p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl1),
  inference(forward_demodulation,[],[f277,f180])).
fof(f277,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1))),s(sK0,sK4('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl1),
  inference(unit_resulting_resolution,[],[f201,f169])).
fof(f169,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK4(X1,X0)))),s(X0,sK5(X1,X0))))) | p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f107])).
fof(f319,plain,(
  ~p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl21),
  inference(cnf_transformation,[],[f319_D])).
fof(f319_D,plain,(
  ~p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f317,plain,(
  ~p(s(bool,'const.relation.RINTER_4'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl1),
  inference(forward_demodulation,[],[f278,f180])).
fof(f278,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),'const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1))),s(sK0,sK3('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0)))),s(sK0,sK5('const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)),sK0))))) | $spl1),
  inference(unit_resulting_resolution,[],[f201,f170])).
fof(f170,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK3(X1,X0)))),s(X0,sK5(X1,X0))))) | p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f107])).
fof(f314,plain,(
  s(bool,'T_0') != s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1))))) | $spl17),
  inference(cnf_transformation,[],[f314_D])).
fof(f314_D,plain,(
  s(bool,'T_0') != s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1))))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f279,plain,(
  s(bool,'T_0') != s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1))))) | $spl1),
  inference(unit_resulting_resolution,[],[f201,f129])).
fof(f280,plain,(
  s(bool,'T_0') != s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1))))) | $spl1),
  inference(unit_resulting_resolution,[],[f201,f131])).
fof(f281,plain,(
  s(bool,'F_0') = s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1))))) | $spl1),
  inference(unit_resulting_resolution,[],[f201,f134])).
fof(f282,plain,(
  s(bool,'F_0') = s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1))))) | $spl1),
  inference(unit_resulting_resolution,[],[f201,f136])).
fof(f289,plain,(
  s(bool,'T_0') != s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1))))) | $spl1),
  inference(unit_resulting_resolution,[],[f127,f126,f126,f201,f196])).
fof(f290,plain,(
  s(bool,'T_0') != s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1))))) | $spl1),
  inference(unit_resulting_resolution,[],[f127,f125,f126,f201,f196])).
fof(f291,plain,(
  s(bool,'T_0') != s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1))))) | $spl1),
  inference(unit_resulting_resolution,[],[f127,f126,f125,f201,f196])).
fof(f292,plain,(
  s(bool,'T_0') != s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1))))) | $spl1),
  inference(unit_resulting_resolution,[],[f127,f125,f125,f201,f196])).
fof(f295,plain,(
  s(bool,'T_0') != s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1))))) | $spl1),
  inference(unit_resulting_resolution,[],[f126,f201,f126,f201,f196])).
fof(f296,plain,(
  s(bool,'T_0') != s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1))))) | $spl1),
  inference(unit_resulting_resolution,[],[f125,f201,f126,f201,f196])).
fof(f299,plain,(
  s(bool,'T_0') != s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1))))) | $spl1),
  inference(unit_resulting_resolution,[],[f126,f201,f125,f201,f196])).
fof(f300,plain,(
  s(bool,'T_0') != s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1))))) | $spl1),
  inference(unit_resulting_resolution,[],[f125,f201,f125,f201,f196])).
fof(f301,plain,(
  s(bool,'T_0') != s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1))))) | $spl1),
  inference(unit_resulting_resolution,[],[f127,f126,f201,f199])).
fof(f302,plain,(
  s(bool,'T_0') != s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1))))) | $spl1),
  inference(unit_resulting_resolution,[],[f127,f125,f201,f199])).
fof(f305,plain,(
  s(bool,'T_0') != s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1))))) | $spl1),
  inference(unit_resulting_resolution,[],[f126,f201,f201,f199])).
fof(f306,plain,(
  s(bool,'T_0') != s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1))))) | $spl1),
  inference(unit_resulting_resolution,[],[f125,f201,f201,f199])).
fof(f272,plain,(
  s(bool,'T_0') = s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK2))) | $spl14),
  inference(cnf_transformation,[],[f272_D])).
fof(f272_D,plain,(
  s(bool,'T_0') = s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK2))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f273,plain,(
  s(bool,'T_0') = s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK2))) | $spl4),
  inference(global_subsumption,[],[f124,f123,f122,f205,f261,f260,f257,f256,f255,f249,f248,f247,f246,f266,f267,f268])).
fof(f268,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK2))) | $spl4),
  inference(superposition,[],[f205,f128])).
fof(f267,plain,(
  s(bool,'T_0') = s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK2))) | $spl4),
  inference(resolution,[],[f205,f130])).
fof(f266,plain,(
  s(bool,'T_0') = s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK2))) | $spl4),
  inference(resolution,[],[f205,f132])).
fof(f246,plain,(
  s(bool,'T_0') = s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK2))) | $spl4),
  inference(unit_resulting_resolution,[],[f205,f130])).
fof(f247,plain,(
  s(bool,'T_0') = s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK2))) | $spl4),
  inference(unit_resulting_resolution,[],[f205,f132])).
fof(f270,plain,(
  s(bool,'F_0') != s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK2))) | $spl13),
  inference(cnf_transformation,[],[f270_D])).
fof(f270_D,plain,(
  s(bool,'F_0') != s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK2))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f248,plain,(
  s(bool,'F_0') != s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK2))) | $spl4),
  inference(unit_resulting_resolution,[],[f205,f133])).
fof(f249,plain,(
  s(bool,'F_0') != s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK2))) | $spl4),
  inference(unit_resulting_resolution,[],[f205,f135])).
fof(f255,plain,(
  s(bool,'F_0') != s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK2))) | $spl4),
  inference(unit_resulting_resolution,[],[f127,f127,f126,f205,f196])).
fof(f256,plain,(
  s(bool,'F_0') != s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK2))) | $spl4),
  inference(unit_resulting_resolution,[],[f127,f127,f125,f205,f196])).
fof(f257,plain,(
  s(bool,'F_0') != s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK2))) | $spl4),
  inference(unit_resulting_resolution,[],[f127,f127,f205,f205,f196])).
fof(f260,plain,(
  s(bool,'F_0') != s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK2))) | $spl4),
  inference(unit_resulting_resolution,[],[f127,f127,f205,f205,f196])).
fof(f261,plain,(
  s(bool,'F_0') != s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK2))) | $spl4),
  inference(unit_resulting_resolution,[],[f127,f127,f205,f199])).
fof(f234,plain,(
  s(bool,'T_0') = s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK1))) | $spl10),
  inference(cnf_transformation,[],[f234_D])).
fof(f234_D,plain,(
  s(bool,'T_0') = s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK1))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f235,plain,(
  s(bool,'T_0') = s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK1))) | $spl2),
  inference(global_subsumption,[],[f124,f123,f122,f203,f221,f220,f219,f218,f217,f216,f215,f209,f208,f207,f206,f226,f227,f228])).
fof(f228,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK1))) | $spl2),
  inference(superposition,[],[f203,f128])).
fof(f227,plain,(
  s(bool,'T_0') = s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK1))) | $spl2),
  inference(resolution,[],[f203,f130])).
fof(f226,plain,(
  s(bool,'T_0') = s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK1))) | $spl2),
  inference(resolution,[],[f203,f132])).
fof(f206,plain,(
  s(bool,'T_0') = s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK1))) | $spl2),
  inference(unit_resulting_resolution,[],[f203,f130])).
fof(f207,plain,(
  s(bool,'T_0') = s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK1))) | $spl2),
  inference(unit_resulting_resolution,[],[f203,f132])).
fof(f230,plain,(
  s(bool,'F_0') != s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK1))) | $spl7),
  inference(cnf_transformation,[],[f230_D])).
fof(f230_D,plain,(
  s(bool,'F_0') != s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK1))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f208,plain,(
  s(bool,'F_0') != s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK1))) | $spl2),
  inference(unit_resulting_resolution,[],[f203,f133])).
fof(f209,plain,(
  s(bool,'F_0') != s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK1))) | $spl2),
  inference(unit_resulting_resolution,[],[f203,f135])).
fof(f215,plain,(
  s(bool,'F_0') != s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK1))) | $spl2),
  inference(unit_resulting_resolution,[],[f127,f127,f126,f203,f196])).
fof(f216,plain,(
  s(bool,'F_0') != s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK1))) | $spl2),
  inference(unit_resulting_resolution,[],[f127,f127,f125,f203,f196])).
fof(f217,plain,(
  s(bool,'F_0') != s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK1))) | $spl2),
  inference(unit_resulting_resolution,[],[f127,f127,f203,f203,f196])).
fof(f232,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl9),
  inference(cnf_transformation,[],[f232_D])).
fof(f232_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f218,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f127,f127,f126,f203,f196])).
fof(f219,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f127,f127,f125,f203,f196])).
fof(f220,plain,(
  s(bool,'F_0') != s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK1))) | $spl2),
  inference(unit_resulting_resolution,[],[f127,f127,f203,f203,f196])).
fof(f221,plain,(
  s(bool,'F_0') != s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK1))) | $spl2),
  inference(unit_resulting_resolution,[],[f127,f127,f203,f199])).
fof(f205,plain,(
  p(s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK2)))) | $spl4),
  inference(cnf_transformation,[],[f205_D])).
fof(f205_D,plain,(
  p(s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK2)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f122,plain,(
  p(s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK2))))),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  p(s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK2)))) & p(s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK1)))) & ~p(s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f66])).
fof(f66,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X2)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) & ~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),'const.relation.RINTER_2'(s(fun(X0,fun(X0,bool)),X2),s(fun(X0,fun(X0,bool)),X1)))))))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X2)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))) & ~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),'const.relation.RINTER_2'(s(fun(X0,fun(X0,bool)),X2),s(fun(X0,fun(X0,bool)),X1)))))))),
  inference(ennf_transformation,[],[f31])).
fof(f31,plain,(
  ~! [X0,X1,X2] : ((p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X2)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))) => p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),'const.relation.RINTER_2'(s(fun(X0,fun(X0,bool)),X2),s(fun(X0,fun(X0,bool)),X1)))))))),
  inference(flattening,[],[f30])).
fof(f30,plain,(
  ~! [X0] : ! [X1,X2] : ((p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X2)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))) => p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),'const.relation.RINTER_2'(s(fun(X0,fun(X0,bool)),X2),s(fun(X0,fun(X0,bool)),X1)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : ((p(s(bool,'const.relation.transitive_1'(s(fun(X6,fun(X6,bool)),X8)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X6,fun(X6,bool)),X7))))) => p(s(bool,'const.relation.transitive_1'(s(fun(X6,fun(X6,bool)),'const.relation.RINTER_2'(s(fun(X6,fun(X6,bool)),X8),s(fun(X6,fun(X6,bool)),X7)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : ((p(s(bool,'const.relation.transitive_1'(s(fun(X6,fun(X6,bool)),X8)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X6,fun(X6,bool)),X7))))) => p(s(bool,'const.relation.transitive_1'(s(fun(X6,fun(X6,bool)),'const.relation.RINTER_2'(s(fun(X6,fun(X6,bool)),X8),s(fun(X6,fun(X6,bool)),X7)))))))),
  file('thm.relation.transitive_RINTER','thm.relation.transitive_RINTER')).
fof(f203,plain,(
  p(s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK1)))) | $spl2),
  inference(cnf_transformation,[],[f203_D])).
fof(f203_D,plain,(
  p(s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK1)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f123,plain,(
  p(s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),sK1))))),
  inference(cnf_transformation,[],[f80])).
fof(f201,plain,(
  ~p(s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)))))) | $spl1),
  inference(cnf_transformation,[],[f201_D])).
fof(f201_D,plain,(
  ~p(s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1)))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f124,plain,(
  ~p(s(bool,'const.relation.transitive_1'(s(fun(sK0,fun(sK0,bool)),'const.relation.RINTER_2'(s(fun(sK0,fun(sK0,bool)),sK2),s(fun(sK0,fun(sK0,bool)),sK1))))))),
  inference(cnf_transformation,[],[f80])).
% SZS output end Proof for thm.relation
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 100
% Passive clauses: 116
% Generated clauses: 4885
% Final active clauses: 67
% Final passive clauses: 49
% Input formulas: 28
% Initial clauses: 76
% 
% Duplicate literals: 30
% Trivial inequalities: 30
% Fw demodulations: 1283
% Bw demodulations: 27
% Global subsumptions: 21
% 
% Simple tautologies: 78
% Equational tautologies: 2
% Forward subsumptions: 2402
% Backward subsumptions: 2
% 
% Binary resolution: 292
% Unit resulting resolution: 2366
% Forward superposition: 58
% Backward superposition: 635
% 
% Split clauses: 26
% Split components: 52
% SAT solver clauses: 2055
% SAT solver unit clauses: 6
% SAT solver binary clauses: 891
% 
% Sat splits: 26
% 
% Memory used [KB]: 3454
% Time elapsed: 0.207 s
% ------------------------------
% ------------------------------
% Success in time 18.234 s
