Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.ringNorm
% (61922)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 129
% Passive clauses: 2183
% Generated clauses: 23341
% Final active clauses: 124
% Final passive clauses: 1941
% Input formulas: 22
% Initial clauses: 23
% 
% Duplicate literals: 1528
% Trivial inequalities: 58
% Fw subsumption resolutions: 2160
% Fw demodulations: 3911
% Bw demodulations: 3
% 
% Simple tautologies: 203
% Equational tautologies: 19
% Forward subsumptions: 12081
% 
% Binary resolution: 13
% Forward superposition: 8884
% Backward superposition: 6339
% Self superposition: 173
% Equality factoring: 36
% Equality resolution: 5
% 
% Split clauses: 774
% Split components: 1553
% SAT solver clauses: 1258
% SAT solver unit clauses: 1
% SAT solver binary clauses: 20
% 
% Sat splits: 774
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1258
% TWLsolver calls for satisfiability: 85
% 
% Memory used [KB]: 15095
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ringNorm
% (61940)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 89
% Passive clauses: 1321
% Generated clauses: 12980
% Final active clauses: 84
% Final passive clauses: 1221
% Input formulas: 22
% Initial clauses: 23
% 
% Duplicate literals: 566
% Trivial inequalities: 52
% Fw subsumption resolutions: 554
% Fw demodulations: 1378
% Bw demodulations: 12
% Condensations: 677
% 
% Simple tautologies: 128
% Equational tautologies: 18
% Forward subsumptions: 6396
% Fw demodulations to eq. taut.: 2
% 
% Binary resolution: 11
% Forward superposition: 4514
% Backward superposition: 3943
% Self superposition: 178
% Equality factoring: 45
% Equality resolution: 4
% 
% Split clauses: 119
% Split components: 243
% SAT solver clauses: 1912
% SAT solver unit clauses: 40
% SAT solver binary clauses: 25
% 
% Sat splits: 119
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 1912
% TWLsolver calls for satisfiability: 110
% 
% Memory used [KB]: 17142
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.ringNorm
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ringNorm
% SZS output start Proof for thm.ringNorm
fof(f67878,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f67877])).
fof(f67877,plain,(
  s(sK0,'const.quote.varmap_find_2'(s('type.quote.index',sK3),s('type.quote.varmap'(sK0),sK2))) != s(sK0,'const.quote.varmap_find_2'(s('type.quote.index',sK3),s('type.quote.varmap'(sK0),sK2)))),
  inference(superposition,[],[f68,f7229])).
fof(f7229,plain,(
  ( ! [X10,X8,X11,X9] : (s(X8,'const.quote.varmap_find_2'(s('type.quote.index',X11),s('type.quote.varmap'(X8),X10))) = s(X8,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X8),X9),s('type.quote.varmap'(X8),X10),s('type.canonical.spolynom'(X8),'const.canonical.SPvar_1'(s('type.quote.index',X11)))))) )),
  inference(superposition,[],[f1979,f77])).
fof(f77,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.quote.varmap_find_2'(s('type.quote.index',X3),s('type.quote.varmap'(X0),X2))) = s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),'const.canonical.SPvar_1'(s('type.quote.index',X3)))))) )),
  inference(cnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.quote.varmap_find_2'(s('type.quote.index',X3),s('type.quote.varmap'(X0),X2))) = s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),'const.canonical.SPvar_1'(s('type.quote.index',X3)))))),
  inference(flattening,[],[f35])).
fof(f35,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.quote.varmap_find_2'(s('type.quote.index',X3),s('type.quote.varmap'(X0),X2))) = s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),'const.canonical.SPvar_1'(s('type.quote.index',X3)))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : ! [X12,X8,X9] : s(X6,'const.quote.varmap_find_2'(s('type.quote.index',X9),s('type.quote.varmap'(X6),X8))) = s(X6,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X6),X12),s('type.quote.varmap'(X6),X8),s('type.canonical.spolynom'(X6),'const.canonical.SPvar_1'(s('type.quote.index',X9)))))),
  file('thm.ringNorm.interp_sp_def|split|1','thm.canonical.interp_sp_def|split|1')).
fof(f1979,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))) = s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3)))) )),
  inference(forward_demodulation,[],[f1911,f83])).
fof(f83,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))) = s(X0,happ(s(fun('type.canonical.spolynom'(X0),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.ringNorm.r_interp_sp_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.canonical.spolynom'(X0),X3)))) )),
  inference(cnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))) = s(X0,happ(s(fun('type.canonical.spolynom'(X0),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.ringNorm.r_interp_sp_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.canonical.spolynom'(X0),X3)))),
  inference(flattening,[],[f47])).
fof(f47,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))) = s(X0,happ(s(fun('type.canonical.spolynom'(X0),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.ringNorm.r_interp_sp_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.canonical.spolynom'(X0),X3)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X20,X21,X22] : s(X6,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X6),X20),s('type.quote.varmap'(X6),X21),s('type.canonical.spolynom'(X6),X22))) = s(X6,happ(s(fun('type.canonical.spolynom'(X6),X6),happ(s(fun('type.quote.varmap'(X6),fun('type.canonical.spolynom'(X6),X6)),'const.ringNorm.r_interp_sp_1'(s('type.ring.ring'(X6),X20))),s('type.quote.varmap'(X6),X21))),s('type.canonical.spolynom'(X6),X22)))),
  file('thm.ringNorm.interp_sp_def|split|1','ARITY_1_3_const.ringNorm.r_interp_sp')).
fof(f1911,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))) = s(X0,happ(s(fun('type.canonical.spolynom'(X0),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.ringNorm.r_interp_sp_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.canonical.spolynom'(X0),X3)))) )),
  inference(superposition,[],[f82,f72])).
fof(f72,plain,(
  ( ! [X0,X1] : (s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.ringNorm.r_interp_sp_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.canonical.interp_sp_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f28])).
fof(f28,plain,(
  ! [X0,X1] : s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.ringNorm.r_interp_sp_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.canonical.interp_sp_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f27])).
fof(f27,plain,(
  ! [X0] : ! [X1] : s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.ringNorm.r_interp_sp_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.canonical.interp_sp_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6] : ! [X7] : s(fun('type.quote.varmap'(X6),fun('type.canonical.spolynom'(X6),X6)),'const.ringNorm.r_interp_sp_1'(s('type.ring.ring'(X6),X7))) = s(fun('type.quote.varmap'(X6),fun('type.canonical.spolynom'(X6),X6)),'const.canonical.interp_sp_1'(s('type.semi_ring.semi_ring'(X6),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X6),X7)))))),
  file('thm.ringNorm.interp_sp_def|split|1','thm.ringNorm.r_interp_sp_def')).
fof(f82,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))) = s(X0,happ(s(fun('type.canonical.spolynom'(X0),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.canonical.interp_sp_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.canonical.spolynom'(X0),X3)))) )),
  inference(cnf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))) = s(X0,happ(s(fun('type.canonical.spolynom'(X0),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.canonical.interp_sp_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.canonical.spolynom'(X0),X3)))),
  inference(flattening,[],[f45])).
fof(f45,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))) = s(X0,happ(s(fun('type.canonical.spolynom'(X0),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.canonical.interp_sp_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.canonical.spolynom'(X0),X3)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6] : ! [X20,X21,X22] : s(X6,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X6),X20),s('type.quote.varmap'(X6),X21),s('type.canonical.spolynom'(X6),X22))) = s(X6,happ(s(fun('type.canonical.spolynom'(X6),X6),happ(s(fun('type.quote.varmap'(X6),fun('type.canonical.spolynom'(X6),X6)),'const.canonical.interp_sp_1'(s('type.semi_ring.semi_ring'(X6),X20))),s('type.quote.varmap'(X6),X21))),s('type.canonical.spolynom'(X6),X22)))),
  file('thm.ringNorm.interp_sp_def|split|1','ARITY_1_3_const.canonical.interp_sp')).
fof(f68,plain,(
  s(sK0,'const.quote.varmap_find_2'(s('type.quote.index',sK3),s('type.quote.varmap'(sK0),sK2))) != s(sK0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(sK0),sK1),s('type.quote.varmap'(sK0),sK2),s('type.canonical.spolynom'(sK0),'const.canonical.SPvar_1'(s('type.quote.index',sK3)))))),
  inference(cnf_transformation,[],[f65])).
fof(f65,plain,(
  s(sK0,'const.quote.varmap_find_2'(s('type.quote.index',sK3),s('type.quote.varmap'(sK0),sK2))) != s(sK0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(sK0),sK1),s('type.quote.varmap'(sK0),sK2),s('type.canonical.spolynom'(sK0),'const.canonical.SPvar_1'(s('type.quote.index',sK3)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f63])).
fof(f63,plain,(
  ? [X0,X1,X2,X3] : s(X0,'const.quote.varmap_find_2'(s('type.quote.index',X3),s('type.quote.varmap'(X0),X2))) != s(X0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),'const.canonical.SPvar_1'(s('type.quote.index',X3)))))),
  inference(ennf_transformation,[],[f25])).
fof(f25,plain,(
  ~! [X0,X1,X2,X3] : s(X0,'const.quote.varmap_find_2'(s('type.quote.index',X3),s('type.quote.varmap'(X0),X2))) = s(X0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),'const.canonical.SPvar_1'(s('type.quote.index',X3)))))),
  inference(flattening,[],[f24])).
fof(f24,plain,(
  ~! [X0] : ! [X1,X2,X3] : s(X0,'const.quote.varmap_find_2'(s('type.quote.index',X3),s('type.quote.varmap'(X0),X2))) = s(X0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),'const.canonical.SPvar_1'(s('type.quote.index',X3)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : s(X6,'const.quote.varmap_find_2'(s('type.quote.index',X9),s('type.quote.varmap'(X6),X8))) = s(X6,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X6),X7),s('type.quote.varmap'(X6),X8),s('type.canonical.spolynom'(X6),'const.canonical.SPvar_1'(s('type.quote.index',X9)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : s(X6,'const.quote.varmap_find_2'(s('type.quote.index',X9),s('type.quote.varmap'(X6),X8))) = s(X6,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X6),X7),s('type.quote.varmap'(X6),X8),s('type.canonical.spolynom'(X6),'const.canonical.SPvar_1'(s('type.quote.index',X9)))))),
  file('thm.ringNorm.interp_sp_def|split|1','thm.ringNorm.interp_sp_def|split|1')).
% SZS output end Proof for thm.ringNorm
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 221
% Passive clauses: 6559
% Generated clauses: 76455
% Final active clauses: 115
% Final passive clauses: 1459
% Input formulas: 22
% Initial clauses: 23
% Discarded non-redundant clauses: 889
% 
% Duplicate literals: 2522
% Trivial inequalities: 540
% Fw subsumption resolutions: 3565
% Fw demodulations: 11308
% Bw demodulations: 3565
% Condensations: 4627
% Global subsumptions: 39
% 
% Simple tautologies: 677
% Equational tautologies: 203
% Forward subsumptions: 46217
% Fw demodulations to eq. taut.: 12
% Bw demodulations to eq. taut.: 405
% 
% Unit resulting resolution: 794
% Forward superposition: 21953
% Backward superposition: 17565
% Self superposition: 458
% Equality factoring: 136
% Equality resolution: 40
% 
% Split clauses: 161
% Split components: 328
% SAT solver clauses: 6945
% SAT solver unit clauses: 130
% SAT solver binary clauses: 1764
% 
% Sat splits: 161
% Sat splitting refutations: 23
% 
% Memory used [KB]: 53986
% Time elapsed: 4.325 s
% ------------------------------
% ------------------------------
% Success in time 5.649 s
