Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.arithmetic
% (86538)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 489
% Passive clauses: 5139
% Generated clauses: 33169
% Final active clauses: 215
% Final passive clauses: 2505
% Input formulas: 43
% Initial clauses: 82
% 
% Duplicate literals: 111
% Trivial inequalities: 248
% Fw subsumption resolutions: 1364
% Fw demodulations: 5785
% Bw demodulations: 130
% 
% Simple tautologies: 228
% Equational tautologies: 31
% Forward subsumptions: 20216
% Fw demodulations to eq. taut.: 53
% Bw demodulations to eq. taut.: 4
% 
% Binary resolution: 80
% Factoring: 1
% Forward superposition: 8823
% Backward superposition: 15878
% Self superposition: 112
% Equality factoring: 50
% Equality resolution: 40
% 
% Split clauses: 137
% Split components: 277
% SAT solver clauses: 407
% SAT solver unit clauses: 11
% SAT solver binary clauses: 75
% SAT solver learnt clauses: 3
% SAT solver learnt literals: 4
% 
% Sat splits: 137
% Sat splitting refutations: 78
% 
% TWLsolver clauses: 404
% TWLsolver calls for satisfiability: 114
% 
% Memory used [KB]: 17014
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.arithmetic
% (86555)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 386
% Passive clauses: 2765
% Generated clauses: 25922
% Final active clauses: 189
% Final passive clauses: 1008
% Input formulas: 43
% Initial clauses: 82
% 
% Duplicate literals: 58
% Trivial inequalities: 148
% Fw subsumption resolutions: 569
% Fw demodulations: 3775
% Bw demodulations: 1383
% Condensations: 105
% Global subsumptions: 13
% 
% Simple tautologies: 223
% Equational tautologies: 22
% Forward subsumptions: 14850
% Fw demodulations to eq. taut.: 31
% Bw demodulations to eq. taut.: 61
% 
% Binary resolution: 35
% Factoring: 1
% Forward superposition: 6246
% Backward superposition: 9328
% Self superposition: 83
% Equality factoring: 26
% Equality resolution: 32
% 
% Split clauses: 81
% Split components: 164
% SAT solver clauses: 2764
% SAT solver unit clauses: 109
% SAT solver binary clauses: 66
% SAT solver learnt clauses: 33
% SAT solver learnt literals: 13
% 
% Sat splits: 81
% Sat splitting refutations: 172
% 
% TWLsolver clauses: 2705
% TWLsolver calls for satisfiability: 530
% 
% Memory used [KB]: 14967
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.arithmetic
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.arithmetic
% SZS output start Proof for thm.arithmetic
fof(f69269,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f137,f225,f223,f221,f136,f224,f135,f222,f220,f238,f256,f257,f2617,f2641,f2639,f2753,f2760,f2758,f10947,f11140,f11139,f13250,f21472,f21473,f21480,f21477,f21474,f21475,f21483,f21476,f21485,f21478,f21482,f21515,f21516,f21517,f21518,f21533,f21534,f21535,f21536,f21538,f23488,f23683,f23682,f42149,f42150,f42162,f42160,f42151,f42152,f42164,f42163,f54076,f54078,f54082,f54087,f54090,f54092,f54100,f54110,f54115,f54121,f54127,f54131,f54140,f54144,f61721,f61760,f61758,f61723,f61765,f61763,f64688,f64679,f64690,f64692,f64680,f64703,f64701,f65075,f65235,f65233,f65094,f65240,f65097,f65241,f65099,f65243,f65100,f65244,f65253,f65254,f65262,f65263,f65265,f65256,f65266,f65147,f65278,f65149,f65280,f65151,f65282,f65153,f65284,f65163,f65289,f65169,f65291,f65183,f65296,f65217,f65299,f65218,f65301,f65239,f65367,f65369,f65371,f65374,f65377,f65379,f65382,f65383,f65388,f65391,f65395,f65398,f65427,f65476,f65489,f65494,f65502,f65508,f65521,f65522,f65525,f65531,f65537,f65540,f65543,f65552,f66140,f66253,f66246,f66248,f66148,f66257,f66255,f66258,f69265,f69266,f69268])).
fof(f69268,plain,(
  $false | ($spl0 | $spl3 | $spl14)),
  inference(subsumption_resolution,[],[f223,f66320])).
fof(f66320,plain,(
  s('type.num.num',sK0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | ($spl0 | $spl14)),
  inference(backward_demodulation,[],[f66272,f140])).
fof(f140,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  file('thm.arithmetic.MULT_EQ_|49|','thm.arithmetic.ONE')).
fof(f66272,plain,(
  s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl14)),
  inference(backward_demodulation,[],[f66270,f1835])).
fof(f1835,plain,(
  s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(sK0)))) | $spl0),
  inference(unit_resulting_resolution,[],[f1662,f153])).
fof(f153,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f55])).
fof(f55,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | ? [X1] : s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X8] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X8) | ? [X9] : s('type.num.num',X8) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))))),
  file('thm.arithmetic.MULT_EQ_|49|','thm.arithmetic.num_CASES')).
fof(f1662,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK0) | $spl0),
  inference(unit_resulting_resolution,[],[f1598,f214])).
fof(f214,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) & ((s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(flattening,[],[f132])).
fof(f132,plain,(
  ! [X0,X1] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1))) & ((s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X8,X9] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X8),s('type.num.num',X9))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X8) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X9)))),
  file('thm.arithmetic.MULT_EQ_|49|','thm.arithmetic.MULT_EQ_|48|')).
fof(f1598,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))) | $spl0),
  inference(superposition,[],[f523,f1543])).
fof(f1543,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl0),
  inference(superposition,[],[f202,f258])).
fof(f258,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK1))) | $spl0),
  inference(backward_demodulation,[],[f140,f220])).
fof(f202,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X8,X9] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X8),s('type.num.num',X9)))))),
  file('thm.arithmetic.MULT_EQ_|49|','thm.arithmetic.MULT_SUC')).
fof(f523,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(unit_resulting_resolution,[],[f195,f211])).
fof(f211,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1))) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(flattening,[],[f130])).
fof(f130,plain,(
  ! [X0,X1] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1))) & ((s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X8,X9] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X9))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X8) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X9)))),
  file('thm.arithmetic.MULT_EQ_|49|','thm.arithmetic.ADD_EQ_|48|')).
fof(f195,plain,(
  ( ! [X0] : (s('type.num.num',X0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : s('type.num.num',X0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),
  inference(flattening,[],[f75])).
fof(f75,plain,(
  ! [X0] : ~s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X9] : ~s('type.num.num',X9) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X9)))),
  file('thm.arithmetic.MULT_EQ_|49|','thm.prim_rec.SUC_ID')).
fof(f66270,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK2(sK0)) | ($spl0 | $spl14)),
  inference(global_subsumption,[],[f13598,f137,f136,f135,f138,f146,f147,f148,f149,f152,f153,f156,f155,f195,f200,f201,f202,f203,f207,f204,f209,f208,f212,f211,f210,f215,f214,f213,f216,f260,f293,f295,f322,f362,f363,f384,f385,f386,f387,f437,f470,f471,f481,f482,f517,f519,f521,f523,f547,f542,f543,f567,f569,f574,f619,f618,f704,f705,f736,f737,f762,f803,f807,f903,f962,f980,f1297,f1347,f1346,f1341,f1340,f1359,f1360,f1441,f1372,f1373,f1445,f1481,f1482,f1545,f1547,f1548,f1561,f1565,f1568,f1570,f1571,f1759,f1760,f1761,f1777,f1781,f1784,f1786,f1787,f2780,f2815,f3001,f3002,f2858,f3100,f3101,f3350,f3351,f3516,f3514,f3566,f3506,f3500,f3611,f3613,f3614,f3616,f3617,f3618,f3619,f3621,f3548,f3550,f3553,f3554,f3555,f3560,f3561,f4792,f4810,f4822,f4826,f4827,f4830,f4831,f4836,f5047,f5065,f5077,f5081,f5082,f5085,f5091,f5481,f5763,f5766,f5858,f6018,f8514,f8532,f8957,f8975,f9481,f9681,f9682,f9485,f9486,f9490,f9491,f9492,f9493,f9501,f9502,f9503,f9504,f9505,f9507,f9508,f9510,f9513,f9536,f9538,f9557,f9558,f9559,f9560,f9685,f9562,f9686,f9687,f9596,f9597,f9688,f9689,f9604,f10111,f10318,f10115,f10117,f10118,f10122,f10123,f10124,f10125,f10133,f10134,f10135,f10136,f10137,f10139,f10140,f10142,f10145,f10170,f10172,f10191,f10192,f10193,f10194,f10322,f10196,f10323,f10324,f10231,f10232,f10325,f10326,f10239,f11115,f11116,f11122,f11123,f11131,f10946,f10948,f10949,f10960,f10962,f10963,f10966,f10970,f10976,f10982,f10983,f11139,f18129,f18130,f18404,f18405,f18411,f18412,f18418,f18080,f18082,f18083,f19001,f19005,f19657,f19661,f21472,f21473,f21474,f21475,f21476,f2816,f2859,f1855,f1854,f226,f160,f159,f40560,f40562,f40563,f40564,f40611,f40614,f40615,f40616,f40620,f40621,f40622,f40623,f40625,f40626,f40627,f40628,f40629,f40630,f40631,f40632,f40633,f40635,f40638,f40679,f40693,f40695,f40697,f40841,f40699,f40854,f40855,f40858,f40859,f41395,f41397,f41398,f41399,f41400,f41449,f41452,f41458,f41459,f41460,f41461,f41463,f41464,f41465,f41466,f41467,f41468,f41469,f41473,f41474,f41477,f41518,f41532,f41534,f41536,f41683,f41538,f42149,f42150,f42151,f42152,f42167,f42474,f42475,f42476,f42477,f42481,f42486,f42495,f42500,f42501,f42502,f42503,f42518,f42520,f42522,f42563,f42565,f42567,f42597,f42598,f42599,f42600,f42601,f42968,f42969,f42970,f42971,f42859,f42861,f42863,f42893,f42894,f42895,f42896,f42897,f42951,f43466,f43467,f43468,f43469,f43401,f43403,f43405,f43436,f43437,f43438,f43439,f43440,f43445,f43628,f43629,f43630,f43631,f43640,f43649,f43654,f43655,f43656,f43657,f43672,f43674,f43676,f43707,f43708,f43709,f43710,f43711,f43726,f43728,f43730,f44127,f44129,f44131,f44162,f44163,f44164,f44165,f44166,f44360,f44361,f44362,f44363,f44264,f44268,f44631,f44633,f44635,f44666,f44667,f44668,f44669,f44670,f44818,f44819,f44820,f44821,f46644,f46645,f46646,f46647,f46321,f46323,f46325,f46356,f46357,f46358,f46359,f46360,f46758,f46759,f46760,f46386,f46391,f47538,f47539,f47540,f47199,f47566,f47567,f47568,f47569,f47278,f47280,f47282,f47313,f47314,f47315,f47316,f47317,f48030,f48032,f48034,f48065,f48066,f48067,f48068,f48069,f48554,f48555,f48556,f48557,f48598,f48599,f48600,f48147,f49330,f49331,f49332,f48949,f48951,f48953,f48984,f48985,f48986,f48987,f48988,f49430,f49431,f49432,f49433,f43611,f41448,f54271,f54272,f54273,f54274,f23411,f23410,f23402,f23398,f23125,f23121,f23106,f23105,f23097,f23093,f22733,f22729,f22714,f22713,f22705,f22701,f21169,f21168,f21160,f21156,f21130,f21126,f20259,f20258,f20250,f20246,f20220,f20216,f19645,f19641,f19626,f19625,f19617,f19613,f18989,f18985,f18970,f18969,f18961,f18957,f18394,f18387,f18384,f18198,f18197,f18085,f18084,f18076,f18072,f18056,f18055,f18047,f18043,f11106,f11098,f11097,f10973,f10969,f10965,f10964,f10955,f10952,f10931,f10930,f10921,f10918,f10190,f10189,f10187,f10186,f10138,f10132,f10131,f10129,f10128,f10110,f10109,f9556,f9555,f9553,f9552,f9506,f9500,f9499,f9497,f9496,f9480,f9479,f8990,f8989,f8982,f8979,f8960,f8959,f8952,f8949,f8517,f8516,f8509,f8506,f8489,f8488,f8481,f8478,f7275,f7258,f7257,f7250,f7247,f6966,f6949,f6948,f6941,f6938,f6584,f6535,f6534,f6527,f6524,f6487,f6438,f6437,f6430,f6427,f6357,f6354,f6351,f6350,f6347,f6346,f6345,f6343,f6338,f6337,f6335,f6331,f6316,f6315,f6313,f6308,f6305,f6169,f6166,f6163,f6162,f6159,f6158,f6157,f6155,f6150,f6149,f6147,f6146,f6143,f6128,f6127,f6125,f6120,f6117,f6001,f6000,f5993,f5990,f5841,f5840,f5833,f5830,f5717,f5716,f5709,f5706,f5432,f5431,f5424,f5421,f5088,f5084,f5080,f5079,f5072,f5071,f5069,f5068,f5050,f5049,f5042,f5039,f4833,f4829,f4825,f4824,f4817,f4816,f4814,f4813,f4795,f4794,f4787,f4784,f4090,f4081,f4077,f4075,f4065,f4064,f4061,f4060,f4059,f4058,f4057,f4056,f4053,f4051,f4049,f4034,f4033,f4019,f4015,f4013,f4003,f4002,f3999,f3998,f3997,f3996,f3995,f3994,f3991,f3989,f3987,f3933,f3932,f3850,f3841,f3837,f3836,f3826,f3825,f3822,f3821,f3820,f3819,f3818,f3817,f3814,f3812,f3810,f3791,f3790,f3780,f3776,f3775,f3765,f3764,f3761,f3760,f3759,f3758,f3757,f3756,f3753,f3751,f3749,f3699,f3620,f3615,f3610,f3603,f3602,f3557,f3552,f3543,f3541,f3518,f3497,f3494,f3490,f3488,f3340,f3338,f3090,f3088,f2513,f2297,f1802,f1741,f1587,f1540,f1368,f1366,f1355,f1353,f575,f538,f536,f514,f477,f475,f381,f380,f379,f378,f228,f227,f140,f56402,f56488,f56510,f57097,f57184,f57207,f57672,f57673,f57674,f57675,f57676,f57678,f57679,f57681,f57682,f57683,f57684,f57685,f57686,f57695,f57696,f57697,f57698,f57699,f57701,f57702,f57704,f57705,f57706,f57707,f57708,f57709,f57953,f57955,f57849,f57850,f57851,f57852,f57853,f57855,f57856,f57858,f57859,f57860,f57861,f57862,f57863,f57872,f57873,f57874,f57875,f57876,f57878,f57879,f57881,f57882,f57883,f57884,f57885,f57886,f58743,f58744,f58745,f58746,f58747,f58749,f58750,f58752,f58753,f58754,f58755,f58756,f58757,f58766,f58767,f58768,f58769,f58770,f58772,f58773,f58775,f58776,f58777,f58778,f58779,f58780,f59025,f59037,f58919,f58920,f58921,f58922,f58923,f58925,f58926,f58928,f58929,f58930,f58931,f58932,f58933,f58942,f58943,f58944,f58945,f58946,f58948,f58949,f58951,f58952,f58953,f58954,f58955,f58956,f61045,f61047,f61062,f61064,f61084,f61086,f61220,f61222,f61237,f61239,f61259,f61261])).
fof(f61261,plain,(
  ( ! [X10,X11,X9] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))))) | s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X11))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X9),s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))))),s('type.num.num',X10)))) )),
  inference(superposition,[],[f381,f1759])).
fof(f61259,plain,(
  ( ! [X4,X5,X3] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))))) | s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X5))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s('type.num.num',X4)))))) )),
  inference(superposition,[],[f381,f1545])).
fof(f61239,plain,(
  ( ! [X26,X27,X25] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X25))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X26))))))) | s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X27))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X25),s('type.num.num','const.num.SUC_1'(s('type.num.num',X26))))),s('type.num.num',X26)))))) )),
  inference(superposition,[],[f381,f1759])).
fof(f61237,plain,(
  ( ! [X21,X19,X20] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X21))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X20))))))) | s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X21))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))),s('type.num.num',X20)))))))) )),
  inference(superposition,[],[f381,f1545])).
fof(f61222,plain,(
  ( ! [X10,X11,X9] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))))) | s('type.num.num',X11) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X9),s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))))),s('type.num.num',X10)))))) )),
  inference(superposition,[],[f381,f1759])).
fof(f61220,plain,(
  ( ! [X4,X5,X3] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))))) | s('type.num.num',X5) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s('type.num.num',X4)))))))) )),
  inference(superposition,[],[f381,f1545])).
fof(f61086,plain,(
  ( ! [X10,X11,X9] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))))) | s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X11),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X9),s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))))),s('type.num.num',X10)))) )),
  inference(superposition,[],[f379,f1759])).
fof(f61084,plain,(
  ( ! [X4,X5,X3] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))))) | s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X5),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s('type.num.num',X4)))))) )),
  inference(superposition,[],[f379,f1545])).
fof(f61064,plain,(
  ( ! [X26,X27,X25] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X25))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X26))))))) | s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X27),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X25),s('type.num.num','const.num.SUC_1'(s('type.num.num',X26))))),s('type.num.num',X26)))))) )),
  inference(superposition,[],[f379,f1759])).
fof(f61062,plain,(
  ( ! [X21,X19,X20] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X21))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X20))))))) | s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X21),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))),s('type.num.num',X20)))))))) )),
  inference(superposition,[],[f379,f1545])).
fof(f61047,plain,(
  ( ! [X10,X11,X9] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))))) | s('type.num.num',X11) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X9),s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))))),s('type.num.num',X10))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f379,f1759])).
fof(f61045,plain,(
  ( ! [X4,X5,X3] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))))) | s('type.num.num',X5) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s('type.num.num',X4))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f379,f1545])).
fof(f58956,plain,(
  ( ! [X288,X290,X289] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X288),s('type.num.num',X288))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X289))),s('type.num.num',X290))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X288))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X289),s('type.num.num',X290))),s('type.num.num',X290)))) )),
  inference(superposition,[],[f1787,f1741])).
fof(f58955,plain,(
  ( ! [X285,X287,X286] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X286),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X287))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X285),s('type.num.num',X285))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X286) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X287),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X285))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X285)))))) )),
  inference(superposition,[],[f1787,f1741])).
fof(f58954,plain,(
  ( ! [X282,X283,X284] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X283),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X282),s('type.num.num',X282))))),s('type.num.num',X284))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X283) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X282))),s('type.num.num',X284))),s('type.num.num',X284)))) )),
  inference(superposition,[],[f1787,f1741])).
fof(f58953,plain,(
  ( ! [X280,X279,X281] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X280))),s('type.num.num',X281))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X279),s('type.num.num',X279))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X280),s('type.num.num',X281))),s('type.num.num',X281))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X279)))) )),
  inference(superposition,[],[f1786,f1741])).
fof(f58952,plain,(
  ( ! [X277,X276,X278] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X277))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X276),s('type.num.num',X276))))),s('type.num.num',X278))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X277),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X276))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X276))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X278)) )),
  inference(superposition,[],[f1786,f1741])).
fof(f58951,plain,(
  ( ! [X273,X275,X274] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X273),s('type.num.num',X273))))),s('type.num.num',X274))),s('type.num.num',X275))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X273))),s('type.num.num',X274))),s('type.num.num',X274))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X275)) )),
  inference(superposition,[],[f1786,f1741])).
fof(f58949,plain,(
  ( ! [X269,X267,X268] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X268),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X269))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X267),s('type.num.num',X267))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X269),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X267))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X267)))))) )),
  inference(superposition,[],[f1784,f1741])).
fof(f58948,plain,(
  ( ! [X265,X264,X266] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X265),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X264),s('type.num.num',X264))))),s('type.num.num',X266))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X264))),s('type.num.num',X266))),s('type.num.num',X266)))) )),
  inference(superposition,[],[f1784,f1741])).
fof(f58946,plain,(
  ( ! [X259,X260,X258] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X259))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X258),s('type.num.num',X258))))),s('type.num.num',X260))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X259),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X258))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X258)))))) )),
  inference(superposition,[],[f1781,f1741])).
fof(f58945,plain,(
  ( ! [X255,X257,X256] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X255),s('type.num.num',X255))))),s('type.num.num',X256))),s('type.num.num',X257))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X255))),s('type.num.num',X256))),s('type.num.num',X256)))) )),
  inference(superposition,[],[f1781,f1741])).
fof(f58944,plain,(
  ( ! [X253,X254,X252] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X252),s('type.num.num',X252))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X253))),s('type.num.num',X254))))) | s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X252))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X253),s('type.num.num',X254))),s('type.num.num',X254)))) )),
  inference(superposition,[],[f1777,f1741])).
fof(f58943,plain,(
  ( ! [X251,X249,X250] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X250))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X251))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X249),s('type.num.num',X249))))))) | s('type.num.num',X250) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X251),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X249))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X249)))))) )),
  inference(superposition,[],[f1777,f1741])).
fof(f58942,plain,(
  ( ! [X246,X248,X247] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X247))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X246),s('type.num.num',X246))))),s('type.num.num',X248))))) | s('type.num.num',X247) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X246))),s('type.num.num',X248))),s('type.num.num',X248)))) )),
  inference(superposition,[],[f1777,f1741])).
fof(f58933,plain,(
  ( ! [X231,X229,X230] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X229),s('type.num.num',X229))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X230),s('type.num.num','const.num.SUC_1'(s('type.num.num',X231))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X229))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X230),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X230),s('type.num.num',X231)))))) )),
  inference(superposition,[],[f1571,f1741])).
fof(f58932,plain,(
  ( ! [X227,X228,X226] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X227),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X228),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X226),s('type.num.num',X226))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X227) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X228),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X228),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X226)))))))) )),
  inference(superposition,[],[f1571,f1741])).
fof(f58931,plain,(
  ( ! [X225,X223,X224] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X224),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X223),s('type.num.num',X223))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X225))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X224) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X223))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X223))),s('type.num.num',X225)))))) )),
  inference(superposition,[],[f1571,f1741])).
fof(f58930,plain,(
  ( ! [X222,X220,X221] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X221),s('type.num.num','const.num.SUC_1'(s('type.num.num',X222))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X220),s('type.num.num',X220))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X221),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X221),s('type.num.num',X222))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X220)))) )),
  inference(superposition,[],[f1570,f1741])).
fof(f58929,plain,(
  ( ! [X218,X219,X217] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X218),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X217),s('type.num.num',X217))))))),s('type.num.num',X219))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X218),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X218),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X217))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X219)) )),
  inference(superposition,[],[f1570,f1741])).
fof(f58928,plain,(
  ( ! [X216,X215,X214] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X214),s('type.num.num',X214))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X215))))),s('type.num.num',X216))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X214))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X214))),s('type.num.num',X215))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X216)) )),
  inference(superposition,[],[f1570,f1741])).
fof(f58926,plain,(
  ( ! [X210,X208,X209] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X209),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X210),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X208),s('type.num.num',X208))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X210),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X210),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X208)))))))) )),
  inference(superposition,[],[f1568,f1741])).
fof(f58925,plain,(
  ( ! [X206,X207,X205] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X206),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X205),s('type.num.num',X205))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X207))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X205))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X205))),s('type.num.num',X207)))))) )),
  inference(superposition,[],[f1568,f1741])).
fof(f58923,plain,(
  ( ! [X200,X199,X201] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X200),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X199),s('type.num.num',X199))))))),s('type.num.num',X201))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X200),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X200),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X199)))))))) )),
  inference(superposition,[],[f1565,f1741])).
fof(f58922,plain,(
  ( ! [X198,X196,X197] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X196),s('type.num.num',X196))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X197))))),s('type.num.num',X198))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X196))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X196))),s('type.num.num',X197)))))) )),
  inference(superposition,[],[f1565,f1741])).
fof(f58921,plain,(
  ( ! [X194,X195,X193] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X193),s('type.num.num',X193))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X194),s('type.num.num','const.num.SUC_1'(s('type.num.num',X195))))))) | s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X193))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X194),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X194),s('type.num.num',X195)))))) )),
  inference(superposition,[],[f1561,f1741])).
fof(f58920,plain,(
  ( ! [X191,X192,X190] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X191))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X192),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X190),s('type.num.num',X190))))))))) | s('type.num.num',X191) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X192),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X192),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X190)))))))) )),
  inference(superposition,[],[f1561,f1741])).
fof(f58919,plain,(
  ( ! [X189,X187,X188] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X188))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X187),s('type.num.num',X187))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X189))))))) | s('type.num.num',X188) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X187))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X187))),s('type.num.num',X189)))))) )),
  inference(superposition,[],[f1561,f1741])).
fof(f59037,plain,(
  ( ! [X36] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X36),s('type.num.num',X36))),s('type.num.num',X36)))))))))) )),
  inference(forward_demodulation,[],[f59036,f201])).
fof(f59036,plain,(
  ( ! [X36] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X36),s('type.num.num',X36))))),s('type.num.num',X36)))))))) )),
  inference(forward_demodulation,[],[f59035,f201])).
fof(f59035,plain,(
  ( ! [X36] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num',X36))),s('type.num.num',X36)))))))) )),
  inference(forward_demodulation,[],[f59034,f201])).
fof(f59034,plain,(
  ( ! [X36] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num',X36))))),s('type.num.num',X36)))))) )),
  inference(forward_demodulation,[],[f58839,f200])).
fof(f58839,plain,(
  ( ! [X36] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))))),s('type.num.num',X36)))))) )),
  inference(superposition,[],[f1759,f1741])).
fof(f59025,plain,(
  ( ! [X26,X27] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X26),s('type.num.num',X26))))))))) | s('type.num.num',X27) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X26)))))) )),
  inference(forward_demodulation,[],[f59024,f201])).
fof(f59024,plain,(
  ( ! [X26,X27] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X26))),s('type.num.num',X26))))))) | s('type.num.num',X27) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X26)))))) )),
  inference(forward_demodulation,[],[f58834,f200])).
fof(f58834,plain,(
  ( ! [X26,X27] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X26))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X26))))))) | s('type.num.num',X27) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X26)))))) )),
  inference(superposition,[],[f1561,f1741])).
fof(f58780,plain,(
  ( ! [X280,X279,X281] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X279))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X280))),s('type.num.num',X281))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X279),s('type.num.num',X279))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X280),s('type.num.num',X281))),s('type.num.num',X281)))) )),
  inference(superposition,[],[f1787,f1741])).
fof(f58779,plain,(
  ( ! [X277,X276,X278] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X277),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X278))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X276))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X277) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X278),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X276),s('type.num.num',X276))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X276),s('type.num.num',X276)))))) )),
  inference(superposition,[],[f1787,f1741])).
fof(f58778,plain,(
  ( ! [X273,X275,X274] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X274),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X273))))),s('type.num.num',X275))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X274) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X273),s('type.num.num',X273))),s('type.num.num',X275))),s('type.num.num',X275)))) )),
  inference(superposition,[],[f1787,f1741])).
fof(f58777,plain,(
  ( ! [X271,X270,X272] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X271))),s('type.num.num',X272))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X270))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X271),s('type.num.num',X272))),s('type.num.num',X272))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X270),s('type.num.num',X270)))) )),
  inference(superposition,[],[f1786,f1741])).
fof(f58776,plain,(
  ( ! [X269,X267,X268] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X268))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X267))))),s('type.num.num',X269))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X268),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X267),s('type.num.num',X267))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X267),s('type.num.num',X267))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X269)) )),
  inference(superposition,[],[f1786,f1741])).
fof(f58775,plain,(
  ( ! [X265,X264,X266] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X264))))),s('type.num.num',X265))),s('type.num.num',X266))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X264),s('type.num.num',X264))),s('type.num.num',X265))),s('type.num.num',X265))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X266)) )),
  inference(superposition,[],[f1786,f1741])).
fof(f58773,plain,(
  ( ! [X259,X260,X258] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X259),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X260))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X258))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X260),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X258),s('type.num.num',X258))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X258),s('type.num.num',X258)))))) )),
  inference(superposition,[],[f1784,f1741])).
fof(f58772,plain,(
  ( ! [X255,X257,X256] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X256),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X255))))),s('type.num.num',X257))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X255),s('type.num.num',X255))),s('type.num.num',X257))),s('type.num.num',X257)))) )),
  inference(superposition,[],[f1784,f1741])).
fof(f58770,plain,(
  ( ! [X251,X249,X250] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X250))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X249))))),s('type.num.num',X251))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X250),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X249),s('type.num.num',X249))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X249),s('type.num.num',X249)))))) )),
  inference(superposition,[],[f1781,f1741])).
fof(f58769,plain,(
  ( ! [X246,X248,X247] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X246))))),s('type.num.num',X247))),s('type.num.num',X248))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X246),s('type.num.num',X246))),s('type.num.num',X247))),s('type.num.num',X247)))) )),
  inference(superposition,[],[f1781,f1741])).
fof(f58768,plain,(
  ( ! [X243,X244,X245] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X244))),s('type.num.num',X245))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X243))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X243),s('type.num.num',X243))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X244),s('type.num.num',X245))),s('type.num.num',X245)))) )),
  inference(superposition,[],[f1777,f1741])).
fof(f58767,plain,(
  ( ! [X241,X242,X240] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X241))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X242))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X240))))))) | s('type.num.num',X241) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X242),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X240),s('type.num.num',X240))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X240),s('type.num.num',X240)))))) )),
  inference(superposition,[],[f1777,f1741])).
fof(f58766,plain,(
  ( ! [X239,X237,X238] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X238))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X237))))),s('type.num.num',X239))))) | s('type.num.num',X238) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X237),s('type.num.num',X237))),s('type.num.num',X239))),s('type.num.num',X239)))) )),
  inference(superposition,[],[f1777,f1741])).
fof(f58757,plain,(
  ( ! [X222,X220,X221] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X220))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X221),s('type.num.num','const.num.SUC_1'(s('type.num.num',X222))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X220),s('type.num.num',X220))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X221),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X221),s('type.num.num',X222)))))) )),
  inference(superposition,[],[f1571,f1741])).
fof(f58756,plain,(
  ( ! [X218,X219,X217] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X218),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X219),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X217))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X218) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X219),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X219),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X217),s('type.num.num',X217)))))))) )),
  inference(superposition,[],[f1571,f1741])).
fof(f58755,plain,(
  ( ! [X216,X215,X214] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X215),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X214))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X216))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X215) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X214),s('type.num.num',X214))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X214),s('type.num.num',X214))),s('type.num.num',X216)))))) )),
  inference(superposition,[],[f1571,f1741])).
fof(f58754,plain,(
  ( ! [X212,X213,X211] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X212),s('type.num.num','const.num.SUC_1'(s('type.num.num',X213))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X211))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X212),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X212),s('type.num.num',X213))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X211),s('type.num.num',X211)))) )),
  inference(superposition,[],[f1570,f1741])).
fof(f58753,plain,(
  ( ! [X210,X208,X209] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X209),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X208))))))),s('type.num.num',X210))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X209),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X209),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X208),s('type.num.num',X208))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X210)) )),
  inference(superposition,[],[f1570,f1741])).
fof(f58752,plain,(
  ( ! [X206,X207,X205] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X205))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X206))))),s('type.num.num',X207))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X205),s('type.num.num',X205))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X205),s('type.num.num',X205))),s('type.num.num',X206))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X207)) )),
  inference(superposition,[],[f1570,f1741])).
fof(f58750,plain,(
  ( ! [X200,X199,X201] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X200),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X201),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X199))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X201),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X201),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X199),s('type.num.num',X199)))))))) )),
  inference(superposition,[],[f1568,f1741])).
fof(f58749,plain,(
  ( ! [X198,X196,X197] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X197),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X196))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X198))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X196),s('type.num.num',X196))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X196),s('type.num.num',X196))),s('type.num.num',X198)))))) )),
  inference(superposition,[],[f1568,f1741])).
fof(f58747,plain,(
  ( ! [X191,X192,X190] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X191),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X190))))))),s('type.num.num',X192))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X191),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X191),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X190),s('type.num.num',X190)))))))) )),
  inference(superposition,[],[f1565,f1741])).
fof(f58746,plain,(
  ( ! [X189,X187,X188] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X187))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X188))))),s('type.num.num',X189))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X187),s('type.num.num',X187))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X187),s('type.num.num',X187))),s('type.num.num',X188)))))) )),
  inference(superposition,[],[f1565,f1741])).
fof(f58745,plain,(
  ( ! [X185,X186,X184] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X185),s('type.num.num','const.num.SUC_1'(s('type.num.num',X186))))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X184))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X184),s('type.num.num',X184))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X185),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X185),s('type.num.num',X186)))))) )),
  inference(superposition,[],[f1561,f1741])).
fof(f58744,plain,(
  ( ! [X182,X183,X181] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X182))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X183),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X181))))))))) | s('type.num.num',X182) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X183),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X183),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X181),s('type.num.num',X181)))))))) )),
  inference(superposition,[],[f1561,f1741])).
fof(f58743,plain,(
  ( ! [X180,X178,X179] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X179))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X178))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X180))))))) | s('type.num.num',X179) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X178),s('type.num.num',X178))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X178),s('type.num.num',X178))),s('type.num.num',X180)))))) )),
  inference(superposition,[],[f1561,f1741])).
fof(f57886,plain,(
  ( ! [X290,X289,X291] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X289),s('type.num.num',X289))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X290))),s('type.num.num',X291))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X289),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X290),s('type.num.num',X291))),s('type.num.num',X291)))) )),
  inference(superposition,[],[f1787,f1540])).
fof(f57885,plain,(
  ( ! [X288,X287,X286] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X287),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X288))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X286),s('type.num.num',X286))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X287) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X288),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X286),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X286),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(superposition,[],[f1787,f1540])).
fof(f57884,plain,(
  ( ! [X285,X283,X284] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X284),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X283),s('type.num.num',X283))))),s('type.num.num',X285))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X284) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X283),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num',X285))),s('type.num.num',X285)))) )),
  inference(superposition,[],[f1787,f1540])).
fof(f57883,plain,(
  ( ! [X280,X282,X281] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X281))),s('type.num.num',X282))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X280),s('type.num.num',X280))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X281),s('type.num.num',X282))),s('type.num.num',X282))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X280),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f1786,f1540])).
fof(f57882,plain,(
  ( ! [X277,X279,X278] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X278))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X277),s('type.num.num',X277))))),s('type.num.num',X279))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X278),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X277),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X277),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X279)) )),
  inference(superposition,[],[f1786,f1540])).
fof(f57881,plain,(
  ( ! [X275,X276,X274] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X274),s('type.num.num',X274))))),s('type.num.num',X275))),s('type.num.num',X276))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X274),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num',X275))),s('type.num.num',X275))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X276)) )),
  inference(superposition,[],[f1786,f1540])).
fof(f57879,plain,(
  ( ! [X269,X268,X270] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X269),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X270))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X268),s('type.num.num',X268))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X270),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X268),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X268),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(superposition,[],[f1784,f1540])).
fof(f57878,plain,(
  ( ! [X265,X267,X266] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X266),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X265),s('type.num.num',X265))))),s('type.num.num',X267))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X265),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num',X267))),s('type.num.num',X267)))) )),
  inference(superposition,[],[f1784,f1540])).
fof(f57876,plain,(
  ( ! [X261,X259,X260] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X260))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X259),s('type.num.num',X259))))),s('type.num.num',X261))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X260),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X259),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X259),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(superposition,[],[f1781,f1540])).
fof(f57875,plain,(
  ( ! [X257,X256,X258] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X256),s('type.num.num',X256))))),s('type.num.num',X257))),s('type.num.num',X258))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X256),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num',X257))),s('type.num.num',X257)))) )),
  inference(superposition,[],[f1781,f1540])).
fof(f57874,plain,(
  ( ! [X255,X253,X254] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X253),s('type.num.num',X253))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X254))),s('type.num.num',X255))))) | s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X253),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X254),s('type.num.num',X255))),s('type.num.num',X255)))) )),
  inference(superposition,[],[f1777,f1540])).
fof(f57873,plain,(
  ( ! [X251,X252,X250] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X251))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X252))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X250),s('type.num.num',X250))))))) | s('type.num.num',X251) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X252),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X250),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X250),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(superposition,[],[f1777,f1540])).
fof(f57872,plain,(
  ( ! [X249,X248,X247] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X248))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X247),s('type.num.num',X247))))),s('type.num.num',X249))))) | s('type.num.num',X248) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X247),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num',X249))),s('type.num.num',X249)))) )),
  inference(superposition,[],[f1777,f1540])).
fof(f57863,plain,(
  ( ! [X231,X230,X232] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X230),s('type.num.num',X230))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X231),s('type.num.num','const.num.SUC_1'(s('type.num.num',X232))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X230),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X231),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X231),s('type.num.num',X232)))))) )),
  inference(superposition,[],[f1571,f1540])).
fof(f57862,plain,(
  ( ! [X229,X227,X228] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X228),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X229),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X227),s('type.num.num',X227))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X228) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X229),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X229),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X227),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))))) )),
  inference(superposition,[],[f1571,f1540])).
fof(f57861,plain,(
  ( ! [X225,X226,X224] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X225),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X224),s('type.num.num',X224))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X226))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X225) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X224),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X224),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num',X226)))))) )),
  inference(superposition,[],[f1571,f1540])).
fof(f57860,plain,(
  ( ! [X222,X223,X221] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X222),s('type.num.num','const.num.SUC_1'(s('type.num.num',X223))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X221),s('type.num.num',X221))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X222),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X222),s('type.num.num',X223))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X221),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f1570,f1540])).
fof(f57859,plain,(
  ( ! [X218,X220,X219] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X219),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X218),s('type.num.num',X218))))))),s('type.num.num',X220))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X219),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X219),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X218),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X220)) )),
  inference(superposition,[],[f1570,f1540])).
fof(f57858,plain,(
  ( ! [X216,X215,X217] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X215),s('type.num.num',X215))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X216))))),s('type.num.num',X217))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X215),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X215),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num',X216))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X217)) )),
  inference(superposition,[],[f1570,f1540])).
fof(f57856,plain,(
  ( ! [X210,X211,X209] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X210),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X211),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X209),s('type.num.num',X209))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X211),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X211),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X209),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))))) )),
  inference(superposition,[],[f1568,f1540])).
fof(f57855,plain,(
  ( ! [X208,X206,X207] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X207),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X206),s('type.num.num',X206))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X208))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X206),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X206),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num',X208)))))) )),
  inference(superposition,[],[f1568,f1540])).
fof(f57853,plain,(
  ( ! [X202,X200,X201] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X201),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X200),s('type.num.num',X200))))))),s('type.num.num',X202))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X201),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X201),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X200),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))))) )),
  inference(superposition,[],[f1565,f1540])).
fof(f57852,plain,(
  ( ! [X198,X199,X197] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X197),s('type.num.num',X197))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X198))))),s('type.num.num',X199))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X197),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X197),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num',X198)))))) )),
  inference(superposition,[],[f1565,f1540])).
fof(f57851,plain,(
  ( ! [X196,X194,X195] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X194),s('type.num.num',X194))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X195),s('type.num.num','const.num.SUC_1'(s('type.num.num',X196))))))) | s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X194),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X195),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X195),s('type.num.num',X196)))))) )),
  inference(superposition,[],[f1561,f1540])).
fof(f57850,plain,(
  ( ! [X191,X192,X193] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X192))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X193),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X191),s('type.num.num',X191))))))))) | s('type.num.num',X192) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X193),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X193),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X191),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))))) )),
  inference(superposition,[],[f1561,f1540])).
fof(f57849,plain,(
  ( ! [X189,X190,X188] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X189))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X188),s('type.num.num',X188))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X190))))))) | s('type.num.num',X189) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X188),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X188),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num',X190)))))) )),
  inference(superposition,[],[f1561,f1540])).
fof(f57955,plain,(
  ( ! [X28,X27] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',X27))))))))) | s('type.num.num',X28) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X27),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(forward_demodulation,[],[f57954,f201])).
fof(f57954,plain,(
  ( ! [X28,X27] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))),s('type.num.num',X27))))))) | s('type.num.num',X28) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X27),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(forward_demodulation,[],[f57759,f200])).
fof(f57759,plain,(
  ( ! [X28,X27] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))))))) | s('type.num.num',X28) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X27),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f1777,f1540])).
fof(f57953,plain,(
  ( ! [X24] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X24))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',X24)))))))))))) )),
  inference(forward_demodulation,[],[f57952,f200])).
fof(f57952,plain,(
  ( ! [X24] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X24))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num',X24)))))))))))) )),
  inference(forward_demodulation,[],[f57951,f201])).
fof(f57951,plain,(
  ( ! [X24] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X24))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X24))),s('type.num.num',X24)))))))))) )),
  inference(forward_demodulation,[],[f57950,f200])).
fof(f57950,plain,(
  ( ! [X24] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X24))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X24))),s('type.num.num',X24)))))))))) )),
  inference(forward_demodulation,[],[f57756,f200])).
fof(f57756,plain,(
  ( ! [X24] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X24))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X24),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X24))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X24)))))))))) )),
  inference(superposition,[],[f1545,f1540])).
fof(f57709,plain,(
  ( ! [X277,X276,X278] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X276),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X277))),s('type.num.num',X278))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X276),s('type.num.num',X276))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X277),s('type.num.num',X278))),s('type.num.num',X278)))) )),
  inference(superposition,[],[f1787,f1540])).
fof(f57708,plain,(
  ( ! [X273,X275,X274] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X274),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X275))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X273),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X274) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X275),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X273),s('type.num.num',X273))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X273),s('type.num.num',X273)))))) )),
  inference(superposition,[],[f1787,f1540])).
fof(f57707,plain,(
  ( ! [X271,X270,X272] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X271),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X270),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.num.num',X272))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X271) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X270),s('type.num.num',X270))),s('type.num.num',X272))),s('type.num.num',X272)))) )),
  inference(superposition,[],[f1787,f1540])).
fof(f57706,plain,(
  ( ! [X269,X267,X268] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X268))),s('type.num.num',X269))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X267),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X268),s('type.num.num',X269))),s('type.num.num',X269))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X267),s('type.num.num',X267)))) )),
  inference(superposition,[],[f1786,f1540])).
fof(f57705,plain,(
  ( ! [X265,X264,X266] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X265))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X264),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.num.num',X266))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X265),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X264),s('type.num.num',X264))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X264),s('type.num.num',X264))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X266)) )),
  inference(superposition,[],[f1786,f1540])).
fof(f57704,plain,(
  ( ! [X261,X263,X262] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X261),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.num.num',X262))),s('type.num.num',X263))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X261),s('type.num.num',X261))),s('type.num.num',X262))),s('type.num.num',X262))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X263)) )),
  inference(superposition,[],[f1786,f1540])).
fof(f57702,plain,(
  ( ! [X255,X257,X256] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X256),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X257))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X255),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X257),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X255),s('type.num.num',X255))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X255),s('type.num.num',X255)))))) )),
  inference(superposition,[],[f1784,f1540])).
fof(f57701,plain,(
  ( ! [X253,X254,X252] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X253),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X252),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.num.num',X254))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X252),s('type.num.num',X252))),s('type.num.num',X254))),s('type.num.num',X254)))) )),
  inference(superposition,[],[f1784,f1540])).
fof(f57699,plain,(
  ( ! [X246,X248,X247] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X247))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X246),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.num.num',X248))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X247),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X246),s('type.num.num',X246))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X246),s('type.num.num',X246)))))) )),
  inference(superposition,[],[f1781,f1540])).
fof(f57698,plain,(
  ( ! [X243,X244,X245] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X243),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.num.num',X244))),s('type.num.num',X245))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X243),s('type.num.num',X243))),s('type.num.num',X244))),s('type.num.num',X244)))) )),
  inference(superposition,[],[f1781,f1540])).
fof(f57697,plain,(
  ( ! [X241,X242,X240] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X241))),s('type.num.num',X242))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X240),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X240),s('type.num.num',X240))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X241),s('type.num.num',X242))),s('type.num.num',X242)))) )),
  inference(superposition,[],[f1777,f1540])).
fof(f57696,plain,(
  ( ! [X239,X237,X238] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X238))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X239))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X237),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))))) | s('type.num.num',X238) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X239),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X237),s('type.num.num',X237))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X237),s('type.num.num',X237)))))) )),
  inference(superposition,[],[f1777,f1540])).
fof(f57695,plain,(
  ( ! [X235,X236,X234] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X235))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X234),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.num.num',X236))))) | s('type.num.num',X235) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X234),s('type.num.num',X234))),s('type.num.num',X236))),s('type.num.num',X236)))) )),
  inference(superposition,[],[f1777,f1540])).
fof(f57686,plain,(
  ( ! [X218,X219,X217] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X217),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X218),s('type.num.num','const.num.SUC_1'(s('type.num.num',X219))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X217),s('type.num.num',X217))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X218),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X218),s('type.num.num',X219)))))) )),
  inference(superposition,[],[f1571,f1540])).
fof(f57685,plain,(
  ( ! [X216,X215,X214] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X215),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X216),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X214),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X215) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X216),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X216),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X214),s('type.num.num',X214)))))))) )),
  inference(superposition,[],[f1571,f1540])).
fof(f57684,plain,(
  ( ! [X212,X213,X211] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X212),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X211),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X213))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X212) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X211),s('type.num.num',X211))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X211),s('type.num.num',X211))),s('type.num.num',X213)))))) )),
  inference(superposition,[],[f1571,f1540])).
fof(f57683,plain,(
  ( ! [X210,X208,X209] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X209),s('type.num.num','const.num.SUC_1'(s('type.num.num',X210))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X208),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X209),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X209),s('type.num.num',X210))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X208),s('type.num.num',X208)))) )),
  inference(superposition,[],[f1570,f1540])).
fof(f57682,plain,(
  ( ! [X206,X207,X205] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X206),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X205),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))))),s('type.num.num',X207))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X206),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X206),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X205),s('type.num.num',X205))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X207)) )),
  inference(superposition,[],[f1570,f1540])).
fof(f57681,plain,(
  ( ! [X204,X202,X203] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X202),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X203))))),s('type.num.num',X204))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X202),s('type.num.num',X202))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X202),s('type.num.num',X202))),s('type.num.num',X203))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X204)) )),
  inference(superposition,[],[f1570,f1540])).
fof(f57679,plain,(
  ( ! [X198,X196,X197] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X197),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X198),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X196),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X198),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X198),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X196),s('type.num.num',X196)))))))) )),
  inference(superposition,[],[f1568,f1540])).
fof(f57678,plain,(
  ( ! [X194,X195,X193] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X194),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X193),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X195))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X193),s('type.num.num',X193))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X193),s('type.num.num',X193))),s('type.num.num',X195)))))) )),
  inference(superposition,[],[f1568,f1540])).
fof(f57676,plain,(
  ( ! [X189,X187,X188] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X188),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X187),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))))),s('type.num.num',X189))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X188),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X188),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X187),s('type.num.num',X187)))))))) )),
  inference(superposition,[],[f1565,f1540])).
fof(f57675,plain,(
  ( ! [X185,X186,X184] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X184),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X185))))),s('type.num.num',X186))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X184),s('type.num.num',X184))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X184),s('type.num.num',X184))),s('type.num.num',X185)))))) )),
  inference(superposition,[],[f1565,f1540])).
fof(f57674,plain,(
  ( ! [X182,X183,X181] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X182),s('type.num.num','const.num.SUC_1'(s('type.num.num',X183))))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X181),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X181),s('type.num.num',X181))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X182),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X182),s('type.num.num',X183)))))) )),
  inference(superposition,[],[f1561,f1540])).
fof(f57673,plain,(
  ( ! [X180,X178,X179] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X179))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X180),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X178),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))))))) | s('type.num.num',X179) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X180),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X180),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X178),s('type.num.num',X178)))))))) )),
  inference(superposition,[],[f1561,f1540])).
fof(f57672,plain,(
  ( ! [X177,X175,X176] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X176))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X175),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X177))))))) | s('type.num.num',X176) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X175),s('type.num.num',X175))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X175),s('type.num.num',X175))),s('type.num.num',X177)))))) )),
  inference(superposition,[],[f1561,f1540])).
fof(f57207,plain,(
  ( ! [X237,X238,X236] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X236))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X237))),s('type.num.num',X238))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X236),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X237),s('type.num.num',X238))),s('type.num.num',X238)))) )),
  inference(superposition,[],[f1777,f1802])).
fof(f57184,plain,(
  ( ! [X177,X178,X179] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X177))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X178),s('type.num.num','const.num.SUC_1'(s('type.num.num',X179))))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X177),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X178),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X178),s('type.num.num',X179)))))) )),
  inference(superposition,[],[f1561,f1802])).
fof(f57097,plain,(
  ( ! [X15] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X15))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X15),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))))) )),
  inference(superposition,[],[f1759,f1802])).
fof(f56510,plain,(
  ( ! [X237,X235,X236] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X235))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X236))),s('type.num.num',X237))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X235))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X236),s('type.num.num',X237))),s('type.num.num',X237)))) )),
  inference(superposition,[],[f1777,f1587])).
fof(f56488,plain,(
  ( ! [X177,X178,X179] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X177))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X178),s('type.num.num','const.num.SUC_1'(s('type.num.num',X179))))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X177))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X178),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X178),s('type.num.num',X179)))))) )),
  inference(superposition,[],[f1561,f1587])).
fof(f56402,plain,(
  ( ! [X17] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X17)))))))) )),
  inference(superposition,[],[f1545,f1587])).
fof(f227,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(forward_demodulation,[],[f150,f140])).
fof(f150,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X8] : s('type.num.num',X8) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X8),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  file('thm.arithmetic.MULT_EQ_|49|','thm.arithmetic.MULT_RIGHT_|49|')).
fof(f228,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))) )),
  inference(forward_demodulation,[],[f151,f140])).
fof(f151,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X8] : s('type.num.num',X8) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X8)))),
  file('thm.arithmetic.MULT_EQ_|49|','thm.arithmetic.MULT_LEFT_|49|')).
fof(f378,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | s('type.num.num',X1) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f208,f147])).
fof(f379,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))) | s('type.num.num',X3) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f208,f227])).
fof(f380,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))) | s('type.num.num',X5) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X4)))) )),
  inference(superposition,[],[f208,f149])).
fof(f381,plain,(
  ( ! [X6,X7] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))) | s('type.num.num',X7) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X6)))) )),
  inference(superposition,[],[f208,f228])).
fof(f475,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f210,f227])).
fof(f477,plain,(
  ( ! [X6,X7] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X6)))) )),
  inference(superposition,[],[f210,f228])).
fof(f514,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(forward_demodulation,[],[f513,f146])).
fof(f513,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))))))) )),
  inference(forward_demodulation,[],[f502,f201])).
fof(f502,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f471,f208])).
fof(f536,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X2))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f211,f227])).
fof(f538,plain,(
  ( ! [X6,X7] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X6)))) )),
  inference(superposition,[],[f211,f228])).
fof(f575,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))))) )),
  inference(unit_resulting_resolution,[],[f514,f208])).
fof(f1353,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X3))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X3)) )),
  inference(superposition,[],[f213,f227])).
fof(f1355,plain,(
  ( ! [X6,X7] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X6),s('type.num.num',X7))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X7)) )),
  inference(superposition,[],[f213,f228])).
fof(f1366,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X3),s('type.num.num',X2))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X3) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f213,f227])).
fof(f1368,plain,(
  ( ! [X6,X7] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X7),s('type.num.num',X6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X7) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X6)))) )),
  inference(superposition,[],[f213,f228])).
fof(f1540,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f202,f227])).
fof(f1587,plain,(
  ( ! [X3] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X3)))) )),
  inference(forward_demodulation,[],[f1542,f228])).
fof(f1542,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X3))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3)))))) )),
  inference(superposition,[],[f202,f228])).
fof(f1741,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X3))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X3)))) )),
  inference(superposition,[],[f203,f228])).
fof(f1802,plain,(
  ( ! [X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(forward_demodulation,[],[f1739,f227])).
fof(f1739,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f203,f227])).
fof(f2297,plain,(
  ( ! [X30,X29] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X30))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X29))))) | s('type.num.num',X30) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X29)))) )),
  inference(superposition,[],[f208,f1587])).
fof(f2513,plain,(
  ( ! [X30,X29] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X30))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X29))))) | s('type.num.num',X30) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X29),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f208,f1802])).
fof(f3088,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f1548,f227])).
fof(f3090,plain,(
  ( ! [X6,X7] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X6)))) )),
  inference(superposition,[],[f1548,f228])).
fof(f3338,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s('type.num.num',X2))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f1760,f227])).
fof(f3340,plain,(
  ( ! [X6,X7] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))),s('type.num.num',X6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X6)))) )),
  inference(superposition,[],[f1760,f228])).
fof(f3488,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0))))) )),
  inference(unit_resulting_resolution,[],[f521,f1587,f387])).
fof(f3490,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))) )),
  inference(unit_resulting_resolution,[],[f547,f1587,f387])).
fof(f3494,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(unit_resulting_resolution,[],[f547,f1802,f387])).
fof(f3497,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(unit_resulting_resolution,[],[f521,f1802,f387])).
fof(f3518,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(unit_resulting_resolution,[],[f1347,f228,f387])).
fof(f3541,plain,(
  ( ! [X2,X3] : (s('type.num.num',X3) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))) | s('type.num.num',sK2(X3)) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X3)) )),
  inference(superposition,[],[f387,f227])).
fof(f3543,plain,(
  ( ! [X6,X7] : (s('type.num.num',X7) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))) | s('type.num.num',sK2(X7)) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X7)) )),
  inference(superposition,[],[f387,f228])).
fof(f3552,plain,(
  ( ! [X21,X20] : (s('type.num.num',X21) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X20))))) | s('type.num.num',sK2(X21)) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X20),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X21)) )),
  inference(superposition,[],[f387,f1802])).
fof(f3557,plain,(
  ( ! [X33,X34] : (s('type.num.num',X34) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X33))))) | s('type.num.num',sK2(X34)) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X33))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X34)) )),
  inference(superposition,[],[f387,f1587])).
fof(f3602,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X2) | s('type.num.num',X2) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))) | s('type.num.num',X3) = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f3520,f227])).
fof(f3520,plain,(
  ( ! [X2,X3] : (s('type.num.num',X2) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))) | s('type.num.num',X3) = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f387,f227])).
fof(f3603,plain,(
  ( ! [X6,X7] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X6) | s('type.num.num',X6) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))) | s('type.num.num',X7) = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X6))))) )),
  inference(forward_demodulation,[],[f3522,f228])).
fof(f3522,plain,(
  ( ! [X6,X7] : (s('type.num.num',X6) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))) | s('type.num.num',X7) = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X6)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X6)))) )),
  inference(superposition,[],[f387,f228])).
fof(f3610,plain,(
  ( ! [X11] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))) | s('type.num.num',X11) = s('type.num.num',sK2('const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) )),
  inference(subsumption_resolution,[],[f3609,f195])).
fof(f3609,plain,(
  ( ! [X11] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))) | s('type.num.num',X11) = s('type.num.num',sK2('const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) )),
  inference(forward_demodulation,[],[f3526,f140])).
fof(f3526,plain,(
  ( ! [X11] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))) | s('type.num.num',X11) = s('type.num.num',sK2('const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) )),
  inference(superposition,[],[f387,f140])).
fof(f3615,plain,(
  ( ! [X21,X20] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X20))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X21))) | s('type.num.num',X21) = s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num',X20),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(subsumption_resolution,[],[f3531,f523])).
fof(f3531,plain,(
  ( ! [X21,X20] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X20))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X21))) | s('type.num.num',X21) = s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num',X20),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X20),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f387,f1802])).
fof(f3620,plain,(
  ( ! [X33,X34] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X33))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X34))) | s('type.num.num',X34) = s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X33))))) )),
  inference(subsumption_resolution,[],[f3536,f471])).
fof(f3536,plain,(
  ( ! [X33,X34] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X33))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X34))) | s('type.num.num',X34) = s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X33)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X33)))) )),
  inference(superposition,[],[f387,f1587])).
fof(f3699,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(unit_resulting_resolution,[],[f1540,f209])).
fof(f3749,plain,(
  ( ! [X28,X27] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X27),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.num.num',X28) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',X27)))) )),
  inference(superposition,[],[f208,f1540])).
fof(f3751,plain,(
  ( ! [X31,X32] : (s('type.num.num',X32) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X31),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X32))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X31),s('type.num.num',X31)))))) )),
  inference(superposition,[],[f209,f1540])).
fof(f3753,plain,(
  ( ! [X35,X36] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X35),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num',X36))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X35),s('type.num.num',X35)))) )),
  inference(superposition,[],[f210,f1540])).
fof(f3756,plain,(
  ( ! [X41,X42] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X42),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X41),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X41),s('type.num.num',X41)))) )),
  inference(superposition,[],[f211,f1540])).
fof(f3757,plain,(
  ( ! [X43,X44] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X43),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X44) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X43),s('type.num.num',X43))),s('type.num.num',X44)))) )),
  inference(superposition,[],[f212,f1540])).
fof(f3758,plain,(
  ( ! [X45,X46] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X45),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num',X46))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X45),s('type.num.num',X45))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X46)) )),
  inference(superposition,[],[f213,f1540])).
fof(f3759,plain,(
  ( ! [X47,X48] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X48),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X47),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X48) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X47),s('type.num.num',X47)))) )),
  inference(superposition,[],[f213,f1540])).
fof(f3760,plain,(
  ( ! [X50,X49] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X49),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X49),s('type.num.num',X49))),s('type.num.num',X50)))) )),
  inference(superposition,[],[f214,f1540])).
fof(f3761,plain,(
  ( ! [X52,X51] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X51),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X52),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X51),s('type.num.num',X51)))))) )),
  inference(superposition,[],[f215,f1540])).
fof(f3764,plain,(
  ( ! [X57,X56] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X57))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X56),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.num.num',X57) = s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num',X56),s('type.num.num',X56)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X56),s('type.num.num',X56)))) )),
  inference(superposition,[],[f387,f1540])).
fof(f3765,plain,(
  ( ! [X59,X58] : (s('type.num.num',X59) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X58),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.num.num',sK2(X59)) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X58),s('type.num.num',X58))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X59)) )),
  inference(superposition,[],[f387,f1540])).
fof(f3775,plain,(
  ( ! [X73] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X73),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X73),s('type.num.num',X73)))))) )),
  inference(superposition,[],[f1297,f1540])).
fof(f3776,plain,(
  ( ! [X74,X75] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X74),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X75))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X74),s('type.num.num',X74)))) )),
  inference(superposition,[],[f1548,f1540])).
fof(f3780,plain,(
  ( ! [X81,X82] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X82))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X81),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X81),s('type.num.num',X81)))) )),
  inference(superposition,[],[f1760,f1540])).
fof(f3790,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f203,f1540])).
fof(f3791,plain,(
  ( ! [X2] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X2))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(superposition,[],[f202,f1540])).
fof(f3810,plain,(
  ( ! [X26,X27] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X26),s('type.num.num',X26))))) | s('type.num.num',X27) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X26),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f208,f1540])).
fof(f3812,plain,(
  ( ! [X30,X31] : (s('type.num.num',X31) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X30),s('type.num.num',X30))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X31))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X30),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(superposition,[],[f209,f1540])).
fof(f3814,plain,(
  ( ! [X35,X34] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X34),s('type.num.num',X34))),s('type.num.num',X35))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X34),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f210,f1540])).
fof(f3817,plain,(
  ( ! [X41,X40] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X41),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X40))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X40),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f211,f1540])).
fof(f3818,plain,(
  ( ! [X43,X42] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X42),s('type.num.num',X42))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X43) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X42),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num',X43)))) )),
  inference(superposition,[],[f212,f1540])).
fof(f3819,plain,(
  ( ! [X45,X44] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X44),s('type.num.num',X44))),s('type.num.num',X45))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X44),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X45)) )),
  inference(superposition,[],[f213,f1540])).
fof(f3820,plain,(
  ( ! [X47,X46] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X47),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X46),s('type.num.num',X46))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X47) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X46),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f213,f1540])).
fof(f3821,plain,(
  ( ! [X48,X49] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X48),s('type.num.num',X48))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X48),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num',X49)))) )),
  inference(superposition,[],[f214,f1540])).
fof(f3822,plain,(
  ( ! [X50,X51] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X50),s('type.num.num',X50))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X51),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X50),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(superposition,[],[f215,f1540])).
fof(f3825,plain,(
  ( ! [X56,X55] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X56))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X55),s('type.num.num',X55))) | s('type.num.num',X56) = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num',X55),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X55),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f387,f1540])).
fof(f3826,plain,(
  ( ! [X57,X58] : (s('type.num.num',X58) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X57),s('type.num.num',X57))))) | s('type.num.num',sK2(X58)) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X57),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X58)) )),
  inference(superposition,[],[f387,f1540])).
fof(f3836,plain,(
  ( ! [X72] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X72),s('type.num.num',X72))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X72),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(superposition,[],[f1297,f1540])).
fof(f3837,plain,(
  ( ! [X74,X73] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X73),s('type.num.num',X73))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X74))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X73),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f1548,f1540])).
fof(f3841,plain,(
  ( ! [X80,X81] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X81))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X80),s('type.num.num',X80))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X80),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f1760,f1540])).
fof(f3850,plain,(
  ( ! [X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1))))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(forward_demodulation,[],[f3700,f201])).
fof(f3700,plain,(
  ( ! [X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f1540,f200])).
fof(f3932,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f1741,f209])).
fof(f3933,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X1)))) )),
  inference(superposition,[],[f1741,f1540])).
fof(f3987,plain,(
  ( ! [X28,X29] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X29))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X28))))) | s('type.num.num',X29) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X28),s('type.num.num',X28)))) )),
  inference(superposition,[],[f208,f1741])).
fof(f3989,plain,(
  ( ! [X33,X32] : (s('type.num.num',X33) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X32))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X33))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X32)))))) )),
  inference(superposition,[],[f209,f1741])).
fof(f3991,plain,(
  ( ! [X37,X36] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X36))),s('type.num.num',X37))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X36),s('type.num.num',X36)))) )),
  inference(superposition,[],[f210,f1741])).
fof(f3994,plain,(
  ( ! [X43,X42] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X43),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X42))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X42),s('type.num.num',X42)))) )),
  inference(superposition,[],[f211,f1741])).
fof(f3995,plain,(
  ( ! [X45,X44] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X44))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X45) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X44),s('type.num.num',X44))),s('type.num.num',X45)))) )),
  inference(superposition,[],[f212,f1741])).
fof(f3996,plain,(
  ( ! [X47,X46] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X46))),s('type.num.num',X47))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X46),s('type.num.num',X46))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X47)) )),
  inference(superposition,[],[f213,f1741])).
fof(f3997,plain,(
  ( ! [X48,X49] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X49),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X48))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X49) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X48),s('type.num.num',X48)))) )),
  inference(superposition,[],[f213,f1741])).
fof(f3998,plain,(
  ( ! [X50,X51] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X50))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X50),s('type.num.num',X50))),s('type.num.num',X51)))) )),
  inference(superposition,[],[f214,f1741])).
fof(f3999,plain,(
  ( ! [X52,X53] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X52))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X53),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X52),s('type.num.num',X52)))))) )),
  inference(superposition,[],[f215,f1741])).
fof(f4002,plain,(
  ( ! [X57,X58] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X58))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X57))) | s('type.num.num',X58) = s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num',X57),s('type.num.num',X57)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X57),s('type.num.num',X57)))) )),
  inference(superposition,[],[f387,f1741])).
fof(f4003,plain,(
  ( ! [X59,X60] : (s('type.num.num',X60) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X59))))) | s('type.num.num',sK2(X60)) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X59),s('type.num.num',X59))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X60)) )),
  inference(superposition,[],[f387,f1741])).
fof(f4013,plain,(
  ( ! [X74] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X74))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X74),s('type.num.num',X74)))))) )),
  inference(superposition,[],[f1297,f1741])).
fof(f4015,plain,(
  ( ! [X76,X77] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X76))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X77))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X76),s('type.num.num',X76)))) )),
  inference(superposition,[],[f1548,f1741])).
fof(f4019,plain,(
  ( ! [X83,X84] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X84))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X83))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X83),s('type.num.num',X83)))) )),
  inference(superposition,[],[f1760,f1741])).
fof(f4033,plain,(
  ( ! [X2] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X2))),s('type.num.num',X2))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num',X2)))) )),
  inference(superposition,[],[f203,f1741])).
fof(f4034,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X3)))))) )),
  inference(superposition,[],[f202,f1741])).
fof(f4049,plain,(
  ( ! [X26,X27] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X26),s('type.num.num',X26))))) | s('type.num.num',X27) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X26)))) )),
  inference(superposition,[],[f208,f1741])).
fof(f4051,plain,(
  ( ! [X30,X31] : (s('type.num.num',X31) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X30),s('type.num.num',X30))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X31))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X30)))))) )),
  inference(superposition,[],[f209,f1741])).
fof(f4053,plain,(
  ( ! [X35,X34] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X34),s('type.num.num',X34))),s('type.num.num',X35))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X34)))) )),
  inference(superposition,[],[f210,f1741])).
fof(f4056,plain,(
  ( ! [X41,X40] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X41),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X40))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X40)))) )),
  inference(superposition,[],[f211,f1741])).
fof(f4057,plain,(
  ( ! [X43,X42] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X42),s('type.num.num',X42))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X43) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X42))),s('type.num.num',X43)))) )),
  inference(superposition,[],[f212,f1741])).
fof(f4058,plain,(
  ( ! [X45,X44] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X44),s('type.num.num',X44))),s('type.num.num',X45))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X44))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X45)) )),
  inference(superposition,[],[f213,f1741])).
fof(f4059,plain,(
  ( ! [X47,X46] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X47),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X46),s('type.num.num',X46))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X47) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X46)))) )),
  inference(superposition,[],[f213,f1741])).
fof(f4060,plain,(
  ( ! [X48,X49] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X48),s('type.num.num',X48))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X48))),s('type.num.num',X49)))) )),
  inference(superposition,[],[f214,f1741])).
fof(f4061,plain,(
  ( ! [X50,X51] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X50),s('type.num.num',X50))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X51),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X50)))))) )),
  inference(superposition,[],[f215,f1741])).
fof(f4064,plain,(
  ( ! [X56,X55] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X56))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X55),s('type.num.num',X55))) | s('type.num.num',X56) = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X55)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X55)))) )),
  inference(superposition,[],[f387,f1741])).
fof(f4065,plain,(
  ( ! [X57,X58] : (s('type.num.num',X58) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X57),s('type.num.num',X57))))) | s('type.num.num',sK2(X58)) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X57))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X58)) )),
  inference(superposition,[],[f387,f1741])).
fof(f4075,plain,(
  ( ! [X72] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X72),s('type.num.num',X72))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X72)))))) )),
  inference(superposition,[],[f1297,f1741])).
fof(f4077,plain,(
  ( ! [X74,X75] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X74),s('type.num.num',X74))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X75))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X74)))) )),
  inference(superposition,[],[f1548,f1741])).
fof(f4081,plain,(
  ( ! [X81,X82] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X82))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X81),s('type.num.num',X81))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X81)))) )),
  inference(superposition,[],[f1760,f1741])).
fof(f4090,plain,(
  ( ! [X2] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X2))))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2)))))) )),
  inference(forward_demodulation,[],[f3934,f201])).
fof(f3934,plain,(
  ( ! [X2] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.num.num',X2))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2)))))) )),
  inference(superposition,[],[f1741,f200])).
fof(f4784,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X4))))) | s('type.num.num',X5) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(superposition,[],[f385,f1540])).
fof(f4787,plain,(
  ( ! [X10,X11] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',X10))))) | s('type.num.num',X11) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X10)))))) )),
  inference(superposition,[],[f385,f1741])).
fof(f4794,plain,(
  ( ! [X19,X20] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X20))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X19))))) | s('type.num.num',X20) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',X19)))))) )),
  inference(superposition,[],[f385,f1741])).
fof(f4795,plain,(
  ( ! [X21,X22] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X22))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X21),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.num.num',X22) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',X21)))))) )),
  inference(superposition,[],[f385,f1540])).
fof(f4813,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X3))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f385,f227])).
fof(f4814,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X4))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X5))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f385,f1540])).
fof(f4816,plain,(
  ( ! [X8,X9] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X9))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X8)))) )),
  inference(superposition,[],[f385,f228])).
fof(f4817,plain,(
  ( ! [X10,X11] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',X10))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X11))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X10)))) )),
  inference(superposition,[],[f385,f1741])).
fof(f4824,plain,(
  ( ! [X19,X20] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X20))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X19))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X20))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',X19)))) )),
  inference(superposition,[],[f385,f1741])).
fof(f4825,plain,(
  ( ! [X21,X22] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X22))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X21),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X22))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',X21)))) )),
  inference(superposition,[],[f385,f1540])).
fof(f4829,plain,(
  ( ! [X31,X32] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X32))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X31))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X32))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X31),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f385,f1802])).
fof(f4833,plain,(
  ( ! [X41,X42] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X42))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X41))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X42))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X41)))) )),
  inference(superposition,[],[f385,f1587])).
fof(f5039,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X4))))) | s('type.num.num',X5) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f386,f1540])).
fof(f5042,plain,(
  ( ! [X10,X11] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',X10))))) | s('type.num.num',X11) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X10))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f386,f1741])).
fof(f5049,plain,(
  ( ! [X19,X20] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X20))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X19))))) | s('type.num.num',X20) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',X19))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f386,f1741])).
fof(f5050,plain,(
  ( ! [X21,X22] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X22))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X21),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.num.num',X22) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',X21))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f386,f1540])).
fof(f5068,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f386,f227])).
fof(f5069,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X4))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f386,f1540])).
fof(f5071,plain,(
  ( ! [X8,X9] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X8)))) )),
  inference(superposition,[],[f386,f228])).
fof(f5072,plain,(
  ( ! [X10,X11] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',X10))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X10)))) )),
  inference(superposition,[],[f386,f1741])).
fof(f5079,plain,(
  ( ! [X19,X20] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X20))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X19))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X20),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',X19)))) )),
  inference(superposition,[],[f386,f1741])).
fof(f5080,plain,(
  ( ! [X21,X22] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X22))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X21),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X22),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',X21)))) )),
  inference(superposition,[],[f386,f1540])).
fof(f5084,plain,(
  ( ! [X31,X32] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X32))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X31))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X31),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f386,f1802])).
fof(f5088,plain,(
  ( ! [X41,X42] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X42))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X41))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X42),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X41)))) )),
  inference(superposition,[],[f386,f1587])).
fof(f5421,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X4))),s('type.num.num',X5))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(superposition,[],[f481,f1540])).
fof(f5424,plain,(
  ( ! [X10,X11] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',X10))),s('type.num.num',X11))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X10)))))) )),
  inference(superposition,[],[f481,f1741])).
fof(f5431,plain,(
  ( ! [X19,X20] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X19))),s('type.num.num',X20))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',X19)))))) )),
  inference(superposition,[],[f481,f1741])).
fof(f5432,plain,(
  ( ! [X21,X22] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X21),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num',X22))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',X21)))))) )),
  inference(superposition,[],[f481,f1540])).
fof(f5706,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X4))),s('type.num.num',X5))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f482,f1540])).
fof(f5709,plain,(
  ( ! [X10,X11] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',X10))),s('type.num.num',X11))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X10))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f482,f1741])).
fof(f5716,plain,(
  ( ! [X19,X20] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X19))),s('type.num.num',X20))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',X19))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f482,f1741])).
fof(f5717,plain,(
  ( ! [X21,X22] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X21),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num',X22))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',X21))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f482,f1540])).
fof(f5830,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X4))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(superposition,[],[f542,f1540])).
fof(f5833,plain,(
  ( ! [X10,X11] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',X10))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X10)))))) )),
  inference(superposition,[],[f542,f1741])).
fof(f5840,plain,(
  ( ! [X19,X20] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X20),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X19))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',X19)))))) )),
  inference(superposition,[],[f542,f1741])).
fof(f5841,plain,(
  ( ! [X21,X22] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X22),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X21),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',X21)))))) )),
  inference(superposition,[],[f542,f1540])).
fof(f5990,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X4))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f543,f1540])).
fof(f5993,plain,(
  ( ! [X10,X11] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',X10))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X10))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f543,f1741])).
fof(f6000,plain,(
  ( ! [X19,X20] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X20),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X19))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',X19))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f543,f1741])).
fof(f6001,plain,(
  ( ! [X21,X22] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X22),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X21),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',X21))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f543,f1540])).
fof(f6117,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X4))))) | s('type.num.num',X5) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f379,f1540])).
fof(f6120,plain,(
  ( ! [X10,X11] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',X10))))) | s('type.num.num',X11) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X10))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f379,f1741])).
fof(f6125,plain,(
  ( ! [X17,X16] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X16))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))) | s('type.num.num',X17) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X16)))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X16)) )),
  inference(superposition,[],[f379,f153])).
fof(f6127,plain,(
  ( ! [X19,X20] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X20))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X19))))) | s('type.num.num',X20) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',X19))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f379,f1741])).
fof(f6128,plain,(
  ( ! [X21,X22] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X22))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X21),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.num.num',X22) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',X21))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f379,f1540])).
fof(f6143,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | s('type.num.num',X1) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK2(X0)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(superposition,[],[f379,f153])).
fof(f6146,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))) | s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f379,f227])).
fof(f6147,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X4))))) | s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X5),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f379,f1540])).
fof(f6149,plain,(
  ( ! [X8,X9] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))) | s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X9),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X8)))) )),
  inference(superposition,[],[f379,f228])).
fof(f6150,plain,(
  ( ! [X10,X11] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',X10))))) | s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X11),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X10)))) )),
  inference(superposition,[],[f379,f1741])).
fof(f6155,plain,(
  ( ! [X17,X16] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X16))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X16)))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X17),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X16)) )),
  inference(superposition,[],[f379,f153])).
fof(f6157,plain,(
  ( ! [X19,X20] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X20))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X19))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',X19))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X20),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f379,f1741])).
fof(f6158,plain,(
  ( ! [X21,X22] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X22))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X21),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',X21))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X22),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f379,f1540])).
fof(f6159,plain,(
  ( ! [X24,X23,X25] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X25))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X23),s('type.num.num','const.num.SUC_1'(s('type.num.num',X24))))))) | s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X25),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X23),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X23),s('type.num.num',X24)))))) )),
  inference(superposition,[],[f379,f202])).
fof(f6162,plain,(
  ( ! [X31,X32] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X32))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X31))))) | s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X32),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X31),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f379,f1802])).
fof(f6163,plain,(
  ( ! [X35,X33,X34] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X35))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X33))),s('type.num.num',X34))))) | s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X35),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X33),s('type.num.num',X34))),s('type.num.num',X34)))) )),
  inference(superposition,[],[f379,f203])).
fof(f6166,plain,(
  ( ! [X41,X42] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X42))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X41))))) | s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X42),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X41)))) )),
  inference(superposition,[],[f379,f1587])).
fof(f6169,plain,(
  ( ! [X45,X46] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X45))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X46))) | s('type.num.num',sK2('const.num.SUC_1'(s('type.num.num',X45)))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X46),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f379,f962])).
fof(f6305,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X4))))) | s('type.num.num',X5) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(superposition,[],[f381,f1540])).
fof(f6308,plain,(
  ( ! [X10,X11] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',X10))))) | s('type.num.num',X11) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X10)))))) )),
  inference(superposition,[],[f381,f1741])).
fof(f6313,plain,(
  ( ! [X17,X16] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X16))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))) | s('type.num.num',X17) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X16)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X16)) )),
  inference(superposition,[],[f381,f153])).
fof(f6315,plain,(
  ( ! [X19,X20] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X20))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X19))))) | s('type.num.num',X20) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',X19)))))) )),
  inference(superposition,[],[f381,f1741])).
fof(f6316,plain,(
  ( ! [X21,X22] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X22))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X21),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.num.num',X22) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',X21)))))) )),
  inference(superposition,[],[f381,f1540])).
fof(f6331,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | s('type.num.num',X1) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK2(X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(superposition,[],[f381,f153])).
fof(f6335,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X4))))) | s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X5))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f381,f1540])).
fof(f6337,plain,(
  ( ! [X8,X9] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))) | s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X8))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X9)))) )),
  inference(superposition,[],[f381,f228])).
fof(f6338,plain,(
  ( ! [X10,X11] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',X10))))) | s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X11))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X10)))) )),
  inference(superposition,[],[f381,f1741])).
fof(f6343,plain,(
  ( ! [X17,X16] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X16))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X16)))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X17))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X16)) )),
  inference(superposition,[],[f381,f153])).
fof(f6345,plain,(
  ( ! [X19,X20] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X20))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X19))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',X19))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X20)))) )),
  inference(superposition,[],[f381,f1741])).
fof(f6346,plain,(
  ( ! [X21,X22] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X22))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X21),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',X21))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X22)))) )),
  inference(superposition,[],[f381,f1540])).
fof(f6347,plain,(
  ( ! [X24,X23,X25] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X25))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X23),s('type.num.num','const.num.SUC_1'(s('type.num.num',X24))))))) | s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X25))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X23),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X23),s('type.num.num',X24)))))) )),
  inference(superposition,[],[f381,f202])).
fof(f6350,plain,(
  ( ! [X31,X32] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X32))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X31))))) | s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X32))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X31),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f381,f1802])).
fof(f6351,plain,(
  ( ! [X35,X33,X34] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X35))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X33))),s('type.num.num',X34))))) | s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X35))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X33),s('type.num.num',X34))),s('type.num.num',X34)))) )),
  inference(superposition,[],[f381,f203])).
fof(f6354,plain,(
  ( ! [X41,X42] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X42))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X41))))) | s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X42))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X41)))) )),
  inference(superposition,[],[f381,f1587])).
fof(f6357,plain,(
  ( ! [X45,X46] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X45))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X46))) | s('type.num.num',sK2('const.num.SUC_1'(s('type.num.num',X45)))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X46)))) )),
  inference(superposition,[],[f381,f962])).
fof(f6427,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X4))),s('type.num.num',X5))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f475,f1540])).
fof(f6430,plain,(
  ( ! [X10,X11] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',X10))),s('type.num.num',X11))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X10))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f475,f1741])).
fof(f6437,plain,(
  ( ! [X19,X20] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X19))),s('type.num.num',X20))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',X19))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f475,f1741])).
fof(f6438,plain,(
  ( ! [X21,X22] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X21),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num',X22))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',X21))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f475,f1540])).
fof(f6487,plain,(
  ( ! [X8,X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num',X9))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X8),s('type.num.num',X9))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f475,f203])).
fof(f6524,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X4))),s('type.num.num',X5))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(superposition,[],[f477,f1540])).
fof(f6527,plain,(
  ( ! [X10,X11] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',X10))),s('type.num.num',X11))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X10)))))) )),
  inference(superposition,[],[f477,f1741])).
fof(f6534,plain,(
  ( ! [X19,X20] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X19))),s('type.num.num',X20))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',X19)))))) )),
  inference(superposition,[],[f477,f1741])).
fof(f6535,plain,(
  ( ! [X21,X22] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X21),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num',X22))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',X21)))))) )),
  inference(superposition,[],[f477,f1540])).
fof(f6584,plain,(
  ( ! [X8,X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num',X9))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X8),s('type.num.num',X9)))))) )),
  inference(superposition,[],[f477,f203])).
fof(f6938,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X4))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f536,f1540])).
fof(f6941,plain,(
  ( ! [X10,X11] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',X10))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X10))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f536,f1741])).
fof(f6948,plain,(
  ( ! [X19,X20] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X20),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X19))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',X19))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f536,f1741])).
fof(f6949,plain,(
  ( ! [X21,X22] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X22),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X21),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',X21))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f536,f1540])).
fof(f6966,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X3))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f536,f202])).
fof(f7247,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X4))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(superposition,[],[f538,f1540])).
fof(f7250,plain,(
  ( ! [X10,X11] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',X10))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X10)))))) )),
  inference(superposition,[],[f538,f1741])).
fof(f7257,plain,(
  ( ! [X19,X20] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X20),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X19))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',X19)))))) )),
  inference(superposition,[],[f538,f1741])).
fof(f7258,plain,(
  ( ! [X21,X22] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X22),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X21),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',X21)))))) )),
  inference(superposition,[],[f538,f1540])).
fof(f7275,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X3)))))) )),
  inference(superposition,[],[f538,f202])).
fof(f8478,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X4))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num',X5)))) )),
  inference(superposition,[],[f1547,f1540])).
fof(f8481,plain,(
  ( ! [X10,X11] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',X10))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X10))),s('type.num.num',X11)))) )),
  inference(superposition,[],[f1547,f1741])).
fof(f8488,plain,(
  ( ! [X19,X20] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X19))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X20))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',X19))),s('type.num.num',X20)))) )),
  inference(superposition,[],[f1547,f1741])).
fof(f8489,plain,(
  ( ! [X21,X22] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X21),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X22))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',X21))),s('type.num.num',X22)))) )),
  inference(superposition,[],[f1547,f1540])).
fof(f8506,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X5),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X4))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(superposition,[],[f1547,f1540])).
fof(f8509,plain,(
  ( ! [X10,X11] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X11),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',X10))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X10)))))) )),
  inference(superposition,[],[f1547,f1741])).
fof(f8516,plain,(
  ( ! [X19,X20] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X20),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X19))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X20),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',X19)))))) )),
  inference(superposition,[],[f1547,f1741])).
fof(f8517,plain,(
  ( ! [X21,X22] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X22),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X21),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X22),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',X21)))))) )),
  inference(superposition,[],[f1547,f1540])).
fof(f8949,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X4))))),s('type.num.num',X5))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num',X5)))) )),
  inference(superposition,[],[f1761,f1540])).
fof(f8952,plain,(
  ( ! [X10,X11] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',X10))))),s('type.num.num',X11))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X10))),s('type.num.num',X11)))) )),
  inference(superposition,[],[f1761,f1741])).
fof(f8959,plain,(
  ( ! [X19,X20] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X19))))),s('type.num.num',X20))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',X19))),s('type.num.num',X20)))) )),
  inference(superposition,[],[f1761,f1741])).
fof(f8960,plain,(
  ( ! [X21,X22] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X21),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.num.num',X22))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',X21))),s('type.num.num',X22)))) )),
  inference(superposition,[],[f1761,f1540])).
fof(f8979,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X4))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(superposition,[],[f1761,f1540])).
fof(f8982,plain,(
  ( ! [X10,X11] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',X10))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X10)))))) )),
  inference(superposition,[],[f1761,f1741])).
fof(f8989,plain,(
  ( ! [X19,X20] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X20))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X19))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X20),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',X19)))))) )),
  inference(superposition,[],[f1761,f1741])).
fof(f8990,plain,(
  ( ! [X21,X22] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X22))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X21),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X22),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',X21)))))) )),
  inference(superposition,[],[f1761,f1540])).
fof(f9479,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X0)),s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',sK2(X0)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(superposition,[],[f762,f1741])).
fof(f9480,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X1)),s('type.num.num',X1))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK2(X1)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) )),
  inference(superposition,[],[f762,f1540])).
fof(f9496,plain,(
  ( ! [X30,X28,X29] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X30))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X28),s('type.num.num',X29))) | s('type.num.num',X30) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X28),s('type.num.num',sK2(X29)))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X29)) )),
  inference(superposition,[],[f379,f762])).
fof(f9497,plain,(
  ( ! [X33,X31,X32] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X33))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X31),s('type.num.num',X32))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X31),s('type.num.num',sK2(X32)))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X33),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X32)) )),
  inference(superposition,[],[f379,f762])).
fof(f9499,plain,(
  ( ! [X39,X37,X38] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X39))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X37),s('type.num.num',X38))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X37),s('type.num.num',sK2(X38)))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X39))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X38)) )),
  inference(superposition,[],[f381,f762])).
fof(f9500,plain,(
  ( ! [X41,X42,X40] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X42))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))) | s('type.num.num',X42) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',sK2(X41)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X41)) )),
  inference(superposition,[],[f381,f762])).
fof(f9506,plain,(
  ( ! [X59,X58] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X58),s('type.num.num',X59))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X59)) )),
  inference(superposition,[],[f514,f762])).
fof(f9552,plain,(
  ( ! [X185,X186,X184] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X186))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X184),s('type.num.num',X185))))) | s('type.num.num',X186) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X184),s('type.num.num',sK2(X185)))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X185)) )),
  inference(superposition,[],[f379,f762])).
fof(f9553,plain,(
  ( ! [X189,X187,X188] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X189))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X187),s('type.num.num',X188))))) | s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X189),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X187),s('type.num.num',sK2(X188)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X188)) )),
  inference(superposition,[],[f379,f762])).
fof(f9555,plain,(
  ( ! [X194,X195,X193] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X195))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X193),s('type.num.num',X194))))) | s('type.num.num',X195) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X193),s('type.num.num',sK2(X194)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X194)) )),
  inference(superposition,[],[f381,f762])).
fof(f9556,plain,(
  ( ! [X198,X196,X197] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X198))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X196),s('type.num.num',X197))))) | s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X198))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X196),s('type.num.num',sK2(X197)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X197)) )),
  inference(superposition,[],[f381,f762])).
fof(f10109,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',sK2(X0)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',sK2(X0)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(superposition,[],[f903,f1741])).
fof(f10110,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',sK2(X1)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK2(X1)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) )),
  inference(superposition,[],[f903,f1540])).
fof(f10128,plain,(
  ( ! [X30,X31,X32] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X32))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X30),s('type.num.num',X31))) | s('type.num.num',X32) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X30)),s('type.num.num',X31))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X30)) )),
  inference(superposition,[],[f379,f903])).
fof(f10129,plain,(
  ( ! [X35,X33,X34] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X35))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X33),s('type.num.num',X34))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X33)),s('type.num.num',X34))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X35),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X33)) )),
  inference(superposition,[],[f379,f903])).
fof(f10131,plain,(
  ( ! [X39,X41,X40] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X41))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X39),s('type.num.num',X40))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X39)),s('type.num.num',X40))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X41))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X39)) )),
  inference(superposition,[],[f381,f903])).
fof(f10132,plain,(
  ( ! [X43,X44,X42] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X44))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X42),s('type.num.num',X43))) | s('type.num.num',X44) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X42)),s('type.num.num',X43))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X42)) )),
  inference(superposition,[],[f381,f903])).
fof(f10138,plain,(
  ( ! [X61,X60] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X60),s('type.num.num',X61))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X60)) )),
  inference(superposition,[],[f514,f903])).
fof(f10186,plain,(
  ( ! [X191,X192,X190] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X192))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X190),s('type.num.num',X191))))) | s('type.num.num',X192) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X190)),s('type.num.num',X191))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X190)) )),
  inference(superposition,[],[f379,f903])).
fof(f10187,plain,(
  ( ! [X194,X195,X193] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X195))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X193),s('type.num.num',X194))))) | s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X195),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X193)),s('type.num.num',X194))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X193)) )),
  inference(superposition,[],[f379,f903])).
fof(f10189,plain,(
  ( ! [X200,X199,X201] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X201))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X199),s('type.num.num',X200))))) | s('type.num.num',X201) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X199)),s('type.num.num',X200))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X199)) )),
  inference(superposition,[],[f381,f903])).
fof(f10190,plain,(
  ( ! [X204,X202,X203] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X204))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X202),s('type.num.num',X203))))) | s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X204))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X202)),s('type.num.num',X203))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X202)) )),
  inference(superposition,[],[f381,f903])).
fof(f10918,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X4))))) | s('type.num.num',X5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f384,f1540])).
fof(f10921,plain,(
  ( ! [X10,X11] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',X10))))) | s('type.num.num',X11) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X10)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X10)))) )),
  inference(superposition,[],[f384,f1741])).
fof(f10930,plain,(
  ( ! [X26,X25] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X26))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X25))))) | s('type.num.num',X26) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num',X25),s('type.num.num',X25)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X25),s('type.num.num',X25)))) )),
  inference(superposition,[],[f384,f1741])).
fof(f10931,plain,(
  ( ! [X28,X27] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X27),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.num.num',X28) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',X27)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',X27)))) )),
  inference(superposition,[],[f384,f1540])).
fof(f10952,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X4))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X5)))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X5)) )),
  inference(superposition,[],[f384,f1540])).
fof(f10955,plain,(
  ( ! [X10,X11] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',X10))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X11)))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X10))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X11)) )),
  inference(superposition,[],[f384,f1741])).
fof(f10964,plain,(
  ( ! [X26,X25] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X26))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X25))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X26)))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X25),s('type.num.num',X25))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X26)) )),
  inference(superposition,[],[f384,f1741])).
fof(f10965,plain,(
  ( ! [X28,X27] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X27),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X28)))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',X27))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X28)) )),
  inference(superposition,[],[f384,f1540])).
fof(f10969,plain,(
  ( ! [X37,X38] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X38))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X37))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X38)))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X37),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X38)) )),
  inference(superposition,[],[f384,f1802])).
fof(f10973,plain,(
  ( ! [X47,X48] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X48))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X47))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X48)))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X47))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X48)) )),
  inference(superposition,[],[f384,f1587])).
fof(f11097,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X2) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))) | s('type.num.num',X3) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) )),
  inference(forward_demodulation,[],[f10917,f227])).
fof(f10917,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))) | s('type.num.num',X3) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f384,f227])).
fof(f11098,plain,(
  ( ! [X8,X9] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X8) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))) | s('type.num.num',X9) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X8))))))) )),
  inference(forward_demodulation,[],[f10920,f228])).
fof(f10920,plain,(
  ( ! [X8,X9] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))) | s('type.num.num',X9) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X8)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X8)))) )),
  inference(superposition,[],[f384,f228])).
fof(f11106,plain,(
  ( ! [X15] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X15))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num',X15) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) )),
  inference(subsumption_resolution,[],[f11105,f195])).
fof(f11105,plain,(
  ( ! [X15] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X15))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num',X15) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) )),
  inference(forward_demodulation,[],[f10925,f140])).
fof(f10925,plain,(
  ( ! [X15] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X15))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num',X15) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) )),
  inference(superposition,[],[f384,f140])).
fof(f18043,plain,(
  ( ! [X4,X5] : (s('type.num.num',X5) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X4))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f437,f1540])).
fof(f18047,plain,(
  ( ! [X12,X11] : (s('type.num.num',X12) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',X11))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X12))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X11)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X11)))) )),
  inference(superposition,[],[f437,f1741])).
fof(f18055,plain,(
  ( ! [X26,X25] : (s('type.num.num',X26) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X25))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X26))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num',X25),s('type.num.num',X25)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X25),s('type.num.num',X25)))) )),
  inference(superposition,[],[f437,f1741])).
fof(f18056,plain,(
  ( ! [X28,X27] : (s('type.num.num',X28) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X27),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',X27)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',X27)))) )),
  inference(superposition,[],[f437,f1540])).
fof(f18072,plain,(
  ( ! [X4,X5] : (s('type.num.num',X5) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X4))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X5)))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X5)) )),
  inference(superposition,[],[f437,f1540])).
fof(f18076,plain,(
  ( ! [X12,X11] : (s('type.num.num',X12) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',X11))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X12)))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X11))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12)) )),
  inference(superposition,[],[f437,f1741])).
fof(f18084,plain,(
  ( ! [X26,X25] : (s('type.num.num',X26) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X25))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X26)))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X25),s('type.num.num',X25))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X26)) )),
  inference(superposition,[],[f437,f1741])).
fof(f18085,plain,(
  ( ! [X28,X27] : (s('type.num.num',X28) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X27),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X28)))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',X27))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X28)) )),
  inference(superposition,[],[f437,f1540])).
fof(f18197,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0))))))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))))) )),
  inference(forward_demodulation,[],[f17930,f3566])).
fof(f17930,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))))))) )),
  inference(unit_resulting_resolution,[],[f1741,f521,f437])).
fof(f18198,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0))))))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(forward_demodulation,[],[f17929,f3566])).
fof(f17929,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))))))) )),
  inference(unit_resulting_resolution,[],[f1540,f521,f437])).
fof(f18384,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X2) | s('type.num.num',X2) != s('type.num.num',X3) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))))) )),
  inference(forward_demodulation,[],[f18042,f227])).
fof(f18042,plain,(
  ( ! [X2,X3] : (s('type.num.num',X2) != s('type.num.num',X3) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f437,f227])).
fof(f18387,plain,(
  ( ! [X10,X9] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X9) | s('type.num.num',X9) != s('type.num.num',X10) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X9))))))))) )),
  inference(forward_demodulation,[],[f18046,f228])).
fof(f18046,plain,(
  ( ! [X10,X9] : (s('type.num.num',X9) != s('type.num.num',X10) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X9)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X9)))) )),
  inference(superposition,[],[f437,f228])).
fof(f18394,plain,(
  ( ! [X15] : (s('type.num.num',X15) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) )),
  inference(subsumption_resolution,[],[f18393,f195])).
fof(f18393,plain,(
  ( ! [X15] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num',X15) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) )),
  inference(forward_demodulation,[],[f18050,f140])).
fof(f18050,plain,(
  ( ! [X15] : (s('type.num.num',X15) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X15))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) )),
  inference(superposition,[],[f437,f140])).
fof(f18957,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X4))),s('type.num.num',X5))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X5)) )),
  inference(superposition,[],[f1359,f1540])).
fof(f18961,plain,(
  ( ! [X12,X11] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',X11))),s('type.num.num',X12))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X11))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12)) )),
  inference(superposition,[],[f1359,f1741])).
fof(f18969,plain,(
  ( ! [X26,X25] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X25))),s('type.num.num',X26))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X25),s('type.num.num',X25))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X26)) )),
  inference(superposition,[],[f1359,f1741])).
fof(f18970,plain,(
  ( ! [X28,X27] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X27),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num',X28))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',X27))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X28)) )),
  inference(superposition,[],[f1359,f1540])).
fof(f18985,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X3),s('type.num.num',X2))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X3))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f1359,f227])).
fof(f18989,plain,(
  ( ! [X10,X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X10),s('type.num.num',X9))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X10))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X9)))) )),
  inference(superposition,[],[f1359,f228])).
fof(f19613,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X4))),s('type.num.num',X5))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X5)) )),
  inference(superposition,[],[f1360,f1540])).
fof(f19617,plain,(
  ( ! [X12,X11] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',X11))),s('type.num.num',X12))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X11))),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12)) )),
  inference(superposition,[],[f1360,f1741])).
fof(f19625,plain,(
  ( ! [X26,X25] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X25))),s('type.num.num',X26))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X25),s('type.num.num',X25))),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X26)) )),
  inference(superposition,[],[f1360,f1741])).
fof(f19626,plain,(
  ( ! [X28,X27] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X27),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num',X28))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',X27))),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X28)) )),
  inference(superposition,[],[f1360,f1540])).
fof(f19641,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X3),s('type.num.num',X2))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f1360,f227])).
fof(f19645,plain,(
  ( ! [X10,X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X10),s('type.num.num',X9))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X9)))) )),
  inference(superposition,[],[f1360,f228])).
fof(f20216,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X3))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X3)))) )),
  inference(superposition,[],[f1372,f227])).
fof(f20220,plain,(
  ( ! [X10,X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X9),s('type.num.num',X10))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X9))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X10)))) )),
  inference(superposition,[],[f1372,f228])).
fof(f20246,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X4))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X5) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(superposition,[],[f1372,f1540])).
fof(f20250,plain,(
  ( ! [X12,X11] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X12),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',X11))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X11)))))) )),
  inference(superposition,[],[f1372,f1741])).
fof(f20258,plain,(
  ( ! [X26,X25] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X26),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X25))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X26) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X25),s('type.num.num',X25)))))) )),
  inference(superposition,[],[f1372,f1741])).
fof(f20259,plain,(
  ( ! [X28,X27] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X28),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X27),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X28) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',X27)))))) )),
  inference(superposition,[],[f1372,f1540])).
fof(f21126,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X3))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1373,f227])).
fof(f21130,plain,(
  ( ! [X10,X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X9),s('type.num.num',X10))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X9))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1373,f228])).
fof(f21156,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X4))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X5) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1373,f1540])).
fof(f21160,plain,(
  ( ! [X12,X11] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X12),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',X11))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X11))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1373,f1741])).
fof(f21168,plain,(
  ( ! [X26,X25] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X26),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X25))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X26) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X25),s('type.num.num',X25))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1373,f1741])).
fof(f21169,plain,(
  ( ! [X28,X27] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X28),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X27),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X28) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',X27))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1373,f1540])).
fof(f22701,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X4))),s('type.num.num',X5))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X5)) )),
  inference(superposition,[],[f1353,f1540])).
fof(f22705,plain,(
  ( ! [X12,X11] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',X11))),s('type.num.num',X12))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X11))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12)) )),
  inference(superposition,[],[f1353,f1741])).
fof(f22713,plain,(
  ( ! [X26,X25] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X25))),s('type.num.num',X26))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X25),s('type.num.num',X25))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X26)) )),
  inference(superposition,[],[f1353,f1741])).
fof(f22714,plain,(
  ( ! [X28,X27] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X27),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num',X28))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',X27))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X28)) )),
  inference(superposition,[],[f1353,f1540])).
fof(f22729,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X3),s('type.num.num',X2))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f1353,f227])).
fof(f22733,plain,(
  ( ! [X10,X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X10),s('type.num.num',X9))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X10),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X9)))) )),
  inference(superposition,[],[f1353,f228])).
fof(f23093,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X4))),s('type.num.num',X5))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X5)) )),
  inference(superposition,[],[f1355,f1540])).
fof(f23097,plain,(
  ( ! [X12,X11] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',X11))),s('type.num.num',X12))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X11))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12)) )),
  inference(superposition,[],[f1355,f1741])).
fof(f23105,plain,(
  ( ! [X26,X25] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X25))),s('type.num.num',X26))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X25),s('type.num.num',X25))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X26)) )),
  inference(superposition,[],[f1355,f1741])).
fof(f23106,plain,(
  ( ! [X28,X27] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X27),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num',X28))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',X27))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X28)) )),
  inference(superposition,[],[f1355,f1540])).
fof(f23121,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X3),s('type.num.num',X2))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X3))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f1355,f227])).
fof(f23125,plain,(
  ( ! [X10,X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X10),s('type.num.num',X9))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X10))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X9)))) )),
  inference(superposition,[],[f1355,f228])).
fof(f23398,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X4))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X5) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f1366,f1540])).
fof(f23402,plain,(
  ( ! [X12,X11] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X12),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',X11))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X11))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f1366,f1741])).
fof(f23410,plain,(
  ( ! [X26,X25] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X26),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X25))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X26) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X25),s('type.num.num',X25))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f1366,f1741])).
fof(f23411,plain,(
  ( ! [X28,X27] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X28),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X27),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X28) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num',X27))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f1366,f1540])).
fof(f54274,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(backward_demodulation,[],[f54273,f41448])).
fof(f54273,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(forward_demodulation,[],[f40607,f148])).
fof(f40607,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f1545,f147])).
fof(f54272,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))))) )),
  inference(forward_demodulation,[],[f40610,f54271])).
fof(f40610,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))))) )),
  inference(superposition,[],[f1545,f146])).
fof(f54271,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(forward_demodulation,[],[f41402,f146])).
fof(f41402,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f1759,f149])).
fof(f41448,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f1759,f148])).
fof(f43611,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))) )),
  inference(unit_resulting_resolution,[],[f547,f1777])).
fof(f49433,plain,(
  ( ! [X50,X48,X49] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X49),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X50))),s('type.num.num',X48))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X49) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X48)) )),
  inference(subsumption_resolution,[],[f49007,f547])).
fof(f49007,plain,(
  ( ! [X50,X48,X49] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X49),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X50))),s('type.num.num',X48))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X49) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X50),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X48)))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X48)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X48)) )),
  inference(superposition,[],[f1787,f153])).
fof(f49432,plain,(
  ( ! [X43,X41,X42,X40] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X42),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X43))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X42) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X40)) )),
  inference(subsumption_resolution,[],[f49005,f547])).
fof(f49005,plain,(
  ( ! [X43,X41,X42,X40] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X42),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X43))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X42) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X43),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X40)),s('type.num.num',X41))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X40)),s('type.num.num',X41))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X40)) )),
  inference(superposition,[],[f1787,f903])).
fof(f49431,plain,(
  ( ! [X39,X37,X38,X36] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X38),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X39))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X37))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X38)) )),
  inference(subsumption_resolution,[],[f49004,f547])).
fof(f49004,plain,(
  ( ! [X39,X37,X38,X36] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X38),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X39))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X37))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X38) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X39),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X36),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num',X37))))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X36),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num',X37)))))))))) )),
  inference(superposition,[],[f1787,f1545])).
fof(f49430,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X34),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X35))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X33))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X34) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X33)) )),
  inference(subsumption_resolution,[],[f49003,f547])).
fof(f49003,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X34),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X35))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X33))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X34) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X35),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',sK2(X33)))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',sK2(X33)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X33)) )),
  inference(superposition,[],[f1787,f762])).
fof(f48988,plain,(
  ( ! [X21,X22,X20] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X21),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X20),s('type.num.num',X22))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK2(X20)),s('type.num.num',X22))),s('type.num.num',X22))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X20)) )),
  inference(superposition,[],[f1787,f153])).
fof(f48987,plain,(
  ( ! [X19,X17,X18,X16] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X18),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X16))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))))),s('type.num.num',X19))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X16),s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))))),s('type.num.num',X17))),s('type.num.num',X19))),s('type.num.num',X19)))) )),
  inference(superposition,[],[f1787,f1759])).
fof(f48986,plain,(
  ( ! [X14,X12,X15,X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X14),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num',X13))),s('type.num.num',X15))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X14) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X12)),s('type.num.num',X13))),s('type.num.num',X15))),s('type.num.num',X15))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12)) )),
  inference(superposition,[],[f1787,f903])).
fof(f48985,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X10),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))))),s('type.num.num',X11))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num',X9))))),s('type.num.num',X11))),s('type.num.num',X11)))) )),
  inference(superposition,[],[f1787,f1545])).
fof(f48984,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X5))),s('type.num.num',X7))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X6) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK2(X5)))),s('type.num.num',X7))),s('type.num.num',X7))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X5)) )),
  inference(superposition,[],[f1787,f762])).
fof(f48953,plain,(
  ( ! [X50,X48,X49] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X49),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X48))),s('type.num.num',X50))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X49) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X48)))),s('type.num.num',X50))),s('type.num.num',X50))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X48)) )),
  inference(superposition,[],[f1787,f153])).
fof(f48951,plain,(
  ( ! [X43,X41,X42,X40] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X42),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))))),s('type.num.num',X43))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X42) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X40)),s('type.num.num',X41))))),s('type.num.num',X43))),s('type.num.num',X43))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X40)) )),
  inference(superposition,[],[f1787,f903])).
fof(f48949,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X34),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X33))))),s('type.num.num',X35))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X34) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',sK2(X33)))))),s('type.num.num',X35))),s('type.num.num',X35))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X33)) )),
  inference(superposition,[],[f1787,f762])).
fof(f49332,plain,(
  ( ! [X43,X41,X42,X40] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X42))),s('type.num.num',X43))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X42),s('type.num.num',X43))),s('type.num.num',X43))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X40)) )),
  inference(subsumption_resolution,[],[f48905,f574])).
fof(f48905,plain,(
  ( ! [X43,X41,X42,X40] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X42))),s('type.num.num',X43))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X40)),s('type.num.num',X41))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X42),s('type.num.num',X43))),s('type.num.num',X43))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X40)) )),
  inference(superposition,[],[f1787,f903])).
fof(f49331,plain,(
  ( ! [X39,X37,X38,X36] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X37))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X38))),s('type.num.num',X39))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X38),s('type.num.num',X39))),s('type.num.num',X39)))) )),
  inference(subsumption_resolution,[],[f48904,f574])).
fof(f48904,plain,(
  ( ! [X39,X37,X38,X36] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X37))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X38))),s('type.num.num',X39))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X36),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num',X37))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X38),s('type.num.num',X39))),s('type.num.num',X39)))) )),
  inference(superposition,[],[f1787,f1545])).
fof(f49330,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X33))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X34))),s('type.num.num',X35))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X34),s('type.num.num',X35))),s('type.num.num',X35))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X33)) )),
  inference(subsumption_resolution,[],[f48903,f574])).
fof(f48903,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X33))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X34))),s('type.num.num',X35))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',sK2(X33)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X34),s('type.num.num',X35))),s('type.num.num',X35))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X33)) )),
  inference(superposition,[],[f1787,f762])).
fof(f48147,plain,(
  ( ! [X78,X76,X79,X77] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X78))),s('type.num.num',X79))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X76))),s('type.num.num',X77))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X78),s('type.num.num',X79))),s('type.num.num',X79))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X76),s('type.num.num',X77))),s('type.num.num',X77)))) )),
  inference(superposition,[],[f1786,f203])).
fof(f48600,plain,(
  ( ! [X43,X41,X42,X40] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X42))),s('type.num.num',X43))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X42),s('type.num.num',X43))),s('type.num.num',X43))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X40)) )),
  inference(subsumption_resolution,[],[f48136,f574])).
fof(f48136,plain,(
  ( ! [X43,X41,X42,X40] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X42))),s('type.num.num',X43))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X42),s('type.num.num',X43))),s('type.num.num',X43))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X40)),s('type.num.num',X41))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X40)) )),
  inference(superposition,[],[f1786,f903])).
fof(f48599,plain,(
  ( ! [X39,X37,X38,X36] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X38))),s('type.num.num',X39))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X37))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X38),s('type.num.num',X39))),s('type.num.num',X39)))) )),
  inference(subsumption_resolution,[],[f48135,f574])).
fof(f48135,plain,(
  ( ! [X39,X37,X38,X36] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X38))),s('type.num.num',X39))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X37))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X38),s('type.num.num',X39))),s('type.num.num',X39))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X36),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num',X37)))))))) )),
  inference(superposition,[],[f1786,f1545])).
fof(f48598,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X34))),s('type.num.num',X35))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X33))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X34),s('type.num.num',X35))),s('type.num.num',X35))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X33)) )),
  inference(subsumption_resolution,[],[f48134,f574])).
fof(f48134,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X34))),s('type.num.num',X35))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X33))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X34),s('type.num.num',X35))),s('type.num.num',X35))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',sK2(X33)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X33)) )),
  inference(superposition,[],[f1786,f762])).
fof(f48557,plain,(
  ( ! [X50,X48,X49] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X49))),s('type.num.num',X48))),s('type.num.num',X50))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X50) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X48)) )),
  inference(subsumption_resolution,[],[f48088,f547])).
fof(f48088,plain,(
  ( ! [X50,X48,X49] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X49))),s('type.num.num',X48))),s('type.num.num',X50))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X49),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X48)))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X48)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X50) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X48)) )),
  inference(superposition,[],[f1786,f153])).
fof(f48556,plain,(
  ( ! [X43,X41,X42,X40] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X42))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))))),s('type.num.num',X43))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X43) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X40)) )),
  inference(subsumption_resolution,[],[f48086,f547])).
fof(f48086,plain,(
  ( ! [X43,X41,X42,X40] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X42))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))))),s('type.num.num',X43))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X42),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X40)),s('type.num.num',X41))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X40)),s('type.num.num',X41))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X43) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X40)) )),
  inference(superposition,[],[f1786,f903])).
fof(f48555,plain,(
  ( ! [X39,X37,X38,X36] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X38))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X37))))))),s('type.num.num',X39))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X39)) )),
  inference(subsumption_resolution,[],[f48085,f547])).
fof(f48085,plain,(
  ( ! [X39,X37,X38,X36] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X38))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X37))))))),s('type.num.num',X39))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X38),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X36),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num',X37))))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X36),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num',X37))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X39)) )),
  inference(superposition,[],[f1786,f1545])).
fof(f48554,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X34))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X33))))),s('type.num.num',X35))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X35) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X33)) )),
  inference(subsumption_resolution,[],[f48084,f547])).
fof(f48084,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X34))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X33))))),s('type.num.num',X35))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X34),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',sK2(X33)))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',sK2(X33)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X35) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X33)) )),
  inference(superposition,[],[f1786,f762])).
fof(f48069,plain,(
  ( ! [X21,X22,X20] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X20),s('type.num.num',X21))),s('type.num.num',X22))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK2(X20)),s('type.num.num',X21))),s('type.num.num',X21))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X20)) )),
  inference(superposition,[],[f1786,f153])).
fof(f48068,plain,(
  ( ! [X19,X17,X18,X16] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X16))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))))),s('type.num.num',X18))),s('type.num.num',X19))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X16),s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))))),s('type.num.num',X17))),s('type.num.num',X18))),s('type.num.num',X18))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X19)) )),
  inference(superposition,[],[f1786,f1759])).
fof(f48067,plain,(
  ( ! [X14,X12,X15,X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num',X13))),s('type.num.num',X14))),s('type.num.num',X15))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X12)),s('type.num.num',X13))),s('type.num.num',X14))),s('type.num.num',X14))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X15) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12)) )),
  inference(superposition,[],[f1786,f903])).
fof(f48066,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))))),s('type.num.num',X10))),s('type.num.num',X11))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num',X9))))),s('type.num.num',X10))),s('type.num.num',X10))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X11)) )),
  inference(superposition,[],[f1786,f1545])).
fof(f48065,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X5))),s('type.num.num',X6))),s('type.num.num',X7))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK2(X5)))),s('type.num.num',X6))),s('type.num.num',X6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X7) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X5)) )),
  inference(superposition,[],[f1786,f762])).
fof(f48034,plain,(
  ( ! [X50,X48,X49] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X48))),s('type.num.num',X49))),s('type.num.num',X50))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X48)))),s('type.num.num',X49))),s('type.num.num',X49))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X50) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X48)) )),
  inference(superposition,[],[f1786,f153])).
fof(f48032,plain,(
  ( ! [X43,X41,X42,X40] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))))),s('type.num.num',X42))),s('type.num.num',X43))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X40)),s('type.num.num',X41))))),s('type.num.num',X42))),s('type.num.num',X42))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X43) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X40)) )),
  inference(superposition,[],[f1786,f903])).
fof(f48030,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X33))))),s('type.num.num',X34))),s('type.num.num',X35))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',sK2(X33)))))),s('type.num.num',X34))),s('type.num.num',X34))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X35) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X33)) )),
  inference(superposition,[],[f1786,f762])).
fof(f47317,plain,(
  ( ! [X21,X22,X20] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X21),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X22),s('type.num.num',X20))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X22),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X22),s('type.num.num',sK2(X20)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X20)) )),
  inference(superposition,[],[f1571,f153])).
fof(f47316,plain,(
  ( ! [X19,X17,X18,X16] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X18),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X19),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X16))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X19),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X16),s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))))),s('type.num.num',X17)))))))) )),
  inference(superposition,[],[f1571,f1759])).
fof(f47315,plain,(
  ( ! [X14,X12,X15,X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X14),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X15),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num',X13))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X14) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X15),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X12)),s('type.num.num',X13))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12)) )),
  inference(superposition,[],[f1571,f903])).
fof(f47314,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X10),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num',X9)))))))))) )),
  inference(superposition,[],[f1571,f1545])).
fof(f47313,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X5))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X6) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK2(X5)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X5)) )),
  inference(superposition,[],[f1571,f762])).
fof(f47282,plain,(
  ( ! [X50,X48,X49] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X49),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X50),s('type.num.num','const.num.SUC_1'(s('type.num.num',X48))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X49) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X50),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X50),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X48)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X48)) )),
  inference(superposition,[],[f1571,f153])).
fof(f47280,plain,(
  ( ! [X43,X41,X42,X40] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X42),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X43),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X42) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X43),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X43),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X40)),s('type.num.num',X41))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X40)) )),
  inference(superposition,[],[f1571,f903])).
fof(f47278,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X34),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X35),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X33))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X34) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X35),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X35),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',sK2(X33)))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X33)) )),
  inference(superposition,[],[f1571,f762])).
fof(f47569,plain,(
  ( ! [X50,X48,X49] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X49),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X48),s('type.num.num','const.num.SUC_1'(s('type.num.num',X50))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X49) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X48)) )),
  inference(subsumption_resolution,[],[f47236,f521])).
fof(f47236,plain,(
  ( ! [X50,X48,X49] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X49),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X48),s('type.num.num','const.num.SUC_1'(s('type.num.num',X50))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X49) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X48)))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X48)))),s('type.num.num',X50))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X48)) )),
  inference(superposition,[],[f1571,f153])).
fof(f47568,plain,(
  ( ! [X43,X41,X42,X40] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X42),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X43))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X42) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X40)) )),
  inference(subsumption_resolution,[],[f47234,f521])).
fof(f47234,plain,(
  ( ! [X43,X41,X42,X40] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X42),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X43))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X42) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X40)),s('type.num.num',X41))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X40)),s('type.num.num',X41))))),s('type.num.num',X43))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X40)) )),
  inference(superposition,[],[f1571,f903])).
fof(f47567,plain,(
  ( ! [X39,X37,X38,X36] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X38),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X37))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X39))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X38)) )),
  inference(subsumption_resolution,[],[f47233,f521])).
fof(f47233,plain,(
  ( ! [X39,X37,X38,X36] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X38),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X37))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X39))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X38) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X36),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num',X37))))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X36),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num',X37))))))),s('type.num.num',X39)))))) )),
  inference(superposition,[],[f1571,f1545])).
fof(f47566,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X34),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X33))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X35))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X34) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X33)) )),
  inference(subsumption_resolution,[],[f47232,f521])).
fof(f47232,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X34),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X33))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X35))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X34) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',sK2(X33)))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',sK2(X33)))))),s('type.num.num',X35))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X33)) )),
  inference(superposition,[],[f1571,f762])).
fof(f47199,plain,(
  ( ! [X78,X76,X79,X77] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X76))),s('type.num.num',X77))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X78),s('type.num.num','const.num.SUC_1'(s('type.num.num',X79))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X76),s('type.num.num',X77))),s('type.num.num',X77))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X78),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X78),s('type.num.num',X79)))))) )),
  inference(superposition,[],[f1571,f203])).
fof(f47540,plain,(
  ( ! [X43,X41,X42,X40] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X42),s('type.num.num','const.num.SUC_1'(s('type.num.num',X43))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X42),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X42),s('type.num.num',X43))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X40)) )),
  inference(subsumption_resolution,[],[f47188,f574])).
fof(f47188,plain,(
  ( ! [X43,X41,X42,X40] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X42),s('type.num.num','const.num.SUC_1'(s('type.num.num',X43))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X40)),s('type.num.num',X41))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X42),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X42),s('type.num.num',X43))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X40)) )),
  inference(superposition,[],[f1571,f903])).
fof(f47539,plain,(
  ( ! [X39,X37,X38,X36] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X37))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X38),s('type.num.num','const.num.SUC_1'(s('type.num.num',X39))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X38),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X38),s('type.num.num',X39)))))) )),
  inference(subsumption_resolution,[],[f47187,f574])).
fof(f47187,plain,(
  ( ! [X39,X37,X38,X36] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X37))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X38),s('type.num.num','const.num.SUC_1'(s('type.num.num',X39))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X36),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num',X37))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X38),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X38),s('type.num.num',X39)))))) )),
  inference(superposition,[],[f1571,f1545])).
fof(f47538,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X33))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X34),s('type.num.num','const.num.SUC_1'(s('type.num.num',X35))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X34),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X34),s('type.num.num',X35))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X33)) )),
  inference(subsumption_resolution,[],[f47186,f574])).
fof(f47186,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X33))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X34),s('type.num.num','const.num.SUC_1'(s('type.num.num',X35))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',sK2(X33)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X34),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X34),s('type.num.num',X35))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X33)) )),
  inference(superposition,[],[f1571,f762])).
fof(f46391,plain,(
  ( ! [X78,X76,X79,X77] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X78),s('type.num.num','const.num.SUC_1'(s('type.num.num',X79))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X76))),s('type.num.num',X77))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X78),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X78),s('type.num.num',X79))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X76),s('type.num.num',X77))),s('type.num.num',X77)))) )),
  inference(superposition,[],[f1570,f203])).
fof(f46386,plain,(
  ( ! [X61,X59,X62,X60] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X61),s('type.num.num','const.num.SUC_1'(s('type.num.num',X62))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X59),s('type.num.num','const.num.SUC_1'(s('type.num.num',X60))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X61),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X61),s('type.num.num',X62))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X59),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X59),s('type.num.num',X60)))))) )),
  inference(superposition,[],[f1570,f202])).
fof(f46760,plain,(
  ( ! [X43,X41,X42,X40] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X42),s('type.num.num','const.num.SUC_1'(s('type.num.num',X43))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X42),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X42),s('type.num.num',X43))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X40)) )),
  inference(subsumption_resolution,[],[f46380,f574])).
fof(f46380,plain,(
  ( ! [X43,X41,X42,X40] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X42),s('type.num.num','const.num.SUC_1'(s('type.num.num',X43))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X42),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X42),s('type.num.num',X43))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X40)),s('type.num.num',X41))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X40)) )),
  inference(superposition,[],[f1570,f903])).
fof(f46759,plain,(
  ( ! [X39,X37,X38,X36] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X38),s('type.num.num','const.num.SUC_1'(s('type.num.num',X39))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X37))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X38),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X38),s('type.num.num',X39)))))) )),
  inference(subsumption_resolution,[],[f46379,f574])).
fof(f46379,plain,(
  ( ! [X39,X37,X38,X36] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X38),s('type.num.num','const.num.SUC_1'(s('type.num.num',X39))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X37))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X38),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X38),s('type.num.num',X39))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X36),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num',X37)))))))) )),
  inference(superposition,[],[f1570,f1545])).
fof(f46758,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X34),s('type.num.num','const.num.SUC_1'(s('type.num.num',X35))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X33))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X34),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X34),s('type.num.num',X35))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X33)) )),
  inference(subsumption_resolution,[],[f46378,f574])).
fof(f46378,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X34),s('type.num.num','const.num.SUC_1'(s('type.num.num',X35))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X33))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X34),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X34),s('type.num.num',X35))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',sK2(X33)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X33)) )),
  inference(superposition,[],[f1570,f762])).
fof(f46360,plain,(
  ( ! [X21,X22,X20] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X21),s('type.num.num',X20))),s('type.num.num',X22))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X21),s('type.num.num',sK2(X20)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X20)) )),
  inference(superposition,[],[f1570,f153])).
fof(f46359,plain,(
  ( ! [X19,X17,X18,X16] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X18),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X16))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))))))),s('type.num.num',X19))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X18),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X18),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X16),s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))))),s('type.num.num',X17))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X19)) )),
  inference(superposition,[],[f1570,f1759])).
fof(f46358,plain,(
  ( ! [X14,X12,X15,X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X14),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num',X13))))),s('type.num.num',X15))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X14),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X14),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X12)),s('type.num.num',X13))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X15) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12)) )),
  inference(superposition,[],[f1570,f903])).
fof(f46357,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X10),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))))))),s('type.num.num',X11))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X10),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num',X9))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X11)) )),
  inference(superposition,[],[f1570,f1545])).
fof(f46356,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X5))))),s('type.num.num',X7))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK2(X5)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X7) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X5)) )),
  inference(superposition,[],[f1570,f762])).
fof(f46325,plain,(
  ( ! [X50,X48,X49] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X49),s('type.num.num','const.num.SUC_1'(s('type.num.num',X48))))),s('type.num.num',X50))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X49),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X49),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X48)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X50) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X48)) )),
  inference(superposition,[],[f1570,f153])).
fof(f46323,plain,(
  ( ! [X43,X41,X42,X40] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X42),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))))))),s('type.num.num',X43))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X42),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X42),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X40)),s('type.num.num',X41))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X43) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X40)) )),
  inference(superposition,[],[f1570,f903])).
fof(f46321,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X34),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X33))))))),s('type.num.num',X35))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X34),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X34),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',sK2(X33)))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X35) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X33)) )),
  inference(superposition,[],[f1570,f762])).
fof(f46647,plain,(
  ( ! [X50,X48,X49] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X48),s('type.num.num','const.num.SUC_1'(s('type.num.num',X49))))),s('type.num.num',X50))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X50) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X48)) )),
  inference(subsumption_resolution,[],[f46279,f521])).
fof(f46279,plain,(
  ( ! [X50,X48,X49] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X48),s('type.num.num','const.num.SUC_1'(s('type.num.num',X49))))),s('type.num.num',X50))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X48)))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X48)))),s('type.num.num',X49))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X50) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X48)) )),
  inference(superposition,[],[f1570,f153])).
fof(f46646,plain,(
  ( ! [X43,X41,X42,X40] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X42))))),s('type.num.num',X43))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X43) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X40)) )),
  inference(subsumption_resolution,[],[f46277,f521])).
fof(f46277,plain,(
  ( ! [X43,X41,X42,X40] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X42))))),s('type.num.num',X43))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X40)),s('type.num.num',X41))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X40)),s('type.num.num',X41))))),s('type.num.num',X42))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X43) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X40)) )),
  inference(superposition,[],[f1570,f903])).
fof(f46645,plain,(
  ( ! [X39,X37,X38,X36] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X37))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X38))))),s('type.num.num',X39))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X39)) )),
  inference(subsumption_resolution,[],[f46276,f521])).
fof(f46276,plain,(
  ( ! [X39,X37,X38,X36] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X37))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X38))))),s('type.num.num',X39))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X36),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num',X37))))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X36),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num',X37))))))),s('type.num.num',X38))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X39)) )),
  inference(superposition,[],[f1570,f1545])).
fof(f46644,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X33))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X34))))),s('type.num.num',X35))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X35) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X33)) )),
  inference(subsumption_resolution,[],[f46275,f521])).
fof(f46275,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X33))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X34))))),s('type.num.num',X35))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',sK2(X33)))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',sK2(X33)))))),s('type.num.num',X34))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X35) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X33)) )),
  inference(superposition,[],[f1570,f762])).
fof(f44821,plain,(
  ( ! [X50,X48,X49] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X49),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X50))),s('type.num.num',X48))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X48)) )),
  inference(subsumption_resolution,[],[f44689,f547])).
fof(f44689,plain,(
  ( ! [X50,X48,X49] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X49),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X50))),s('type.num.num',X48))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X50),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X48)))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X48)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X48)) )),
  inference(superposition,[],[f1784,f153])).
fof(f44820,plain,(
  ( ! [X43,X41,X42,X40] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X42),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X43))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X40)) )),
  inference(subsumption_resolution,[],[f44687,f547])).
fof(f44687,plain,(
  ( ! [X43,X41,X42,X40] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X42),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X43))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X43),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X40)),s('type.num.num',X41))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X40)),s('type.num.num',X41))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X40)) )),
  inference(superposition,[],[f1784,f903])).
fof(f44819,plain,(
  ( ! [X39,X37,X38,X36] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X38),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X39))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X37)))))))))) )),
  inference(subsumption_resolution,[],[f44686,f547])).
fof(f44686,plain,(
  ( ! [X39,X37,X38,X36] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X38),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X39))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X37))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X39),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X36),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num',X37))))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X36),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num',X37)))))))))) )),
  inference(superposition,[],[f1784,f1545])).
fof(f44818,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X34),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X35))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X33))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X33)) )),
  inference(subsumption_resolution,[],[f44685,f547])).
fof(f44685,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X34),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X35))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X33))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X35),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',sK2(X33)))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',sK2(X33)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X33)) )),
  inference(superposition,[],[f1784,f762])).
fof(f44670,plain,(
  ( ! [X21,X22,X20] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X20),s('type.num.num',X22))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK2(X20)),s('type.num.num',X22))),s('type.num.num',X22))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X20)) )),
  inference(superposition,[],[f1784,f153])).
fof(f44669,plain,(
  ( ! [X19,X17,X18,X16] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X18),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X16))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))))),s('type.num.num',X19))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X16),s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))))),s('type.num.num',X17))),s('type.num.num',X19))),s('type.num.num',X19)))) )),
  inference(superposition,[],[f1784,f1759])).
fof(f44668,plain,(
  ( ! [X14,X12,X15,X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X14),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num',X13))),s('type.num.num',X15))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X12)),s('type.num.num',X13))),s('type.num.num',X15))),s('type.num.num',X15))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12)) )),
  inference(superposition,[],[f1784,f903])).
fof(f44667,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))))),s('type.num.num',X11))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num',X9))))),s('type.num.num',X11))),s('type.num.num',X11)))) )),
  inference(superposition,[],[f1784,f1545])).
fof(f44666,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X5))),s('type.num.num',X7))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK2(X5)))),s('type.num.num',X7))),s('type.num.num',X7))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X5)) )),
  inference(superposition,[],[f1784,f762])).
fof(f44635,plain,(
  ( ! [X50,X48,X49] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X49),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X48))),s('type.num.num',X50))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X48)))),s('type.num.num',X50))),s('type.num.num',X50))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X48)) )),
  inference(superposition,[],[f1784,f153])).
fof(f44633,plain,(
  ( ! [X43,X41,X42,X40] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X42),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))))),s('type.num.num',X43))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X40)),s('type.num.num',X41))))),s('type.num.num',X43))),s('type.num.num',X43))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X40)) )),
  inference(superposition,[],[f1784,f903])).
fof(f44631,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X34),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X33))))),s('type.num.num',X35))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',sK2(X33)))))),s('type.num.num',X35))),s('type.num.num',X35))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X33)) )),
  inference(superposition,[],[f1784,f762])).
fof(f44268,plain,(
  ( ! [X10,X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))),s('type.num.num',X10))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X9),s('type.num.num',X10))),s('type.num.num',X10)))) )),
  inference(superposition,[],[f1781,f148])).
fof(f44264,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))),s('type.num.num',X1))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X1)))) )),
  inference(superposition,[],[f1781,f203])).
fof(f44363,plain,(
  ( ! [X50,X48,X49] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X49))),s('type.num.num',X48))),s('type.num.num',X50))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X48)) )),
  inference(subsumption_resolution,[],[f44185,f547])).
fof(f44185,plain,(
  ( ! [X50,X48,X49] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X49))),s('type.num.num',X48))),s('type.num.num',X50))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X49),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X48)))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X48)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X48)) )),
  inference(superposition,[],[f1781,f153])).
fof(f44362,plain,(
  ( ! [X43,X41,X42,X40] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X42))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))))),s('type.num.num',X43))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X40)) )),
  inference(subsumption_resolution,[],[f44183,f547])).
fof(f44183,plain,(
  ( ! [X43,X41,X42,X40] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X42))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))))),s('type.num.num',X43))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X42),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X40)),s('type.num.num',X41))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X40)),s('type.num.num',X41))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X40)) )),
  inference(superposition,[],[f1781,f903])).
fof(f44361,plain,(
  ( ! [X39,X37,X38,X36] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X38))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X37))))))),s('type.num.num',X39)))) )),
  inference(subsumption_resolution,[],[f44182,f547])).
fof(f44182,plain,(
  ( ! [X39,X37,X38,X36] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X38))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X37))))))),s('type.num.num',X39))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X38),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X36),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num',X37))))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X36),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num',X37)))))))))) )),
  inference(superposition,[],[f1781,f1545])).
fof(f44360,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X34))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X33))))),s('type.num.num',X35))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X33)) )),
  inference(subsumption_resolution,[],[f44181,f547])).
fof(f44181,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X34))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X33))))),s('type.num.num',X35))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X34),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',sK2(X33)))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',sK2(X33)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X33)) )),
  inference(superposition,[],[f1781,f762])).
fof(f44166,plain,(
  ( ! [X21,X22,X20] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X20),s('type.num.num',X21))),s('type.num.num',X22))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK2(X20)),s('type.num.num',X21))),s('type.num.num',X21))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X20)) )),
  inference(superposition,[],[f1781,f153])).
fof(f44165,plain,(
  ( ! [X19,X17,X18,X16] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X16))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))))),s('type.num.num',X18))),s('type.num.num',X19))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X16),s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))))),s('type.num.num',X17))),s('type.num.num',X18))),s('type.num.num',X18)))) )),
  inference(superposition,[],[f1781,f1759])).
fof(f44164,plain,(
  ( ! [X14,X12,X15,X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num',X13))),s('type.num.num',X14))),s('type.num.num',X15))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X12)),s('type.num.num',X13))),s('type.num.num',X14))),s('type.num.num',X14))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12)) )),
  inference(superposition,[],[f1781,f903])).
fof(f44163,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))))),s('type.num.num',X10))),s('type.num.num',X11))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num',X9))))),s('type.num.num',X10))),s('type.num.num',X10)))) )),
  inference(superposition,[],[f1781,f1545])).
fof(f44162,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X5))),s('type.num.num',X6))),s('type.num.num',X7))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK2(X5)))),s('type.num.num',X6))),s('type.num.num',X6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X5)) )),
  inference(superposition,[],[f1781,f762])).
fof(f44131,plain,(
  ( ! [X50,X48,X49] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X48))),s('type.num.num',X49))),s('type.num.num',X50))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X48)))),s('type.num.num',X49))),s('type.num.num',X49))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X48)) )),
  inference(superposition,[],[f1781,f153])).
fof(f44129,plain,(
  ( ! [X43,X41,X42,X40] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))))),s('type.num.num',X42))),s('type.num.num',X43))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X40)),s('type.num.num',X41))))),s('type.num.num',X42))),s('type.num.num',X42))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X40)) )),
  inference(superposition,[],[f1781,f903])).
fof(f44127,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X33))))),s('type.num.num',X34))),s('type.num.num',X35))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',sK2(X33)))))),s('type.num.num',X34))),s('type.num.num',X34))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X33)) )),
  inference(superposition,[],[f1781,f762])).
fof(f43730,plain,(
  ( ! [X50,X48,X49] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X49))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X50))),s('type.num.num',X48))))) | s('type.num.num',X49) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X50),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X48)))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X48)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X48)) )),
  inference(superposition,[],[f1777,f153])).
fof(f43728,plain,(
  ( ! [X43,X41,X42,X40] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X42))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X43))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))))))) | s('type.num.num',X42) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X43),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X40)),s('type.num.num',X41))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X40)),s('type.num.num',X41))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X40)) )),
  inference(superposition,[],[f1777,f903])).
fof(f43726,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X34))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X35))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X33))))))) | s('type.num.num',X34) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X35),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',sK2(X33)))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',sK2(X33)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X33)) )),
  inference(superposition,[],[f1777,f762])).
fof(f43711,plain,(
  ( ! [X21,X22,X20] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X21))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X20),s('type.num.num',X22))))) | s('type.num.num',X21) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK2(X20)),s('type.num.num',X22))),s('type.num.num',X22))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X20)) )),
  inference(superposition,[],[f1777,f153])).
fof(f43710,plain,(
  ( ! [X19,X17,X18,X16] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X18))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X16))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))))),s('type.num.num',X19))))) | s('type.num.num',X18) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X16),s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))))),s('type.num.num',X17))),s('type.num.num',X19))),s('type.num.num',X19)))) )),
  inference(superposition,[],[f1777,f1759])).
fof(f43709,plain,(
  ( ! [X14,X12,X15,X13] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X14))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num',X13))),s('type.num.num',X15))))) | s('type.num.num',X14) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X12)),s('type.num.num',X13))),s('type.num.num',X15))),s('type.num.num',X15))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12)) )),
  inference(superposition,[],[f1777,f903])).
fof(f43708,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))))),s('type.num.num',X11))))) | s('type.num.num',X10) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num',X9))))),s('type.num.num',X11))),s('type.num.num',X11)))) )),
  inference(superposition,[],[f1777,f1545])).
fof(f43707,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X5))),s('type.num.num',X7))))) | s('type.num.num',X6) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK2(X5)))),s('type.num.num',X7))),s('type.num.num',X7))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X5)) )),
  inference(superposition,[],[f1777,f762])).
fof(f43676,plain,(
  ( ! [X50,X48,X49] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X49))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X48))),s('type.num.num',X50))))) | s('type.num.num',X49) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X48)))),s('type.num.num',X50))),s('type.num.num',X50))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X48)) )),
  inference(superposition,[],[f1777,f153])).
fof(f43674,plain,(
  ( ! [X43,X41,X42,X40] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X42))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))))),s('type.num.num',X43))))) | s('type.num.num',X42) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X40)),s('type.num.num',X41))))),s('type.num.num',X43))),s('type.num.num',X43))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X40)) )),
  inference(superposition,[],[f1777,f903])).
fof(f43672,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X34))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X33))))),s('type.num.num',X35))))) | s('type.num.num',X34) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',sK2(X33)))))),s('type.num.num',X35))),s('type.num.num',X35))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X33)) )),
  inference(superposition,[],[f1777,f762])).
fof(f43657,plain,(
  ( ! [X19,X17,X18,X16] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X16))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X18))),s('type.num.num',X19))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X18),s('type.num.num',X19))),s('type.num.num',X19))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X16),s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))))),s('type.num.num',X17)))) )),
  inference(superposition,[],[f1777,f1759])).
fof(f43656,plain,(
  ( ! [X14,X12,X15,X13] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num',X13))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X14))),s('type.num.num',X15))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X12)),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X14),s('type.num.num',X15))),s('type.num.num',X15))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12)) )),
  inference(superposition,[],[f1777,f903])).
fof(f43655,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))),s('type.num.num',X11))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X10),s('type.num.num',X11))),s('type.num.num',X11))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num',X9)))))) )),
  inference(superposition,[],[f1777,f1545])).
fof(f43654,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X5))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.num.num',X7))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK2(X5)))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X6),s('type.num.num',X7))),s('type.num.num',X7))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X5)) )),
  inference(superposition,[],[f1777,f762])).
fof(f43649,plain,(
  ( ! [X111,X112,X113] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X111))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X112))),s('type.num.num',X113))))) | s('type.num.num',sK2('const.num.SUC_1'(s('type.num.num',X111)))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X112),s('type.num.num',X113))),s('type.num.num',X113)))) )),
  inference(superposition,[],[f1777,f962])).
fof(f43640,plain,(
  ( ! [X78,X76,X79,X77] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X76))),s('type.num.num',X77))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X78))),s('type.num.num',X79))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X76),s('type.num.num',X77))),s('type.num.num',X77))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X78),s('type.num.num',X79))),s('type.num.num',X79)))) )),
  inference(superposition,[],[f1777,f203])).
fof(f43631,plain,(
  ( ! [X47,X45,X46,X44] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X46))),s('type.num.num',X47))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X44))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45))))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X46),s('type.num.num',X47))),s('type.num.num',X47))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45))))),s('type.num.num',X45)))))) )),
  inference(superposition,[],[f1777,f1759])).
fof(f43630,plain,(
  ( ! [X43,X41,X42,X40] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X42))),s('type.num.num',X43))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X40)),s('type.num.num',X41))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X42),s('type.num.num',X43))),s('type.num.num',X43))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X40)) )),
  inference(superposition,[],[f1777,f903])).
fof(f43629,plain,(
  ( ! [X39,X37,X38,X36] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X38))),s('type.num.num',X39))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X37))))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X38),s('type.num.num',X39))),s('type.num.num',X39))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X36),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num',X37)))))))) )),
  inference(superposition,[],[f1777,f1545])).
fof(f43628,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X33))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X34))),s('type.num.num',X35))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',sK2(X33)))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X34),s('type.num.num',X35))),s('type.num.num',X35))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X33)) )),
  inference(superposition,[],[f1777,f762])).
fof(f43445,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1)))))) )),
  inference(superposition,[],[f1568,f202])).
fof(f43440,plain,(
  ( ! [X21,X22,X20] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X22),s('type.num.num',X20))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X22),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X22),s('type.num.num',sK2(X20)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X20)) )),
  inference(superposition,[],[f1568,f153])).
fof(f43439,plain,(
  ( ! [X19,X17,X18,X16] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X18),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X19),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X16))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X19),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X16),s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))))),s('type.num.num',X17)))))))) )),
  inference(superposition,[],[f1568,f1759])).
fof(f43438,plain,(
  ( ! [X14,X12,X15,X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X14),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X15),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num',X13))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X15),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X12)),s('type.num.num',X13))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12)) )),
  inference(superposition,[],[f1568,f903])).
fof(f43437,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num',X9)))))))))) )),
  inference(superposition,[],[f1568,f1545])).
fof(f43436,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X5))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK2(X5)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X5)) )),
  inference(superposition,[],[f1568,f762])).
fof(f43405,plain,(
  ( ! [X50,X48,X49] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X49),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X50),s('type.num.num','const.num.SUC_1'(s('type.num.num',X48))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X50),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X50),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X48)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X48)) )),
  inference(superposition,[],[f1568,f153])).
fof(f43403,plain,(
  ( ! [X43,X41,X42,X40] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X42),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X43),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X43),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X43),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X40)),s('type.num.num',X41))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X40)) )),
  inference(superposition,[],[f1568,f903])).
fof(f43401,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X34),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X35),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X33))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X35),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X35),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',sK2(X33)))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X33)) )),
  inference(superposition,[],[f1568,f762])).
fof(f43469,plain,(
  ( ! [X50,X48,X49] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X49),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X48),s('type.num.num','const.num.SUC_1'(s('type.num.num',X50))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X48)) )),
  inference(subsumption_resolution,[],[f43359,f521])).
fof(f43359,plain,(
  ( ! [X50,X48,X49] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X49),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X48),s('type.num.num','const.num.SUC_1'(s('type.num.num',X50))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X48)))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X48)))),s('type.num.num',X50))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X48)) )),
  inference(superposition,[],[f1568,f153])).
fof(f43468,plain,(
  ( ! [X43,X41,X42,X40] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X42),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X43))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X40)) )),
  inference(subsumption_resolution,[],[f43357,f521])).
fof(f43357,plain,(
  ( ! [X43,X41,X42,X40] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X42),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X43))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X40)),s('type.num.num',X41))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X40)),s('type.num.num',X41))))),s('type.num.num',X43))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X40)) )),
  inference(superposition,[],[f1568,f903])).
fof(f43467,plain,(
  ( ! [X39,X37,X38,X36] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X38),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X37))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X39)))))))) )),
  inference(subsumption_resolution,[],[f43356,f521])).
fof(f43356,plain,(
  ( ! [X39,X37,X38,X36] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X38),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X37))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X39))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X36),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num',X37))))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X36),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num',X37))))))),s('type.num.num',X39)))))) )),
  inference(superposition,[],[f1568,f1545])).
fof(f43466,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X34),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X33))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X35))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X33)) )),
  inference(subsumption_resolution,[],[f43355,f521])).
fof(f43355,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X34),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X33))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X35))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',sK2(X33)))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',sK2(X33)))))),s('type.num.num',X35))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X33)) )),
  inference(superposition,[],[f1568,f762])).
fof(f42951,plain,(
  ( ! [X10,X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X9),s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X9),s('type.num.num',X10)))))) )),
  inference(superposition,[],[f1565,f148])).
fof(f42897,plain,(
  ( ! [X21,X22,X20] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X21),s('type.num.num',X20))),s('type.num.num',X22))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X21),s('type.num.num',sK2(X20)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X20)) )),
  inference(superposition,[],[f1565,f153])).
fof(f42896,plain,(
  ( ! [X19,X17,X18,X16] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X18),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X16))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))))))),s('type.num.num',X19))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X18),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X18),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X16),s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))))),s('type.num.num',X17)))))))) )),
  inference(superposition,[],[f1565,f1759])).
fof(f42895,plain,(
  ( ! [X14,X12,X15,X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X14),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num',X13))))),s('type.num.num',X15))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X14),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X14),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X12)),s('type.num.num',X13))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12)) )),
  inference(superposition,[],[f1565,f903])).
fof(f42894,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X10),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))))))),s('type.num.num',X11))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X10),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num',X9)))))))))) )),
  inference(superposition,[],[f1565,f1545])).
fof(f42893,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X5))))),s('type.num.num',X7))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK2(X5)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X5)) )),
  inference(superposition,[],[f1565,f762])).
fof(f42863,plain,(
  ( ! [X50,X48,X49] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X49),s('type.num.num','const.num.SUC_1'(s('type.num.num',X48))))),s('type.num.num',X50))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X49),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X49),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X48)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X48)) )),
  inference(superposition,[],[f1565,f153])).
fof(f42861,plain,(
  ( ! [X43,X41,X42,X40] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X42),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))))))),s('type.num.num',X43))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X42),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X42),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X40)),s('type.num.num',X41))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X40)) )),
  inference(superposition,[],[f1565,f903])).
fof(f42859,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X34),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X33))))))),s('type.num.num',X35))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X34),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X34),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',sK2(X33)))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X33)) )),
  inference(superposition,[],[f1565,f762])).
fof(f42971,plain,(
  ( ! [X50,X48,X49] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X48),s('type.num.num','const.num.SUC_1'(s('type.num.num',X49))))),s('type.num.num',X50))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X48)) )),
  inference(subsumption_resolution,[],[f42818,f521])).
fof(f42818,plain,(
  ( ! [X50,X48,X49] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X48),s('type.num.num','const.num.SUC_1'(s('type.num.num',X49))))),s('type.num.num',X50))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X48)))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X48)))),s('type.num.num',X49))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X48)) )),
  inference(superposition,[],[f1565,f153])).
fof(f42970,plain,(
  ( ! [X43,X41,X42,X40] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X42))))),s('type.num.num',X43))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X40)) )),
  inference(subsumption_resolution,[],[f42816,f521])).
fof(f42816,plain,(
  ( ! [X43,X41,X42,X40] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X42))))),s('type.num.num',X43))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X40)),s('type.num.num',X41))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X40)),s('type.num.num',X41))))),s('type.num.num',X42))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X40)) )),
  inference(superposition,[],[f1565,f903])).
fof(f42969,plain,(
  ( ! [X39,X37,X38,X36] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X37))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X38))))),s('type.num.num',X39)))) )),
  inference(subsumption_resolution,[],[f42815,f521])).
fof(f42815,plain,(
  ( ! [X39,X37,X38,X36] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X37))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X38))))),s('type.num.num',X39))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X36),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num',X37))))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X36),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num',X37))))))),s('type.num.num',X38)))))) )),
  inference(superposition,[],[f1565,f1545])).
fof(f42968,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X33))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X34))))),s('type.num.num',X35))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X33)) )),
  inference(subsumption_resolution,[],[f42814,f521])).
fof(f42814,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X33))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X34))))),s('type.num.num',X35))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',sK2(X33)))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',sK2(X33)))))),s('type.num.num',X34))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X33)) )),
  inference(superposition,[],[f1565,f762])).
fof(f42601,plain,(
  ( ! [X21,X22,X20] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X21))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X22),s('type.num.num',X20))))) | s('type.num.num',X21) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X22),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X22),s('type.num.num',sK2(X20)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X20)) )),
  inference(superposition,[],[f1561,f153])).
fof(f42600,plain,(
  ( ! [X19,X17,X18,X16] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X18))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X19),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X16))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))))))))) | s('type.num.num',X18) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X19),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X16),s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))))),s('type.num.num',X17)))))))) )),
  inference(superposition,[],[f1561,f1759])).
fof(f42599,plain,(
  ( ! [X14,X12,X15,X13] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X14))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X15),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num',X13))))))) | s('type.num.num',X14) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X15),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X12)),s('type.num.num',X13))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12)) )),
  inference(superposition,[],[f1561,f903])).
fof(f42598,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))))))))) | s('type.num.num',X10) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num',X9)))))))))) )),
  inference(superposition,[],[f1561,f1545])).
fof(f42597,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X5))))))) | s('type.num.num',X6) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK2(X5)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X5)) )),
  inference(superposition,[],[f1561,f762])).
fof(f42567,plain,(
  ( ! [X50,X48,X49] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X49))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X50),s('type.num.num','const.num.SUC_1'(s('type.num.num',X48))))))) | s('type.num.num',X49) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X50),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X50),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X48)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X48)) )),
  inference(superposition,[],[f1561,f153])).
fof(f42565,plain,(
  ( ! [X43,X41,X42,X40] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X42))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X43),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))))))))) | s('type.num.num',X42) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X43),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X43),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X40)),s('type.num.num',X41))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X40)) )),
  inference(superposition,[],[f1561,f903])).
fof(f42563,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X34))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X35),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X33))))))))) | s('type.num.num',X34) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X35),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X35),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',sK2(X33)))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X33)) )),
  inference(superposition,[],[f1561,f762])).
fof(f42522,plain,(
  ( ! [X50,X48,X49] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X49))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X48),s('type.num.num','const.num.SUC_1'(s('type.num.num',X50))))))) | s('type.num.num',X49) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X48)))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X48)))),s('type.num.num',X50))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X48)) )),
  inference(superposition,[],[f1561,f153])).
fof(f42520,plain,(
  ( ! [X43,X41,X42,X40] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X42))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X43))))))) | s('type.num.num',X42) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X40)),s('type.num.num',X41))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X40)),s('type.num.num',X41))))),s('type.num.num',X43))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X40)) )),
  inference(superposition,[],[f1561,f903])).
fof(f42518,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X34))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X33))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X35))))))) | s('type.num.num',X34) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',sK2(X33)))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',sK2(X33)))))),s('type.num.num',X35))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X33)) )),
  inference(superposition,[],[f1561,f762])).
fof(f42503,plain,(
  ( ! [X19,X17,X18,X16] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X16))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X18),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X18),s('type.num.num',X19))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X16),s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))))),s('type.num.num',X17)))) )),
  inference(superposition,[],[f1561,f1759])).
fof(f42502,plain,(
  ( ! [X14,X12,X15,X13] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num',X13))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',X15))))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X12)),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X14),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X14),s('type.num.num',X15))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12)) )),
  inference(superposition,[],[f1561,f903])).
fof(f42501,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X10),s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X10),s('type.num.num',X11))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num',X9)))))) )),
  inference(superposition,[],[f1561,f1545])).
fof(f42500,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X5))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK2(X5)))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X6),s('type.num.num',X7))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X5)) )),
  inference(superposition,[],[f1561,f762])).
fof(f42495,plain,(
  ( ! [X111,X109,X110] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X109))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X110),s('type.num.num','const.num.SUC_1'(s('type.num.num',X111))))))) | s('type.num.num',sK2('const.num.SUC_1'(s('type.num.num',X109)))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X110),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X110),s('type.num.num',X111)))))) )),
  inference(superposition,[],[f1561,f962])).
fof(f42486,plain,(
  ( ! [X78,X76,X79,X77] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X78),s('type.num.num','const.num.SUC_1'(s('type.num.num',X79))))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X76))),s('type.num.num',X77))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X78),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X78),s('type.num.num',X79))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X76),s('type.num.num',X77))),s('type.num.num',X77)))) )),
  inference(superposition,[],[f1561,f203])).
fof(f42481,plain,(
  ( ! [X61,X59,X62,X60] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X59),s('type.num.num','const.num.SUC_1'(s('type.num.num',X60))))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X61),s('type.num.num','const.num.SUC_1'(s('type.num.num',X62))))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X59),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X59),s('type.num.num',X60))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X61),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X61),s('type.num.num',X62)))))) )),
  inference(superposition,[],[f1561,f202])).
fof(f42477,plain,(
  ( ! [X47,X45,X46,X44] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X46),s('type.num.num','const.num.SUC_1'(s('type.num.num',X47))))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X44))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45))))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X46),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X46),s('type.num.num',X47))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45))))),s('type.num.num',X45)))))) )),
  inference(superposition,[],[f1561,f1759])).
fof(f42476,plain,(
  ( ! [X43,X41,X42,X40] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X42),s('type.num.num','const.num.SUC_1'(s('type.num.num',X43))))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X40)),s('type.num.num',X41))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X42),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X42),s('type.num.num',X43))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X40)) )),
  inference(superposition,[],[f1561,f903])).
fof(f42475,plain,(
  ( ! [X39,X37,X38,X36] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X38),s('type.num.num','const.num.SUC_1'(s('type.num.num',X39))))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X37))))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X38),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X38),s('type.num.num',X39))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X36),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num',X37)))))))) )),
  inference(superposition,[],[f1561,f1545])).
fof(f42474,plain,(
  ( ! [X35,X33,X34,X32] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X33))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X34),s('type.num.num','const.num.SUC_1'(s('type.num.num',X35))))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',sK2(X33)))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X34),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X34),s('type.num.num',X35))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X33)) )),
  inference(superposition,[],[f1561,f762])).
fof(f42167,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK3(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,f(s(X0,sK3(X1,X2,bool,X0))))) )),
  inference(condensation,[],[f42166])).
fof(f42166,plain,(
  ( ! [X10,X8,X11,X9] : (s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK3(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,f(s(X8,sK3(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(subsumption_resolution,[],[f42156,f152])).
fof(f42156,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') != s(bool,X11) | s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK3(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,f(s(X8,sK3(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f1855,f293])).
fof(f41538,plain,(
  ( ! [X231,X230,X232] : (s('type.num.num',X232) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X230))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X231))))))) | s('type.num.num',sK2(X232)) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X230),s('type.num.num','const.num.SUC_1'(s('type.num.num',X231))))),s('type.num.num',X231))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X232)) )),
  inference(superposition,[],[f387,f1759])).
fof(f41683,plain,(
  ( ! [X229,X227,X228] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X229))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X227))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X228))))) | s('type.num.num',X229) = s('type.num.num',sK2('const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X227),s('type.num.num','const.num.SUC_1'(s('type.num.num',X228))))),s('type.num.num',X228))))))) )),
  inference(subsumption_resolution,[],[f41537,f574])).
fof(f41537,plain,(
  ( ! [X229,X227,X228] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X229))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X227))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X228))))) | s('type.num.num',X229) = s('type.num.num',sK2('const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X227),s('type.num.num','const.num.SUC_1'(s('type.num.num',X228))))),s('type.num.num',X228)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X227),s('type.num.num','const.num.SUC_1'(s('type.num.num',X228))))),s('type.num.num',X228)))))) )),
  inference(superposition,[],[f387,f1759])).
fof(f41536,plain,(
  ( ! [X225,X226,X224] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X226))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X224))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X225))))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X226),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X224),s('type.num.num','const.num.SUC_1'(s('type.num.num',X225))))),s('type.num.num',X225)))))) )),
  inference(superposition,[],[f386,f1759])).
fof(f41534,plain,(
  ( ! [X218,X220,X219] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X220))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X218))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X219))))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X220))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X218),s('type.num.num','const.num.SUC_1'(s('type.num.num',X219))))),s('type.num.num',X219)))))) )),
  inference(superposition,[],[f385,f1759])).
fof(f41532,plain,(
  ( ! [X212,X213,X214] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X214))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X212))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X213))))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X214)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X212),s('type.num.num','const.num.SUC_1'(s('type.num.num',X213))))),s('type.num.num',X213))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X214)) )),
  inference(superposition,[],[f384,f1759])).
fof(f41518,plain,(
  ( ! [X169,X170,X168] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X170))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X168))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X169))))))) | s('type.num.num',X170) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X168),s('type.num.num','const.num.SUC_1'(s('type.num.num',X169))))),s('type.num.num',X169)))))) )),
  inference(superposition,[],[f208,f1759])).
fof(f41477,plain,(
  ( ! [X68,X69,X67] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X67))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X68))))),s('type.num.num',X69))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X67),s('type.num.num','const.num.SUC_1'(s('type.num.num',X68))))),s('type.num.num',X68))),s('type.num.num',X69)))) )),
  inference(superposition,[],[f1761,f1759])).
fof(f41474,plain,(
  ( ! [X59,X60,X58] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X60),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X58))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X59))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X60),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X58),s('type.num.num','const.num.SUC_1'(s('type.num.num',X59))))),s('type.num.num',X59)))))) )),
  inference(superposition,[],[f1547,f1759])).
fof(f41473,plain,(
  ( ! [X57,X56,X55] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X55))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X56))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X57))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X55),s('type.num.num','const.num.SUC_1'(s('type.num.num',X56))))),s('type.num.num',X56))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X55))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X56))))),s('type.num.num',X57)))))))) )),
  inference(superposition,[],[f1545,f1759])).
fof(f41469,plain,(
  ( ! [X45,X46,X44] : (s('type.num.num',X46) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X44))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45))))) | s('type.num.num',sK2(X46)) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45))))),s('type.num.num',X45))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X46)) )),
  inference(superposition,[],[f387,f1759])).
fof(f41468,plain,(
  ( ! [X43,X41,X42] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X43))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X41))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X42))))) | s('type.num.num',X43) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X41),s('type.num.num','const.num.SUC_1'(s('type.num.num',X42))))),s('type.num.num',X42))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f386,f1759])).
fof(f41467,plain,(
  ( ! [X39,X38,X40] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X40))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X38))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X39))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X38),s('type.num.num','const.num.SUC_1'(s('type.num.num',X39))))),s('type.num.num',X39)))) )),
  inference(superposition,[],[f386,f1759])).
fof(f41466,plain,(
  ( ! [X37,X35,X36] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X37))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X35))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))))) | s('type.num.num',X37) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X35),s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))))),s('type.num.num',X36)))))) )),
  inference(superposition,[],[f385,f1759])).
fof(f41465,plain,(
  ( ! [X33,X34,X32] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X34))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X32))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X33))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X34))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X32),s('type.num.num','const.num.SUC_1'(s('type.num.num',X33))))),s('type.num.num',X33)))) )),
  inference(superposition,[],[f385,f1759])).
fof(f41464,plain,(
  ( ! [X30,X31,X29] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X31))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X29))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X30))))) | s('type.num.num',X31) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X29),s('type.num.num','const.num.SUC_1'(s('type.num.num',X30))))),s('type.num.num',X30)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X29),s('type.num.num','const.num.SUC_1'(s('type.num.num',X30))))),s('type.num.num',X30)))) )),
  inference(superposition,[],[f384,f1759])).
fof(f41463,plain,(
  ( ! [X28,X26,X27] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X26))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X28)))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X26),s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))))),s('type.num.num',X27))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X28)) )),
  inference(superposition,[],[f384,f1759])).
fof(f41461,plain,(
  ( ! [X21,X22,X20] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X22))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X20))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X21))))) | s('type.num.num',X22) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X20),s('type.num.num','const.num.SUC_1'(s('type.num.num',X21))))),s('type.num.num',X21)))) )),
  inference(superposition,[],[f208,f1759])).
fof(f41460,plain,(
  ( ! [X19,X17,X18] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X18))))),s('type.num.num',X19))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X17),s('type.num.num','const.num.SUC_1'(s('type.num.num',X18))))),s('type.num.num',X18))),s('type.num.num',X19)))))) )),
  inference(superposition,[],[f201,f1759])).
fof(f41459,plain,(
  ( ! [X14,X15,X16] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X16),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X14))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X15))))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X16),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X14),s('type.num.num','const.num.SUC_1'(s('type.num.num',X15))))),s('type.num.num',X15)))))))) )),
  inference(superposition,[],[f200,f1759])).
fof(f41458,plain,(
  ( ! [X12,X13] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X12))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X13))))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X12),s('type.num.num','const.num.SUC_1'(s('type.num.num',X13))))),s('type.num.num',X13)))) )),
  inference(superposition,[],[f195,f1759])).
fof(f41452,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X3)))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X3)))))),s('type.num.num',X3))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X3)) )),
  inference(superposition,[],[f1759,f762])).
fof(f41449,plain,(
  ( ! [X2,X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))))),s('type.num.num',X2)))))))) )),
  inference(superposition,[],[f1759,f200])).
fof(f41400,plain,(
  ( ! [X14,X12,X13] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X14))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X12))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X13))))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X14),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X12))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X13))))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X12),s('type.num.num','const.num.SUC_1'(s('type.num.num',X13))))),s('type.num.num',X13)))))))) )),
  inference(superposition,[],[f1759,f1759])).
fof(f41399,plain,(
  ( ! [X10,X11,X9] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',X10))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',X10))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X9)),s('type.num.num',X10))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X9)) )),
  inference(superposition,[],[f1759,f903])).
fof(f41398,plain,(
  ( ! [X6,X8,X7] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X8),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.num.num',X7)))))))))) )),
  inference(superposition,[],[f1759,f1545])).
fof(f41397,plain,(
  ( ! [X4,X5,X3] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X4))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X4))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',sK2(X4)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X4)) )),
  inference(superposition,[],[f1759,f762])).
fof(f41395,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num',sK2(X0)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(superposition,[],[f1759,f153])).
fof(f40859,plain,(
  ( ! [X349,X347,X348] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X349))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X347))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X348)))))))) )),
  inference(subsumption_resolution,[],[f40738,f574])).
fof(f40738,plain,(
  ( ! [X349,X347,X348] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X349))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X347))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X348))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X347),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X347))),s('type.num.num',X348)))))))) )),
  inference(superposition,[],[f1760,f1545])).
fof(f40858,plain,(
  ( ! [X339,X340,X338] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X338))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X339))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X340)))))) )),
  inference(subsumption_resolution,[],[f40735,f574])).
fof(f40735,plain,(
  ( ! [X339,X340,X338] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X338))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X339))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X340))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X338),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X338))),s('type.num.num',X339)))))))) )),
  inference(superposition,[],[f1548,f1545])).
fof(f40855,plain,(
  ( ! [X329,X331,X328,X330] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X328))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X329))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X330),s('type.num.num',X331)))))))) )),
  inference(subsumption_resolution,[],[f40732,f574])).
fof(f40732,plain,(
  ( ! [X329,X331,X328,X330] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X328))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X329))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X330),s('type.num.num',X331))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X328),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X328))),s('type.num.num',X329)))))))) )),
  inference(superposition,[],[f1445,f1545])).
fof(f40854,plain,(
  ( ! [X316,X318,X317,X319] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X318),s('type.num.num',X319))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X316))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X317)))))))) )),
  inference(subsumption_resolution,[],[f40729,f574])).
fof(f40729,plain,(
  ( ! [X316,X318,X317,X319] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X318),s('type.num.num',X319))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X316))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X317))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X316),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X316))),s('type.num.num',X317)))))))) )),
  inference(superposition,[],[f1441,f1545])).
fof(f40699,plain,(
  ( ! [X229,X227,X228] : (s('type.num.num',X229) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X227))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X228))))))) | s('type.num.num',sK2(X229)) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X227),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X227))),s('type.num.num',X228))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X229)) )),
  inference(superposition,[],[f387,f1545])).
fof(f40841,plain,(
  ( ! [X225,X226,X224] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X226))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X224))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X225))))) | s('type.num.num',X226) = s('type.num.num',sK2('const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X224),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X224))),s('type.num.num',X225))))))))) )),
  inference(subsumption_resolution,[],[f40698,f574])).
fof(f40698,plain,(
  ( ! [X225,X226,X224] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X226))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X224))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X225))))) | s('type.num.num',X226) = s('type.num.num',sK2('const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X224),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X224))),s('type.num.num',X225)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X224),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X224))),s('type.num.num',X225)))))))) )),
  inference(superposition,[],[f387,f1545])).
fof(f40697,plain,(
  ( ! [X222,X223,X221] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X223))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X221))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X222))))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X223),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X221),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X221))),s('type.num.num',X222)))))))) )),
  inference(superposition,[],[f386,f1545])).
fof(f40695,plain,(
  ( ! [X216,X215,X217] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X217))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X215))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X216))))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X217))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X215),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X215))),s('type.num.num',X216)))))))) )),
  inference(superposition,[],[f385,f1545])).
fof(f40693,plain,(
  ( ! [X210,X211,X209] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X211))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X209))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X210))))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X211)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X209),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X209))),s('type.num.num',X210))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X211)) )),
  inference(superposition,[],[f384,f1545])).
fof(f40679,plain,(
  ( ! [X167,X165,X166] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X167))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X165))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X166))))))) | s('type.num.num',X167) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X165),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X165))),s('type.num.num',X166)))))))) )),
  inference(superposition,[],[f208,f1545])).
fof(f40638,plain,(
  ( ! [X66,X64,X65] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X64))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X65))))),s('type.num.num',X66))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X64),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X64))),s('type.num.num',X65))))),s('type.num.num',X66)))) )),
  inference(superposition,[],[f1761,f1545])).
fof(f40635,plain,(
  ( ! [X57,X56,X55] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X57),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X55))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X56))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X57),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X55),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X55))),s('type.num.num',X56)))))))) )),
  inference(superposition,[],[f1547,f1545])).
fof(f40633,plain,(
  ( ! [X52,X50,X51] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X52),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X50))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X51)))))))) )),
  inference(superposition,[],[f547,f1545])).
fof(f40632,plain,(
  ( ! [X47,X48,X49] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X47))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X48))))),s('type.num.num',X49)))) )),
  inference(superposition,[],[f521,f1545])).
fof(f40631,plain,(
  ( ! [X45,X46,X44] : (s('type.num.num',X46) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X44))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45))))) | s('type.num.num',sK2(X46)) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X44),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X44))),s('type.num.num',X45))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X46)) )),
  inference(superposition,[],[f387,f1545])).
fof(f40630,plain,(
  ( ! [X43,X41,X42] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X43))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X41))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X42))))) | s('type.num.num',X43) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X41),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X41))),s('type.num.num',X42))))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f386,f1545])).
fof(f40629,plain,(
  ( ! [X39,X38,X40] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X40))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X38))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X39))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X38),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X38))),s('type.num.num',X39)))))) )),
  inference(superposition,[],[f386,f1545])).
fof(f40628,plain,(
  ( ! [X37,X35,X36] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X37))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X35))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))))) | s('type.num.num',X37) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X35),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X35))),s('type.num.num',X36)))))))) )),
  inference(superposition,[],[f385,f1545])).
fof(f40627,plain,(
  ( ! [X33,X34,X32] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X34))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X32))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X33))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X34))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X32))),s('type.num.num',X33)))))) )),
  inference(superposition,[],[f385,f1545])).
fof(f40626,plain,(
  ( ! [X30,X31,X29] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X31))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X29))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X30))))) | s('type.num.num',X31) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num',X29),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X29))),s('type.num.num',X30)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X29),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X29))),s('type.num.num',X30)))))) )),
  inference(superposition,[],[f384,f1545])).
fof(f40625,plain,(
  ( ! [X28,X26,X27] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X26))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X28)))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X26),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X26))),s('type.num.num',X27))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X28)) )),
  inference(superposition,[],[f384,f1545])).
fof(f40623,plain,(
  ( ! [X21,X22,X20] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X22))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X20))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X21))))) | s('type.num.num',X22) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X20),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X20))),s('type.num.num',X21)))))) )),
  inference(superposition,[],[f208,f1545])).
fof(f40622,plain,(
  ( ! [X19,X17,X18] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X18))))),s('type.num.num',X19))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X17),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))),s('type.num.num',X18))))),s('type.num.num',X19)))))) )),
  inference(superposition,[],[f201,f1545])).
fof(f40621,plain,(
  ( ! [X14,X15,X16] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X16),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X14))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X15))))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X16),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X14),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X14))),s('type.num.num',X15)))))))))) )),
  inference(superposition,[],[f200,f1545])).
fof(f40620,plain,(
  ( ! [X12,X13] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X12))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X13))))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X12))),s('type.num.num',X13)))))) )),
  inference(superposition,[],[f195,f1545])).
fof(f40616,plain,(
  ( ! [X4,X5,X3] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))))),s('type.num.num',X5))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X5)) )),
  inference(superposition,[],[f1441,f1545])).
fof(f40615,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X2)) )),
  inference(superposition,[],[f1445,f1545])).
fof(f40614,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X2)))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X2)))),s('type.num.num',X3))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X2)) )),
  inference(superposition,[],[f1545,f903])).
fof(f40611,plain,(
  ( ! [X2,X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))),s('type.num.num',X2)))))))))) )),
  inference(superposition,[],[f1545,f201])).
fof(f40564,plain,(
  ( ! [X10,X11,X9] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',X10))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X9)),s('type.num.num',X10))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',X10))),s('type.num.num',X11))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X9)) )),
  inference(superposition,[],[f1545,f903])).
fof(f40563,plain,(
  ( ! [X6,X8,X7] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.num.num',X7))))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))),s('type.num.num',X8)))))))) )),
  inference(superposition,[],[f1545,f1545])).
fof(f40562,plain,(
  ( ! [X4,X5,X3] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X4))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',sK2(X4)))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X4))),s('type.num.num',X5))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X4)) )),
  inference(superposition,[],[f1545,f762])).
fof(f40560,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X0)),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(superposition,[],[f1545,f153])).
fof(f159,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f59])).
fof(f59,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.arithmetic.MULT_EQ_|49|','thm.bool.EQ_CLAUSES|split|2')).
fof(f160,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f104])).
fof(f226,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f142,f138])).
fof(f142,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f99])).
fof(f99,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f46])).
fof(f46,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f36])).
fof(f36,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.arithmetic.MULT_EQ_|49|','thm.bool.NOT_CLAUSES|split|2')).
fof(f1854,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK3(X0,X1,X2,bool))) )),
  inference(superposition,[],[f216,f152])).
fof(f1855,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK3(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK3(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f216,f152])).
fof(f2859,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,'T_0') != s(bool,f(s(X39,sK3(X40,X41,bool,X39)))) | s(fun(X39,bool),X40) = s(fun(X39,bool),X41) | s(bool,X42) = s(bool,g(s(X39,sK3(X40,X41,bool,X39)))) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f216,f293])).
fof(f2816,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,X44) != s(bool,f(s(X41,sK3(X42,X43,bool,X41)))) | s(fun(X41,bool),X42) = s(fun(X41,bool),X43) | s(bool,'T_0') = s(bool,X44) | s(bool,'T_0') = s(bool,g(s(X41,sK3(X42,X43,bool,X41))))) )),
  inference(superposition,[],[f216,f293])).
fof(f19661,plain,(
  ( ! [X43,X42] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X43),s('type.num.num',X42))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X43),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X42)))) )),
  inference(superposition,[],[f1360,f146])).
fof(f19657,plain,(
  ( ! [X33,X32] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X33),s('type.num.num',X32))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X33),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1360,f148])).
fof(f19005,plain,(
  ( ! [X43,X42] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X43),s('type.num.num',X42))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X43))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X42)))) )),
  inference(superposition,[],[f1359,f146])).
fof(f19001,plain,(
  ( ! [X33,X32] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X33),s('type.num.num',X32))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X33))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1359,f148])).
fof(f18083,plain,(
  ( ! [X24,X23,X22] : (s('type.num.num',X24) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X22),s('type.num.num',X23))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X24)))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X22)),s('type.num.num',X23))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22)) )),
  inference(superposition,[],[f437,f903])).
fof(f18082,plain,(
  ( ! [X21,X19,X20] : (s('type.num.num',X21) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',X20))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X21)))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',sK2(X20)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X20)) )),
  inference(superposition,[],[f437,f762])).
fof(f18080,plain,(
  ( ! [X17,X16] : (s('type.num.num',X16) != s('type.num.num',X17) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X16)))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X17)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X17) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X16)) )),
  inference(superposition,[],[f437,f153])).
fof(f18418,plain,(
  ( ! [X50,X51] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X50) | s('type.num.num',X50) != s('type.num.num',X51) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X51))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(sK2('const.num.SUC_1'(s('type.num.num',X50)))))))))) )),
  inference(forward_demodulation,[],[f18066,f962])).
fof(f18066,plain,(
  ( ! [X50,X51] : (s('type.num.num',X50) != s('type.num.num',X51) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X51))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(sK2('const.num.SUC_1'(s('type.num.num',X50))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK2('const.num.SUC_1'(s('type.num.num',X50))))) )),
  inference(superposition,[],[f437,f962])).
fof(f18412,plain,(
  ( ! [X43,X42] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X42) | s('type.num.num',X42) != s('type.num.num',X43) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X43))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X42))))))))) )),
  inference(forward_demodulation,[],[f18062,f146])).
fof(f18062,plain,(
  ( ! [X43,X42] : (s('type.num.num',X42) != s('type.num.num',X43) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X43))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X42)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X42)))) )),
  inference(superposition,[],[f437,f146])).
fof(f18411,plain,(
  ( ! [X39,X41,X40] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X39))),s('type.num.num',X40))) | s('type.num.num',X41) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X39))),s('type.num.num',X40))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X41))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X39),s('type.num.num',X40))),s('type.num.num',X40))))))))) )),
  inference(forward_demodulation,[],[f18061,f203])).
fof(f18061,plain,(
  ( ! [X39,X41,X40] : (s('type.num.num',X41) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X39))),s('type.num.num',X40))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X41))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X39),s('type.num.num',X40))),s('type.num.num',X40)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X39),s('type.num.num',X40))),s('type.num.num',X40)))) )),
  inference(superposition,[],[f437,f203])).
fof(f18405,plain,(
  ( ! [X33,X32] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X32) | s('type.num.num',X32) != s('type.num.num',X33) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X33))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num','const.num.|48|_0'))))))))) )),
  inference(forward_demodulation,[],[f18058,f148])).
fof(f18058,plain,(
  ( ! [X33,X32] : (s('type.num.num',X32) != s('type.num.num',X33) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X33))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num','const.num.|48|_0')))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f437,f148])).
fof(f18404,plain,(
  ( ! [X30,X31,X29] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X29),s('type.num.num','const.num.SUC_1'(s('type.num.num',X30))))) | s('type.num.num',X31) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X29),s('type.num.num','const.num.SUC_1'(s('type.num.num',X30))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X31))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num',X29),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X29),s('type.num.num',X30))))))))))) )),
  inference(forward_demodulation,[],[f18057,f202])).
fof(f18057,plain,(
  ( ! [X30,X31,X29] : (s('type.num.num',X31) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X29),s('type.num.num','const.num.SUC_1'(s('type.num.num',X30))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X31))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num',X29),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X29),s('type.num.num',X30)))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X29),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X29),s('type.num.num',X30)))))) )),
  inference(superposition,[],[f437,f202])).
fof(f18130,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) )),
  inference(forward_demodulation,[],[f18003,f3500])).
fof(f18003,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))) != s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))) )),
  inference(unit_resulting_resolution,[],[f547,f195,f437])).
fof(f18129,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) )),
  inference(forward_demodulation,[],[f18008,f3506])).
fof(f18008,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1))) != s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1))))) )),
  inference(unit_resulting_resolution,[],[f521,f195,f437])).
fof(f10983,plain,(
  ( ! [X6,X8,X7] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X8)))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X6)),s('type.num.num',X7))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X8) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X6)) )),
  inference(superposition,[],[f384,f903])).
fof(f10982,plain,(
  ( ! [X4,X5,X3] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X4))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X5)))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',sK2(X4)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X5) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X4)) )),
  inference(superposition,[],[f384,f762])).
fof(f10976,plain,(
  ( ! [X52,X51] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X51))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X52))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X52)))) = s('type.num.num',sK2('const.num.SUC_1'(s('type.num.num',X51)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X52)) )),
  inference(superposition,[],[f384,f962])).
fof(f10970,plain,(
  ( ! [X39,X41,X40] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X41))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X39))),s('type.num.num',X40))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X41)))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X39),s('type.num.num',X40))),s('type.num.num',X40))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X41)) )),
  inference(superposition,[],[f384,f203])).
fof(f10966,plain,(
  ( ! [X30,X31,X29] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X31))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X29),s('type.num.num','const.num.SUC_1'(s('type.num.num',X30))))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X31)))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X29),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X29),s('type.num.num',X30))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X31)) )),
  inference(superposition,[],[f384,f202])).
fof(f10963,plain,(
  ( ! [X24,X23,X22] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X24))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X22),s('type.num.num',X23))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X24)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X22)),s('type.num.num',X23))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22)) )),
  inference(superposition,[],[f384,f903])).
fof(f10962,plain,(
  ( ! [X21,X19,X20] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X21))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',X20))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X21)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',sK2(X20)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X20)) )),
  inference(superposition,[],[f384,f762])).
fof(f10960,plain,(
  ( ! [X17,X16] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X16))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X16)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X17)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X17) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X16)) )),
  inference(superposition,[],[f384,f153])).
fof(f10949,plain,(
  ( ! [X6,X8,X7] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7))) | s('type.num.num',X8) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num',sK2(X6)),s('type.num.num',X7)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X6)),s('type.num.num',X7))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X6)) )),
  inference(superposition,[],[f384,f903])).
fof(f10948,plain,(
  ( ! [X4,X5,X3] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X4))) | s('type.num.num',X5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',sK2(X4))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',sK2(X4)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X4)) )),
  inference(superposition,[],[f384,f762])).
fof(f10946,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | s('type.num.num',X1) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(sK2(X0))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK2(X0)) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(superposition,[],[f384,f153])).
fof(f11131,plain,(
  ( ! [X52,X51] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X51) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X51))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X52))) | s('type.num.num',X52) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(sK2('const.num.SUC_1'(s('type.num.num',X51)))))))) )),
  inference(forward_demodulation,[],[f10942,f962])).
fof(f10942,plain,(
  ( ! [X52,X51] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X51))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X52))) | s('type.num.num',X52) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(sK2('const.num.SUC_1'(s('type.num.num',X51))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK2('const.num.SUC_1'(s('type.num.num',X51))))) )),
  inference(superposition,[],[f384,f962])).
fof(f11123,plain,(
  ( ! [X43,X42] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X42) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X42))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X43))) | s('type.num.num',X43) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X42))))))) )),
  inference(forward_demodulation,[],[f10937,f146])).
fof(f10937,plain,(
  ( ! [X43,X42] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X42))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X43))) | s('type.num.num',X43) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X42)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X42)))) )),
  inference(superposition,[],[f384,f146])).
fof(f11122,plain,(
  ( ! [X39,X41,X40] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X39))),s('type.num.num',X40))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X41))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X39))),s('type.num.num',X40))))) | s('type.num.num',X41) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X39),s('type.num.num',X40))),s('type.num.num',X40))))))) )),
  inference(forward_demodulation,[],[f10936,f203])).
fof(f10936,plain,(
  ( ! [X39,X41,X40] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X41))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X39))),s('type.num.num',X40))))) | s('type.num.num',X41) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X39),s('type.num.num',X40))),s('type.num.num',X40)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X39),s('type.num.num',X40))),s('type.num.num',X40)))) )),
  inference(superposition,[],[f384,f203])).
fof(f11116,plain,(
  ( ! [X33,X32] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X32) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X32))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X33))) | s('type.num.num',X33) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f10933,f148])).
fof(f10933,plain,(
  ( ! [X33,X32] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X32))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X33))) | s('type.num.num',X33) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f384,f148])).
fof(f11115,plain,(
  ( ! [X30,X31,X29] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X29),s('type.num.num','const.num.SUC_1'(s('type.num.num',X30))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X31))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X29),s('type.num.num','const.num.SUC_1'(s('type.num.num',X30))))))) | s('type.num.num',X31) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num',X29),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X29),s('type.num.num',X30))))))))) )),
  inference(forward_demodulation,[],[f10932,f202])).
fof(f10932,plain,(
  ( ! [X30,X31,X29] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X31))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X29),s('type.num.num','const.num.SUC_1'(s('type.num.num',X30))))))) | s('type.num.num',X31) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num',X29),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X29),s('type.num.num',X30)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X29),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X29),s('type.num.num',X30)))))) )),
  inference(superposition,[],[f384,f202])).
fof(f10239,plain,(
  ( ! [X346,X347,X348] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X346),s('type.num.num',X347))))),s('type.num.num',X348))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X346)),s('type.num.num',X347))))),s('type.num.num',X348))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X346)) )),
  inference(superposition,[],[f1761,f903])).
fof(f10326,plain,(
  ( ! [X344,X343,X345] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X345))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X343),s('type.num.num',X344))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X343)) )),
  inference(subsumption_resolution,[],[f10238,f574])).
fof(f10238,plain,(
  ( ! [X344,X343,X345] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X345))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X343),s('type.num.num',X344))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X343)),s('type.num.num',X344))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X343)) )),
  inference(superposition,[],[f1760,f903])).
fof(f10325,plain,(
  ( ! [X331,X332,X330] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X330),s('type.num.num',X331))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X332))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X330)) )),
  inference(subsumption_resolution,[],[f10233,f574])).
fof(f10233,plain,(
  ( ! [X331,X332,X330] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X330),s('type.num.num',X331))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X332))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X330)),s('type.num.num',X331))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X330)) )),
  inference(superposition,[],[f1548,f903])).
fof(f10232,plain,(
  ( ! [X327,X329,X328] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X329),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X327),s('type.num.num',X328))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X329),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X327)),s('type.num.num',X328))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X327)) )),
  inference(superposition,[],[f1547,f903])).
fof(f10231,plain,(
  ( ! [X325,X324,X326] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X324),s('type.num.num',X325))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X326))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X324)),s('type.num.num',X325))))),s('type.num.num',X326))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X324)) )),
  inference(superposition,[],[f1547,f903])).
fof(f10324,plain,(
  ( ! [X321,X318,X320,X319] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X318),s('type.num.num',X319))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X320),s('type.num.num',X321))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X318)) )),
  inference(subsumption_resolution,[],[f10229,f574])).
fof(f10229,plain,(
  ( ! [X321,X318,X320,X319] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X318),s('type.num.num',X319))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X320),s('type.num.num',X321))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X318)),s('type.num.num',X319))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X318)) )),
  inference(superposition,[],[f1445,f903])).
fof(f10323,plain,(
  ( ! [X308,X306,X309,X307] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X308),s('type.num.num',X309))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X306),s('type.num.num',X307))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X306)) )),
  inference(subsumption_resolution,[],[f10226,f574])).
fof(f10226,plain,(
  ( ! [X308,X306,X309,X307] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X308),s('type.num.num',X309))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X306),s('type.num.num',X307))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X306)),s('type.num.num',X307))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X306)) )),
  inference(superposition,[],[f1441,f903])).
fof(f10196,plain,(
  ( ! [X222,X220,X221] : (s('type.num.num',X222) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X220),s('type.num.num',X221))))) | s('type.num.num',sK2(X222)) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X220)),s('type.num.num',X221))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X222) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X220)) )),
  inference(superposition,[],[f387,f903])).
fof(f10322,plain,(
  ( ! [X218,X219,X217] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X219))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X217),s('type.num.num',X218))) | s('type.num.num',X219) = s('type.num.num',sK2('const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X217)),s('type.num.num',X218)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X217)) )),
  inference(subsumption_resolution,[],[f10195,f574])).
fof(f10195,plain,(
  ( ! [X218,X219,X217] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X219))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X217),s('type.num.num',X218))) | s('type.num.num',X219) = s('type.num.num',sK2('const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X217)),s('type.num.num',X218)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X217)),s('type.num.num',X218))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X217)) )),
  inference(superposition,[],[f387,f903])).
fof(f10194,plain,(
  ( ! [X216,X215,X214] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X216))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X214),s('type.num.num',X215))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X216),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X214)),s('type.num.num',X215))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X214)) )),
  inference(superposition,[],[f386,f903])).
fof(f10193,plain,(
  ( ! [X212,X213,X211] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X213))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X211),s('type.num.num',X212))))) | s('type.num.num',X213) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X211)),s('type.num.num',X212))))),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X211)) )),
  inference(superposition,[],[f386,f903])).
fof(f10192,plain,(
  ( ! [X210,X208,X209] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X210))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X208),s('type.num.num',X209))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X210))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X208)),s('type.num.num',X209))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X208)) )),
  inference(superposition,[],[f385,f903])).
fof(f10191,plain,(
  ( ! [X206,X207,X205] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X207))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X205),s('type.num.num',X206))))) | s('type.num.num',X207) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X205)),s('type.num.num',X206))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X205)) )),
  inference(superposition,[],[f385,f903])).
fof(f10172,plain,(
  ( ! [X149,X150,X148] : (s('type.num.num',X150) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X148),s('type.num.num',X149))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X150))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X148)),s('type.num.num',X149))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X148)) )),
  inference(superposition,[],[f209,f903])).
fof(f10170,plain,(
  ( ! [X144,X142,X143] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X144))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X142),s('type.num.num',X143))))) | s('type.num.num',X144) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X142)),s('type.num.num',X143))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X142)) )),
  inference(superposition,[],[f208,f903])).
fof(f10145,plain,(
  ( ! [X80,X81,X79] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X79),s('type.num.num',X80))),s('type.num.num',X81))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X79)),s('type.num.num',X80))),s('type.num.num',X81))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X79)) )),
  inference(superposition,[],[f1761,f903])).
fof(f10142,plain,(
  ( ! [X70,X72,X71] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X72),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X70),s('type.num.num',X71))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X72),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X70)),s('type.num.num',X71))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X70)) )),
  inference(superposition,[],[f1547,f903])).
fof(f10140,plain,(
  ( ! [X66,X67,X65] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X67),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X65),s('type.num.num',X66))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X65)) )),
  inference(superposition,[],[f547,f903])).
fof(f10139,plain,(
  ( ! [X64,X62,X63] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X62),s('type.num.num',X63))),s('type.num.num',X64))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X62)) )),
  inference(superposition,[],[f521,f903])).
fof(f10137,plain,(
  ( ! [X59,X57,X58] : (s('type.num.num',X59) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X57),s('type.num.num',X58))) | s('type.num.num',sK2(X59)) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X57)),s('type.num.num',X58))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X59) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X57)) )),
  inference(superposition,[],[f387,f903])).
fof(f10136,plain,(
  ( ! [X54,X56,X55] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X56))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X54),s('type.num.num',X55))) | s('type.num.num',X56) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X54)),s('type.num.num',X55))),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X54)) )),
  inference(superposition,[],[f386,f903])).
fof(f10135,plain,(
  ( ! [X52,X53,X51] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X53))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X51),s('type.num.num',X52))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X53),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X51)),s('type.num.num',X52))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X51)) )),
  inference(superposition,[],[f386,f903])).
fof(f10134,plain,(
  ( ! [X50,X48,X49] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X50))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X48),s('type.num.num',X49))) | s('type.num.num',X50) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X48)),s('type.num.num',X49))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X48)) )),
  inference(superposition,[],[f385,f903])).
fof(f10133,plain,(
  ( ! [X47,X45,X46] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X47))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X45),s('type.num.num',X46))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X47))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X45)),s('type.num.num',X46))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X45)) )),
  inference(superposition,[],[f385,f903])).
fof(f10125,plain,(
  ( ! [X23,X21,X22] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X23))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',X22))) | s('type.num.num',X23) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X21)),s('type.num.num',X22))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X21)) )),
  inference(superposition,[],[f208,f903])).
fof(f10124,plain,(
  ( ! [X19,X20,X18] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X18),s('type.num.num',X19))),s('type.num.num',X20))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X18)),s('type.num.num',X19))),s('type.num.num',X20))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X18)) )),
  inference(superposition,[],[f201,f903])).
fof(f10123,plain,(
  ( ! [X17,X15,X16] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X17),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num',X16))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X17),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X15)),s('type.num.num',X16))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X15)) )),
  inference(superposition,[],[f200,f903])).
fof(f10122,plain,(
  ( ! [X14,X13] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X13),s('type.num.num',X14))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X13)),s('type.num.num',X14))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13)) )),
  inference(superposition,[],[f195,f903])).
fof(f10118,plain,(
  ( ! [X6,X4,X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X5))),s('type.num.num',X6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X6) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X4)) )),
  inference(superposition,[],[f1441,f903])).
fof(f10117,plain,(
  ( ! [X2,X3,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X3) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) )),
  inference(superposition,[],[f1445,f903])).
fof(f10115,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',sK2(X3)))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X2)),s('type.num.num',X3))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X2) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X3)) )),
  inference(superposition,[],[f903,f762])).
fof(f10318,plain,(
  ( ! [X6,X5] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num',X6))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X5)),s('type.num.num',X6))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X5)) )),
  inference(forward_demodulation,[],[f10113,f200])).
fof(f10113,plain,(
  ( ! [X6,X5] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X5)),s('type.num.num',X6))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X5)) )),
  inference(superposition,[],[f903,f200])).
fof(f10111,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK2(X2)),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK2(X2)),s('type.num.num',X3))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X2)) )),
  inference(superposition,[],[f903,f202])).
fof(f9604,plain,(
  ( ! [X337,X339,X338] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X337),s('type.num.num',X338))))),s('type.num.num',X339))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X337),s('type.num.num',sK2(X338)))))),s('type.num.num',X339))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X338)) )),
  inference(superposition,[],[f1761,f762])).
fof(f9689,plain,(
  ( ! [X335,X334,X336] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X336))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X334),s('type.num.num',X335))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X335)) )),
  inference(subsumption_resolution,[],[f9603,f574])).
fof(f9603,plain,(
  ( ! [X335,X334,X336] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X336))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X334),s('type.num.num',X335))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X334),s('type.num.num',sK2(X335)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X335)) )),
  inference(superposition,[],[f1760,f762])).
fof(f9688,plain,(
  ( ! [X321,X323,X322] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X321),s('type.num.num',X322))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X323))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X322)) )),
  inference(subsumption_resolution,[],[f9598,f574])).
fof(f9598,plain,(
  ( ! [X321,X323,X322] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X321),s('type.num.num',X322))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X323))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X321),s('type.num.num',sK2(X322)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X322)) )),
  inference(superposition,[],[f1548,f762])).
fof(f9597,plain,(
  ( ! [X318,X320,X319] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X320),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X318),s('type.num.num',X319))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X320),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X318),s('type.num.num',sK2(X319)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X319)) )),
  inference(superposition,[],[f1547,f762])).
fof(f9596,plain,(
  ( ! [X315,X316,X317] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X315),s('type.num.num',X316))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X317))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X315),s('type.num.num',sK2(X316)))))),s('type.num.num',X317))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X316)) )),
  inference(superposition,[],[f1547,f762])).
fof(f9687,plain,(
  ( ! [X310,X312,X309,X311] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X309),s('type.num.num',X310))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X311),s('type.num.num',X312))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X310)) )),
  inference(subsumption_resolution,[],[f9594,f574])).
fof(f9594,plain,(
  ( ! [X310,X312,X309,X311] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X309),s('type.num.num',X310))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X311),s('type.num.num',X312))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X309),s('type.num.num',sK2(X310)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X310)) )),
  inference(superposition,[],[f1445,f762])).
fof(f9686,plain,(
  ( ! [X300,X298,X297,X299] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X299),s('type.num.num',X300))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X297),s('type.num.num',X298))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X298)) )),
  inference(subsumption_resolution,[],[f9591,f574])).
fof(f9591,plain,(
  ( ! [X300,X298,X297,X299] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X299),s('type.num.num',X300))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X297),s('type.num.num',X298))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X297),s('type.num.num',sK2(X298)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X298)) )),
  inference(superposition,[],[f1441,f762])).
fof(f9562,plain,(
  ( ! [X216,X215,X214] : (s('type.num.num',X216) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X214),s('type.num.num',X215))))) | s('type.num.num',sK2(X216)) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X214),s('type.num.num',sK2(X215)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X216) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X215)) )),
  inference(superposition,[],[f387,f762])).
fof(f9685,plain,(
  ( ! [X212,X213,X211] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X213))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X211),s('type.num.num',X212))) | s('type.num.num',X213) = s('type.num.num',sK2('const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X211),s('type.num.num',sK2(X212))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X212)) )),
  inference(subsumption_resolution,[],[f9561,f574])).
fof(f9561,plain,(
  ( ! [X212,X213,X211] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X213))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X211),s('type.num.num',X212))) | s('type.num.num',X213) = s('type.num.num',sK2('const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X211),s('type.num.num',sK2(X212))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X211),s('type.num.num',sK2(X212)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X212)) )),
  inference(superposition,[],[f387,f762])).
fof(f9560,plain,(
  ( ! [X210,X208,X209] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X210))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X208),s('type.num.num',X209))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X210),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X208),s('type.num.num',sK2(X209)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X209)) )),
  inference(superposition,[],[f386,f762])).
fof(f9559,plain,(
  ( ! [X206,X207,X205] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X207))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X205),s('type.num.num',X206))))) | s('type.num.num',X207) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X205),s('type.num.num',sK2(X206)))))),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X206)) )),
  inference(superposition,[],[f386,f762])).
fof(f9558,plain,(
  ( ! [X204,X202,X203] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X204))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X202),s('type.num.num',X203))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X204))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X202),s('type.num.num',sK2(X203)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X203)) )),
  inference(superposition,[],[f385,f762])).
fof(f9557,plain,(
  ( ! [X200,X199,X201] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X201))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X199),s('type.num.num',X200))))) | s('type.num.num',X201) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X199),s('type.num.num',sK2(X200)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X200)) )),
  inference(superposition,[],[f385,f762])).
fof(f9538,plain,(
  ( ! [X144,X142,X143] : (s('type.num.num',X144) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X142),s('type.num.num',X143))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X144))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X142),s('type.num.num',sK2(X143)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X143)) )),
  inference(superposition,[],[f209,f762])).
fof(f9536,plain,(
  ( ! [X138,X136,X137] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X138))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X136),s('type.num.num',X137))))) | s('type.num.num',X138) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X136),s('type.num.num',sK2(X137)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X137)) )),
  inference(superposition,[],[f208,f762])).
fof(f9513,plain,(
  ( ! [X78,X79,X77] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X77),s('type.num.num',X78))),s('type.num.num',X79))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X77),s('type.num.num',sK2(X78)))),s('type.num.num',X79))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X78)) )),
  inference(superposition,[],[f1761,f762])).
fof(f9510,plain,(
  ( ! [X70,X68,X69] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X70),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X68),s('type.num.num',X69))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X70),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X68),s('type.num.num',sK2(X69)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X69)) )),
  inference(superposition,[],[f1547,f762])).
fof(f9508,plain,(
  ( ! [X64,X65,X63] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X65),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X63),s('type.num.num',X64))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X64)) )),
  inference(superposition,[],[f547,f762])).
fof(f9507,plain,(
  ( ! [X61,X62,X60] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X60),s('type.num.num',X61))),s('type.num.num',X62))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X61)) )),
  inference(superposition,[],[f521,f762])).
fof(f9505,plain,(
  ( ! [X57,X56,X55] : (s('type.num.num',X57) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X55),s('type.num.num',X56))) | s('type.num.num',sK2(X57)) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X55),s('type.num.num',sK2(X56)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X57) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X56)) )),
  inference(superposition,[],[f387,f762])).
fof(f9504,plain,(
  ( ! [X54,X52,X53] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X54))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X52),s('type.num.num',X53))) | s('type.num.num',X54) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X52),s('type.num.num',sK2(X53)))),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X53)) )),
  inference(superposition,[],[f386,f762])).
fof(f9503,plain,(
  ( ! [X50,X51,X49] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X51))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X49),s('type.num.num',X50))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X51),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X49),s('type.num.num',sK2(X50)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X50)) )),
  inference(superposition,[],[f386,f762])).
fof(f9502,plain,(
  ( ! [X47,X48,X46] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X48))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X46),s('type.num.num',X47))) | s('type.num.num',X48) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X46),s('type.num.num',sK2(X47)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X47)) )),
  inference(superposition,[],[f385,f762])).
fof(f9501,plain,(
  ( ! [X45,X43,X44] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X45))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X43),s('type.num.num',X44))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X45))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X43),s('type.num.num',sK2(X44)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X44)) )),
  inference(superposition,[],[f385,f762])).
fof(f9493,plain,(
  ( ! [X21,X19,X20] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X21))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',X20))) | s('type.num.num',X21) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',sK2(X20)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X20)) )),
  inference(superposition,[],[f208,f762])).
fof(f9492,plain,(
  ( ! [X17,X18,X16] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X16),s('type.num.num',X17))),s('type.num.num',X18))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X16),s('type.num.num',sK2(X17)))),s('type.num.num',X18))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X17)) )),
  inference(superposition,[],[f201,f762])).
fof(f9491,plain,(
  ( ! [X14,X15,X13] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X13),s('type.num.num',X14))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X13),s('type.num.num',sK2(X14)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X14)) )),
  inference(superposition,[],[f200,f762])).
fof(f9490,plain,(
  ( ! [X12,X11] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',X12))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',sK2(X12)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12)) )),
  inference(superposition,[],[f195,f762])).
fof(f9486,plain,(
  ( ! [X4,X5,X3] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X4))),s('type.num.num',X5))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X5) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X4)) )),
  inference(superposition,[],[f1441,f762])).
fof(f9485,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X2) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) )),
  inference(superposition,[],[f1445,f762])).
fof(f9682,plain,(
  ( ! [X7] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X7)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X7)) )),
  inference(forward_demodulation,[],[f9484,f1587])).
fof(f9484,plain,(
  ( ! [X7] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X7)))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X7))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X7)) )),
  inference(superposition,[],[f762,f1587])).
fof(f9681,plain,(
  ( ! [X6,X5] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num',X6))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num',sK2(X6)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X6)) )),
  inference(forward_demodulation,[],[f9483,f201])).
fof(f9483,plain,(
  ( ! [X6,X5] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.num.num',X6))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num',sK2(X6)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X6)) )),
  inference(superposition,[],[f762,f201])).
fof(f9481,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.num.num',sK2(X3)))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',sK2(X3)))),s('type.num.num',X3))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X3)) )),
  inference(superposition,[],[f762,f203])).
fof(f8975,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK2(X0)),s('type.num.num',X1))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(superposition,[],[f1761,f153])).
fof(f8957,plain,(
  ( ! [X17,X16] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X16))),s('type.num.num',X17))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X16)))),s('type.num.num',X17))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X16)) )),
  inference(superposition,[],[f1761,f153])).
fof(f8532,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X0))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',sK2(X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(superposition,[],[f1547,f153])).
fof(f8514,plain,(
  ( ! [X17,X16] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X17),s('type.num.num','const.num.SUC_1'(s('type.num.num',X16))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X17),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X16)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X16)) )),
  inference(superposition,[],[f1547,f153])).
fof(f6018,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X3))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f543,f202])).
fof(f5858,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X3)))))) )),
  inference(superposition,[],[f542,f202])).
fof(f5766,plain,(
  ( ! [X8,X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num',X9))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X8),s('type.num.num',X9))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f482,f203])).
fof(f5763,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X4) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f482,f148])).
fof(f5481,plain,(
  ( ! [X8,X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num',X9))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X8),s('type.num.num',X9)))))) )),
  inference(superposition,[],[f481,f203])).
fof(f5091,plain,(
  ( ! [X45,X46] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X45))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X46))) | s('type.num.num',sK2('const.num.SUC_1'(s('type.num.num',X45)))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X46),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f386,f962])).
fof(f5085,plain,(
  ( ! [X35,X33,X34] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X35))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X33))),s('type.num.num',X34))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X35),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X33),s('type.num.num',X34))),s('type.num.num',X34)))) )),
  inference(superposition,[],[f386,f203])).
fof(f5082,plain,(
  ( ! [X26,X27] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X26))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X26),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X27),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f386,f148])).
fof(f5081,plain,(
  ( ! [X24,X23,X25] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X25))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X23),s('type.num.num','const.num.SUC_1'(s('type.num.num',X24))))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X25),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X23),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X23),s('type.num.num',X24)))))) )),
  inference(superposition,[],[f386,f202])).
fof(f5077,plain,(
  ( ! [X17,X16] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X16))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X16)))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X17),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X16)) )),
  inference(superposition,[],[f386,f153])).
fof(f5065,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | s('type.num.num',X1) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X0)),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(superposition,[],[f386,f153])).
fof(f5047,plain,(
  ( ! [X17,X16] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X16))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))) | s('type.num.num',X17) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X16)))),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X16)) )),
  inference(superposition,[],[f386,f153])).
fof(f4836,plain,(
  ( ! [X45,X46] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X45))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X46))) | s('type.num.num',sK2('const.num.SUC_1'(s('type.num.num',X45)))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X46)))) )),
  inference(superposition,[],[f385,f962])).
fof(f4831,plain,(
  ( ! [X37,X36] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X37))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X36))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X37)))) )),
  inference(superposition,[],[f385,f146])).
fof(f4830,plain,(
  ( ! [X35,X33,X34] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X35))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X33))),s('type.num.num',X34))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X35))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X33),s('type.num.num',X34))),s('type.num.num',X34)))) )),
  inference(superposition,[],[f385,f203])).
fof(f4827,plain,(
  ( ! [X26,X27] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X26))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X26),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X27)))) )),
  inference(superposition,[],[f385,f148])).
fof(f4826,plain,(
  ( ! [X24,X23,X25] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X25))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X23),s('type.num.num','const.num.SUC_1'(s('type.num.num',X24))))))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X25))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X23),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X23),s('type.num.num',X24)))))) )),
  inference(superposition,[],[f385,f202])).
fof(f4822,plain,(
  ( ! [X17,X16] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X16))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X16)))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X17))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X16)) )),
  inference(superposition,[],[f385,f153])).
fof(f4810,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | s('type.num.num',X1) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK2(X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(superposition,[],[f385,f153])).
fof(f4792,plain,(
  ( ! [X17,X16] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X16))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))) | s('type.num.num',X17) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X16)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X16)) )),
  inference(superposition,[],[f385,f153])).
fof(f3561,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num',X1) | s('type.num.num',sK2(X0)) = s('type.num.num',sK2(X1)) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(superposition,[],[f387,f153])).
fof(f3560,plain,(
  ( ! [X37,X38] : (s('type.num.num',X38) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X37))) | s('type.num.num',sK2(X38)) = s('type.num.num',sK2('const.num.SUC_1'(s('type.num.num',X37)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X38)) )),
  inference(superposition,[],[f387,f962])).
fof(f3555,plain,(
  ( ! [X28,X29] : (s('type.num.num',X29) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))) | s('type.num.num',sK2(X29)) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X28))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X29)) )),
  inference(superposition,[],[f387,f146])).
fof(f3554,plain,(
  ( ! [X26,X27,X25] : (s('type.num.num',X27) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X25))),s('type.num.num',X26))))) | s('type.num.num',sK2(X27)) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X25),s('type.num.num',X26))),s('type.num.num',X26))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X27)) )),
  inference(superposition,[],[f387,f203])).
fof(f3553,plain,(
  ( ! [X24,X23,X22] : (s('type.num.num',X24) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X22),s('type.num.num','const.num.SUC_1'(s('type.num.num',X23))))))) | s('type.num.num',sK2(X24)) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X22),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X22),s('type.num.num',X23))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X24)) )),
  inference(superposition,[],[f387,f202])).
fof(f3550,plain,(
  ( ! [X15,X16] : (s('type.num.num',X16) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X15))) | s('type.num.num',sK2(X16)) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X16)) )),
  inference(superposition,[],[f387,f148])).
fof(f3548,plain,(
  ( ! [X12,X13] : (s('type.num.num',X13) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X12))) | s('type.num.num',sK2(X13)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X12)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12)) )),
  inference(superposition,[],[f387,f153])).
fof(f3621,plain,(
  ( ! [X37,X38] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X37) | s('type.num.num',X37) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X38))) | s('type.num.num',X38) = s('type.num.num',sK2(sK2('const.num.SUC_1'(s('type.num.num',X37)))))) )),
  inference(forward_demodulation,[],[f3539,f962])).
fof(f3539,plain,(
  ( ! [X37,X38] : (s('type.num.num',X37) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X38))) | s('type.num.num',X38) = s('type.num.num',sK2(sK2('const.num.SUC_1'(s('type.num.num',X37))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK2('const.num.SUC_1'(s('type.num.num',X37))))) )),
  inference(superposition,[],[f387,f962])).
fof(f3619,plain,(
  ( ! [X30,X31,X32] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X32))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X30),s('type.num.num',X31))))) | s('type.num.num',X32) = s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X30))),s('type.num.num',X31))))) )),
  inference(subsumption_resolution,[],[f3535,f521])).
fof(f3535,plain,(
  ( ! [X30,X31,X32] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X32))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X30),s('type.num.num',X31))))) | s('type.num.num',X32) = s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X30))),s('type.num.num',X31)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X30))),s('type.num.num',X31)))) )),
  inference(superposition,[],[f387,f201])).
fof(f3618,plain,(
  ( ! [X28,X29] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X28) | s('type.num.num',X28) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X29))) | s('type.num.num',X29) = s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X28))))) )),
  inference(forward_demodulation,[],[f3534,f146])).
fof(f3534,plain,(
  ( ! [X28,X29] : (s('type.num.num',X28) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X29))) | s('type.num.num',X29) = s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X28)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X28)))) )),
  inference(superposition,[],[f387,f146])).
fof(f3617,plain,(
  ( ! [X26,X27,X25] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X25))),s('type.num.num',X26))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X25))),s('type.num.num',X26))) | s('type.num.num',X27) = s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X25),s('type.num.num',X26))),s('type.num.num',X26))))) )),
  inference(forward_demodulation,[],[f3533,f203])).
fof(f3533,plain,(
  ( ! [X26,X27,X25] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X25))),s('type.num.num',X26))) | s('type.num.num',X27) = s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X25),s('type.num.num',X26))),s('type.num.num',X26)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X25),s('type.num.num',X26))),s('type.num.num',X26)))) )),
  inference(superposition,[],[f387,f203])).
fof(f3616,plain,(
  ( ! [X24,X23,X22] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X22),s('type.num.num','const.num.SUC_1'(s('type.num.num',X23))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X24))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X22),s('type.num.num','const.num.SUC_1'(s('type.num.num',X23))))) | s('type.num.num',X24) = s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num',X22),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X22),s('type.num.num',X23))))))) )),
  inference(forward_demodulation,[],[f3532,f202])).
fof(f3532,plain,(
  ( ! [X24,X23,X22] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X24))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X22),s('type.num.num','const.num.SUC_1'(s('type.num.num',X23))))) | s('type.num.num',X24) = s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num',X22),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X22),s('type.num.num',X23)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X22),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X22),s('type.num.num',X23)))))) )),
  inference(superposition,[],[f387,f202])).
fof(f3614,plain,(
  ( ! [X19,X17,X18] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X17),s('type.num.num',X18))))) | s('type.num.num',X19) = s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num',X17),s('type.num.num','const.num.SUC_1'(s('type.num.num',X18))))))) )),
  inference(subsumption_resolution,[],[f3530,f547])).
fof(f3530,plain,(
  ( ! [X19,X17,X18] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X17),s('type.num.num',X18))))) | s('type.num.num',X19) = s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num',X17),s('type.num.num','const.num.SUC_1'(s('type.num.num',X18)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X17),s('type.num.num','const.num.SUC_1'(s('type.num.num',X18)))))) )),
  inference(superposition,[],[f387,f200])).
fof(f3613,plain,(
  ( ! [X15,X16] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X15) | s('type.num.num',X15) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X16))) | s('type.num.num',X16) = s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num','const.num.|48|_0'))))) )),
  inference(forward_demodulation,[],[f3529,f148])).
fof(f3529,plain,(
  ( ! [X15,X16] : (s('type.num.num',X15) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X16))) | s('type.num.num',X16) = s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f387,f148])).
fof(f3611,plain,(
  ( ! [X12,X13] : (s('type.num.num',X12) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X13))) | s('type.num.num',X13) = s('type.num.num',sK2('const.num.SUC_1'(s('type.num.num',sK2(X12))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12)) )),
  inference(subsumption_resolution,[],[f3527,f574])).
fof(f3527,plain,(
  ( ! [X12,X13] : (s('type.num.num',X12) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X13))) | s('type.num.num',X13) = s('type.num.num',sK2('const.num.SUC_1'(s('type.num.num',sK2(X12))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X12)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12)) )),
  inference(superposition,[],[f387,f153])).
fof(f3500,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))) )),
  inference(unit_resulting_resolution,[],[f547,f200,f387])).
fof(f3506,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1))))) )),
  inference(unit_resulting_resolution,[],[f521,f201,f387])).
fof(f3566,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))) )),
  inference(forward_demodulation,[],[f3508,f200])).
fof(f3508,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))) = s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))) )),
  inference(unit_resulting_resolution,[],[f547,f201,f387])).
fof(f3514,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))) )),
  inference(unit_resulting_resolution,[],[f547,f146,f387])).
fof(f3516,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0'))))) )),
  inference(unit_resulting_resolution,[],[f521,f148,f387])).
fof(f3351,plain,(
  ( ! [X26,X27,X25] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X25))),s('type.num.num',X26))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X25),s('type.num.num',X26))),s('type.num.num',X26)))) )),
  inference(superposition,[],[f1760,f203])).
fof(f3350,plain,(
  ( ! [X24,X23,X22] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X24))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X22),s('type.num.num','const.num.SUC_1'(s('type.num.num',X23))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X22),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X22),s('type.num.num',X23)))))) )),
  inference(superposition,[],[f1760,f202])).
fof(f3101,plain,(
  ( ! [X26,X27,X25] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X25))),s('type.num.num',X26))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X25),s('type.num.num',X26))),s('type.num.num',X26)))) )),
  inference(superposition,[],[f1548,f203])).
fof(f3100,plain,(
  ( ! [X24,X23,X22] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X22),s('type.num.num','const.num.SUC_1'(s('type.num.num',X23))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X24))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X22),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X22),s('type.num.num',X23)))))) )),
  inference(superposition,[],[f1548,f202])).
fof(f2858,plain,(
  ( ! [X37,X35,X38,X36] : (s(X37,f(s(bool,'T_0'))) != s(X37,g(s(bool,'T_0'))) | s(fun(bool,X37),X35) = s(fun(bool,X37),X36) | s(bool,X38) = s(bool,sK3(X35,X36,X37,bool)) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f216,f293])).
fof(f3002,plain,(
  ( ! [X28,X26,X27] : (s(bool,X26) != s(bool,X27) | s(bool,'T_0') = s(bool,X27) | s(bool,X26) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f2856,f138])).
fof(f2856,plain,(
  ( ! [X28,X26,X27] : (~p(s(bool,'T_0')) | s(bool,X26) != s(bool,X27) | s(bool,'T_0') = s(bool,X27) | s(bool,X26) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f295,f293])).
fof(f3001,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,X24)) | s(bool,X23) = s(bool,X24) | s(bool,X23) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(subsumption_resolution,[],[f2855,f138])).
fof(f2855,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'T_0')) | ~p(s(bool,X24)) | s(bool,X23) = s(bool,X24) | s(bool,X23) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f207,f293])).
fof(f2815,plain,(
  ( ! [X39,X37,X38,X40] : (s(X39,f(s(bool,X40))) != s(X39,g(s(bool,X40))) | s(fun(bool,X39),X37) = s(fun(bool,X39),X38) | s(bool,'T_0') = s(bool,X40) | s(bool,'T_0') = s(bool,sK3(X37,X38,X39,bool))) )),
  inference(superposition,[],[f216,f293])).
fof(f2780,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f293,f293])).
fof(f1787,plain,(
  ( ! [X70,X68,X69] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X70),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X68))),s('type.num.num',X69))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X70) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X68),s('type.num.num',X69))),s('type.num.num',X69)))) )),
  inference(superposition,[],[f213,f203])).
fof(f1786,plain,(
  ( ! [X66,X67,X65] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X65))),s('type.num.num',X66))),s('type.num.num',X67))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X65),s('type.num.num',X66))),s('type.num.num',X66))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X67)) )),
  inference(superposition,[],[f213,f203])).
fof(f1784,plain,(
  ( ! [X61,X59,X60] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X61),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X59))),s('type.num.num',X60))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X59),s('type.num.num',X60))),s('type.num.num',X60)))) )),
  inference(superposition,[],[f211,f203])).
fof(f1781,plain,(
  ( ! [X52,X50,X51] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X50))),s('type.num.num',X51))),s('type.num.num',X52))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X50),s('type.num.num',X51))),s('type.num.num',X51)))) )),
  inference(superposition,[],[f210,f203])).
fof(f1777,plain,(
  ( ! [X39,X38,X40] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X40))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X38))),s('type.num.num',X39))))) | s('type.num.num',X40) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X38),s('type.num.num',X39))),s('type.num.num',X39)))) )),
  inference(superposition,[],[f208,f203])).
fof(f1761,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.num.num',X3))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num',X3)))) )),
  inference(superposition,[],[f210,f203])).
fof(f1760,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) )),
  inference(superposition,[],[f211,f203])).
fof(f1759,plain,(
  ( ! [X4,X3] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X3),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))))),s('type.num.num',X4)))))) )),
  inference(superposition,[],[f203,f200])).
fof(f1571,plain,(
  ( ! [X64,X62,X63] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X64),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X62),s('type.num.num','const.num.SUC_1'(s('type.num.num',X63))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X64) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X62),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X62),s('type.num.num',X63)))))) )),
  inference(superposition,[],[f213,f202])).
fof(f1570,plain,(
  ( ! [X61,X59,X60] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X59),s('type.num.num','const.num.SUC_1'(s('type.num.num',X60))))),s('type.num.num',X61))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X59),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X59),s('type.num.num',X60))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X61)) )),
  inference(superposition,[],[f213,f202])).
fof(f1568,plain,(
  ( ! [X54,X55,X53] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X55),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X53),s('type.num.num','const.num.SUC_1'(s('type.num.num',X54))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X53),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X53),s('type.num.num',X54)))))) )),
  inference(superposition,[],[f211,f202])).
fof(f1565,plain,(
  ( ! [X45,X46,X44] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X44),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45))))),s('type.num.num',X46))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X44),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X44),s('type.num.num',X45)))))) )),
  inference(superposition,[],[f210,f202])).
fof(f1561,plain,(
  ( ! [X33,X34,X32] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X34))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X32),s('type.num.num','const.num.SUC_1'(s('type.num.num',X33))))))) | s('type.num.num',X34) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X32),s('type.num.num',X33)))))) )),
  inference(superposition,[],[f208,f202])).
fof(f1548,plain,(
  ( ! [X4,X3] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X3),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X3)) )),
  inference(superposition,[],[f210,f202])).
fof(f1547,plain,(
  ( ! [X2,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X2)))) )),
  inference(superposition,[],[f211,f202])).
fof(f1545,plain,(
  ( ! [X4,X3] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s('type.num.num',X4)))))))) )),
  inference(superposition,[],[f202,f201])).
fof(f1482,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(forward_demodulation,[],[f1473,f227])).
fof(f1473,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(unit_resulting_resolution,[],[f1347,f363])).
fof(f1481,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),
  inference(forward_demodulation,[],[f1474,f227])).
fof(f1474,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),
  inference(unit_resulting_resolution,[],[f1347,f153])).
fof(f1445,plain,(
  ( ! [X17,X18,X16] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X18),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X16),s('type.num.num',X17))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X18)) )),
  inference(subsumption_resolution,[],[f1374,f547])).
fof(f1374,plain,(
  ( ! [X17,X18,X16] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X18),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X16),s('type.num.num',X17))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X16),s('type.num.num','const.num.SUC_1'(s('type.num.num',X17)))))) )),
  inference(superposition,[],[f213,f200])).
fof(f1373,plain,(
  ( ! [X14,X15] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X15),s('type.num.num',X14))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X15) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X14),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f213,f148])).
fof(f1372,plain,(
  ( ! [X12,X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X13),s('type.num.num',X12))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X12)))) )),
  inference(superposition,[],[f213,f146])).
fof(f1441,plain,(
  ( ! [X17,X18,X16] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X16),s('type.num.num',X17))))),s('type.num.num',X18))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X18)) )),
  inference(subsumption_resolution,[],[f1361,f547])).
fof(f1361,plain,(
  ( ! [X17,X18,X16] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X16),s('type.num.num',X17))))),s('type.num.num',X18))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X16),s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X18)) )),
  inference(superposition,[],[f213,f200])).
fof(f1360,plain,(
  ( ! [X14,X15] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X14),s('type.num.num',X15))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X14),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X15)) )),
  inference(superposition,[],[f213,f148])).
fof(f1359,plain,(
  ( ! [X12,X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X12),s('type.num.num',X13))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X12))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13)) )),
  inference(superposition,[],[f213,f146])).
fof(f1340,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f574,f574,f213])).
fof(f1341,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(unit_resulting_resolution,[],[f195,f574,f213])).
fof(f1346,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f574,f195,f213])).
fof(f1347,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f195,f195,f213])).
fof(f1297,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) )),
  inference(forward_demodulation,[],[f1178,f147])).
fof(f1178,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))),s('type.num.num',X1)))) )),
  inference(trivial_inequality_removal,[],[f1162])).
fof(f1162,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))),s('type.num.num',X1)))) )),
  inference(superposition,[],[f212,f147])).
fof(f980,plain,(
  ( ! [X6,X7] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))) | s('type.num.num',X7) = s('type.num.num',sK2('const.num.SUC_1'(s('type.num.num',X6))))) )),
  inference(superposition,[],[f208,f619])).
fof(f962,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num',sK2('const.num.SUC_1'(s('type.num.num',X0))))) )),
  inference(unit_resulting_resolution,[],[f619,f208])).
fof(f903,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(X0)),s('type.num.num',X1))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(superposition,[],[f201,f153])).
fof(f807,plain,(
  ( ! [X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | s('type.num.num',X1) = s('type.num.num',sK2('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) )),
  inference(superposition,[],[f208,f362])).
fof(f803,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK2('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(unit_resulting_resolution,[],[f362,f208])).
fof(f762,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',sK2(X0)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(superposition,[],[f200,f153])).
fof(f737,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))) != s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))) )),
  inference(forward_demodulation,[],[f710,f200])).
fof(f710,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))) != s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))) )),
  inference(unit_resulting_resolution,[],[f547,f363])).
fof(f736,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))))) )),
  inference(forward_demodulation,[],[f711,f200])).
fof(f711,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))))) )),
  inference(unit_resulting_resolution,[],[f547,f153])).
fof(f705,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))) != s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1))))) )),
  inference(forward_demodulation,[],[f678,f201])).
fof(f678,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1))) != s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1))))) )),
  inference(unit_resulting_resolution,[],[f521,f363])).
fof(f704,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1))))))) )),
  inference(forward_demodulation,[],[f679,f201])).
fof(f679,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1))))))) )),
  inference(unit_resulting_resolution,[],[f521,f153])).
fof(f618,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) != s('type.num.num',sK2('const.num.SUC_1'(s('type.num.num',X0))))) )),
  inference(unit_resulting_resolution,[],[f574,f363])).
fof(f619,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.num.SUC_1'(s('type.num.num',X0))))))) )),
  inference(unit_resulting_resolution,[],[f574,f153])).
fof(f574,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f514,f209])).
fof(f569,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) != s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f568,f148])).
fof(f568,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) != s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f550,f200])).
fof(f550,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(unit_resulting_resolution,[],[f523,f363])).
fof(f567,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) )),
  inference(forward_demodulation,[],[f566,f148])).
fof(f566,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) )),
  inference(forward_demodulation,[],[f551,f200])).
fof(f551,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) )),
  inference(unit_resulting_resolution,[],[f523,f153])).
fof(f543,plain,(
  ( ! [X14,X15] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num',X14))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X14),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f211,f148])).
fof(f542,plain,(
  ( ! [X12,X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X13),s('type.num.num',X12))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X12)))) )),
  inference(superposition,[],[f211,f146])).
fof(f547,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) )),
  inference(forward_demodulation,[],[f546,f146])).
fof(f546,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))))))) )),
  inference(forward_demodulation,[],[f522,f201])).
fof(f522,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f471,f211])).
fof(f521,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) )),
  inference(forward_demodulation,[],[f520,f146])).
fof(f520,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))),s('type.num.num',X1)))) )),
  inference(forward_demodulation,[],[f498,f201])).
fof(f498,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0))),s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f471,f210])).
fof(f519,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) != s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0))))) )),
  inference(forward_demodulation,[],[f518,f146])).
fof(f518,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) != s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0))))) )),
  inference(forward_demodulation,[],[f499,f201])).
fof(f499,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0))) != s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0))))) )),
  inference(unit_resulting_resolution,[],[f471,f363])).
fof(f517,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0))))))) )),
  inference(forward_demodulation,[],[f516,f146])).
fof(f516,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0))))))) )),
  inference(forward_demodulation,[],[f500,f201])).
fof(f500,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0))))))) )),
  inference(unit_resulting_resolution,[],[f471,f153])).
fof(f482,plain,(
  ( ! [X14,X15] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X14),s('type.num.num',X15))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X14),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f210,f148])).
fof(f481,plain,(
  ( ! [X12,X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num',X13))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X12)))) )),
  inference(superposition,[],[f210,f146])).
fof(f471,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f195,f210])).
fof(f470,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num',sK2('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(unit_resulting_resolution,[],[f195,f363])).
fof(f437,plain,(
  ( ! [X10,X11] : (s('type.num.num',X10) != s('type.num.num',X11) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X10)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) )),
  inference(superposition,[],[f209,f153])).
fof(f387,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | s('type.num.num',X1) = s('type.num.num',sK2(X0)) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(superposition,[],[f208,f153])).
fof(f386,plain,(
  ( ! [X14,X15] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X14))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X15))) | s('type.num.num',X15) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X14),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f208,f148])).
fof(f385,plain,(
  ( ! [X12,X13] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X12))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X13))) | s('type.num.num',X13) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X12)))) )),
  inference(superposition,[],[f208,f146])).
fof(f384,plain,(
  ( ! [X10,X11] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))) | s('type.num.num',X11) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(X10)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X10)) )),
  inference(superposition,[],[f208,f153])).
fof(f363,plain,(
  ( ! [X0] : (s('type.num.num',X0) != s('type.num.num',sK2(X0)) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(superposition,[],[f195,f153])).
fof(f362,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(unit_resulting_resolution,[],[f195,f153])).
fof(f322,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f319,f226])).
fof(f319,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f204,f152])).
fof(f295,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f159,f152])).
fof(f293,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f152,f152])).
fof(f260,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f226,f155])).
fof(f216,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK3(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK3(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK3(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK3(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f95])).
fof(f95,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f89])).
fof(f89,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.arithmetic.MULT_EQ_|49|',aHL_EXT)).
fof(f213,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f133])).
fof(f215,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f133])).
fof(f210,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f131])).
fof(f212,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f131])).
fof(f208,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1] : ((s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.num.num',X0) != s('type.num.num',X1) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X8,X9] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))) <=> s('type.num.num',X8) = s('type.num.num',X9))),
  file('thm.arithmetic.MULT_EQ_|49|','thm.prim_rec.INV_SUC_EQ')).
fof(f209,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num',X1) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f129])).
fof(f204,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f93])).
fof(f93,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X10,X11] : ((p(s(bool,X10)) => p(s(bool,X11))) => ((p(s(bool,X11)) => p(s(bool,X10))) => s(bool,X10) = s(bool,X11)))),
  file('thm.arithmetic.MULT_EQ_|49|','thm.bool.IMP_ANTISYM_AX')).
fof(f207,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f94])).
fof(f203,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X1)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X8,X9] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num',X9))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X8),s('type.num.num',X9))),s('type.num.num',X9)))),
  file('thm.arithmetic.MULT_EQ_|49|','thm.arithmetic.MULT|split|1')).
fof(f201,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X9,X8] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X9))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num',X9)))),
  file('thm.arithmetic.MULT_EQ_|49|','thm.arithmetic.ADD_CLAUSES|split|2')).
fof(f200,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0,X1] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X9,X8] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X9))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9)))))),
  file('thm.arithmetic.MULT_EQ_|49|','thm.arithmetic.ADD_CLAUSES|split|3')).
fof(f155,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.arithmetic.MULT_EQ_|49|','thm.bool.EQ_CLAUSES|split|0')).
fof(f156,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f102])).
fof(f152,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.arithmetic.MULT_EQ_|49|',aHL_BOOL_CASES)).
fof(f149,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X8] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X8)))),
  file('thm.arithmetic.MULT_EQ_|49|','thm.arithmetic.MULT_CLAUSES|split|0')).
fof(f148,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X8] : s('type.num.num',X8) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.|48|_0')))),
  file('thm.arithmetic.MULT_EQ_|49|','thm.arithmetic.ADD_|48|')).
fof(f147,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X8] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X8),s('type.num.num','const.num.|48|_0')))),
  file('thm.arithmetic.MULT_EQ_|49|','thm.arithmetic.MULT_|48|')).
fof(f146,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X8] : s('type.num.num',X8) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X8)))),
  file('thm.arithmetic.MULT_EQ_|49|','thm.arithmetic.ADD_CLAUSES|split|0')).
fof(f138,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.arithmetic.MULT_EQ_|49|',aHL_TRUTH)).
fof(f13598,plain,(
  ( ! [X5] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))) | s('type.num.num',X5) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(sK2(sK0))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK2(sK0))) ) | $spl0),
  inference(superposition,[],[f384,f1835])).
fof(f69266,plain,(
  $false | ($spl0 | $spl14 | $spl97)),
  inference(subsumption_resolution,[],[f66319,f149])).
fof(f66319,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK1))) | ($spl0 | $spl14 | $spl97)),
  inference(backward_demodulation,[],[f66270,f65301])).
fof(f69265,plain,(
  $false | ($spl0 | $spl14 | $spl95)),
  inference(subsumption_resolution,[],[f66318,f147])).
fof(f66318,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK1),s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl14 | $spl95)),
  inference(backward_demodulation,[],[f66270,f65299])).
fof(f66258,plain,(
  $false | ($spl0 | $spl16 | $spl64)),
  inference(subsumption_resolution,[],[f66151,f258])).
fof(f66151,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK1))) | ($spl0 | $spl16 | $spl64)),
  inference(superposition,[],[f64596,f65911])).
fof(f65911,plain,(
  s('type.num.num',sK1) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(sK1)))) | $spl64),
  inference(unit_resulting_resolution,[],[f65764,f153])).
fof(f65764,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK1) | $spl64),
  inference(unit_resulting_resolution,[],[f147,f65240,f212])).
fof(f64596,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) ) | ($spl0 | $spl16)),
  inference(superposition,[],[f64368,f202])).
fof(f64368,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',X0)))) ) | ($spl0 | $spl16)),
  inference(subsumption_resolution,[],[f64363,f1662])).
fof(f64363,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',X0))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK0)) ) | ($spl0 | $spl16)),
  inference(superposition,[],[f28304,f903])).
fof(f28304,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(sK0)),s('type.num.num',X0)))))) ) | ($spl0 | $spl16)),
  inference(unit_resulting_resolution,[],[f574,f23494,f384])).
fof(f23494,plain,(
  ( ! [X17] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(sK0)),s('type.num.num',X17)))) ) | ($spl0 | $spl16)),
  inference(superposition,[],[f521,f22867])).
fof(f22867,plain,(
  s('type.num.num',sK2(sK0)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(sK2(sK0))))) | ($spl0 | $spl16)),
  inference(unit_resulting_resolution,[],[f1835,f11140])).
fof(f66255,plain,(
  s('type.num.num',sK2(sK1)) = s('type.num.num',sK2(sK2(sK0))) | $spl102),
  inference(cnf_transformation,[],[f66255_D])).
fof(f66255_D,plain,(
  s('type.num.num',sK2(sK1)) = s('type.num.num',sK2(sK2(sK0))) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f66257,plain,(
  s('type.num.num',sK1) != s('type.num.num',sK2(sK0)) | $spl105),
  inference(cnf_transformation,[],[f66257_D])).
fof(f66257_D,plain,(
  s('type.num.num',sK1) != s('type.num.num',sK2(sK0)) <=> ~$spl105),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl105])])).
fof(f66148,plain,(
  s('type.num.num',sK1) != s('type.num.num',sK2(sK0)) | s('type.num.num',sK2(sK1)) = s('type.num.num',sK2(sK2(sK0))) | ($spl0 | $spl16 | $spl64)),
  inference(superposition,[],[f23479,f65911])).
fof(f23479,plain,(
  ( ! [X2] : (s('type.num.num',sK2(sK0)) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))) | s('type.num.num',X2) = s('type.num.num',sK2(sK2(sK0)))) ) | ($spl0 | $spl16)),
  inference(superposition,[],[f208,f22867])).
fof(f66248,plain,(
  s('type.num.num',sK0) != s('type.num.num',sK1) | $spl101),
  inference(cnf_transformation,[],[f66248_D])).
fof(f66248_D,plain,(
  s('type.num.num',sK0) != s('type.num.num',sK1) <=> ~$spl101),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl101])])).
fof(f66246,plain,(
  s('type.num.num',sK2(sK0)) = s('type.num.num',sK2(sK1)) | $spl98),
  inference(cnf_transformation,[],[f66246_D])).
fof(f66246_D,plain,(
  s('type.num.num',sK2(sK0)) = s('type.num.num',sK2(sK1)) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f66253,plain,(
  s('type.num.num',sK2(sK0)) = s('type.num.num',sK2(sK1)) | s('type.num.num',sK0) != s('type.num.num',sK1) | ($spl0 | $spl64)),
  inference(forward_demodulation,[],[f66147,f962])).
fof(f66147,plain,(
  s('type.num.num',sK0) != s('type.num.num',sK1) | s('type.num.num',sK2(sK1)) = s('type.num.num',sK2('const.num.SUC_1'(s('type.num.num',sK2(sK0))))) | ($spl0 | $spl64)),
  inference(superposition,[],[f3612,f65911])).
fof(f3612,plain,(
  ( ! [X14] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X14))) | s('type.num.num',X14) = s('type.num.num',sK2('const.num.SUC_1'(s('type.num.num',sK2(sK0)))))) ) | $spl0),
  inference(subsumption_resolution,[],[f3528,f574])).
fof(f3528,plain,(
  ( ! [X14] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X14))) | s('type.num.num',X14) = s('type.num.num',sK2('const.num.SUC_1'(s('type.num.num',sK2(sK0))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(sK0))))) ) | $spl0),
  inference(superposition,[],[f387,f1835])).
fof(f66140,plain,(
  s('type.num.num',sK0) != s('type.num.num',sK1) | s('type.num.num',sK2(sK0)) = s('type.num.num',sK2(sK1)) | ($spl0 | $spl64)),
  inference(superposition,[],[f1861,f65911])).
fof(f1861,plain,(
  ( ! [X2] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))) | s('type.num.num',X2) = s('type.num.num',sK2(sK0))) ) | $spl0),
  inference(superposition,[],[f208,f1835])).
fof(f65552,plain,(
  $false | ($spl0 | $spl62)),
  inference(subsumption_resolution,[],[f65551,f195])).
fof(f65551,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl62)),
  inference(forward_demodulation,[],[f65550,f147])).
fof(f65550,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl62)),
  inference(forward_demodulation,[],[f258,f65428])).
fof(f65428,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK1) | ($spl0 | $spl62)),
  inference(forward_demodulation,[],[f64681,f65239])).
fof(f64681,plain,(
  s('type.num.num',sK1) = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1)))) | $spl0),
  inference(equality_resolution,[],[f3605])).
fof(f3605,plain,(
  ( ! [X8] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))) | s('type.num.num',X8) = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1))))) ) | $spl0),
  inference(subsumption_resolution,[],[f3604,f574])).
fof(f3604,plain,(
  ( ! [X8] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))) | s('type.num.num',X8) = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1))))) ) | $spl0),
  inference(forward_demodulation,[],[f3523,f1804])).
fof(f1804,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1))) | $spl0),
  inference(forward_demodulation,[],[f1742,f1587])).
fof(f1742,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK1))) | $spl0),
  inference(superposition,[],[f203,f258])).
fof(f3523,plain,(
  ( ! [X8] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))) | s('type.num.num',X8) = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1)))) ) | $spl0),
  inference(superposition,[],[f387,f1804])).
fof(f65543,plain,(
  $false | ($spl0 | $spl62)),
  inference(subsumption_resolution,[],[f65542,f195])).
fof(f65542,plain,(
  s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) | ($spl0 | $spl62)),
  inference(forward_demodulation,[],[f65541,f227])).
fof(f65541,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl62)),
  inference(forward_demodulation,[],[f1642,f65428])).
fof(f1642,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))) | $spl0),
  inference(forward_demodulation,[],[f1597,f148])).
fof(f1597,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num','const.num.|48|_0'))))) | $spl0),
  inference(superposition,[],[f1543,f200])).
fof(f65540,plain,(
  $false | ($spl0 | $spl62)),
  inference(subsumption_resolution,[],[f65539,f195])).
fof(f65539,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl62)),
  inference(forward_demodulation,[],[f65538,f147])).
fof(f65538,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl62)),
  inference(forward_demodulation,[],[f1804,f65428])).
fof(f65537,plain,(
  $false | ($spl0 | $spl62)),
  inference(subsumption_resolution,[],[f65536,f195])).
fof(f65536,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) | ($spl0 | $spl62)),
  inference(forward_demodulation,[],[f65535,f227])).
fof(f65535,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl62)),
  inference(forward_demodulation,[],[f65534,f1802])).
fof(f65534,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl62)),
  inference(forward_demodulation,[],[f2133,f65428])).
fof(f2133,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))) | $spl0),
  inference(superposition,[],[f202,f1804])).
fof(f65531,plain,(
  $false | ($spl0 | $spl62)),
  inference(subsumption_resolution,[],[f65530,f574])).
fof(f65530,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl62)),
  inference(forward_demodulation,[],[f65529,f147])).
fof(f65529,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl62)),
  inference(forward_demodulation,[],[f2174,f65428])).
fof(f2174,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK1))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num',sK1))) | $spl0),
  inference(forward_demodulation,[],[f2132,f201])).
fof(f2132,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK1))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num',sK1))) | $spl0),
  inference(superposition,[],[f203,f1804])).
fof(f65525,plain,(
  $false | ($spl0 | $spl62)),
  inference(subsumption_resolution,[],[f65524,f195])).
fof(f65524,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK0))))) | ($spl0 | $spl62)),
  inference(forward_demodulation,[],[f65523,f1540])).
fof(f65523,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK0))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl0 | $spl62)),
  inference(forward_demodulation,[],[f2754,f65428])).
fof(f2754,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK0))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))))) | $spl0),
  inference(forward_demodulation,[],[f2707,f200])).
fof(f2707,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))))) | $spl0),
  inference(superposition,[],[f202,f1642])).
fof(f65522,plain,(
  $false | ($spl0 | $spl62)),
  inference(subsumption_resolution,[],[f65520,f195])).
fof(f65520,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',X0)))))) ) | ($spl0 | $spl62)),
  inference(backward_demodulation,[],[f65518,f10302])).
fof(f10302,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),s('type.num.num',X0)))))) ) | $spl0),
  inference(forward_demodulation,[],[f10042,f227])).
fof(f10042,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),s('type.num.num',X0)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1830,f903])).
fof(f1830,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl0),
  inference(unit_resulting_resolution,[],[f195,f1662,f213])).
fof(f65518,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl62)),
  inference(forward_demodulation,[],[f3512,f65428])).
fof(f3512,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1)))))) | $spl0),
  inference(unit_resulting_resolution,[],[f1598,f1642,f387])).
fof(f65521,plain,(
  $false | ($spl0 | $spl62)),
  inference(subsumption_resolution,[],[f65519,f195])).
fof(f65519,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',sK0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',sK0)))))) ) | ($spl0 | $spl62)),
  inference(backward_demodulation,[],[f65518,f9665])).
fof(f9665,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',sK0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))))) ) | $spl0),
  inference(forward_demodulation,[],[f9414,f227])).
fof(f9414,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1830,f762])).
fof(f65508,plain,(
  $false | ($spl0 | $spl62)),
  inference(subsumption_resolution,[],[f65507,f227])).
fof(f65507,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl62)),
  inference(forward_demodulation,[],[f3687,f65428])).
fof(f3687,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))) | $spl0),
  inference(subsumption_resolution,[],[f3632,f1598])).
fof(f3632,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))) | $spl0),
  inference(superposition,[],[f363,f3512])).
fof(f65502,plain,(
  $false | ($spl0 | $spl62)),
  inference(subsumption_resolution,[],[f65501,f195])).
fof(f65501,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) | ($spl0 | $spl62)),
  inference(forward_demodulation,[],[f65500,f227])).
fof(f65500,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl62)),
  inference(forward_demodulation,[],[f65499,f148])).
fof(f65499,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num','const.num.|48|_0'))))))) | ($spl0 | $spl62)),
  inference(forward_demodulation,[],[f41661,f65428])).
fof(f41661,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK1))))))) | $spl0),
  inference(forward_demodulation,[],[f41403,f201])).
fof(f41403,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1))))) | $spl0),
  inference(superposition,[],[f1759,f1642])).
fof(f65494,plain,(
  $false | ($spl0 | $spl62)),
  inference(subsumption_resolution,[],[f65493,f195])).
fof(f65493,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | ($spl0 | $spl62)),
  inference(forward_demodulation,[],[f65492,f227])).
fof(f65492,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl62)),
  inference(forward_demodulation,[],[f65491,f1802])).
fof(f65491,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl62)),
  inference(forward_demodulation,[],[f65490,f146])).
fof(f65490,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))))) | ($spl0 | $spl62)),
  inference(forward_demodulation,[],[f64374,f65428])).
fof(f64374,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK1))))))) | $spl0),
  inference(superposition,[],[f202,f2174])).
fof(f65489,plain,(
  $false | ($spl0 | $spl62)),
  inference(subsumption_resolution,[],[f65488,f574])).
fof(f65488,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl62)),
  inference(forward_demodulation,[],[f65487,f147])).
fof(f65487,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl62)),
  inference(forward_demodulation,[],[f65486,f148])).
fof(f65486,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl62)),
  inference(forward_demodulation,[],[f64375,f65428])).
fof(f64375,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.num.num',sK1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',sK1))))),s('type.num.num',sK1))) | $spl0),
  inference(superposition,[],[f203,f2174])).
fof(f65476,plain,(
  $false | ($spl0 | $spl61 | $spl62)),
  inference(subsumption_resolution,[],[f65454,f147])).
fof(f65454,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl61 | $spl62)),
  inference(backward_demodulation,[],[f65428,f65235])).
fof(f65427,plain,(
  $false | ($spl61 | $spl74)),
  inference(subsumption_resolution,[],[f65426,f147])).
fof(f65426,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num','const.num.|48|_0'))) | ($spl61 | $spl74)),
  inference(forward_demodulation,[],[f65235,f65256])).
fof(f65398,plain,(
  $false | ($spl0 | $spl74)),
  inference(subsumption_resolution,[],[f65397,f574])).
fof(f65397,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl74)),
  inference(forward_demodulation,[],[f65396,f147])).
fof(f65396,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl74)),
  inference(forward_demodulation,[],[f65336,f148])).
fof(f65336,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl74)),
  inference(backward_demodulation,[],[f65256,f64375])).
fof(f65395,plain,(
  $false | ($spl0 | $spl74)),
  inference(subsumption_resolution,[],[f65394,f195])).
fof(f65394,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | ($spl0 | $spl74)),
  inference(forward_demodulation,[],[f65393,f227])).
fof(f65393,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl74)),
  inference(forward_demodulation,[],[f65392,f1802])).
fof(f65392,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl74)),
  inference(forward_demodulation,[],[f65335,f146])).
fof(f65335,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))))) | ($spl0 | $spl74)),
  inference(backward_demodulation,[],[f65256,f64374])).
fof(f65391,plain,(
  $false | ($spl0 | $spl74)),
  inference(subsumption_resolution,[],[f65390,f195])).
fof(f65390,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) | ($spl0 | $spl74)),
  inference(forward_demodulation,[],[f65389,f227])).
fof(f65389,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl74)),
  inference(forward_demodulation,[],[f65330,f148])).
fof(f65330,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num','const.num.|48|_0'))))))) | ($spl0 | $spl74)),
  inference(backward_demodulation,[],[f65256,f41661])).
fof(f65388,plain,(
  $false | ($spl0 | $spl74)),
  inference(subsumption_resolution,[],[f65325,f227])).
fof(f65325,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl74)),
  inference(backward_demodulation,[],[f65256,f3687])).
fof(f65383,plain,(
  $false | ($spl0 | $spl74)),
  inference(subsumption_resolution,[],[f65381,f195])).
fof(f65381,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',X0)))))) ) | ($spl0 | $spl74)),
  inference(backward_demodulation,[],[f65319,f10302])).
fof(f65319,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl74)),
  inference(backward_demodulation,[],[f65256,f3512])).
fof(f65382,plain,(
  $false | ($spl0 | $spl74)),
  inference(subsumption_resolution,[],[f65380,f195])).
fof(f65380,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',sK0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',sK0)))))) ) | ($spl0 | $spl74)),
  inference(backward_demodulation,[],[f65319,f9665])).
fof(f65379,plain,(
  $false | ($spl0 | $spl74)),
  inference(subsumption_resolution,[],[f65378,f195])).
fof(f65378,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK0))))) | ($spl0 | $spl74)),
  inference(forward_demodulation,[],[f65318,f1540])).
fof(f65318,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK0))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl0 | $spl74)),
  inference(backward_demodulation,[],[f65256,f2754])).
fof(f65377,plain,(
  $false | ($spl0 | $spl74)),
  inference(subsumption_resolution,[],[f65376,f574])).
fof(f65376,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl74)),
  inference(forward_demodulation,[],[f65314,f147])).
fof(f65314,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl74)),
  inference(backward_demodulation,[],[f65256,f2174])).
fof(f65374,plain,(
  $false | ($spl0 | $spl74)),
  inference(subsumption_resolution,[],[f65373,f195])).
fof(f65373,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) | ($spl0 | $spl74)),
  inference(forward_demodulation,[],[f65372,f227])).
fof(f65372,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl74)),
  inference(forward_demodulation,[],[f65312,f1802])).
fof(f65312,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl74)),
  inference(backward_demodulation,[],[f65256,f2133])).
fof(f65371,plain,(
  $false | ($spl0 | $spl74)),
  inference(subsumption_resolution,[],[f65370,f195])).
fof(f65370,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl74)),
  inference(forward_demodulation,[],[f65311,f147])).
fof(f65311,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl74)),
  inference(backward_demodulation,[],[f65256,f1804])).
fof(f65369,plain,(
  $false | ($spl0 | $spl74)),
  inference(subsumption_resolution,[],[f65368,f195])).
fof(f65368,plain,(
  s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) | ($spl0 | $spl74)),
  inference(forward_demodulation,[],[f65310,f227])).
fof(f65310,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl74)),
  inference(backward_demodulation,[],[f65256,f1642])).
fof(f65367,plain,(
  $false | ($spl0 | $spl74)),
  inference(subsumption_resolution,[],[f65366,f195])).
fof(f65366,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl74)),
  inference(forward_demodulation,[],[f65303,f147])).
fof(f65303,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl74)),
  inference(backward_demodulation,[],[f65256,f258])).
fof(f65239,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1)))) | $spl62),
  inference(cnf_transformation,[],[f65239_D])).
fof(f65239_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1)))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f65301,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK2(sK0)),s('type.num.num',sK1))) | $spl97),
  inference(cnf_transformation,[],[f65301_D])).
fof(f65301_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK2(sK0)),s('type.num.num',sK1))) <=> ~$spl97),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl97])])).
fof(f65218,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK2(sK0)),s('type.num.num',sK1))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1)))) | ($spl0 | $spl16)),
  inference(superposition,[],[f23499,f64681])).
fof(f23499,plain,(
  ( ! [X21] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK2(sK0)),s('type.num.num',X21))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X21)) ) | ($spl0 | $spl16)),
  inference(superposition,[],[f1760,f22867])).
fof(f65299,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK1),s('type.num.num',sK2(sK0)))) | $spl95),
  inference(cnf_transformation,[],[f65299_D])).
fof(f65299_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK1),s('type.num.num',sK2(sK0)))) <=> ~$spl95),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl95])])).
fof(f65217,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK1),s('type.num.num',sK2(sK0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1)))) | ($spl0 | $spl16)),
  inference(superposition,[],[f23498,f64681])).
fof(f23498,plain,(
  ( ! [X20] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X20),s('type.num.num',sK2(sK0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X20)) ) | ($spl0 | $spl16)),
  inference(superposition,[],[f1548,f22867])).
fof(f65296,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK1),s('type.num.num',sK0))) | $spl93),
  inference(cnf_transformation,[],[f65296_D])).
fof(f65296_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK1),s('type.num.num',sK0))) <=> ~$spl93),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl93])])).
fof(f65183,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK1),s('type.num.num',sK0))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1)))) | $spl0),
  inference(superposition,[],[f1907,f64681])).
fof(f1907,plain,(
  ( ! [X24] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X24),s('type.num.num',sK0))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X24)) ) | $spl0),
  inference(subsumption_resolution,[],[f1890,f574])).
fof(f1890,plain,(
  ( ! [X24] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X24),s('type.num.num',sK0))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(sK0))))) ) | $spl0),
  inference(superposition,[],[f213,f1835])).
fof(f65291,plain,(
  ( ! [X86] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X86))),s('type.num.num',sK1)))) ) | $spl90),
  inference(cnf_transformation,[],[f65291_D])).
fof(f65291_D,plain,(
  ( ! [X86] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X86))),s('type.num.num',sK1)))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f65169,plain,(
  ( ! [X86] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X86))),s('type.num.num',sK1))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1))))) ) | $spl0),
  inference(superposition,[],[f1760,f64681])).
fof(f65289,plain,(
  ( ! [X83] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X83)))))) ) | $spl88),
  inference(cnf_transformation,[],[f65289_D])).
fof(f65289_D,plain,(
  ( ! [X83] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X83)))))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f65163,plain,(
  ( ! [X83] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X83))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1))))) ) | $spl0),
  inference(superposition,[],[f1548,f64681])).
fof(f65284,plain,(
  ( ! [X68] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK1),s('type.num.num',X68))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X68),s('type.num.num','const.num.|48|_0')))) ) | $spl86),
  inference(cnf_transformation,[],[f65284_D])).
fof(f65284_D,plain,(
  ( ! [X68] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK1),s('type.num.num',X68))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X68),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f65153,plain,(
  ( ! [X68] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK1),s('type.num.num',X68))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X68),s('type.num.num','const.num.|48|_0')))) ) | $spl0),
  inference(superposition,[],[f1373,f64681])).
fof(f65282,plain,(
  ( ! [X66] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK1),s('type.num.num',X66))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X66)))) ) | $spl84),
  inference(cnf_transformation,[],[f65282_D])).
fof(f65282_D,plain,(
  ( ! [X66] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK1),s('type.num.num',X66))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X66)))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f65151,plain,(
  ( ! [X66] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK1),s('type.num.num',X66))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X66)))) ) | $spl0),
  inference(superposition,[],[f1372,f64681])).
fof(f65280,plain,(
  ( ! [X64] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X64),s('type.num.num',sK1))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X64),s('type.num.num','const.num.|48|_0')))) ) | $spl82),
  inference(cnf_transformation,[],[f65280_D])).
fof(f65280_D,plain,(
  ( ! [X64] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X64),s('type.num.num',sK1))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X64),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f65149,plain,(
  ( ! [X64] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X64),s('type.num.num',sK1))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X64),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1))))) ) | $spl0),
  inference(superposition,[],[f1360,f64681])).
fof(f65278,plain,(
  ( ! [X62] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X62),s('type.num.num',sK1))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X62)))) ) | $spl80),
  inference(cnf_transformation,[],[f65278_D])).
fof(f65278_D,plain,(
  ( ! [X62] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X62),s('type.num.num',sK1))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X62)))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f65147,plain,(
  ( ! [X62] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X62),s('type.num.num',sK1))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X62))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1))))) ) | $spl0),
  inference(superposition,[],[f1359,f64681])).
fof(f65266,plain,(
  ( ! [X37] : (s('type.num.num',sK1) != s('type.num.num',X37) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X37))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1)))))))))) ) | $spl78),
  inference(cnf_transformation,[],[f65266_D])).
fof(f65266_D,plain,(
  ( ! [X37] : (s('type.num.num',sK1) != s('type.num.num',X37) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X37))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1)))))))))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f65256,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK1) | $spl74),
  inference(cnf_transformation,[],[f65256_D])).
fof(f65256_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK1) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f65265,plain,(
  ( ! [X37] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',sK1) | s('type.num.num',sK1) != s('type.num.num',X37) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X37))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1)))))))))) ) | $spl0),
  inference(forward_demodulation,[],[f65119,f64681])).
fof(f65119,plain,(
  ( ! [X37] : (s('type.num.num',sK1) != s('type.num.num',X37) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X37))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1))))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1))))) ) | $spl0),
  inference(superposition,[],[f437,f64681])).
fof(f65263,plain,(
  ( ! [X35] : (s('type.num.num',sK1) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X35))) | s('type.num.num',X35) = s('type.num.num',sK2(sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1)))))) ) | $spl76),
  inference(cnf_transformation,[],[f65263_D])).
fof(f65263_D,plain,(
  ( ! [X35] : (s('type.num.num',sK1) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X35))) | s('type.num.num',X35) = s('type.num.num',sK2(sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1)))))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f65262,plain,(
  ( ! [X35] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',sK1) | s('type.num.num',sK1) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X35))) | s('type.num.num',X35) = s('type.num.num',sK2(sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1)))))) ) | $spl0),
  inference(forward_demodulation,[],[f65117,f64681])).
fof(f65117,plain,(
  ( ! [X35] : (s('type.num.num',sK1) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X35))) | s('type.num.num',X35) = s('type.num.num',sK2(sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1))))) ) | $spl0),
  inference(superposition,[],[f387,f64681])).
fof(f65254,plain,(
  ( ! [X29] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X29))) | s('type.num.num',X29) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1)))))))) ) | $spl72),
  inference(cnf_transformation,[],[f65254_D])).
fof(f65254_D,plain,(
  ( ! [X29] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X29))) | s('type.num.num',X29) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1)))))))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f65253,plain,(
  ( ! [X29] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',sK1) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X29))) | s('type.num.num',X29) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1)))))))) ) | $spl0),
  inference(forward_demodulation,[],[f65111,f64681])).
fof(f65111,plain,(
  ( ! [X29] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X29))) | s('type.num.num',X29) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1))))) ) | $spl0),
  inference(superposition,[],[f384,f64681])).
fof(f65244,plain,(
  ( ! [X20] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X20),s('type.num.num',sK1))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X20)) ) | $spl70),
  inference(cnf_transformation,[],[f65244_D])).
fof(f65244_D,plain,(
  ( ! [X20] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X20),s('type.num.num',sK1))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X20)) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f65100,plain,(
  ( ! [X20] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X20),s('type.num.num',sK1))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X20) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1))))) ) | $spl0),
  inference(superposition,[],[f213,f64681])).
fof(f65243,plain,(
  ( ! [X19] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK1),s('type.num.num',X19))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X19)) ) | $spl68),
  inference(cnf_transformation,[],[f65243_D])).
fof(f65243_D,plain,(
  ( ! [X19] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK1),s('type.num.num',X19))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X19)) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f65099,plain,(
  ( ! [X19] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK1),s('type.num.num',X19))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X19)) ) | $spl0),
  inference(superposition,[],[f213,f64681])).
fof(f65241,plain,(
  ( ! [X17] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X17),s('type.num.num',sK1)))) ) | $spl66),
  inference(cnf_transformation,[],[f65241_D])).
fof(f65241_D,plain,(
  ( ! [X17] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X17),s('type.num.num',sK1)))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f65097,plain,(
  ( ! [X17] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X17),s('type.num.num',sK1))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1))))) ) | $spl0),
  inference(superposition,[],[f211,f64681])).
fof(f65240,plain,(
  ( ! [X14] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',X14)))) ) | $spl64),
  inference(cnf_transformation,[],[f65240_D])).
fof(f65240_D,plain,(
  ( ! [X14] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',X14)))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f65094,plain,(
  ( ! [X14] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK1),s('type.num.num',X14))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1))))) ) | $spl0),
  inference(superposition,[],[f210,f64681])).
fof(f65233,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1))) | $spl58),
  inference(cnf_transformation,[],[f65233_D])).
fof(f65233_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f65235,plain,(
  s('type.num.num',sK1) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1))) | $spl61),
  inference(cnf_transformation,[],[f65235_D])).
fof(f65235_D,plain,(
  s('type.num.num',sK1) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f65075,plain,(
  s('type.num.num',sK1) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1))) | $spl0),
  inference(superposition,[],[f363,f64681])).
fof(f64701,plain,(
  s('type.num.num',sK2(sK2(sK0))) = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1)))) | $spl54),
  inference(cnf_transformation,[],[f64701_D])).
fof(f64701_D,plain,(
  s('type.num.num',sK2(sK2(sK0))) = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1)))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f64703,plain,(
  s('type.num.num',sK2(sK0)) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))) | $spl57),
  inference(cnf_transformation,[],[f64703_D])).
fof(f64703_D,plain,(
  s('type.num.num',sK2(sK0)) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))) <=> ~$spl57),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl57])])).
fof(f64680,plain,(
  s('type.num.num',sK2(sK0)) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))) | s('type.num.num',sK2(sK2(sK0))) = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1)))) | ($spl0 | $spl16)),
  inference(superposition,[],[f3605,f26394])).
fof(f26394,plain,(
  s('type.num.num',sK2(sK0)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(sK2(sK0))))) | ($spl0 | $spl16)),
  inference(unit_resulting_resolution,[],[f23496,f153])).
fof(f23496,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK2(sK0)) | ($spl0 | $spl16)),
  inference(superposition,[],[f574,f22867])).
fof(f64692,plain,(
  s('type.num.num',sK2(sK0)) = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1)))) | $spl52),
  inference(cnf_transformation,[],[f64692_D])).
fof(f64692_D,plain,(
  s('type.num.num',sK2(sK0)) = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1)))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f64690,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))) | $spl51),
  inference(cnf_transformation,[],[f64690_D])).
fof(f64690_D,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))) <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f64679,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))) | s('type.num.num',sK2(sK0)) = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1)))) | $spl0),
  inference(superposition,[],[f3605,f1835])).
fof(f64688,plain,(
  s('type.num.num',sK2(sK0)) = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1)))) | s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))) | ($spl0 | $spl16)),
  inference(forward_demodulation,[],[f64687,f26394])).
fof(f64687,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(sK2(sK0))))) = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1)))) | ($spl0 | $spl16)),
  inference(forward_demodulation,[],[f64659,f1835])).
fof(f64659,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(sK0)))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(sK2(sK0))))) = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1)))) | ($spl0 | $spl16)),
  inference(superposition,[],[f3605,f26394])).
fof(f61763,plain,(
  s('type.num.num',sK2(sK2(sK0))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))) | $spl46),
  inference(cnf_transformation,[],[f61763_D])).
fof(f61763_D,plain,(
  s('type.num.num',sK2(sK2(sK0))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f61765,plain,(
  s('type.num.num',sK2(sK0)) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) | $spl49),
  inference(cnf_transformation,[],[f61765_D])).
fof(f61765_D,plain,(
  s('type.num.num',sK2(sK0)) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f61723,plain,(
  s('type.num.num',sK2(sK0)) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) | s('type.num.num',sK2(sK2(sK0))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))) | ($spl0 | $spl16)),
  inference(superposition,[],[f23479,f1642])).
fof(f61758,plain,(
  s('type.num.num',sK2(sK2(sK0))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1))) | $spl42),
  inference(cnf_transformation,[],[f61758_D])).
fof(f61758_D,plain,(
  s('type.num.num',sK2(sK2(sK0))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f61760,plain,(
  s('type.num.num',sK2(sK0)) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))) | $spl45),
  inference(cnf_transformation,[],[f61760_D])).
fof(f61760_D,plain,(
  s('type.num.num',sK2(sK0)) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f61721,plain,(
  s('type.num.num',sK2(sK0)) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))) | s('type.num.num',sK2(sK2(sK0))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1))) | ($spl0 | $spl16)),
  inference(superposition,[],[f23479,f1804])).
fof(f54144,plain,(
  $false | ($spl0 | $spl5 | $spl16 | $spl30)),
  inference(subsumption_resolution,[],[f54143,f38813])).
fof(f38813,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) | ($spl0 | $spl16 | $spl30)),
  inference(superposition,[],[f24194,f24881])).
fof(f24881,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',X3))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X3)))) ) | ($spl0 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f23785,f1835])).
fof(f23785,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X3))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(sK0)))),s('type.num.num',X3)))) ) | ($spl0 | $spl16 | $spl30)),
  inference(backward_demodulation,[],[f23740,f1741])).
fof(f23740,plain,(
  s('type.num.num',sK2(sK0)) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl16 | $spl30)),
  inference(backward_demodulation,[],[f23682,f22867])).
fof(f24194,plain,(
  ( ! [X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',X1)))) ) | ($spl0 | $spl16 | $spl30)),
  inference(backward_demodulation,[],[f23784,f1860])).
fof(f1860,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',X1))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(sK0)),s('type.num.num',X1)))))) ) | $spl0),
  inference(superposition,[],[f201,f1835])).
fof(f23784,plain,(
  ( ! [X3] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK2(sK0)),s('type.num.num',X3)))) ) | ($spl0 | $spl16 | $spl30)),
  inference(backward_demodulation,[],[f23740,f1587])).
fof(f54143,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) | ($spl0 | $spl5 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f54142,f54133])).
fof(f54133,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK0))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) | ($spl0 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f54132,f23761])).
fof(f23761,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',sK2(sK0))))) ) | ($spl0 | $spl16 | $spl30)),
  inference(backward_demodulation,[],[f23740,f227])).
fof(f54132,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK0))),s('type.num.num',sK2(sK0)))) | ($spl0 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f54064,f23740])).
fof(f54064,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl16 | $spl30)),
  inference(backward_demodulation,[],[f53968,f39007])).
fof(f39007,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))) | ($spl0 | $spl16 | $spl30)),
  inference(backward_demodulation,[],[f38813,f37971])).
fof(f37971,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))) | ($spl0 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f37970,f25854])).
fof(f25854,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',sK0)))))))) ) | ($spl0 | $spl16 | $spl30)),
  inference(backward_demodulation,[],[f25543,f25537])).
fof(f25537,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',sK0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',sK0)))))))) ) | ($spl0 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f25536,f24191])).
fof(f24191,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',sK0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1)))) ) | ($spl0 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f23783,f1835])).
fof(f23783,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(sK0))))))) ) | ($spl0 | $spl16 | $spl30)),
  inference(backward_demodulation,[],[f23740,f1540])).
fof(f25536,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',sK0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1)))))))) ) | ($spl0 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f23861,f1835])).
fof(f23861,plain,(
  ( ! [X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1))))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(sK0))))))) ) | ($spl0 | $spl16 | $spl30)),
  inference(backward_demodulation,[],[f23740,f3850])).
fof(f25543,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',sK0))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',X1)))) ) | ($spl0 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f23865,f1835])).
fof(f23865,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(sK0)))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(sK0)))),s('type.num.num',X1)))) ) | ($spl0 | $spl16 | $spl30)),
  inference(backward_demodulation,[],[f23740,f3933])).
fof(f37970,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK0))))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))) | ($spl0 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f37969,f28715])).
fof(f28715,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | ($spl0 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f28547,f24893])).
fof(f24893,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',sK0))))))) ) | ($spl0 | $spl16 | $spl30)),
  inference(backward_demodulation,[],[f24882,f2029])).
fof(f2029,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',sK0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',sK0))))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1832,f153])).
fof(f1832,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',sK0)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1662,f211])).
fof(f24882,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',sK0)))) ) | ($spl0 | $spl16 | $spl30)),
  inference(backward_demodulation,[],[f23786,f1859])).
fof(f1859,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',sK0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',sK2(sK0))))))) ) | $spl0),
  inference(superposition,[],[f200,f1835])).
fof(f23786,plain,(
  ( ! [X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',sK2(sK0))))) ) | ($spl0 | $spl16 | $spl30)),
  inference(backward_demodulation,[],[f23740,f1802])).
fof(f28547,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2('const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK0)))))))) | ($spl0 | $spl16 | $spl30)),
  inference(unit_resulting_resolution,[],[f1833,f24191,f437])).
fof(f1833,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',X0)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1662,f210])).
fof(f37969,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))) | ($spl0 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f37824,f24882])).
fof(f37824,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num',sK0))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))) | ($spl0 | $spl16 | $spl30)),
  inference(superposition,[],[f202,f25542])).
fof(f25542,plain,(
  s('type.num.num',sK0) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num',sK1))) | ($spl0 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f25541,f1835])).
fof(f25541,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(sK0)))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num',sK1))) | ($spl0 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f25540,f23763])).
fof(f23763,plain,(
  s('type.num.num',sK2(sK0)) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK1))) | ($spl0 | $spl16 | $spl30)),
  inference(backward_demodulation,[],[f23740,f258])).
fof(f25540,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK1))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num',sK1))) | ($spl0 | $spl16 | $spl30)),
  inference(backward_demodulation,[],[f25539,f24193])).
fof(f24193,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK1),s('type.num.num',sK0))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num',sK1))) | ($spl0 | $spl16 | $spl30)),
  inference(backward_demodulation,[],[f24191,f2174])).
fof(f25539,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',sK0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',X0)))))) ) | ($spl0 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f25538,f24191])).
fof(f25538,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',X0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0)))))) ) | ($spl0 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f23864,f1835])).
fof(f23864,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(sK0)))),s('type.num.num',X0)))))) ) | ($spl0 | $spl16 | $spl30)),
  inference(backward_demodulation,[],[f23740,f3932])).
fof(f53968,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK1) | ($spl0 | $spl16 | $spl30)),
  inference(unit_resulting_resolution,[],[f23763,f53049])).
fof(f53049,plain,(
  ( ! [X11] : (s('type.num.num',sK2(sK0)) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',X11))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X11)) ) | ($spl0 | $spl16 | $spl30)),
  inference(superposition,[],[f52624,f153])).
fof(f52624,plain,(
  ( ! [X2] : (s('type.num.num',sK2(sK0)) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2)))))) ) | ($spl0 | $spl16 | $spl30)),
  inference(superposition,[],[f23774,f25854])).
fof(f23774,plain,(
  ( ! [X0] : (s('type.num.num',sK2(sK0)) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) ) | ($spl0 | $spl16 | $spl30)),
  inference(backward_demodulation,[],[f23740,f514])).
fof(f54142,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) | ($spl0 | $spl5 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f54141,f23761])).
fof(f54141,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK2(sK0)))))) | ($spl0 | $spl5 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f54069,f23740])).
fof(f54069,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl0 | $spl5 | $spl16 | $spl30)),
  inference(backward_demodulation,[],[f53968,f40818])).
fof(f40818,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))))) | ($spl0 | $spl5 | $spl16 | $spl30)),
  inference(backward_demodulation,[],[f40817,f28719])).
fof(f28719,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))))))))) | ($spl0 | $spl5 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f28716,f25854])).
fof(f28716,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK0))))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))))))))) | ($spl0 | $spl5 | $spl16 | $spl30)),
  inference(backward_demodulation,[],[f28715,f24190])).
fof(f24190,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))))))))) | ($spl0 | $spl5 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f23778,f1835])).
fof(f23778,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))))))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(sK0)))))))))))) | ($spl0 | $spl5 | $spl16 | $spl30)),
  inference(backward_demodulation,[],[f23740,f1308])).
fof(f1308,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))))))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))))))) | $spl5),
  inference(unit_resulting_resolution,[],[f469,f208])).
fof(f469,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))))) | $spl5),
  inference(unit_resulting_resolution,[],[f408,f208])).
fof(f408,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | $spl5),
  inference(unit_resulting_resolution,[],[f406,f208])).
fof(f406,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl5),
  inference(unit_resulting_resolution,[],[f376,f208])).
fof(f376,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl5),
  inference(unit_resulting_resolution,[],[f330,f208])).
fof(f330,plain,(
  s('type.num.num',sK1) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl5),
  inference(superposition,[],[f225,f140])).
fof(f40817,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))))))) | ($spl0 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f40816,f24194])).
fof(f40816,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK1))))))) | $spl0),
  inference(forward_demodulation,[],[f40606,f200])).
fof(f40606,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))))) | $spl0),
  inference(superposition,[],[f1545,f1804])).
fof(f54140,plain,(
  $false | ($spl0 | $spl16 | $spl30)),
  inference(subsumption_resolution,[],[f54139,f195])).
fof(f54139,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) | ($spl0 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f54138,f23761])).
fof(f54138,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK2(sK0)))) | ($spl0 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f54137,f1835])).
fof(f54137,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK2(sK0)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(sK0)))))))) | ($spl0 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f54068,f23740])).
fof(f54068,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | ($spl0 | $spl16 | $spl30)),
  inference(backward_demodulation,[],[f53968,f40817])).
fof(f54131,plain,(
  $false | ($spl0 | $spl16 | $spl30)),
  inference(subsumption_resolution,[],[f54130,f1662])).
fof(f54130,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK0) | ($spl0 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f54062,f147])).
fof(f54062,plain,(
  s('type.num.num',sK0) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK0))),s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl16 | $spl30)),
  inference(backward_demodulation,[],[f53968,f39004])).
fof(f39004,plain,(
  s('type.num.num',sK0) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK0))),s('type.num.num',sK1))) | ($spl0 | $spl16 | $spl30)),
  inference(backward_demodulation,[],[f38813,f25542])).
fof(f54127,plain,(
  $false | ($spl0 | $spl9 | $spl16 | $spl30)),
  inference(subsumption_resolution,[],[f54126,f38813])).
fof(f54126,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) | ($spl0 | $spl9 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f54125,f1835])).
fof(f54125,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(sK0)))))))) | ($spl0 | $spl9 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f54060,f23740])).
fof(f54060,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | ($spl0 | $spl9 | $spl16 | $spl30)),
  inference(backward_demodulation,[],[f53968,f39002])).
fof(f39002,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))))))) | ($spl0 | $spl9 | $spl16 | $spl30)),
  inference(backward_demodulation,[],[f38813,f5791])).
fof(f5791,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))))))) | $spl9),
  inference(unit_resulting_resolution,[],[f2701,f208])).
fof(f2701,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))))) | $spl9),
  inference(unit_resulting_resolution,[],[f2641,f208])).
fof(f54121,plain,(
  $false | ($spl0 | $spl16 | $spl30)),
  inference(subsumption_resolution,[],[f54120,f574])).
fof(f54120,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f54057,f147])).
fof(f54057,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK0))))),s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl16 | $spl30)),
  inference(backward_demodulation,[],[f53968,f37973])).
fof(f37973,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK0))))),s('type.num.num',sK1))) | ($spl0 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f37972,f24194])).
fof(f37972,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK1))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK0))))),s('type.num.num',sK1))) | ($spl0 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f37825,f28715])).
fof(f37825,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK1))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.num.num',sK1))) | ($spl0 | $spl16 | $spl30)),
  inference(superposition,[],[f203,f25542])).
fof(f54115,plain,(
  $false | ($spl0 | $spl16 | $spl30)),
  inference(subsumption_resolution,[],[f54114,f195])).
fof(f54114,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK0))))) | ($spl0 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f54113,f1835])).
fof(f54113,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK0))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(sK0)))))) | ($spl0 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f54052,f23740])).
fof(f54052,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK0))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl0 | $spl16 | $spl30)),
  inference(backward_demodulation,[],[f53968,f24192])).
fof(f24192,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK0))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))))) | ($spl0 | $spl16 | $spl30)),
  inference(backward_demodulation,[],[f24191,f2754])).
fof(f54110,plain,(
  $false | ($spl0 | $spl16 | $spl30)),
  inference(subsumption_resolution,[],[f54109,f23496])).
fof(f54109,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK2(sK0)) | ($spl0 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f54048,f147])).
fof(f54048,plain,(
  s('type.num.num',sK2(sK0)) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl16 | $spl30)),
  inference(backward_demodulation,[],[f53968,f23763])).
fof(f54100,plain,(
  $false | ($spl0 | $spl16 | $spl30)),
  inference(subsumption_resolution,[],[f54099,f23761])).
fof(f54099,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK2(sK0)))) | ($spl0 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f54042,f23740])).
fof(f54042,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl16 | $spl30)),
  inference(backward_demodulation,[],[f53968,f3687])).
fof(f54092,plain,(
  $false | ($spl0 | $spl16 | $spl30)),
  inference(subsumption_resolution,[],[f54091,f29918])).
fof(f29918,plain,(
  s('type.num.num',sK0) != s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK2(sK0))))) | ($spl0 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f29846,f23761])).
fof(f29846,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK2(sK0)))) != s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK2(sK0))))) | ($spl0 | $spl16)),
  inference(unit_resulting_resolution,[],[f195,f23502,f437])).
fof(f23502,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK2(sK0)))) | ($spl0 | $spl16)),
  inference(superposition,[],[f1824,f22867])).
fof(f1824,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f574,f1662,f213])).
fof(f54091,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK2(sK0))))) | ($spl0 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f54038,f23740])).
fof(f54038,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK2('const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl0 | $spl16 | $spl30)),
  inference(backward_demodulation,[],[f53968,f3512])).
fof(f54090,plain,(
  $false | ($spl0 | $spl9 | $spl16 | $spl30)),
  inference(trivial_inequality_removal,[],[f54089])).
fof(f54089,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) | ($spl0 | $spl9 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f54088,f1835])).
fof(f54088,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(sK0)))))) | ($spl0 | $spl9 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f54037,f23740])).
fof(f54037,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl0 | $spl9 | $spl16 | $spl30)),
  inference(backward_demodulation,[],[f53968,f2701])).
fof(f54087,plain,(
  $false | ($spl0 | $spl9 | $spl16 | $spl30)),
  inference(subsumption_resolution,[],[f54086,f1835])).
fof(f54086,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(sK0)))) | ($spl0 | $spl9 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f54036,f23740])).
fof(f54036,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl9 | $spl16 | $spl30)),
  inference(backward_demodulation,[],[f53968,f2641])).
fof(f54082,plain,(
  $false | ($spl0 | $spl16 | $spl30)),
  inference(subsumption_resolution,[],[f54081,f195])).
fof(f54081,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) | ($spl0 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f54080,f23761])).
fof(f54080,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK2(sK0)))) | ($spl0 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f54079,f23786])).
fof(f54079,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK2(sK0)))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK2(sK0)))) | ($spl0 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f54034,f23740])).
fof(f54034,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl16 | $spl30)),
  inference(backward_demodulation,[],[f53968,f2133])).
fof(f54078,plain,(
  $false | ($spl0 | $spl16 | $spl30)),
  inference(subsumption_resolution,[],[f54077,f195])).
fof(f54077,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f54033,f147])).
fof(f54033,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl16 | $spl30)),
  inference(backward_demodulation,[],[f53968,f1804])).
fof(f54076,plain,(
  $false | ($spl0 | $spl16 | $spl30)),
  inference(subsumption_resolution,[],[f54075,f195])).
fof(f54075,plain,(
  s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) | ($spl0 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f54074,f23761])).
fof(f54074,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK2(sK0)))) | ($spl0 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f54032,f23740])).
fof(f54032,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl16 | $spl30)),
  inference(backward_demodulation,[],[f53968,f1642])).
fof(f42163,plain,(
  ( ! [X6,X7] : (s(fun(bool,bool),X6) = s(fun(bool,bool),X7) | s(bool,'T_0') = s(bool,sK3(X6,X7,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X6,X7,bool,bool))))) ) | $spl38),
  inference(cnf_transformation,[],[f42163_D])).
fof(f42163_D,plain,(
  ( ! [X6,X7] : (s(fun(bool,bool),X6) = s(fun(bool,bool),X7) | s(bool,'T_0') = s(bool,sK3(X6,X7,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X6,X7,bool,bool))))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f42164,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,f(s(bool,X8))) | s(bool,'T_0') = s(bool,X8)) ) | $spl40),
  inference(cnf_transformation,[],[f42164_D])).
fof(f42164_D,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,f(s(bool,X8))) | s(bool,'T_0') = s(bool,X8)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f42152,plain,(
  ( ! [X10,X11,X9] : (s(bool,'F_0') != s(bool,f(s(bool,X11))) | s(fun(bool,bool),X9) = s(fun(bool,bool),X10) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X9,X10,bool,bool)))) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,sK3(X9,X10,bool,bool))) )),
  inference(superposition,[],[f1855,f293])).
fof(f42151,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') != s(bool,f(s(bool,X8))) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X6,X7,bool,bool)))) | s(bool,'T_0') = s(bool,sK3(X6,X7,bool,bool)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f1855,f293])).
fof(f42160,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK3(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X0,X1,bool,bool))))) ) | $spl34),
  inference(cnf_transformation,[],[f42160_D])).
fof(f42160_D,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK3(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X0,X1,bool,bool))))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f42162,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl37),
  inference(cnf_transformation,[],[f42162_D])).
fof(f42162_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f42150,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X3,X4,bool,bool)))) | s(bool,X5) = s(bool,sK3(X3,X4,bool,bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1855,f293])).
fof(f42149,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X0,X1,bool,bool)))) | s(bool,X2) = s(bool,sK3(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1855,f293])).
fof(f23682,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK2(sK2(sK0))) | $spl30),
  inference(cnf_transformation,[],[f23682_D])).
fof(f23682_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK2(sK2(sK0))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f23683,plain,(
  ( ! [X11] : (s('type.num.num',sK2(sK0)) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))) | s('type.num.num',X11) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(sK2(sK2(sK0))))))) ) | $spl32),
  inference(cnf_transformation,[],[f23683_D])).
fof(f23683_D,plain,(
  ( ! [X11] : (s('type.num.num',sK2(sK0)) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))) | s('type.num.num',X11) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(sK2(sK2(sK0))))))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f23488,plain,(
  ( ! [X11] : (s('type.num.num',sK2(sK0)) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))) | s('type.num.num',X11) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(sK2(sK2(sK0)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK2(sK2(sK0)))) ) | ($spl0 | $spl16)),
  inference(superposition,[],[f384,f22867])).
fof(f21538,plain,(
  $false | ($spl23 | $spl26)),
  inference(subsumption_resolution,[],[f21537,f138])).
fof(f21537,plain,(
  ~p(s(bool,'T_0')) | ($spl23 | $spl26)),
  inference(forward_demodulation,[],[f21528,f21519])).
fof(f21519,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl23 | $spl26)),
  inference(subsumption_resolution,[],[f21514,f21483])).
fof(f21514,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl23),
  inference(trivial_inequality_removal,[],[f21508])).
fof(f21508,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl23),
  inference(superposition,[],[f21480,f293])).
fof(f21528,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl23 | $spl26)),
  inference(backward_demodulation,[],[f21519,f21501])).
fof(f21501,plain,(
  ~p(s(bool,f(s(bool,'F_0')))) | $spl23),
  inference(unit_resulting_resolution,[],[f138,f21480,f207])).
fof(f21536,plain,(
  $false | ($spl23 | $spl26 | $spl29)),
  inference(subsumption_resolution,[],[f21527,f21519])).
fof(f21527,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl23 | $spl26 | $spl29)),
  inference(backward_demodulation,[],[f21519,f21485])).
fof(f21535,plain,(
  $false | ($spl23 | $spl26)),
  inference(subsumption_resolution,[],[f21525,f21519])).
fof(f21525,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl23 | $spl26)),
  inference(backward_demodulation,[],[f21519,f21480])).
fof(f21534,plain,(
  $false | ($spl23 | $spl26)),
  inference(subsumption_resolution,[],[f21522,f138])).
fof(f21522,plain,(
  ~p(s(bool,'T_0')) | ($spl23 | $spl26)),
  inference(backward_demodulation,[],[f21519,f226])).
fof(f21533,plain,(
  $false | ($spl23 | $spl26)),
  inference(global_subsumption,[],[f137,f136,f135,f138,f226,f146,f147,f148,f149,f152,f153,f156,f155,f160,f159,f195,f200,f201,f202,f203,f207,f204,f209,f208,f212,f211,f210,f215,f214,f213,f216,f260,f293,f295,f322,f362,f363,f384,f385,f386,f387,f437,f470,f471,f481,f482,f517,f519,f521,f523,f547,f542,f543,f567,f569,f574,f619,f618,f704,f705,f736,f737,f762,f803,f807,f903,f962,f980,f1297,f1347,f1346,f1341,f1340,f1359,f1360,f1441,f1372,f1373,f1445,f1481,f1482,f1545,f1547,f1548,f1561,f1565,f1568,f1570,f1571,f1759,f1760,f1761,f1777,f1781,f1784,f1786,f1787,f1854,f1855,f2780,f2815,f2816,f3001,f3002,f2858,f2859,f3100,f3101,f3350,f3351,f3516,f3514,f3566,f3506,f3500,f3611,f3613,f3614,f3616,f3617,f3618,f3619,f3621,f3548,f3550,f3553,f3554,f3555,f3560,f3561,f4792,f4810,f4822,f4826,f4827,f4830,f4831,f4836,f5047,f5065,f5077,f5081,f5082,f5085,f5091,f5481,f5763,f5766,f5858,f6018,f8514,f8532,f8957,f8975,f9481,f9681,f9682,f9485,f9486,f9490,f9491,f9492,f9493,f9501,f9502,f9503,f9504,f9505,f9507,f9508,f9510,f9513,f9536,f9538,f9557,f9558,f9559,f9560,f9685,f9562,f9686,f9687,f9596,f9597,f9688,f9689,f9604,f10111,f10318,f10115,f10117,f10118,f10122,f10123,f10124,f10125,f10133,f10134,f10135,f10136,f10137,f10139,f10140,f10142,f10145,f10170,f10172,f10191,f10192,f10193,f10194,f10322,f10196,f10323,f10324,f10231,f10232,f10325,f10326,f10239,f11115,f11116,f11122,f11123,f11131,f10946,f10948,f10949,f10960,f10962,f10963,f10966,f10970,f10976,f10982,f10983,f11106,f11098,f11097,f10973,f10969,f10965,f10964,f10955,f10952,f10931,f10930,f10921,f10918,f10190,f10189,f10187,f10186,f10138,f10132,f10131,f10129,f10128,f10110,f10109,f9556,f9555,f9553,f9552,f9506,f9500,f9499,f9497,f9496,f9480,f9479,f8990,f8989,f8982,f8979,f8960,f8959,f8952,f8949,f8517,f8516,f8509,f8506,f8489,f8488,f8481,f8478,f7275,f7258,f7257,f7250,f7247,f6966,f6949,f6948,f6941,f6938,f6584,f6535,f6534,f6527,f6524,f6487,f6438,f6437,f6430,f6427,f6357,f6354,f6351,f6350,f6347,f6346,f6345,f6343,f6338,f6337,f6335,f6331,f6316,f6315,f6313,f6308,f6305,f6169,f6166,f6163,f6162,f6159,f6158,f6157,f6155,f6150,f6149,f6147,f6146,f6143,f6128,f6127,f6125,f6120,f6117,f6001,f6000,f5993,f5990,f5841,f5840,f5833,f5830,f5717,f5716,f5709,f5706,f5432,f5431,f5424,f5421,f5088,f5084,f5080,f5079,f5072,f5071,f5069,f5068,f5050,f5049,f5042,f5039,f4833,f4829,f4825,f4824,f4817,f4816,f4814,f4813,f4795,f4794,f4787,f4784,f4090,f4081,f4077,f4075,f4065,f4064,f4061,f4060,f4059,f4058,f4057,f4056,f4053,f4051,f4049,f4034,f4033,f4019,f4015,f4013,f4003,f4002,f3999,f3998,f3997,f3996,f3995,f3994,f3991,f3989,f3987,f3933,f3932,f3850,f3841,f3837,f3836,f3826,f3825,f3822,f3821,f3820,f3819,f3818,f3817,f3814,f3812,f3810,f3791,f3790,f3780,f3776,f3775,f3765,f3764,f3761,f3760,f3759,f3758,f3757,f3756,f3753,f3751,f3749,f3699,f3620,f3615,f3610,f3603,f3602,f3557,f3552,f3543,f3541,f3518,f3497,f3494,f3490,f3488,f3340,f3338,f3090,f3088,f2513,f2297,f1802,f1741,f1587,f1540,f1368,f1366,f1355,f1353,f575,f538,f536,f514,f477,f475,f381,f380,f379,f378,f228,f227,f140,f18129,f18130,f18197,f18198,f18384,f18043,f18387,f18047,f18394,f18055,f18056,f18404,f18405,f18411,f18412,f18418,f18072,f18076,f18080,f18082,f18083,f18084,f18085,f18957,f18961,f18969,f18970,f18985,f18989,f19001,f19005,f19613,f19617,f19625,f19626,f19641,f19645,f19657,f19661,f20216,f20220,f20246,f20250,f20258,f20259,f21126,f21130,f21156,f21160,f21168,f21169,f21472,f21473,f21474,f21475,f21476,f21483,f21480,f21501,f21502,f21519,f21532])).
fof(f21532,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl23 | $spl26)),
  inference(subsumption_resolution,[],[f21520,f21519])).
fof(f21520,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl23 | $spl26)),
  inference(backward_demodulation,[],[f21519,f159])).
fof(f21502,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl23),
  inference(superposition,[],[f21480,f152])).
fof(f21518,plain,(
  $false | ($spl23 | $spl29)),
  inference(subsumption_resolution,[],[f21486,f21485])).
fof(f21486,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | $spl23),
  inference(unit_resulting_resolution,[],[f260,f21480,f293])).
fof(f21517,plain,(
  $false | ($spl23 | $spl29)),
  inference(subsumption_resolution,[],[f21489,f21485])).
fof(f21489,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | $spl23),
  inference(unit_resulting_resolution,[],[f260,f21480,f293])).
fof(f21516,plain,(
  $false | ($spl23 | $spl29)),
  inference(subsumption_resolution,[],[f21492,f21485])).
fof(f21492,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | $spl23),
  inference(unit_resulting_resolution,[],[f226,f21480,f322])).
fof(f21515,plain,(
  $false | ($spl23 | $spl29)),
  inference(subsumption_resolution,[],[f21495,f21485])).
fof(f21495,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | $spl23),
  inference(unit_resulting_resolution,[],[f21480,f152])).
fof(f21482,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl24),
  inference(cnf_transformation,[],[f21482_D])).
fof(f21482_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f21478,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK3(X0,X1,bool,bool))) ) | $spl20),
  inference(cnf_transformation,[],[f21478_D])).
fof(f21478_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK3(X0,X1,bool,bool))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f21485,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl29),
  inference(cnf_transformation,[],[f21485_D])).
fof(f21485_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f21476,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK3(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1854,f152])).
fof(f21483,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl26),
  inference(cnf_transformation,[],[f21483_D])).
fof(f21483_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f21475,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK3(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1854,f293])).
fof(f21474,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK3(X7,X8,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1854,f293])).
fof(f21477,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl18),
  inference(cnf_transformation,[],[f21477_D])).
fof(f21477_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f21480,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl23),
  inference(cnf_transformation,[],[f21480_D])).
fof(f21480_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f21473,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK3(X3,X4,bool,bool)) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1854,f293])).
fof(f21472,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK3(X0,X1,bool,bool)) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1854,f293])).
fof(f13250,plain,(
  $false | ($spl0 | $spl5 | $spl14)),
  inference(subsumption_resolution,[],[f13152,f11193])).
fof(f11193,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))) | ($spl0 | $spl5 | $spl14)),
  inference(backward_demodulation,[],[f11152,f406])).
fof(f11152,plain,(
  s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl14)),
  inference(backward_demodulation,[],[f11139,f1835])).
fof(f13152,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))) | ($spl0 | $spl14)),
  inference(backward_demodulation,[],[f11214,f12917])).
fof(f12917,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',sK0))))) | ($spl0 | $spl14)),
  inference(backward_demodulation,[],[f11185,f2754])).
fof(f11185,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',X0)))) ) | ($spl0 | $spl14)),
  inference(backward_demodulation,[],[f11152,f228])).
fof(f11214,plain,(
  ( ! [X3] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK0),s('type.num.num',X3)))) ) | ($spl0 | $spl14)),
  inference(backward_demodulation,[],[f11152,f1587])).
fof(f11139,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK2(sK0)) | $spl14),
  inference(cnf_transformation,[],[f11139_D])).
fof(f11139_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK2(sK0)) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f11140,plain,(
  ( ! [X2] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))) | s('type.num.num',X2) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(sK2(sK0)))))) ) | $spl16),
  inference(cnf_transformation,[],[f11140_D])).
fof(f11140_D,plain,(
  ( ! [X2] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))) | s('type.num.num',X2) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(sK2(sK0)))))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f10947,plain,(
  ( ! [X2] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))) | s('type.num.num',X2) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2(sK2(sK0))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK2(sK0))) ) | $spl0),
  inference(superposition,[],[f384,f1835])).
fof(f2758,plain,(
  s('type.num.num',sK2(sK0)) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))) | $spl10),
  inference(cnf_transformation,[],[f2758_D])).
fof(f2758_D,plain,(
  s('type.num.num',sK2(sK0)) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f2760,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) | $spl13),
  inference(cnf_transformation,[],[f2760_D])).
fof(f2760_D,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f2753,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) | s('type.num.num',sK2(sK0)) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))) | $spl0),
  inference(superposition,[],[f1861,f1642])).
fof(f2639,plain,(
  s('type.num.num',sK2(sK0)) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1))) | $spl6),
  inference(cnf_transformation,[],[f2639_D])).
fof(f2639_D,plain,(
  s('type.num.num',sK2(sK0)) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f2641,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))) | $spl9),
  inference(cnf_transformation,[],[f2641_D])).
fof(f2641_D,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f2617,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))) | s('type.num.num',sK2(sK0)) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',sK1))) | $spl0),
  inference(superposition,[],[f1861,f1804])).
fof(f257,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f253,f242])).
fof(f242,plain,(
  s('type.num.num',sK0) != s('type.num.num',sK1) | ($spl2 | $spl5)),
  inference(forward_demodulation,[],[f225,f222])).
fof(f253,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) | ($spl0 | $spl2)),
  inference(superposition,[],[f240,f243])).
fof(f243,plain,(
  s('type.num.num',sK0) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK1))) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f220,f222])).
fof(f240,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',X0)))) ) | $spl2),
  inference(backward_demodulation,[],[f239,f228])).
fof(f239,plain,(
  s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl2),
  inference(forward_demodulation,[],[f140,f222])).
fof(f256,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f252,f242])).
fof(f252,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) | ($spl0 | $spl2)),
  inference(superposition,[],[f243,f240])).
fof(f238,plain,(
  $false | ($spl1 | $spl2 | $spl4)),
  inference(subsumption_resolution,[],[f237,f233])).
fof(f233,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',sK0)))) ) | ($spl2 | $spl4)),
  inference(backward_demodulation,[],[f232,f231])).
fof(f231,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',sK1)))) ) | $spl4),
  inference(backward_demodulation,[],[f229,f227])).
fof(f229,plain,(
  s('type.num.num',sK1) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl4),
  inference(backward_demodulation,[],[f224,f140])).
fof(f232,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK1) | ($spl2 | $spl4)),
  inference(backward_demodulation,[],[f222,f224])).
fof(f237,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK0))) | ($spl1 | $spl2 | $spl4)),
  inference(forward_demodulation,[],[f236,f232])).
fof(f236,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK1))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f221,f222])).
fof(f220,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK1))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl0),
  inference(cnf_transformation,[],[f220_D])).
fof(f220_D,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK1))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f222,plain,(
  s('type.num.num',sK0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl2),
  inference(cnf_transformation,[],[f222_D])).
fof(f222_D,plain,(
  s('type.num.num',sK0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f135,plain,(
  s('type.num.num',sK0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK1))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f98])).
fof(f98,plain,(
  (s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK1))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | (s('type.num.num',sK0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) & s('type.num.num',sK1) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) & (s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK1))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s('type.num.num',sK0) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s('type.num.num',sK1) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f97])).
fof(f97,plain,(
  ? [X0,X1] : ((s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))) | (s('type.num.num',X0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) & s('type.num.num',X1) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) & (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))) | s('type.num.num',X0) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s('type.num.num',X1) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(flattening,[],[f96])).
fof(f96,plain,(
  ? [X0,X1] : ((s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))) | (s('type.num.num',X0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) & s('type.num.num',X1) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) & (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))) | (s('type.num.num',X0) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s('type.num.num',X1) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),
  inference(nnf_transformation,[],[f91])).
fof(f91,plain,(
  ? [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))) <~> (s('type.num.num',X0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) & s('type.num.num',X1) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(ennf_transformation,[],[f45])).
fof(f45,plain,(
  ~! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))) <=> (s('type.num.num',X0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) & s('type.num.num',X1) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X6),s('type.num.num',X7))) <=> (s('type.num.num',X6) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) & s('type.num.num',X7) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X6),s('type.num.num',X7))) <=> (s('type.num.num',X6) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) & s('type.num.num',X7) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  file('thm.arithmetic.MULT_EQ_|49|','thm.arithmetic.MULT_EQ_|49|')).
fof(f224,plain,(
  s('type.num.num',sK1) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl4),
  inference(cnf_transformation,[],[f224_D])).
fof(f224_D,plain,(
  s('type.num.num',sK1) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f136,plain,(
  s('type.num.num',sK1) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK1))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f98])).
fof(f221,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK1))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl1),
  inference(cnf_transformation,[],[f221_D])).
fof(f221_D,plain,(
  s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK1))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f223,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl3),
  inference(cnf_transformation,[],[f223_D])).
fof(f223_D,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f225,plain,(
  s('type.num.num',sK1) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl5),
  inference(cnf_transformation,[],[f225_D])).
fof(f225_D,plain,(
  s('type.num.num',sK1) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f137,plain,(
  s('type.num.num',sK1) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s('type.num.num',sK0) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num',sK1))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f98])).
% SZS output end Proof for thm.arithmetic
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 466
% Passive clauses: 9798
% Generated clauses: 74185
% Final active clauses: 61
% Final passive clauses: 179
% Input formulas: 43
% Initial clauses: 82
% Discarded non-redundant clauses: 1036
% 
% Duplicate literals: 53
% Trivial inequalities: 105
% Fw subsumption resolutions: 758
% Fw demodulations: 18393
% Bw demodulations: 5403
% Condensations: 181
% Global subsumptions: 2
% 
% Simple tautologies: 150
% Equational tautologies: 393
% Forward subsumptions: 43151
% Fw demodulations to eq. taut.: 1103
% Bw demodulations to eq. taut.: 41
% 
% Unit resulting resolution: 17878
% Factoring: 1
% Forward superposition: 13247
% Backward superposition: 12710
% Self superposition: 91
% Equality factoring: 26
% Equality resolution: 34
% 
% Split clauses: 41
% Split components: 88
% SAT solver clauses: 9941
% SAT solver unit clauses: 157
% SAT solver binary clauses: 5013
% 
% Sat splits: 41
% Sat splitting refutations: 57
% 
% Memory used [KB]: 37867
% Time elapsed: 4.141 s
% ------------------------------
% ------------------------------
% Success in time 5.469 s
