Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.real_sigma
% (114584)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 139
% Passive clauses: 1233
% Generated clauses: 16381
% Final active clauses: 130
% Final passive clauses: 956
% Input formulas: 27
% Initial clauses: 59
% 
% Duplicate literals: 899
% Trivial inequalities: 393
% Fw subsumption resolutions: 2721
% Fw demodulations: 1486
% Bw demodulations: 6
% 
% Simple tautologies: 449
% Equational tautologies: 54
% Forward subsumptions: 7196
% Fw demodulations to eq. taut.: 2
% 
% Binary resolution: 147
% Forward superposition: 4443
% Backward superposition: 5496
% Self superposition: 232
% Equality factoring: 171
% Equality resolution: 48
% 
% Split clauses: 1424
% Split components: 2962
% SAT solver clauses: 1999
% SAT solver unit clauses: 49
% SAT solver binary clauses: 271
% 
% Sat splits: 1424
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 1998
% TWLsolver calls for satisfiability: 139
% 
% Memory used [KB]: 8443
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.real_sigma
% (114597)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 132
% Passive clauses: 418
% Generated clauses: 13308
% Final active clauses: 130
% Final passive clauses: 261
% Input formulas: 27
% Initial clauses: 59
% 
% Duplicate literals: 685
% Trivial inequalities: 256
% Fw subsumption resolutions: 1596
% Fw demodulations: 1160
% Bw demodulations: 26
% Condensations: 1525
% Global subsumptions: 4
% 
% Simple tautologies: 303
% Equational tautologies: 24
% Forward subsumptions: 6326
% 
% Binary resolution: 27
% Forward superposition: 3754
% Backward superposition: 3579
% Self superposition: 114
% Equality factoring: 48
% Equality resolution: 59
% 
% Split clauses: 156
% Split components: 321
% SAT solver clauses: 870
% SAT solver unit clauses: 67
% SAT solver binary clauses: 29
% SAT solver learnt clauses: 8
% 
% Sat splits: 156
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 854
% TWLsolver calls for satisfiability: 307
% 
% Memory used [KB]: 10746
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 285 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.real_sigma
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.real_sigma
% SZS output start Proof for thm.real_sigma
fof(f1664,plain,(
  $false),
  inference(subsumption_resolution,[],[f1654,f100])).
fof(f100,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK0,'type.realax.real'),'_dst_x0_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(sK0,bool),sK1)))),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) & s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK0,'type.realax.real'),'_dst_x0_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(sK0,bool),sK1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f60])).
fof(f60,plain,(
  ? [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) & s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),'_dst_x0_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(X0,bool),X1))))),
  inference(ennf_transformation,[],[f30])).
fof(f30,plain,(
  ~! [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),'_dst_x0_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(X0,bool),X1))))),
  inference(flattening,[],[f29])).
fof(f29,plain,(
  ~! [X0] : ! [X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),'_dst_x0_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(X0,bool),X1))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X7)))) => s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X6,'type.realax.real'),'_dst_x0_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(X6,bool),X7))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X7)))) => s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X6,'type.realax.real'),'_dst_x0_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(X6,bool),X7))))),
  file('thm.real_sigma.REAL_SUM_IMAGE_|48|','thm.real_sigma.REAL_SUM_IMAGE_|48|')).
fof(f1654,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK0,'type.realax.real'),'_dst_x0_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s(fun(sK0,bool),sK1)))),
  inference(superposition,[],[f1638,f836])).
fof(f836,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(forward_demodulation,[],[f835,f772])).
fof(f772,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f771,f142])).
fof(f142,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X22,X23] : ! [X24,X25] : s(X23,X24) = s(X23,happ(s(fun(X22,X23),'_dst_x0_1'(s(X23,X24))),s(X22,X25)))),
  file('thm.real_sigma.REAL_SUM_IMAGE_|48|','_dst_x0_ax')).
fof(f771,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',happ(s(fun(X0,'type.num.num'),'_dst_x0_1'(s('type.num.num','const.num.|48|_0'))),s(X0,X1)))))) )),
  inference(forward_demodulation,[],[f766,f142])).
fof(f766,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',happ(s(fun(X0,'type.num.num'),'_dst_x0_1'(s('type.num.num','const.num.|48|_0'))),s(X0,X1))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',happ(s(fun(X2,'type.num.num'),'_dst_x0_1'(s('type.num.num','const.num.|48|_0'))),s(X2,X3)))))))) )),
  inference(unit_resulting_resolution,[],[f142,f142,f136])).
fof(f136,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0,X1] : ((s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1))) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0)))))))),
  inference(flattening,[],[f89])).
fof(f89,plain,(
  ! [X0,X1] : ((s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1))) & ((s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0)))))))),
  inference(nnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X19,X8] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X8))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X19))))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X19) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X8)))),
  file('thm.real_sigma.REAL_SUM_IMAGE_|48|','thm.real.eq_ints|split|1')).
fof(f835,plain,(
  ( ! [X0] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(forward_demodulation,[],[f831,f104])).
fof(f104,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X8] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X8),s('type.num.num','const.num.|48|_0')))),
  file('thm.real_sigma.REAL_SUM_IMAGE_|48|','thm.arithmetic.MULT_|48|')).
fof(f831,plain,(
  ( ! [X0] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f129,f772])).
fof(f129,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X0))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0)))))))) )),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1] : s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X0))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0)))))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X17,X18] : s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X18),s('type.num.num',X17))))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X18))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X17)))))))),
  file('thm.real_sigma.REAL_SUM_IMAGE_|48|','thm.real.mult_ints|split|2')).
fof(f1638,plain,(
  ( ! [X0] : (s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK0,'type.realax.real'),'_dst_x0_1'(s('type.realax.real',X0))),s(fun(sK0,bool),sK1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.realax.real',X0)))) )),
  inference(unit_resulting_resolution,[],[f142,f1534])).
fof(f1534,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X0) != s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),X1),s(sK0,sK2(X0,X1,sK1,sK0)))) | s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK0,'type.realax.real'),X1),s(fun(sK0,bool),sK1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.realax.real',X0)))) )),
  inference(subsumption_resolution,[],[f1524,f101])).
fof(f101,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f8])).
fof(f8,axiom,(
  p(s(bool,'T_0'))),
  file('thm.real_sigma.REAL_SUM_IMAGE_|48|','thm.bool.TRUTH')).
fof(f1524,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.realax.real',X0) != s('type.realax.real',happ(s(fun(sK0,'type.realax.real'),X1),s(sK0,sK2(X0,X1,sK1,sK0)))) | s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(sK0,'type.realax.real'),X1),s(fun(sK0,bool),sK1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))),s('type.realax.real',X0)))) )),
  inference(superposition,[],[f121,f183])).
fof(f183,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))),
  inference(subsumption_resolution,[],[f180,f103])).
fof(f103,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f31])).
fof(f31,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.real_sigma.REAL_SUM_IMAGE_|48|',aHL_FALSITY)).
fof(f180,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))),
  inference(superposition,[],[f99,f105])).
fof(f105,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.real_sigma.REAL_SUM_IMAGE_|48|',aHL_BOOL_CASES)).
fof(f99,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1))))),
  inference(cnf_transformation,[],[f72])).
fof(f121,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | s('type.realax.real',X3) != s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,sK2(X3,X2,X1,X0)))) | s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.realax.real',X3)))) )),
  inference(cnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | ! [X2,X3] : (s('type.realax.real',X3) != s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,sK2(X3,X2,X1,X0)))) | s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.realax.real',X3)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f66])).
fof(f66,plain,(
  ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | ! [X2,X3] : (? [X4] : s('type.realax.real',X3) != s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X4))) | s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.realax.real',X3)))))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2,X3] : (! [X4] : s('type.realax.real',X3) = s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X4))) => s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.realax.real',X3)))))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2,X3] : (! [X4] : s('type.realax.real',X3) = s('type.realax.real',happ(s(fun(X0,'type.realax.real'),X2),s(X0,X4))) => s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X0,'type.realax.real'),X2),s(fun(X0,bool),X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))),s('type.realax.real',X3)))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6] : ! [X20] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X20)))) => ! [X21,X10] : (! [X11] : s('type.realax.real',X10) = s('type.realax.real',happ(s(fun(X6,'type.realax.real'),X21),s(X6,X11))) => s('type.realax.real','const.real_sigma.REAL_SUM_IMAGE_2'(s(fun(X6,'type.realax.real'),X21),s(fun(X6,bool),X20))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),X20))))),s('type.realax.real',X10)))))),
  file('thm.real_sigma.REAL_SUM_IMAGE_|48|','thm.real_sigma.REAL_SUM_IMAGE_FINITE_CONST')).
% SZS output end Proof for thm.real_sigma
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 49
% Passive clauses: 99
% Generated clauses: 1522
% Final active clauses: 47
% Final passive clauses: 48
% Input formulas: 27
% Initial clauses: 59
% 
% Duplicate literals: 134
% Trivial inequalities: 29
% Fw subsumption resolutions: 62
% Fw demodulations: 130
% Bw demodulations: 4
% Condensations: 241
% 
% Simple tautologies: 67
% Equational tautologies: 31
% Forward subsumptions: 702
% Fw demodulations to eq. taut.: 18
% 
% Unit resulting resolution: 74
% Forward superposition: 394
% Backward superposition: 288
% Self superposition: 55
% Equality factoring: 40
% Equality resolution: 13
% 
% Split clauses: 1
% Split components: 2
% SAT solver clauses: 103
% SAT solver unit clauses: 21
% SAT solver binary clauses: 32
% 
% Sat splits: 1
% 
% Memory used [KB]: 1918
% Time elapsed: 0.101 s
% ------------------------------
% ------------------------------
% Success in time 1.518 s
