Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.integer
% (30637)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1068
% Passive clauses: 7486
% Generated clauses: 34235
% Final active clauses: 919
% Final passive clauses: 5947
% Input formulas: 261
% Initial clauses: 509
% 
% Duplicate literals: 246
% Trivial inequalities: 98
% Fw subsumption resolutions: 1788
% Fw demodulations: 4548
% Bw demodulations: 137
% 
% Simple tautologies: 341
% Equational tautologies: 30
% Forward subsumptions: 19135
% Fw demodulations to eq. taut.: 4
% 
% Binary resolution: 7073
% Factoring: 710
% Forward superposition: 10090
% Backward superposition: 8668
% Self superposition: 52
% Equality factoring: 2
% Equality resolution: 21
% 
% Split clauses: 718
% Split components: 1438
% SAT solver clauses: 886
% SAT solver unit clauses: 17
% SAT solver binary clauses: 68
% 
% Sat splits: 718
% Sat splitting refutations: 7
% 
% TWLsolver clauses: 883
% TWLsolver calls for satisfiability: 435
% 
% Memory used [KB]: 15095
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.integer
% (30651)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 644
% Passive clauses: 3130
% Generated clauses: 26577
% Final active clauses: 566
% Final passive clauses: 2194
% Input formulas: 261
% Initial clauses: 509
% 
% Duplicate literals: 97
% Trivial inequalities: 56
% Fw subsumption resolutions: 839
% Fw demodulations: 3914
% Bw demodulations: 335
% Condensations: 119
% Global subsumptions: 24
% 
% Simple tautologies: 315
% Equational tautologies: 11
% Forward subsumptions: 14244
% Fw demodulations to eq. taut.: 21
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 2599
% Factoring: 273
% Forward superposition: 8349
% Backward superposition: 6110
% Self superposition: 94
% Equality factoring: 2
% Equality resolution: 15
% 
% Split clauses: 20
% Split components: 40
% SAT solver clauses: 4298
% SAT solver unit clauses: 639
% SAT solver binary clauses: 157
% SAT solver learnt clauses: 49
% SAT solver learnt literals: 13
% 
% Sat splits: 20
% Sat splitting refutations: 26
% 
% TWLsolver clauses: 4201
% TWLsolver calls for satisfiability: 3251
% 
% Memory used [KB]: 23027
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.integer
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.integer
% SZS output start Proof for thm.integer
fof(f90904,plain,(
  $false),
  inference(subsumption_resolution,[],[f90903,f53302])).
fof(f53302,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int',sK39(sK12,'const.integer.int_|48|_0'))),
  inference(unit_resulting_resolution,[],[f53241,f1442])).
fof(f1442,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int',X0) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))) )),
  inference(forward_demodulation,[],[f1441,f864])).
fof(f864,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),
  inference(cnf_transformation,[],[f250])).
fof(f250,axiom,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),
  file('thm.integer.INT_MOD|48|','thm.integer.INT_|48|')).
fof(f1441,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f976,f864])).
fof(f976,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f707])).
fof(f707,plain,(
  ! [X0] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))),
  inference(nnf_transformation,[],[f357])).
fof(f357,plain,(
  ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) <=> s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X17] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X17))) <=> s('type.integer.int',X17) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.integer.INT_MOD|48|','thm.integer.INT_NEG_EQ|48|')).
fof(f53241,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK39(sK12,'const.integer.int_|48|_0'))))),
  inference(backward_demodulation,[],[f53237,f8428])).
fof(f8428,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))),
  inference(unit_resulting_resolution,[],[f1392,f1444])).
fof(f1444,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',X0)) )),
  inference(forward_demodulation,[],[f1443,f864])).
fof(f1443,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',X0) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))) )),
  inference(forward_demodulation,[],[f975,f864])).
fof(f975,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f707])).
fof(f1392,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))),
  inference(backward_demodulation,[],[f864,f859])).
fof(f859,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))),
  inference(cnf_transformation,[],[f690])).
fof(f690,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK12)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12])],[f599])).
fof(f599,plain,(
  ? [X0] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))),
  inference(ennf_transformation,[],[f264])).
fof(f264,plain,(
  ~! [X0] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))),
  inference(flattening,[],[f263])).
fof(f263,plain,(
  ~! [X0] : (~s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : (~s('type.integer.int',X6) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X6))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : (~s('type.integer.int',X6) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X6))))),
  file('thm.integer.INT_MOD|48|','thm.integer.INT_MOD|48|')).
fof(f53237,plain,(
  ( ! [X0] : (s('type.integer.int',sK39(sK12,X0)) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int',sK12)))) )),
  inference(forward_demodulation,[],[f52995,f50552])).
fof(f50552,plain,(
  ( ! [X0] : (s('type.integer.int',sK39(sK12,X0)) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X0),s('type.integer.int',sK12))),s('type.integer.int',sK12)))))) )),
  inference(unit_resulting_resolution,[],[f1391,f1504])).
fof(f1504,plain,(
  ( ! [X0,X1] : (s('type.integer.int',sK39(X1,X0)) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int',X1))))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',X1)) )),
  inference(forward_demodulation,[],[f1220,f864])).
fof(f1220,plain,(
  ( ! [X0,X1] : (s('type.integer.int',sK39(X1,X0)) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int',X1))))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f806])).
fof(f806,plain,(
  ! [X0,X1] : (s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',sK39(X1,X0)) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int',X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK39])],[f652])).
fof(f652,plain,(
  ! [X0,X1] : ? [X2] : (s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X2) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int',X1))))))),
  inference(ennf_transformation,[],[f532])).
fof(f532,plain,(
  ! [X0,X1] : ? [X2] : (s('type.integer.int',X1) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.integer.int',X2) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int',X1))))))),
  inference(flattening,[],[f531])).
fof(f531,plain,(
  ! [X0,X1] : ? [X2] : (~s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.integer.int',X2) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int',X1))))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X7,X8] : ? [X12] : (~s('type.integer.int',X8) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.integer.int',X12) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X7),s('type.integer.int',X8))),s('type.integer.int',X8))))))),
  file('thm.integer.INT_MOD|48|','thm.integer.integer_266_prove_ltb')).
fof(f1391,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int',sK12)),
  inference(backward_demodulation,[],[f864,f858])).
fof(f858,plain,(
  s('type.integer.int',sK12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),
  inference(cnf_transformation,[],[f690])).
fof(f52995,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int',sK12))) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X0),s('type.integer.int',sK12))),s('type.integer.int',sK12)))))) )),
  inference(unit_resulting_resolution,[],[f1391,f1464])).
fof(f1464,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int',X1))))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',X1)) )),
  inference(forward_demodulation,[],[f1125,f864])).
fof(f1125,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int',X1))))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f634])).
fof(f634,plain,(
  ! [X0,X1] : (s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int',X1))))))),
  inference(ennf_transformation,[],[f479])).
fof(f479,plain,(
  ! [X0,X1] : (s('type.integer.int',X1) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int',X1))))))),
  inference(flattening,[],[f478])).
fof(f478,plain,(
  ! [X0,X1] : (~s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int',X1))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7,X8] : (~s('type.integer.int',X8) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X7),s('type.integer.int',X8))) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_div_2'(s('type.integer.int',X7),s('type.integer.int',X8))),s('type.integer.int',X8))))))),
  file('thm.integer.INT_MOD|48|','thm.integer.int_mod')).
fof(f90903,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK39(sK12,'const.integer.int_|48|_0'))),
  inference(forward_demodulation,[],[f90902,f13778])).
fof(f13778,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_|48|_0')))) )),
  inference(forward_demodulation,[],[f13724,f1404])).
fof(f1404,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_|48|_0')))) )),
  inference(forward_demodulation,[],[f899,f864])).
fof(f899,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(cnf_transformation,[],[f294])).
fof(f294,plain,(
  ! [X0] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f69])).
fof(f69,axiom,(
  ! [X17] : s('type.integer.int',X17) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X17),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.integer.INT_MOD|48|','thm.integer.INT_ADD_RID')).
fof(f13724,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_|48|_0'))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_|48|_0')))) )),
  inference(superposition,[],[f1067,f1393])).
fof(f1393,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_|48|_0')))),
  inference(forward_demodulation,[],[f867,f864])).
fof(f867,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(cnf_transformation,[],[f144])).
fof(f144,axiom,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.integer.INT_MOD|48|','thm.integer.INT_NEG_|48|')).
fof(f1067,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) )),
  inference(cnf_transformation,[],[f418])).
fof(f418,plain,(
  ! [X0,X1] : s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X17,X20] : s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X17),s('type.integer.int',X20))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X17),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X20)))))),
  file('thm.integer.INT_MOD|48|','thm.integer.int_sub')).
fof(f90902,plain,(
  s('type.integer.int',sK39(sK12,'const.integer.int_|48|_0')) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))),
  inference(forward_demodulation,[],[f90901,f1409])).
fof(f1409,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',X0)))) )),
  inference(forward_demodulation,[],[f914,f864])).
fof(f914,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f309])).
fof(f309,plain,(
  ! [X0] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X17] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X17)))),
  file('thm.integer.INT_MOD|48|','thm.integer.INT_MUL_LZERO')).
fof(f90901,plain,(
  s('type.integer.int',sK39(sK12,'const.integer.int_|48|_0')) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))))),
  inference(subsumption_resolution,[],[f90685,f1391])).
fof(f90685,plain,(
  s('type.integer.int',sK39(sK12,'const.integer.int_|48|_0')) = s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12))))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK12)),
  inference(superposition,[],[f1504,f12269])).
fof(f12269,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_div_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK12)))),
  inference(forward_demodulation,[],[f12210,f889])).
fof(f889,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))) )),
  inference(cnf_transformation,[],[f284])).
fof(f284,plain,(
  ! [X0] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X17] : s('type.integer.int',X17) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X17)))))),
  file('thm.integer.INT_MOD|48|','thm.integer.INT_NEGNEG')).
fof(f12210,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_div_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))))),
  inference(unit_resulting_resolution,[],[f8451,f1427])).
fof(f1427,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_div_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',X0))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',X0)) )),
  inference(forward_demodulation,[],[f1426,f864])).
fof(f1426,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_div_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f950,f864])).
fof(f950,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_div_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f612])).
fof(f612,plain,(
  ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_div_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))),
  inference(ennf_transformation,[],[f344])).
fof(f344,plain,(
  ! [X0] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_div_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))),
  inference(flattening,[],[f343])).
fof(f343,plain,(
  ! [X0] : (~s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_div_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0))))),
  inference(rectify,[],[f231])).
fof(f231,axiom,(
  ! [X11] : (~s('type.integer.int',X11) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_div_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X11))))),
  file('thm.integer.INT_MOD|48|','thm.integer.INT_DIV_|48|')).
fof(f8451,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))))),
  inference(unit_resulting_resolution,[],[f8427,f1444])).
fof(f8427,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK12)))),
  inference(unit_resulting_resolution,[],[f1391,f1444])).
% SZS output end Proof for thm.integer
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1289
% Passive clauses: 12466
% Generated clauses: 93941
% Final active clauses: 1057
% Final passive clauses: 5121
% Input formulas: 261
% Initial clauses: 509
% 
% Duplicate literals: 19
% Trivial inequalities: 189
% Fw subsumption resolutions: 2834
% Fw demodulations: 18219
% Bw demodulations: 4330
% Condensations: 69
% 
% Simple tautologies: 135
% Equational tautologies: 108
% Forward subsumptions: 56385
% Fw demodulations to eq. taut.: 586
% Bw demodulations to eq. taut.: 60
% 
% Unit resulting resolution: 18512
% Factoring: 7
% Forward superposition: 21178
% Backward superposition: 23849
% Self superposition: 222
% Equality factoring: 2
% Equality resolution: 28
% 
% Split clauses: 22
% Split components: 44
% SAT solver clauses: 15233
% SAT solver unit clauses: 6793
% SAT solver binary clauses: 6286
% 
% Sat splits: 22
% Sat splitting refutations: 2721
% 
% Memory used [KB]: 50532
% Time elapsed: 3.145 s
% ------------------------------
% ------------------------------
% Success in time 4.51 s
