Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.pred_set
% (56343)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 111
% Passive clauses: 3969
% Generated clauses: 15793
% Final active clauses: 97
% Final passive clauses: 3667
% Input formulas: 37
% Initial clauses: 78
% 
% Duplicate literals: 579
% Trivial inequalities: 224
% Fw subsumption resolutions: 2221
% Fw demodulations: 1157
% Bw demodulations: 12
% 
% Simple tautologies: 349
% Equational tautologies: 82
% Forward subsumptions: 5958
% 
% Binary resolution: 412
% Factoring: 2
% Forward superposition: 5681
% Backward superposition: 5095
% Self superposition: 253
% Equality factoring: 194
% Equality resolution: 30
% 
% Split clauses: 419
% Split components: 848
% SAT solver clauses: 1235
% SAT solver unit clauses: 44
% SAT solver binary clauses: 453
% 
% Sat splits: 419
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1234
% TWLsolver calls for satisfiability: 170
% 
% Memory used [KB]: 9594
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (56362)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 147
% Passive clauses: 1396
% Generated clauses: 14372
% Final active clauses: 136
% Final passive clauses: 1151
% Input formulas: 37
% Initial clauses: 78
% 
% Duplicate literals: 464
% Trivial inequalities: 191
% Fw subsumption resolutions: 1202
% Fw demodulations: 2173
% Bw demodulations: 106
% Condensations: 598
% Global subsumptions: 12
% 
% Simple tautologies: 242
% Equational tautologies: 29
% Forward subsumptions: 5891
% Fw demodulations to eq. taut.: 4
% 
% Binary resolution: 627
% Factoring: 1
% Forward superposition: 3637
% Backward superposition: 3750
% Self superposition: 58
% Equality factoring: 42
% Equality resolution: 25
% 
% Split clauses: 151
% Split components: 320
% SAT solver clauses: 2274
% SAT solver unit clauses: 96
% SAT solver binary clauses: 51
% SAT solver learnt clauses: 24
% 
% Sat splits: 151
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 2226
% TWLsolver calls for satisfiability: 411
% 
% Memory used [KB]: 13432
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% (56381)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 657
% Passive clauses: 8025
% Generated clauses: 124651
% Final active clauses: 39
% Final passive clauses: 89
% Input formulas: 37
% Initial clauses: 78
% 
% Duplicate literals: 4018
% Trivial inequalities: 1011
% Fw subsumption resolutions: 8083
% Fw demodulations: 23034
% Bw demodulations: 5498
% Condensations: 7829
% Global subsumptions: 2785
% 
% Simple tautologies: 1123
% Equational tautologies: 515
% Forward subsumptions: 65041
% Fw demodulations to eq. taut.: 26
% Bw demodulations to eq. taut.: 260
% 
% Unit resulting resolution: 5554
% Factoring: 1
% Forward superposition: 24676
% Backward superposition: 26137
% Self superposition: 568
% Equality factoring: 307
% Equality resolution: 83
% 
% Split clauses: 192
% Split components: 394
% SAT solver clauses: 14203
% SAT solver unit clauses: 186
% SAT solver binary clauses: 1470
% 
% Sat splits: 192
% Sat splitting refutations: 3006
% 
% Memory used [KB]: 70105
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f31397,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f206,f207,f208,f209,f210,f211,f212,f213,f214,f215,f358,f379,f380,f384,f365,f366,f434,f437,f436,f454,f474,f476,f477,f2017,f2069,f2070,f2071,f2072,f2857,f2919,f2921,f2966,f2967,f2968,f3086,f3087,f3651,f3650,f3649,f3639,f3632,f3631,f3630,f3690,f3691,f3692,f3693,f3762,f3764,f3765,f3767,f3078,f3079,f3080,f3081,f3082,f3083,f3084,f3085,f3805,f3804,f3803,f3793,f3786,f3785,f3784,f3842,f3843,f3844,f3845,f3867,f3901,f3902,f3937,f3938,f3943,f3944,f3945,f3946,f3947,f3948,f3952,f4794,f4797,f4788,f4787,f4786,f3951,f3950,f3949,f4872,f4873,f4910,f4911,f382,f4916,f4917,f4918,f4919,f4922,f4923,f4924,f239,f241,f5934,f5988,f6184,f6185,f6178,f6177,f6176,f4799,f6242,f6243,f6244,f6245,f6246,f6776,f6778,f6770,f6769,f6768,f7784,f7839,f8239,f8241,f8231,f8230,f8229,f8714,f8716,f8706,f8705,f8704,f9548,f9602,f9986,f9988,f9987,f10422,f10476,f10855,f10909,f12232,f12286,f12471,f12473,f12472,f12957,f13013,f13024,f13023,f13084,f13085,f13087,f14008,f14013,f14012,f14010,f14110,f14111,f14512,f14572,f14513,f14574,f14514,f14576,f14570,f15639,f15641,f15633,f15632,f15631,f16395,f16447,f16396,f16451,f16785,f16787,f16777,f16776,f16775,f16895,f16897,f16898,f17097,f17127,f17128,f17129,f17130,f17131,f17132,f13153,f17133,f13152,f17134,f17135,f17137,f12843,f17139,f17138,f17140,f17141,f17136,f8762,f17142,f8758,f8757,f8499,f8497,f8495,f8493,f17143,f8463,f8462,f17144,f7414,f17163,f7410,f7409,f7176,f7174,f7172,f7170,f17164,f7128,f7127,f17165,f6948,f6946,f6944,f6942,f6911,f6910,f385,f17166,f6816,f17168,f6815,f6814,f386,f17167,f4925,f4921,f4920,f17235,f17236,f17272,f17273,f381,f17278,f17279,f234,f17280,f17281,f235,f232,f17282,f17283,f233,f17287,f237,f243,f18621,f18626,f17253,f473,f18015,f18001,f378,f17907,f435,f377,f19217,f20036,f20037,f4795,f20386,f20387,f6777,f20380,f20379,f20378,f20706,f20707,f15640,f20700,f4798,f20699,f20698,f475,f20826,f21235,f2911,f21605,f21606,f8240,f21597,f21596,f21595,f8243,f22018,f22019,f8715,f22010,f22009,f22008,f8717,f22635,f13012,f23055,f2912,f22686,f24549,f24603,f24734,f24736,f24735,f25613,f25667,f26128,f9603,f26182,f27060,f10477,f27114,f27367,f27369,f27368,f29058,f29059,f16786,f29050,f8242,f29049,f29048,f6779,f8244,f30622,f30624,f383,f30623,f30806,f30858,f30807,f16449,f30860,f30896,f30898,f4796,f30897,f30999,f31000,f31097,f31098,f31392])).
fof(f31392,plain,(
  $false | ($spl0 | $spl6 | $spl12)),
  inference(subsumption_resolution,[],[f31391,f198])).
fof(f198,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.SUBSET_INSERT_DELETE',aHL_TRUTH)).
fof(f31391,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl6 | $spl12)),
  inference(forward_demodulation,[],[f31390,f18237])).
fof(f18237,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)),sK2,sK0)),s(fun(sK0,bool),sK2))) | $spl6),
  inference(resolution,[],[f1173,f148])).
fof(f148,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X9] : (s(bool,'T_0') = s(bool,X9) <=> p(s(bool,X9)))),
  file('thm.pred_set.SUBSET_INSERT_DELETE','thm.bool.EQ_CLAUSES|split|1')).
fof(f1173,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)),sK2,sK0)),s(fun(sK0,bool),sK2)))) | $spl6),
  inference(subsumption_resolution,[],[f1079,f191])).
fof(f191,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0] : (~p(s(bool,'F_0')) | p(s(bool,X0)))),
  inference(ennf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : (p(s(bool,'F_0')) => p(s(bool,X0)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X9] : (p(s(bool,'F_0')) => p(s(bool,X9)))),
  file('thm.pred_set.SUBSET_INSERT_DELETE','thm.bool.FALSITY')).
fof(f1079,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)),sK2,sK0)),s(fun(sK0,bool),sK2)))) | $spl6),
  inference(superposition,[],[f158,f236])).
fof(f236,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | $spl6),
  inference(cnf_transformation,[],[f236_D])).
fof(f236_D,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f158,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))))) & ((p(s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(nnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(ennf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f64])).
fof(f64,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : ! [X8,X9] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X9)))) <=> ! [X7] : (p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X8)))) => p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X9))))))),
  file('thm.pred_set.SUBSET_INSERT_DELETE','thm.pred_set.SUBSET_DEF')).
fof(f31390,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)),sK2,sK0)),s(fun(sK0,bool),sK2)))) | ($spl0 | $spl12)),
  inference(subsumption_resolution,[],[f31389,f28549])).
fof(f28549,plain,(
  s(sK0,sK1) != s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)),sK2,sK0)) | $spl12),
  inference(resolution,[],[f17878,f135])).
fof(f135,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) & ((s(X0,X1) != s(X0,X2) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))))))),
  inference(flattening,[],[f90])).
fof(f90,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))))) & ((s(X0,X1) != s(X0,X2) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))))))),
  inference(nnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) <=> (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))))),
  inference(flattening,[],[f44])).
fof(f44,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) <=> (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6] : ! [X7,X10,X8] : (p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X10),s(fun(X6,bool),X8)))))) <=> (s(X6,X7) = s(X6,X10) | p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X8))))))),
  file('thm.pred_set.SUBSET_INSERT_DELETE','thm.pred_set.IN_INSERT')).
fof(f17878,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))))) | $spl12),
  inference(resolution,[],[f3935,f159])).
fof(f159,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f103])).
fof(f3935,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))))) | $spl12),
  inference(trivial_inequality_removal,[],[f3927])).
fof(f3927,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))))) | $spl12),
  inference(superposition,[],[f2917,f150])).
fof(f150,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X9] : (s(bool,'T_0') = s(bool,X9) <=> p(s(bool,X9)))),
  file('thm.pred_set.SUBSET_INSERT_DELETE','thm.bool.EQ_CLAUSES|split|0')).
fof(f2917,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | $spl12),
  inference(forward_demodulation,[],[f127,f243])).
fof(f127,plain,(
  s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))))),
  inference(cnf_transformation,[],[f88])).
fof(f88,plain,(
  s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f82])).
fof(f82,plain,(
  ? [X0,X1,X2,X3] : s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X3))))) != s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X2),s(X0,X1))),s(fun(X0,bool),X3)))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ~! [X0,X1,X2,X3] : s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X3))))) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X2),s(X0,X1))),s(fun(X0,bool),X3)))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ~! [X0] : ! [X1,X2,X3] : s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X3))))) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X2),s(X0,X1))),s(fun(X0,bool),X3)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X8),s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X7),s(fun(X6,bool),X9))))) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),'const.pred_set.DELETE_2'(s(fun(X6,bool),X8),s(X6,X7))),s(fun(X6,bool),X9)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X8),s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X7),s(fun(X6,bool),X9))))) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),'const.pred_set.DELETE_2'(s(fun(X6,bool),X8),s(X6,X7))),s(fun(X6,bool),X9)))),
  file('thm.pred_set.SUBSET_INSERT_DELETE','thm.pred_set.SUBSET_INSERT_DELETE')).
fof(f31389,plain,(
  s(sK0,sK1) = s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)),sK2,sK0)) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)),sK2,sK0)),s(fun(sK0,bool),sK2)))) | ($spl0 | $spl12)),
  inference(subsumption_resolution,[],[f31364,f192])).
fof(f192,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pred_set.SUBSET_INSERT_DELETE',aHL_FALSITY)).
fof(f31364,plain,(
  p(s(bool,'F_0')) | s(sK0,sK1) = s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)),sK2,sK0)) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)),sK2,sK0)),s(fun(sK0,bool),sK2)))) | ($spl0 | $spl12)),
  inference(superposition,[],[f31263,f28621])).
fof(f28621,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)),sK2,sK0)),s(fun(sK0,bool),sK3))) | $spl12),
  inference(resolution,[],[f28550,f144])).
fof(f144,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X9] : (s(bool,'F_0') = s(bool,X9) <=> ~p(s(bool,X9)))),
  file('thm.pred_set.SUBSET_INSERT_DELETE','thm.bool.EQ_CLAUSES|split|3')).
fof(f28550,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)),sK2,sK0)),s(fun(sK0,bool),sK3)))) | $spl12),
  inference(resolution,[],[f17878,f136])).
fof(f136,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f91])).
fof(f31263,plain,(
  ( ! [X3] : (p(s(bool,'const.bool.IN_2'(s(sK0,X3),s(fun(sK0,bool),sK3)))) | s(sK0,sK1) = s(sK0,X3) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X3),s(fun(sK0,bool),sK2))))) ) | $spl0),
  inference(resolution,[],[f17304,f139])).
fof(f139,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3)))))) | s(X0,X2) = s(X0,X3) | ~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) & s(X0,X2) != s(X0,X3))) & (~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,X3) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3))))))))),
  inference(flattening,[],[f92])).
fof(f92,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) & s(X0,X2) != s(X0,X3))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,X3)) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3))))))))),
  inference(nnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) & s(X0,X2) != s(X0,X3)))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) & ~s(X0,X2) = s(X0,X3)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X8,X7,X10] : (p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),'const.pred_set.DELETE_2'(s(fun(X6,bool),X8),s(X6,X10)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X8)))) & ~s(X6,X7) = s(X6,X10)))),
  file('thm.pred_set.SUBSET_INSERT_DELETE','thm.pred_set.IN_DELETE')).
fof(f17304,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))) | p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK3))))) ) | $spl0),
  inference(resolution,[],[f233,f157])).
fof(f157,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f103])).
fof(f31098,plain,(
  $false | ($spl6 | $spl48)),
  inference(subsumption_resolution,[],[f31094,f236])).
fof(f31094,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | $spl48),
  inference(equality_resolution,[],[f7839])).
fof(f31097,plain,(
  $false | ($spl6 | $spl12 | $spl16 | $spl48)),
  inference(subsumption_resolution,[],[f31096,f236])).
fof(f31096,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | ($spl12 | $spl16 | $spl48)),
  inference(forward_demodulation,[],[f31095,f24142])).
fof(f24142,plain,(
  s(sK0,sK1) = s(sK0,sK4(sK3,sK2,sK0)) | ($spl12 | $spl16)),
  inference(subsumption_resolution,[],[f24141,f199])).
fof(f199,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f141])).
fof(f141,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6] : ! [X7] : (s(X6,X7) = s(X6,X7) <=> p(s(bool,'T_0')))),
  file('thm.pred_set.SUBSET_INSERT_DELETE','thm.bool.REFL_CLAUSE')).
fof(f24141,plain,(
  ~p(s(bool,'T_0')) | s(sK0,sK1) = s(sK0,sK4(sK3,sK2,sK0)) | ($spl12 | $spl16)),
  inference(forward_demodulation,[],[f24082,f19304])).
fof(f19304,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),sK2))) | $spl16),
  inference(resolution,[],[f19232,f148])).
fof(f19232,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),sK2)))) | $spl16),
  inference(subsumption_resolution,[],[f19137,f191])).
fof(f19137,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),sK2)))) | $spl16),
  inference(superposition,[],[f158,f378])).
fof(f24082,plain,(
  s(sK0,sK1) = s(sK0,sK4(sK3,sK2,sK0)) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),sK2)))) | ($spl12 | $spl16)),
  inference(resolution,[],[f24077,f139])).
fof(f24077,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))) | ($spl12 | $spl16)),
  inference(subsumption_resolution,[],[f24058,f199])).
fof(f24058,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))) | ($spl12 | $spl16)),
  inference(superposition,[],[f19338,f243])).
fof(f19338,plain,(
  ( ! [X0] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),sK3)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),X0))))) ) | $spl16),
  inference(resolution,[],[f19233,f157])).
fof(f19233,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),sK3)))) | $spl16),
  inference(subsumption_resolution,[],[f19138,f192])).
fof(f19138,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),sK3)))) | $spl16),
  inference(superposition,[],[f159,f378])).
fof(f31095,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),sK3))))) | ($spl12 | $spl16 | $spl48)),
  inference(trivial_inequality_removal,[],[f31089])).
fof(f31089,plain,(
  s(sK0,sK1) != s(sK0,sK1) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),sK3))))) | ($spl12 | $spl16 | $spl48)),
  inference(superposition,[],[f7839,f24142])).
fof(f31000,plain,(
  $false | ($spl12 | $spl42)),
  inference(subsumption_resolution,[],[f30996,f243])).
fof(f30996,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))) | $spl42),
  inference(equality_resolution,[],[f5988])).
fof(f30999,plain,(
  $false | ($spl12 | $spl16 | $spl42)),
  inference(subsumption_resolution,[],[f30998,f243])).
fof(f30998,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))) | ($spl12 | $spl16 | $spl42)),
  inference(forward_demodulation,[],[f30997,f24142])).
fof(f30997,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK4(sK3,sK2,sK0)))),s(fun(sK0,bool),sK3))) | ($spl12 | $spl16 | $spl42)),
  inference(trivial_inequality_removal,[],[f30991])).
fof(f30991,plain,(
  s(sK0,sK1) != s(sK0,sK1) | s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK4(sK3,sK2,sK0)))),s(fun(sK0,bool),sK3))) | ($spl12 | $spl16 | $spl42)),
  inference(superposition,[],[f5988,f24142])).
fof(f30897,plain,(
  ( ! [X17,X15] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))),s(fun(fun(sK0,bool),bool),X17)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X15)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),X15),s(fun(fun(sK0,bool),bool),X17))))))) ) | $spl150),
  inference(cnf_transformation,[],[f30897_D])).
fof(f30897_D,plain,(
  ( ! [X17,X15] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))),s(fun(fun(sK0,bool),bool),X17)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X15)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),X15),s(fun(fun(sK0,bool),bool),X17))))))) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f4796,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,X7)) ) | $spl36),
  inference(cnf_transformation,[],[f4796_D])).
fof(f4796_D,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,X7)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f30898,plain,(
  ( ! [X19,X20,X18] : (s(bool,'T_0') != s(bool,X19) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))),s(fun(fun(sK0,bool),bool),X20)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),X18),s(fun(fun(sK0,bool),bool),X20)))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X18))))) ) | $spl12),
  inference(subsumption_resolution,[],[f30881,f147])).
fof(f147,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f97])).
fof(f30881,plain,(
  ( ! [X19,X20,X18] : (s(bool,'T_0') != s(bool,X19) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))),s(fun(fun(sK0,bool),bool),X20)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),X18),s(fun(fun(sK0,bool),bool),X20)))))) | ~p(s(bool,X19)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X18))))) ) | $spl12),
  inference(superposition,[],[f17252,f131])).
fof(f131,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f83])).
fof(f83,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X11,X12] : ((p(s(bool,X11)) => p(s(bool,X12))) => ((p(s(bool,X12)) => p(s(bool,X11))) => s(bool,X11) = s(bool,X12)))),
  file('thm.pred_set.SUBSET_INSERT_DELETE','thm.bool.IMP_ANTISYM_AX')).
fof(f17252,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X6))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))),s(fun(fun(sK0,bool),bool),X7)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),X6),s(fun(fun(sK0,bool),bool),X7))))))) ) | $spl12),
  inference(superposition,[],[f2917,f134])).
fof(f134,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))) | s(X0,X1) = s(X0,X2) | ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))))) )),
  inference(cnf_transformation,[],[f91])).
fof(f30896,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') != s(bool,X16) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))),s(fun(fun(sK0,bool),bool),X17)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),X15),s(fun(fun(sK0,bool),bool),X17)))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X15))))) ) | $spl12),
  inference(subsumption_resolution,[],[f30880,f147])).
fof(f30880,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') != s(bool,X16) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))),s(fun(fun(sK0,bool),bool),X17)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),X15),s(fun(fun(sK0,bool),bool),X17)))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X15)))) | ~p(s(bool,X16))) ) | $spl12),
  inference(superposition,[],[f17252,f131])).
fof(f30860,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK4(sK3,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)))),s(fun(sK0,bool),sK3))) | $spl149),
  inference(cnf_transformation,[],[f30860_D])).
fof(f30860_D,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK4(sK3,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)))),s(fun(sK0,bool),sK3))) <=> ~$spl149),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl149])])).
fof(f16449,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | $spl98),
  inference(cnf_transformation,[],[f16449_D])).
fof(f16449_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f30807,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK4(sK3,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)))),s(fun(sK0,bool),sK3))) | ($spl6 | $spl12 | $spl16 | $spl82)),
  inference(resolution,[],[f26524,f17208])).
fof(f17208,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X4),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),X5)))))) | p(s(bool,'const.bool.IN_2'(s(sK0,X4),s(fun(sK0,bool),X5)))) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,X4))),s(fun(sK0,bool),sK3)))) ) | $spl6),
  inference(superposition,[],[f1047,f134])).
fof(f1047,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))) | $spl6),
  inference(backward_demodulation,[],[f236,f127])).
fof(f26524,plain,(
  ( ! [X5] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,'const.pred_set.INSERT_2'(s(sK0,X5),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X5),s(fun(sK0,bool),sK2))))))) ) | ($spl12 | $spl16 | $spl82)),
  inference(subsumption_resolution,[],[f26422,f191])).
fof(f26422,plain,(
  ( ! [X5] : (p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,'const.pred_set.INSERT_2'(s(sK0,X5),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X5),s(fun(sK0,bool),sK2))))))) ) | ($spl12 | $spl16 | $spl82)),
  inference(superposition,[],[f158,f26389])).
fof(f26389,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)))) ) | ($spl12 | $spl16 | $spl82)),
  inference(subsumption_resolution,[],[f26360,f199])).
fof(f26360,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)))) ) | ($spl12 | $spl16 | $spl82)),
  inference(superposition,[],[f25677,f19882])).
fof(f19882,plain,(
  ( ! [X15] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X15),s(fun(sK0,bool),sK2)))))) ) | $spl82),
  inference(resolution,[],[f18820,f148])).
fof(f18820,plain,(
  ( ! [X3] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X3),s(fun(sK0,bool),sK2))))))) ) | $spl82),
  inference(resolution,[],[f14400,f136])).
fof(f14400,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)))) | $spl82),
  inference(resolution,[],[f14010,f137])).
fof(f137,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3))))))) )),
  inference(cnf_transformation,[],[f93])).
fof(f25677,plain,(
  ( ! [X9] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),X9)))) | s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X9),s(fun(sK0,bool),sK3)))) ) | ($spl12 | $spl16)),
  inference(resolution,[],[f24150,f144])).
fof(f24150,plain,(
  ( ! [X0] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X0),s(fun(sK0,bool),sK3)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),X0))))) ) | ($spl12 | $spl16)),
  inference(backward_demodulation,[],[f24142,f19338])).
fof(f30858,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK3,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK3))))) | $spl147),
  inference(cnf_transformation,[],[f30858_D])).
fof(f30858_D,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK3,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK3))))) <=> ~$spl147),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl147])])).
fof(f30806,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK3,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK3))))) | ($spl12 | $spl16 | $spl82)),
  inference(resolution,[],[f26524,f17243])).
fof(f17243,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X4),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),X5)))))) | p(s(bool,'const.bool.IN_2'(s(sK0,X4),s(fun(sK0,bool),X5)))) | s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X4),s(fun(sK0,bool),sK3)))))) ) | $spl12),
  inference(superposition,[],[f2917,f134])).
fof(f30623,plain,(
  ( ! [X6,X4] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),X6)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X4),s(fun(sK0,bool),sK3)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),X4),s(fun(fun(sK0,bool),bool),X6))))))) ) | $spl144),
  inference(cnf_transformation,[],[f30623_D])).
fof(f30623_D,plain,(
  ( ! [X6,X4] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),X6)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X4),s(fun(sK0,bool),sK3)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),X4),s(fun(fun(sK0,bool),bool),X6))))))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f383,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl20),
  inference(cnf_transformation,[],[f383_D])).
fof(f383_D,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f30624,plain,(
  ( ! [X8,X7,X9] : (s(bool,'F_0') != s(bool,X8) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),X9)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),X7),s(fun(fun(sK0,bool),bool),X9)))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X7),s(fun(sK0,bool),sK3))))) ) | $spl6),
  inference(subsumption_resolution,[],[f30603,f143])).
fof(f143,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f95])).
fof(f30603,plain,(
  ( ! [X8,X7,X9] : (s(bool,'F_0') != s(bool,X8) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),X9)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),X7),s(fun(fun(sK0,bool),bool),X9)))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X7),s(fun(sK0,bool),sK3)))) | p(s(bool,X8))) ) | $spl6),
  inference(superposition,[],[f17213,f128])).
fof(f128,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f84])).
fof(f17213,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X6),s(fun(sK0,bool),sK3))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),X7)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),X6),s(fun(fun(sK0,bool),bool),X7))))))) ) | $spl6),
  inference(superposition,[],[f1047,f134])).
fof(f30622,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') != s(bool,X5) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),X6)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),X4),s(fun(fun(sK0,bool),bool),X6)))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X4),s(fun(sK0,bool),sK3))))) ) | $spl6),
  inference(subsumption_resolution,[],[f30602,f143])).
fof(f30602,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') != s(bool,X5) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),X6)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),X4),s(fun(fun(sK0,bool),bool),X6)))))) | p(s(bool,X5)) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X4),s(fun(sK0,bool),sK3))))) ) | $spl6),
  inference(superposition,[],[f17213,f128])).
fof(f8244,plain,(
  ( ! [X29] : (p(s(bool,X29))) ) | $spl56),
  inference(cnf_transformation,[],[f8244_D])).
fof(f8244_D,plain,(
  ( ! [X29] : (p(s(bool,X29))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f6779,plain,(
  ( ! [X25] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X25)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X25))))) ) | $spl46),
  inference(cnf_transformation,[],[f6779_D])).
fof(f6779_D,plain,(
  ( ! [X25] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X25)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X25))))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f29048,plain,(
  ( ! [X31,X32] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X32)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X32)))))) | p(s(bool,X31))) ) | ($spl6 | $spl12)),
  inference(trivial_inequality_removal,[],[f29005])).
fof(f29005,plain,(
  ( ! [X31,X32] : (s(bool,X31) != s(bool,X31) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X32)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X32)))))) | p(s(bool,X31))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f28973,f144])).
fof(f28973,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') != s(bool,X10) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X11)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X10),s(fun(bool,bool),X11))))))) ) | ($spl6 | $spl12)),
  inference(forward_demodulation,[],[f28972,f243])).
fof(f28972,plain,(
  ( ! [X10,X11] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X11)))) | s(bool,'F_0') != s(bool,X10) | ~p(s(bool,'const.bool.IN_2'(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X10),s(fun(bool,bool),X11))))))) ) | ($spl6 | $spl12)),
  inference(forward_demodulation,[],[f17231,f243])).
fof(f17231,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') != s(bool,X10) | p(s(bool,'const.bool.IN_2'(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))),s(fun(bool,bool),X11)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X10),s(fun(bool,bool),X11))))))) ) | $spl6),
  inference(superposition,[],[f1047,f134])).
fof(f29049,plain,(
  ( ! [X30,X29] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X30)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X30)))))) | p(s(bool,X29))) ) | ($spl6 | $spl12)),
  inference(trivial_inequality_removal,[],[f29004])).
fof(f29004,plain,(
  ( ! [X30,X29] : (s(bool,X29) != s(bool,X29) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X30)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X30)))))) | p(s(bool,X29))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f28973,f146])).
fof(f146,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X9] : (s(bool,'F_0') = s(bool,X9) <=> ~p(s(bool,X9)))),
  file('thm.pred_set.SUBSET_INSERT_DELETE','thm.bool.EQ_CLAUSES|split|2')).
fof(f8242,plain,(
  ( ! [X27] : (s(bool,'T_0') = s(bool,X27)) ) | $spl52),
  inference(cnf_transformation,[],[f8242_D])).
fof(f8242_D,plain,(
  ( ! [X27] : (s(bool,'T_0') = s(bool,X27)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f29050,plain,(
  ( ! [X28,X27] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X28)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X28)))))) | s(bool,'T_0') = s(bool,X27)) ) | ($spl6 | $spl12)),
  inference(trivial_inequality_removal,[],[f29003])).
fof(f29003,plain,(
  ( ! [X28,X27] : (s(bool,X27) != s(bool,X27) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X28)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X28)))))) | s(bool,'T_0') = s(bool,X27)) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f28973,f151])).
fof(f151,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.SUBSET_INSERT_DELETE',aHL_BOOL_CASES)).
fof(f16786,plain,(
  ( ! [X2,X0] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X2)))) | p(s(bool,X0)) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X2))))))) ) | $spl102),
  inference(cnf_transformation,[],[f16786_D])).
fof(f16786_D,plain,(
  ( ! [X2,X0] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X2)))) | p(s(bool,X0)) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X2))))))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f29059,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X4) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X5)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X3),s(fun(bool,bool),X5)))))) | p(s(bool,X3))) ) | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f28994,f143])).
fof(f28994,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X4) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X5)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X3),s(fun(bool,bool),X5)))))) | p(s(bool,X3)) | p(s(bool,X4))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f28973,f128])).
fof(f29058,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X2)))))) | p(s(bool,X0))) ) | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f28993,f143])).
fof(f28993,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X2)))))) | p(s(bool,X1)) | p(s(bool,X0))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f28973,f128])).
fof(f27368,plain,(
  ( ! [X17,X15] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X15),s(fun(fun(sK0,bool),bool),X17)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X15)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X15),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))),s(fun(fun(sK0,bool),bool),X17))))))) ) | $spl142),
  inference(cnf_transformation,[],[f27368_D])).
fof(f27368_D,plain,(
  ( ! [X17,X15] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X15),s(fun(fun(sK0,bool),bool),X17)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X15)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X15),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))),s(fun(fun(sK0,bool),bool),X17))))))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f27369,plain,(
  ( ! [X19,X20,X18] : (s(bool,'T_0') != s(bool,X19) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X18),s(fun(fun(sK0,bool),bool),X20)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X18),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))),s(fun(fun(sK0,bool),bool),X20)))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X18))))) ) | $spl12),
  inference(subsumption_resolution,[],[f27351,f147])).
fof(f27351,plain,(
  ( ! [X19,X20,X18] : (s(bool,'T_0') != s(bool,X19) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X18),s(fun(fun(sK0,bool),bool),X20)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X18),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))),s(fun(fun(sK0,bool),bool),X20)))))) | ~p(s(bool,X19)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X18))))) ) | $spl12),
  inference(superposition,[],[f17251,f131])).
fof(f17251,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X4))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X4),s(fun(fun(sK0,bool),bool),X5)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X4),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))),s(fun(fun(sK0,bool),bool),X5))))))) ) | $spl12),
  inference(superposition,[],[f2917,f134])).
fof(f27367,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') != s(bool,X16) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X15),s(fun(fun(sK0,bool),bool),X17)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X15),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))),s(fun(fun(sK0,bool),bool),X17)))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X15))))) ) | $spl12),
  inference(subsumption_resolution,[],[f27350,f147])).
fof(f27350,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') != s(bool,X16) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X15),s(fun(fun(sK0,bool),bool),X17)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X15),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))),s(fun(fun(sK0,bool),bool),X17)))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X15)))) | ~p(s(bool,X16))) ) | $spl12),
  inference(superposition,[],[f17251,f131])).
fof(f27114,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),X1))))) | s(fun(sK0,bool),sK3) != s(fun(sK0,bool),X1)) ) | $spl140),
  inference(cnf_transformation,[],[f27114_D])).
fof(f27114_D,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),X1))))) | s(fun(sK0,bool),sK3) != s(fun(sK0,bool),X1)) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f10477,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK3),s(fun(fun(sK0,bool),bool),X0))))) ) | $spl70),
  inference(cnf_transformation,[],[f10477_D])).
fof(f10477_D,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK3),s(fun(fun(sK0,bool),bool),X0))))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f27060,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK3),s(fun(fun(sK0,bool),bool),X0)))) | s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),X1))))) | s(fun(sK0,bool),sK3) != s(fun(sK0,bool),X1)) ) | $spl12),
  inference(resolution,[],[f17248,f135])).
fof(f17248,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK3),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),X6),s(fun(fun(sK0,bool),bool),X7)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK3),s(fun(fun(sK0,bool),bool),X7)))) | s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),X6)))))) ) | $spl12),
  inference(superposition,[],[f2917,f134])).
fof(f26182,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | s(fun(sK0,bool),sK2) != s(fun(sK0,bool),X1)) ) | $spl138),
  inference(cnf_transformation,[],[f26182_D])).
fof(f26182_D,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | s(fun(sK0,bool),sK2) != s(fun(sK0,bool),X1)) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f9603,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),X0))))) ) | $spl64),
  inference(cnf_transformation,[],[f9603_D])).
fof(f9603_D,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),X0))))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f26128,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),X0)))) | s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | s(fun(sK0,bool),sK2) != s(fun(sK0,bool),X1)) ) | $spl12),
  inference(resolution,[],[f17240,f135])).
fof(f17240,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),X6),s(fun(fun(sK0,bool),bool),X7)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),X7)))) | s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X6),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))))) ) | $spl12),
  inference(superposition,[],[f2917,f134])).
fof(f25667,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),X1))) | s(fun(sK0,bool),sK3) != s(fun(sK0,bool),X1)) ) | $spl136),
  inference(cnf_transformation,[],[f25667_D])).
fof(f25667_D,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),X1))) | s(fun(sK0,bool),sK3) != s(fun(sK0,bool),X1)) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f25613,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK3),s(fun(fun(sK0,bool),bool),X0)))) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),X1))) | s(fun(sK0,bool),sK3) != s(fun(sK0,bool),X1)) ) | $spl6),
  inference(resolution,[],[f17217,f135])).
fof(f17217,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK3),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),X6),s(fun(fun(sK0,bool),bool),X7)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK3),s(fun(fun(sK0,bool),bool),X7)))) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),X6)))) ) | $spl6),
  inference(superposition,[],[f1047,f134])).
fof(f24735,plain,(
  ( ! [X4,X2] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X2),s(fun(fun(sK0,bool),bool),X4)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X2),s(fun(sK0,bool),sK3)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X2),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),X4))))))) ) | $spl134),
  inference(cnf_transformation,[],[f24735_D])).
fof(f24735_D,plain,(
  ( ! [X4,X2] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X2),s(fun(fun(sK0,bool),bool),X4)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X2),s(fun(sK0,bool),sK3)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X2),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),X4))))))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f24736,plain,(
  ( ! [X6,X7,X5] : (s(bool,'F_0') != s(bool,X6) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X5),s(fun(fun(sK0,bool),bool),X7)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X5),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),X7)))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X5),s(fun(sK0,bool),sK3))))) ) | $spl6),
  inference(subsumption_resolution,[],[f24715,f143])).
fof(f24715,plain,(
  ( ! [X6,X7,X5] : (s(bool,'F_0') != s(bool,X6) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X5),s(fun(fun(sK0,bool),bool),X7)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X5),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),X7)))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X5),s(fun(sK0,bool),sK3)))) | p(s(bool,X6))) ) | $spl6),
  inference(superposition,[],[f17212,f128])).
fof(f17212,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X4),s(fun(sK0,bool),sK3))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X4),s(fun(fun(sK0,bool),bool),X5)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X4),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),X5))))))) ) | $spl6),
  inference(superposition,[],[f1047,f134])).
fof(f24734,plain,(
  ( ! [X4,X2,X3] : (s(bool,'F_0') != s(bool,X3) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X2),s(fun(fun(sK0,bool),bool),X4)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X2),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),X4)))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X2),s(fun(sK0,bool),sK3))))) ) | $spl6),
  inference(subsumption_resolution,[],[f24714,f143])).
fof(f24714,plain,(
  ( ! [X4,X2,X3] : (s(bool,'F_0') != s(bool,X3) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X2),s(fun(fun(sK0,bool),bool),X4)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X2),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),X4)))))) | p(s(bool,X3)) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X2),s(fun(sK0,bool),sK3))))) ) | $spl6),
  inference(superposition,[],[f17212,f128])).
fof(f24603,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X1),s(sK0,sK1))),s(fun(sK0,bool),sK3))) | s(fun(sK0,bool),sK2) != s(fun(sK0,bool),X1)) ) | $spl132),
  inference(cnf_transformation,[],[f24603_D])).
fof(f24603_D,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X1),s(sK0,sK1))),s(fun(sK0,bool),sK3))) | s(fun(sK0,bool),sK2) != s(fun(sK0,bool),X1)) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f24549,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),X0)))) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X1),s(sK0,sK1))),s(fun(sK0,bool),sK3))) | s(fun(sK0,bool),sK2) != s(fun(sK0,bool),X1)) ) | $spl6),
  inference(resolution,[],[f17205,f135])).
fof(f17205,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),X6),s(fun(fun(sK0,bool),bool),X7)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),X7)))) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X6),s(sK0,sK1))),s(fun(sK0,bool),sK3)))) ) | $spl6),
  inference(superposition,[],[f1047,f134])).
fof(f22686,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X2),s(fun(sK0,bool),sK3))))) | s(sK0,sK1) != s(sK0,X2)) ) | $spl130),
  inference(cnf_transformation,[],[f22686_D])).
fof(f22686_D,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X2),s(fun(sK0,bool),sK3))))) | s(sK0,sK1) != s(sK0,X2)) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f2912,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),X0))))) ) | $spl32),
  inference(cnf_transformation,[],[f2912_D])).
fof(f2912_D,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),X0))))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f23055,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),X0)))) | s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X1),s(fun(sK0,bool),sK3))))) | s(sK0,sK1) != s(sK0,X1)) ) | $spl12),
  inference(resolution,[],[f17244,f135])).
fof(f17244,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X6),s(fun(sK0,bool),X7)))))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),X7)))) | s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X6),s(fun(sK0,bool),sK3)))))) ) | $spl12),
  inference(superposition,[],[f2917,f134])).
fof(f13012,plain,(
  ( ! [X13] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),X13))))) ) | $spl78),
  inference(cnf_transformation,[],[f13012_D])).
fof(f13012_D,plain,(
  ( ! [X13] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),X13))))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f22635,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X2),s(fun(sK0,bool),sK3))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),X3)))) | s(sK0,sK1) != s(sK0,X2)) ) | $spl12),
  inference(resolution,[],[f17242,f138])).
fof(f138,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) != s(X0,X3) | ~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3))))))) )),
  inference(cnf_transformation,[],[f93])).
fof(f17242,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X3),s(sK0,X2)))))) | s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X2),s(fun(sK0,bool),sK3))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),X3))))) ) | $spl12),
  inference(superposition,[],[f2917,f139])).
fof(f8717,plain,(
  ( ! [X28] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X28)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X28))))) ) | $spl60),
  inference(cnf_transformation,[],[f8717_D])).
fof(f8717_D,plain,(
  ( ! [X28] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X28)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X28))))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f22008,plain,(
  ( ! [X31,X32] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X32)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X32)))) | p(s(bool,X31))) ) | ($spl6 | $spl12)),
  inference(trivial_inequality_removal,[],[f21979])).
fof(f21979,plain,(
  ( ! [X31,X32] : (s(bool,X31) != s(bool,X31) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X32)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X32)))) | p(s(bool,X31))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f21867,f144])).
fof(f21867,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,X8) | ~p(s(bool,'const.bool.IN_2'(s(bool,X8),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X9)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X8),s(fun(bool,bool),X9))))) ) | ($spl6 | $spl12)),
  inference(forward_demodulation,[],[f17230,f243])).
fof(f17230,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,X8) | p(s(bool,'const.bool.IN_2'(s(bool,X8),s(fun(bool,bool),X9)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X8),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))),s(fun(bool,bool),X9))))))) ) | $spl6),
  inference(superposition,[],[f1047,f134])).
fof(f22009,plain,(
  ( ! [X30,X29] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X30)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X30)))) | p(s(bool,X29))) ) | ($spl6 | $spl12)),
  inference(trivial_inequality_removal,[],[f21978])).
fof(f21978,plain,(
  ( ! [X30,X29] : (s(bool,X29) != s(bool,X29) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X30)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X30)))) | p(s(bool,X29))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f21867,f146])).
fof(f22010,plain,(
  ( ! [X28,X27] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X28)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X28)))) | s(bool,'T_0') = s(bool,X27)) ) | ($spl6 | $spl12)),
  inference(trivial_inequality_removal,[],[f21977])).
fof(f21977,plain,(
  ( ! [X28,X27] : (s(bool,X27) != s(bool,X27) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X28)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X28)))) | s(bool,'T_0') = s(bool,X27)) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f21867,f151])).
fof(f8715,plain,(
  ( ! [X2,X0] : (~p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X2)))))) | p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X2))))) ) | $spl58),
  inference(cnf_transformation,[],[f8715_D])).
fof(f8715_D,plain,(
  ( ! [X2,X0] : (~p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X2)))))) | p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X2))))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f22019,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X4) | ~p(s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X5)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),X5)))) | p(s(bool,X3))) ) | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f21968,f143])).
fof(f21968,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X4) | ~p(s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X5)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),X5)))) | p(s(bool,X3)) | p(s(bool,X4))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f21867,f128])).
fof(f22018,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,X1) | ~p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X2)))) | p(s(bool,X0))) ) | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f21967,f143])).
fof(f21967,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,X1) | ~p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X2)))) | p(s(bool,X1)) | p(s(bool,X0))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f21867,f128])).
fof(f8243,plain,(
  ( ! [X28] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X28)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X28),s(bool,'T_0'))))))) ) | $spl54),
  inference(cnf_transformation,[],[f8243_D])).
fof(f8243_D,plain,(
  ( ! [X28] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X28)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X28),s(bool,'T_0'))))))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f21595,plain,(
  ( ! [X31,X32] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X32),s(bool,'T_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X32)))) | p(s(bool,X31))) ) | ($spl6 | $spl12)),
  inference(trivial_inequality_removal,[],[f21567])).
fof(f21567,plain,(
  ( ! [X31,X32] : (s(bool,X31) != s(bool,X31) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X32),s(bool,'T_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X32)))) | p(s(bool,X31))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f21489,f144])).
fof(f21489,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X4) | p(s(bool,'const.bool.IN_2'(s(bool,X4),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X5),s(bool,'T_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X4),s(fun(bool,bool),X5))))) ) | ($spl6 | $spl12)),
  inference(forward_demodulation,[],[f17228,f243])).
fof(f17228,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X4) | p(s(bool,'const.bool.IN_2'(s(bool,X4),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X5),s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3)))))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X4),s(fun(bool,bool),X5))))) ) | $spl6),
  inference(superposition,[],[f1047,f139])).
fof(f21596,plain,(
  ( ! [X30,X29] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X30),s(bool,'T_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X30)))) | p(s(bool,X29))) ) | ($spl6 | $spl12)),
  inference(trivial_inequality_removal,[],[f21566])).
fof(f21566,plain,(
  ( ! [X30,X29] : (s(bool,X29) != s(bool,X29) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X30),s(bool,'T_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X30)))) | p(s(bool,X29))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f21489,f146])).
fof(f21597,plain,(
  ( ! [X28,X27] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X28),s(bool,'T_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X28)))) | s(bool,'T_0') = s(bool,X27)) ) | ($spl6 | $spl12)),
  inference(trivial_inequality_removal,[],[f21565])).
fof(f21565,plain,(
  ( ! [X28,X27] : (s(bool,X27) != s(bool,X27) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X28),s(bool,'T_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X28)))) | s(bool,'T_0') = s(bool,X27)) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f21489,f151])).
fof(f8240,plain,(
  ( ! [X2,X0] : (p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X2),s(bool,'T_0')))))) | p(s(bool,X0)) | ~p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X2))))) ) | $spl50),
  inference(cnf_transformation,[],[f8240_D])).
fof(f8240_D,plain,(
  ( ! [X2,X0] : (p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X2),s(bool,'T_0')))))) | p(s(bool,X0)) | ~p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X2))))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f21606,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X4) | p(s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X5),s(bool,'T_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),X5)))) | p(s(bool,X3))) ) | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f21556,f143])).
fof(f21556,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X4) | p(s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X5),s(bool,'T_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),X5)))) | p(s(bool,X3)) | p(s(bool,X4))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f21489,f128])).
fof(f21605,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X2),s(bool,'T_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X2)))) | p(s(bool,X0))) ) | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f21555,f143])).
fof(f21555,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X2),s(bool,'T_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X2)))) | p(s(bool,X1)) | p(s(bool,X0))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f21489,f128])).
fof(f2911,plain,(
  ( ! [X1] : (s(sK0,sK1) != s(sK0,X1) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,X1))),s(fun(sK0,bool),sK3)))) ) | $spl30),
  inference(cnf_transformation,[],[f2911_D])).
fof(f2911_D,plain,(
  ( ! [X1] : (s(sK0,sK1) != s(sK0,X1) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,X1))),s(fun(sK0,bool),sK3)))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f21235,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),X0)))) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,X1))),s(fun(sK0,bool),sK3))) | s(sK0,sK1) != s(sK0,X1)) ) | $spl6),
  inference(resolution,[],[f17209,f135])).
fof(f17209,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X6),s(fun(sK0,bool),X7)))))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),X7)))) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,X6))),s(fun(sK0,bool),sK3)))) ) | $spl6),
  inference(superposition,[],[f1047,f134])).
fof(f20826,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,X2))),s(fun(sK0,bool),sK3))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),X3)))) | s(sK0,sK1) != s(sK0,X2)) ) | $spl6),
  inference(resolution,[],[f17207,f138])).
fof(f17207,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X3),s(sK0,X2)))))) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,X2))),s(fun(sK0,bool),sK3))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),X3))))) ) | $spl6),
  inference(superposition,[],[f1047,f139])).
fof(f475,plain,(
  ( ! [X2] : (~p(s(bool,X2))) ) | $spl28),
  inference(cnf_transformation,[],[f475_D])).
fof(f475_D,plain,(
  ( ! [X2] : (~p(s(bool,X2))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f20698,plain,(
  ( ! [X28,X29] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X29)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X29)))))) | ~p(s(bool,X28))) ) | ($spl6 | $spl12)),
  inference(trivial_inequality_removal,[],[f20670])).
fof(f20670,plain,(
  ( ! [X28,X29] : (s(bool,X28) != s(bool,X28) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X29)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X29)))))) | ~p(s(bool,X28))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f17197,f148])).
fof(f17197,plain,(
  ( ! [X14,X13] : (s(bool,'T_0') != s(bool,X13) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X14)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X13),s(fun(bool,bool),X14))))))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f1756,f134])).
fof(f1756,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ($spl6 | $spl12)),
  inference(backward_demodulation,[],[f243,f1047])).
fof(f20699,plain,(
  ( ! [X26,X27] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X27)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X27)))))) | ~p(s(bool,X26))) ) | ($spl6 | $spl12)),
  inference(trivial_inequality_removal,[],[f20669])).
fof(f20669,plain,(
  ( ! [X26,X27] : (s(bool,X26) != s(bool,X26) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X27)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X27)))))) | ~p(s(bool,X26))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f17197,f150])).
fof(f4798,plain,(
  ( ! [X24] : (s(bool,'F_0') = s(bool,X24)) ) | $spl38),
  inference(cnf_transformation,[],[f4798_D])).
fof(f4798_D,plain,(
  ( ! [X24] : (s(bool,'F_0') = s(bool,X24)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f20700,plain,(
  ( ! [X24,X25] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X25)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X25)))))) | s(bool,'F_0') = s(bool,X24)) ) | ($spl6 | $spl12)),
  inference(trivial_inequality_removal,[],[f20668])).
fof(f20668,plain,(
  ( ! [X24,X25] : (s(bool,X24) != s(bool,X24) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X25)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X25)))))) | s(bool,'F_0') = s(bool,X24)) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f17197,f151])).
fof(f15640,plain,(
  ( ! [X6,X8] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X8)))) | ~p(s(bool,X6)) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X6),s(fun(bool,bool),X8))))))) ) | $spl94),
  inference(cnf_transformation,[],[f15640_D])).
fof(f15640_D,plain,(
  ( ! [X6,X8] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X8)))) | ~p(s(bool,X6)) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X6),s(fun(bool,bool),X8))))))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f20707,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,X10) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X11)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X9),s(fun(bool,bool),X11)))))) | ~p(s(bool,X9))) ) | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f20661,f147])).
fof(f20661,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,X10) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X11)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X9),s(fun(bool,bool),X11)))))) | ~p(s(bool,X10)) | ~p(s(bool,X9))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f17197,f131])).
fof(f20706,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X7) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X8)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X6),s(fun(bool,bool),X8)))))) | ~p(s(bool,X6))) ) | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f20660,f147])).
fof(f20660,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X7) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X8)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X6),s(fun(bool,bool),X8)))))) | ~p(s(bool,X6)) | ~p(s(bool,X7))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f17197,f131])).
fof(f20378,plain,(
  ( ! [X28,X29] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X29)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X29)))))) | ~p(s(bool,X28))) ) | ($spl6 | $spl12)),
  inference(trivial_inequality_removal,[],[f20352])).
fof(f20352,plain,(
  ( ! [X28,X29] : (s(bool,X28) != s(bool,X28) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X29)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X29)))))) | ~p(s(bool,X28))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f17196,f148])).
fof(f17196,plain,(
  ( ! [X12,X11] : (s(bool,'T_0') != s(bool,X11) | p(s(bool,'const.bool.IN_2'(s(bool,X11),s(fun(bool,bool),X12)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X11),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X12))))))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f1756,f134])).
fof(f20379,plain,(
  ( ! [X26,X27] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X27)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X27)))))) | ~p(s(bool,X26))) ) | ($spl6 | $spl12)),
  inference(trivial_inequality_removal,[],[f20351])).
fof(f20351,plain,(
  ( ! [X26,X27] : (s(bool,X26) != s(bool,X26) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X27)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X27)))))) | ~p(s(bool,X26))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f17196,f150])).
fof(f20380,plain,(
  ( ! [X24,X25] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X25)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X25)))))) | s(bool,'F_0') = s(bool,X24)) ) | ($spl6 | $spl12)),
  inference(trivial_inequality_removal,[],[f20350])).
fof(f20350,plain,(
  ( ! [X24,X25] : (s(bool,X24) != s(bool,X24) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X25)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X25)))))) | s(bool,'F_0') = s(bool,X24)) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f17196,f151])).
fof(f6777,plain,(
  ( ! [X6,X8] : (~p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X8)))))) | ~p(s(bool,X6)) | p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),X8))))) ) | $spl44),
  inference(cnf_transformation,[],[f6777_D])).
fof(f6777_D,plain,(
  ( ! [X6,X8] : (~p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X8)))))) | ~p(s(bool,X6)) | p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),X8))))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f20387,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,X10) | p(s(bool,'const.bool.IN_2'(s(bool,X9),s(fun(bool,bool),X11)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X9),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X11)))))) | ~p(s(bool,X9))) ) | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f20343,f147])).
fof(f20343,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,X10) | p(s(bool,'const.bool.IN_2'(s(bool,X9),s(fun(bool,bool),X11)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X9),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X11)))))) | ~p(s(bool,X10)) | ~p(s(bool,X9))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f17196,f131])).
fof(f20386,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X7) | p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),X8)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X8)))))) | ~p(s(bool,X6))) ) | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f20342,f147])).
fof(f20342,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X7) | p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),X8)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X8)))))) | ~p(s(bool,X6)) | ~p(s(bool,X7))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f17196,f131])).
fof(f4795,plain,(
  ( ! [X6,X8] : (p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X8),s(bool,'F_0')))))) | ~p(s(bool,X6)) | ~p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),X8))))) ) | $spl34),
  inference(cnf_transformation,[],[f4795_D])).
fof(f4795_D,plain,(
  ( ! [X6,X8] : (p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X8),s(bool,'F_0')))))) | ~p(s(bool,X6)) | ~p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),X8))))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f20037,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,X10) | p(s(bool,'const.bool.IN_2'(s(bool,X9),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X11),s(bool,'F_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X9),s(fun(bool,bool),X11)))) | ~p(s(bool,X9))) ) | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f19993,f147])).
fof(f19993,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,X10) | p(s(bool,'const.bool.IN_2'(s(bool,X9),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X11),s(bool,'F_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X9),s(fun(bool,bool),X11)))) | ~p(s(bool,X10)) | ~p(s(bool,X9))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f17194,f131])).
fof(f17194,plain,(
  ( ! [X8,X7] : (s(bool,'T_0') != s(bool,X7) | p(s(bool,'const.bool.IN_2'(s(bool,X7),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X8),s(bool,'F_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X7),s(fun(bool,bool),X8))))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f1756,f139])).
fof(f20036,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X7) | p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X8),s(bool,'F_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),X8)))) | ~p(s(bool,X6))) ) | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f19992,f147])).
fof(f19992,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X7) | p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X8),s(bool,'F_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),X8)))) | ~p(s(bool,X6)) | ~p(s(bool,X7))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f17194,f131])).
fof(f19217,plain,(
  $false | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f19116,f192])).
fof(f19116,plain,(
  p(s(bool,'F_0')) | ($spl16 | $spl24)),
  inference(backward_demodulation,[],[f378,f436])).
fof(f377,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)))) | $spl14),
  inference(cnf_transformation,[],[f377_D])).
fof(f377_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f435,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))) | $spl25),
  inference(cnf_transformation,[],[f435_D])).
fof(f435_D,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f17907,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)))) | $spl12),
  inference(superposition,[],[f3935,f140])).
fof(f140,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X3))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) | ! [X3] : s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X3))))))),
  inference(ennf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) => ! [X3] : s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X3))))))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0] : ! [X1,X2] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) => ! [X3] : s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : ! [X7,X8] : (~p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X8)))) => ! [X9] : s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X9))) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X8),s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X7),s(fun(X6,bool),X9))))))),
  file('thm.pred_set.SUBSET_INSERT_DELETE','thm.pred_set.SUBSET_INSERT')).
fof(f378,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))) | $spl16),
  inference(cnf_transformation,[],[f378_D])).
fof(f378_D,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f18001,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)))) | $spl6),
  inference(superposition,[],[f236,f140])).
fof(f18015,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)))) | $spl6),
  inference(superposition,[],[f140,f236])).
fof(f473,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))) | $spl27),
  inference(cnf_transformation,[],[f473_D])).
fof(f473_D,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f17253,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)))) | $spl12),
  inference(superposition,[],[f2917,f140])).
fof(f18626,plain,(
  $false | ($spl17 | $spl27)),
  inference(subsumption_resolution,[],[f473,f431])).
fof(f431,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))) | $spl17),
  inference(trivial_inequality_removal,[],[f426])).
fof(f426,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))) | $spl17),
  inference(superposition,[],[f379,f151])).
fof(f18621,plain,(
  $false | ($spl14 | $spl17 | $spl87)),
  inference(subsumption_resolution,[],[f18620,f198])).
fof(f18620,plain,(
  ~p(s(bool,'T_0')) | ($spl14 | $spl17 | $spl87)),
  inference(forward_demodulation,[],[f18619,f2702])).
fof(f2702,plain,(
  ( ! [X14] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),X14),s(sK0,sK1)),sK2,sK0)),s(fun(sK0,bool),sK2)))) ) | $spl14),
  inference(resolution,[],[f2076,f148])).
fof(f2076,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),X1),s(sK0,sK1)),sK2,sK0)),s(fun(sK0,bool),sK2))))) ) | $spl14),
  inference(resolution,[],[f479,f158])).
fof(f479,plain,(
  ( ! [X1] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X1),s(sK0,sK1))))))) ) | $spl14),
  inference(resolution,[],[f377,f152])).
fof(f152,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X3),s(X0,X1))))))) )),
  inference(cnf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X3),s(X0,X1)))))) | (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))) & (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X3),s(X0,X1))))))))),
  inference(flattening,[],[f99])).
fof(f99,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X3),s(X0,X1)))))) | (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))) & ((p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X3),s(X0,X1))))))))),
  inference(nnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X3),s(X0,X1)))))) <=> (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))))),
  inference(flattening,[],[f58])).
fof(f58,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X3),s(X0,X1)))))) <=> (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6] : ! [X7,X8,X9] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X8),s(fun(X6,bool),'const.pred_set.DELETE_2'(s(fun(X6,bool),X9),s(X6,X7)))))) <=> (~p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X8)))) & p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X9))))))),
  file('thm.pred_set.SUBSET_INSERT_DELETE','thm.pred_set.SUBSET_DELETE')).
fof(f18619,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1)),sK2,sK0)),s(fun(sK0,bool),sK2)))) | ($spl17 | $spl87)),
  inference(subsumption_resolution,[],[f18597,f192])).
fof(f18597,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1)),sK2,sK0)),s(fun(sK0,bool),sK2)))) | ($spl17 | $spl87)),
  inference(superposition,[],[f17351,f14585])).
fof(f14585,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1)),sK2,sK0)),s(fun(sK0,bool),sK3))) | $spl87),
  inference(resolution,[],[f14570,f144])).
fof(f17351,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK3)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2))))) ) | $spl17),
  inference(resolution,[],[f432,f157])).
fof(f432,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))) | $spl17),
  inference(trivial_inequality_removal,[],[f422])).
fof(f422,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))) | $spl17),
  inference(superposition,[],[f379,f146])).
fof(f243,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))) | $spl12),
  inference(cnf_transformation,[],[f243_D])).
fof(f243_D,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f237,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | $spl7),
  inference(cnf_transformation,[],[f237_D])).
fof(f237_D,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f17287,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3)))),
  inference(superposition,[],[f127,f151])).
fof(f233,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3)))) | $spl0),
  inference(cnf_transformation,[],[f233_D])).
fof(f233_D,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f17283,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))))),
  inference(superposition,[],[f127,f146])).
fof(f17282,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))))),
  inference(superposition,[],[f127,f144])).
fof(f232,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3)))) | $spl1),
  inference(cnf_transformation,[],[f232_D])).
fof(f232_D,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f235,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | ~p(s(bool,X2))) ) | $spl4),
  inference(cnf_transformation,[],[f235_D])).
fof(f235_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | ~p(s(bool,X2))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f17281,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | ~p(s(bool,X3)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))))) )),
  inference(superposition,[],[f127,f131])).
fof(f17280,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3)))) | ~p(s(bool,X2))) )),
  inference(superposition,[],[f127,f131])).
fof(f234,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | p(s(bool,X0))) ) | $spl2),
  inference(cnf_transformation,[],[f234_D])).
fof(f234_D,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | p(s(bool,X0))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f17279,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3)))) | p(s(bool,X1))) )),
  inference(superposition,[],[f127,f128])).
fof(f17278,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | p(s(bool,X0)) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))))) )),
  inference(superposition,[],[f127,f128])).
fof(f381,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))))) | $spl19),
  inference(cnf_transformation,[],[f381_D])).
fof(f381_D,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f17273,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))))) ) | $spl12),
  inference(subsumption_resolution,[],[f17257,f148])).
fof(f17257,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X3)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))))) ) | $spl12),
  inference(superposition,[],[f2917,f131])).
fof(f17272,plain,(
  ( ! [X2] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))))) | ~p(s(bool,X2))) ) | $spl12),
  inference(subsumption_resolution,[],[f17256,f148])).
fof(f17256,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))))) | ~p(s(bool,X2))) ) | $spl12),
  inference(superposition,[],[f2917,f131])).
fof(f17236,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))))) ) | $spl6),
  inference(subsumption_resolution,[],[f17219,f143])).
fof(f17219,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3)))) | p(s(bool,X1))) ) | $spl6),
  inference(superposition,[],[f1047,f128])).
fof(f17235,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))))) ) | $spl6),
  inference(subsumption_resolution,[],[f17218,f143])).
fof(f17218,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,X0)) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))))) ) | $spl6),
  inference(superposition,[],[f1047,f128])).
fof(f4920,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))))),
  inference(superposition,[],[f127,f144])).
fof(f4921,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))))),
  inference(superposition,[],[f127,f146])).
fof(f4925,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3)))),
  inference(superposition,[],[f127,f151])).
fof(f17167,plain,(
  ( ! [X49] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X49),s(fun(sK0,bool),sK2))))))) ) | $spl126),
  inference(cnf_transformation,[],[f17167_D])).
fof(f17167_D,plain,(
  ( ! [X49] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X49),s(fun(sK0,bool),sK2))))))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f386,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl23),
  inference(cnf_transformation,[],[f386_D])).
fof(f386_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f6814,plain,(
  ( ! [X14] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X14),s(fun(sK0,bool),sK2))))))) ) | $spl16),
  inference(superposition,[],[f2293,f148])).
fof(f2293,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X13),s(fun(sK0,bool),sK2)))))) ) | $spl16),
  inference(resolution,[],[f965,f143])).
fof(f965,plain,(
  ( ! [X3] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X3),s(fun(sK0,bool),sK2))))))) ) | $spl16),
  inference(resolution,[],[f899,f136])).
fof(f899,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),sK2)))) | $spl16),
  inference(subsumption_resolution,[],[f805,f191])).
fof(f805,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),sK2)))) | $spl16),
  inference(superposition,[],[f158,f378])).
fof(f6815,plain,(
  ( ! [X15] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X15),s(fun(sK0,bool),sK2))))))) ) | $spl16),
  inference(superposition,[],[f2293,f150])).
fof(f17168,plain,(
  ( ! [X16] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X16),s(fun(sK0,bool),sK2)))))) ) | $spl128),
  inference(cnf_transformation,[],[f17168_D])).
fof(f17168_D,plain,(
  ( ! [X16] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X16),s(fun(sK0,bool),sK2)))))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f6816,plain,(
  ( ! [X16] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X16),s(fun(sK0,bool),sK2)))))) ) | $spl16),
  inference(superposition,[],[f2293,f151])).
fof(f17166,plain,(
  ( ! [X51] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X51),s(fun(sK0,bool),sK2))))))) ) | $spl124),
  inference(cnf_transformation,[],[f17166_D])).
fof(f17166_D,plain,(
  ( ! [X51] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X51),s(fun(sK0,bool),sK2))))))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f385,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl22),
  inference(cnf_transformation,[],[f385_D])).
fof(f385_D,plain,(
  s(bool,'T_0') = s(bool,'F_0') <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f6910,plain,(
  ( ! [X13] : (s(bool,'T_0') = s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X13),s(fun(sK0,bool),sK2))))))) ) | $spl16),
  inference(superposition,[],[f2295,f144])).
fof(f2295,plain,(
  ( ! [X15] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X15),s(fun(sK0,bool),sK2)))))) ) | $spl16),
  inference(resolution,[],[f965,f148])).
fof(f6911,plain,(
  ( ! [X14] : (s(bool,'T_0') = s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X14),s(fun(sK0,bool),sK2))))))) ) | $spl16),
  inference(superposition,[],[f2295,f146])).
fof(f6942,plain,(
  ( ! [X45] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X45),s(fun(sK0,bool),sK2))))))) ) | $spl16),
  inference(superposition,[],[f143,f2295])).
fof(f6944,plain,(
  ( ! [X47] : (s(bool,'T_0') = s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X47),s(fun(sK0,bool),sK2))))))) ) | $spl16),
  inference(superposition,[],[f144,f2295])).
fof(f6946,plain,(
  ( ! [X49] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X49),s(fun(sK0,bool),sK2))))))) ) | $spl16),
  inference(superposition,[],[f145,f2295])).
fof(f145,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f96])).
fof(f6948,plain,(
  ( ! [X51] : (s(bool,'T_0') = s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X51),s(fun(sK0,bool),sK2))))))) ) | $spl16),
  inference(superposition,[],[f146,f2295])).
fof(f17165,plain,(
  ( ! [X63] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,X63))))))) ) | $spl122),
  inference(cnf_transformation,[],[f17165_D])).
fof(f17165_D,plain,(
  ( ! [X63] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,X63))))))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f7127,plain,(
  ( ! [X17] : (s(bool,'T_0') = s(bool,'F_0') | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,X17))))))) ) | $spl16),
  inference(superposition,[],[f2347,f148])).
fof(f2347,plain,(
  ( ! [X15] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,X15)))))) ) | $spl16),
  inference(resolution,[],[f1007,f144])).
fof(f1007,plain,(
  ( ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,X3))))))) ) | $spl16),
  inference(resolution,[],[f900,f137])).
fof(f900,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),sK3)))) | $spl16),
  inference(subsumption_resolution,[],[f806,f192])).
fof(f806,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),sK3)))) | $spl16),
  inference(superposition,[],[f159,f378])).
fof(f7128,plain,(
  ( ! [X18] : (s(bool,'T_0') = s(bool,'F_0') | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,X18))))))) ) | $spl16),
  inference(superposition,[],[f2347,f150])).
fof(f17164,plain,(
  ( ! [X16] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,X16))))))) ) | $spl120),
  inference(cnf_transformation,[],[f17164_D])).
fof(f17164_D,plain,(
  ( ! [X16] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,X16))))))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f7170,plain,(
  ( ! [X57] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,X57))))))) ) | $spl16),
  inference(superposition,[],[f147,f2347])).
fof(f7172,plain,(
  ( ! [X59] : (s(bool,'T_0') = s(bool,'F_0') | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,X59))))))) ) | $spl16),
  inference(superposition,[],[f148,f2347])).
fof(f7174,plain,(
  ( ! [X61] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,X61))))))) ) | $spl16),
  inference(superposition,[],[f149,f2347])).
fof(f149,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f98])).
fof(f7176,plain,(
  ( ! [X63] : (s(bool,'T_0') = s(bool,'F_0') | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,X63))))))) ) | $spl16),
  inference(superposition,[],[f150,f2347])).
fof(f7409,plain,(
  ( ! [X15] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,X15))))))) ) | $spl16),
  inference(superposition,[],[f2349,f144])).
fof(f2349,plain,(
  ( ! [X17] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,X17)))))) ) | $spl16),
  inference(resolution,[],[f1007,f147])).
fof(f7410,plain,(
  ( ! [X16] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,X16))))))) ) | $spl16),
  inference(superposition,[],[f2349,f146])).
fof(f17163,plain,(
  ( ! [X20] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,X20)))))) ) | $spl118),
  inference(cnf_transformation,[],[f17163_D])).
fof(f17163_D,plain,(
  ( ! [X20] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,X20)))))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f7414,plain,(
  ( ! [X20] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,X20)))))) ) | $spl16),
  inference(superposition,[],[f2349,f151])).
fof(f17144,plain,(
  ( ! [X42] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X42),s(sK0,sK4(sK3,sK2,sK0)))))))) ) | $spl116),
  inference(cnf_transformation,[],[f17144_D])).
fof(f17144_D,plain,(
  ( ! [X42] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X42),s(sK0,sK4(sK3,sK2,sK0)))))))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f8462,plain,(
  ( ! [X11] : (s(bool,'T_0') = s(bool,'F_0') | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X11),s(sK0,sK4(sK3,sK2,sK0)))))))) ) | $spl16),
  inference(superposition,[],[f4497,f148])).
fof(f4497,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X9),s(sK0,sK4(sK3,sK2,sK0))))))) ) | $spl16),
  inference(resolution,[],[f963,f144])).
fof(f963,plain,(
  ( ! [X1] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X1),s(sK0,sK4(sK3,sK2,sK0)))))))) ) | $spl16),
  inference(resolution,[],[f899,f152])).
fof(f8463,plain,(
  ( ! [X12] : (s(bool,'T_0') = s(bool,'F_0') | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X12),s(sK0,sK4(sK3,sK2,sK0)))))))) ) | $spl16),
  inference(superposition,[],[f4497,f150])).
fof(f17143,plain,(
  ( ! [X10] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X10),s(sK0,sK4(sK3,sK2,sK0)))))))) ) | $spl114),
  inference(cnf_transformation,[],[f17143_D])).
fof(f17143_D,plain,(
  ( ! [X10] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X10),s(sK0,sK4(sK3,sK2,sK0)))))))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f8493,plain,(
  ( ! [X36] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X36),s(sK0,sK4(sK3,sK2,sK0)))))))) ) | $spl16),
  inference(superposition,[],[f147,f4497])).
fof(f8495,plain,(
  ( ! [X38] : (s(bool,'T_0') = s(bool,'F_0') | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X38),s(sK0,sK4(sK3,sK2,sK0)))))))) ) | $spl16),
  inference(superposition,[],[f148,f4497])).
fof(f8497,plain,(
  ( ! [X40] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X40),s(sK0,sK4(sK3,sK2,sK0)))))))) ) | $spl16),
  inference(superposition,[],[f149,f4497])).
fof(f8499,plain,(
  ( ! [X42] : (s(bool,'T_0') = s(bool,'F_0') | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X42),s(sK0,sK4(sK3,sK2,sK0)))))))) ) | $spl16),
  inference(superposition,[],[f150,f4497])).
fof(f8757,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X9),s(sK0,sK4(sK3,sK2,sK0)))))))) ) | $spl16),
  inference(superposition,[],[f4499,f144])).
fof(f4499,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X11),s(sK0,sK4(sK3,sK2,sK0))))))) ) | $spl16),
  inference(resolution,[],[f963,f147])).
fof(f8758,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X10),s(sK0,sK4(sK3,sK2,sK0)))))))) ) | $spl16),
  inference(superposition,[],[f4499,f146])).
fof(f17142,plain,(
  ( ! [X14] : (s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X14),s(sK0,sK4(sK3,sK2,sK0))))))) ) | $spl112),
  inference(cnf_transformation,[],[f17142_D])).
fof(f17142_D,plain,(
  ( ! [X14] : (s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X14),s(sK0,sK4(sK3,sK2,sK0))))))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f8762,plain,(
  ( ! [X14] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X14),s(sK0,sK4(sK3,sK2,sK0))))))) ) | $spl16),
  inference(superposition,[],[f4499,f151])).
fof(f17136,plain,(
  ( ! [X62] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,X62))))))) ) | $spl108),
  inference(cnf_transformation,[],[f17136_D])).
fof(f17136_D,plain,(
  ( ! [X62] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,X62))))))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f17141,plain,(
  ( ! [X16] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,X16))))))) ) | $spl14),
  inference(subsumption_resolution,[],[f12684,f138])).
fof(f12684,plain,(
  ( ! [X16] : (s(bool,'T_0') != s(bool,'F_0') | s(sK0,sK1) = s(sK0,X16) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,X16))))))) ) | $spl14),
  inference(superposition,[],[f8118,f148])).
fof(f8118,plain,(
  ( ! [X18] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,X18))))) | s(sK0,sK1) = s(sK0,X18)) ) | $spl14),
  inference(resolution,[],[f480,f143])).
fof(f480,plain,(
  ( ! [X2] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,X2)))))) | s(sK0,sK1) = s(sK0,X2)) ) | $spl14),
  inference(resolution,[],[f377,f139])).
fof(f17140,plain,(
  ( ! [X17] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,X17))))))) ) | $spl14),
  inference(subsumption_resolution,[],[f12685,f138])).
fof(f12685,plain,(
  ( ! [X17] : (s(bool,'T_0') != s(bool,'F_0') | s(sK0,sK1) = s(sK0,X17) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,X17))))))) ) | $spl14),
  inference(superposition,[],[f8118,f150])).
fof(f17138,plain,(
  ( ! [X0] : (s(sK0,sK1) = s(sK0,X0)) ) | $spl110),
  inference(cnf_transformation,[],[f17138_D])).
fof(f17138_D,plain,(
  ( ! [X0] : (s(sK0,sK1) = s(sK0,X0)) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f17139,plain,(
  ( ! [X18] : (s(bool,'T_0') != s(bool,'F_0') | s(sK0,sK1) = s(sK0,X18)) ) | $spl14),
  inference(subsumption_resolution,[],[f12686,f8118])).
fof(f12686,plain,(
  ( ! [X18] : (s(bool,'T_0') != s(bool,'F_0') | s(sK0,sK1) = s(sK0,X18) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,X18)))))) ) | $spl14),
  inference(superposition,[],[f8118,f151])).
fof(f12843,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(sK0,sK1) = s(sK0,X0)) ) | $spl14),
  inference(duplicate_literal_removal,[],[f12730])).
fof(f12730,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(sK0,sK1) = s(sK0,X0) | s(sK0,sK1) = s(sK0,X0)) ) | $spl14),
  inference(superposition,[],[f8118,f8120])).
fof(f8120,plain,(
  ( ! [X20] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,X20))))) | s(sK0,sK1) = s(sK0,X20)) ) | $spl14),
  inference(resolution,[],[f480,f148])).
fof(f17137,plain,(
  ( ! [X58] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,X58))))))) ) | $spl14),
  inference(subsumption_resolution,[],[f12767,f138])).
fof(f12767,plain,(
  ( ! [X58] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,X58)))))) | s(sK0,sK1) = s(sK0,X58)) ) | $spl14),
  inference(superposition,[],[f143,f8120])).
fof(f17135,plain,(
  ( ! [X62] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,X62))))))) ) | $spl14),
  inference(subsumption_resolution,[],[f12771,f138])).
fof(f12771,plain,(
  ( ! [X62] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,X62)))))) | s(sK0,sK1) = s(sK0,X62)) ) | $spl14),
  inference(superposition,[],[f145,f8120])).
fof(f17134,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X0),s(sK0,sK1))))))) ) | $spl106),
  inference(cnf_transformation,[],[f17134_D])).
fof(f17134_D,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X0),s(sK0,sK1))))))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f13152,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X0),s(sK0,sK1))))))) ) | $spl14),
  inference(superposition,[],[f8268,f2702])).
fof(f8268,plain,(
  ( ! [X12] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(X12,sK2,sK0)),s(fun(sK0,bool),sK2))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),X12))))) ) | $spl14),
  inference(resolution,[],[f5995,f143])).
fof(f5995,plain,(
  ( ! [X2] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4(X2,sK2,sK0)),s(fun(sK0,bool),sK2)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),X2))))) ) | $spl14),
  inference(resolution,[],[f478,f158])).
fof(f478,plain,(
  ( ! [X0] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X0)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),X0))))) ) | $spl14),
  inference(resolution,[],[f377,f157])).
fof(f17133,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,X1))))))) ) | $spl104),
  inference(cnf_transformation,[],[f17133_D])).
fof(f17133_D,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,X1))))))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f13153,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,X1))))))) ) | ($spl14 | $spl16)),
  inference(superposition,[],[f8268,f2753])).
fof(f2753,plain,(
  ( ! [X14] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,X14)),sK2,sK0)),s(fun(sK0,bool),sK2)))) ) | $spl16),
  inference(resolution,[],[f2126,f148])).
fof(f2126,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,X1)),sK2,sK0)),s(fun(sK0,bool),sK2))))) ) | $spl16),
  inference(resolution,[],[f901,f158])).
fof(f901,plain,(
  ( ! [X1] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,X1))))))) ) | $spl16),
  inference(subsumption_resolution,[],[f807,f192])).
fof(f807,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,X1))))))) ) | $spl16),
  inference(superposition,[],[f153,f378])).
fof(f153,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X3),s(X0,X1))))))) )),
  inference(cnf_transformation,[],[f100])).
fof(f17132,plain,(
  ( ! [X16] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),X16))))) ) | $spl14),
  inference(subsumption_resolution,[],[f13162,f5995])).
fof(f13162,plain,(
  ( ! [X16] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),X16)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(X16,sK2,sK0)),s(fun(sK0,bool),sK2))))) ) | $spl14),
  inference(superposition,[],[f8268,f148])).
fof(f17131,plain,(
  ( ! [X17] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),X17))))) ) | $spl14),
  inference(subsumption_resolution,[],[f13163,f5995])).
fof(f13163,plain,(
  ( ! [X17] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),X17)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(X17,sK2,sK0)),s(fun(sK0,bool),sK2))))) ) | $spl14),
  inference(superposition,[],[f8268,f150])).
fof(f17130,plain,(
  ( ! [X18] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),X18))))) ) | $spl14),
  inference(subsumption_resolution,[],[f13164,f8268])).
fof(f13164,plain,(
  ( ! [X18] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),X18)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(X18,sK2,sK0)),s(fun(sK0,bool),sK2)))) ) | $spl14),
  inference(superposition,[],[f8268,f151])).
fof(f17129,plain,(
  ( ! [X16] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),X16))))) ) | $spl14),
  inference(subsumption_resolution,[],[f13394,f5996])).
fof(f5996,plain,(
  ( ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(X3,sK2,sK0)),s(fun(sK0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),X3))))) ) | $spl14),
  inference(resolution,[],[f478,f159])).
fof(f13394,plain,(
  ( ! [X16] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),X16)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(X16,sK2,sK0)),s(fun(sK0,bool),X16))))) ) | $spl14),
  inference(superposition,[],[f8325,f144])).
fof(f8325,plain,(
  ( ! [X24] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(X24,sK2,sK0)),s(fun(sK0,bool),X24))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),X24))))) ) | $spl14),
  inference(resolution,[],[f5996,f147])).
fof(f17128,plain,(
  ( ! [X17] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),X17))))) ) | $spl14),
  inference(subsumption_resolution,[],[f13395,f5996])).
fof(f13395,plain,(
  ( ! [X17] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),X17)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(X17,sK2,sK0)),s(fun(sK0,bool),X17))))) ) | $spl14),
  inference(superposition,[],[f8325,f146])).
fof(f17127,plain,(
  ( ! [X21] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),X21))))) ) | $spl14),
  inference(subsumption_resolution,[],[f13399,f8325])).
fof(f13399,plain,(
  ( ! [X21] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),X21)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(X21,sK2,sK0)),s(fun(sK0,bool),X21)))) ) | $spl14),
  inference(superposition,[],[f8325,f151])).
fof(f17097,plain,(
  $false | ($spl1 | $spl10 | $spl84)),
  inference(subsumption_resolution,[],[f17096,f192])).
fof(f17096,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl10 | $spl84)),
  inference(forward_demodulation,[],[f17095,f5046])).
fof(f5046,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)),sK0)),s(fun(sK0,bool),sK3))) | $spl10),
  inference(resolution,[],[f3694,f144])).
fof(f3694,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)),sK0)),s(fun(sK0,bool),sK3)))) | $spl10),
  inference(subsumption_resolution,[],[f3456,f192])).
fof(f3456,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)),sK0)),s(fun(sK0,bool),sK3)))) | $spl10),
  inference(superposition,[],[f159,f241])).
fof(f17095,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)),sK0)),s(fun(sK0,bool),sK3)))) | ($spl1 | $spl10 | $spl84)),
  inference(subsumption_resolution,[],[f17039,f13958])).
fof(f13958,plain,(
  s(sK0,sK1) != s(sK0,sK4(sK3,'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)),sK0)) | $spl1),
  inference(resolution,[],[f4942,f138])).
fof(f4942,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))) | $spl1),
  inference(resolution,[],[f232,f158])).
fof(f17039,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)),sK0)),s(fun(sK0,bool),sK3)))) | s(sK0,sK1) = s(sK0,sK4(sK3,'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)),sK0)) | ($spl10 | $spl84)),
  inference(resolution,[],[f17035,f134])).
fof(f17035,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))))) | ($spl10 | $spl84)),
  inference(subsumption_resolution,[],[f16999,f199])).
fof(f16999,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))))) | ($spl10 | $spl84)),
  inference(superposition,[],[f4989,f14011])).
fof(f14011,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)),sK0)),s(fun(sK0,bool),sK2))) | $spl84),
  inference(cnf_transformation,[],[f14011_D])).
fof(f14011_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)),sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f4989,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2)))) | p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))))) ) | $spl10),
  inference(resolution,[],[f4906,f157])).
fof(f4906,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))))) | $spl10),
  inference(trivial_inequality_removal,[],[f4896])).
fof(f4896,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))))) | $spl10),
  inference(superposition,[],[f3709,f146])).
fof(f3709,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | $spl10),
  inference(forward_demodulation,[],[f127,f241])).
fof(f16898,plain,(
  $false | ($spl10 | $spl30)),
  inference(subsumption_resolution,[],[f16891,f241])).
fof(f16891,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))) | $spl30),
  inference(equality_resolution,[],[f2911])).
fof(f16897,plain,(
  $false | ($spl8 | $spl10 | $spl14 | $spl16 | $spl30)),
  inference(subsumption_resolution,[],[f16896,f241])).
fof(f16896,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))) | ($spl8 | $spl14 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f16892,f16510])).
fof(f16510,plain,(
  s(sK0,sK1) = s(sK0,sK4(sK3,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)),sK0)) | ($spl8 | $spl14 | $spl16)),
  inference(subsumption_resolution,[],[f16509,f192])).
fof(f16509,plain,(
  p(s(bool,'F_0')) | s(sK0,sK1) = s(sK0,sK4(sK3,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)),sK0)) | ($spl8 | $spl14 | $spl16)),
  inference(forward_demodulation,[],[f16457,f11484])).
fof(f11484,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)),sK0)),s(fun(sK0,bool),sK3))) | ($spl8 | $spl14 | $spl16)),
  inference(resolution,[],[f11411,f144])).
fof(f11411,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)),sK0)),s(fun(sK0,bool),sK3)))) | ($spl8 | $spl14 | $spl16)),
  inference(subsumption_resolution,[],[f11311,f192])).
fof(f11311,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)),sK0)),s(fun(sK0,bool),sK3)))) | ($spl8 | $spl14 | $spl16)),
  inference(superposition,[],[f159,f10978])).
fof(f10978,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))),s(fun(sK0,bool),sK3))) | ($spl8 | $spl14 | $spl16)),
  inference(subsumption_resolution,[],[f10951,f199])).
fof(f10951,plain,(
  ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))),s(fun(sK0,bool),sK3))) | ($spl8 | $spl14 | $spl16)),
  inference(superposition,[],[f10102,f6054])).
fof(f6054,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | ($spl8 | $spl14)),
  inference(resolution,[],[f6040,f148])).
fof(f6040,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))))) | ($spl8 | $spl14)),
  inference(subsumption_resolution,[],[f6021,f199])).
fof(f6021,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))))) | ($spl8 | $spl14)),
  inference(superposition,[],[f478,f238])).
fof(f238,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | $spl8),
  inference(cnf_transformation,[],[f238_D])).
fof(f238_D,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f10102,plain,(
  ( ! [X9] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),X9)))) | s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X9),s(fun(sK0,bool),sK3)))) ) | ($spl8 | $spl16)),
  inference(resolution,[],[f9443,f144])).
fof(f9443,plain,(
  ( ! [X13] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X13),s(fun(sK0,bool),sK3)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),X13))))) ) | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f9341,f192])).
fof(f9341,plain,(
  ( ! [X13] : (p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),X13)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X13),s(fun(sK0,bool),sK3))))) ) | ($spl8 | $spl16)),
  inference(superposition,[],[f157,f9065])).
fof(f9065,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))) | ($spl8 | $spl16)),
  inference(backward_demodulation,[],[f9059,f1011])).
fof(f1011,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),sK3))) | $spl16),
  inference(resolution,[],[f900,f144])).
fof(f9059,plain,(
  s(sK0,sK1) = s(sK0,sK4(sK3,sK2,sK0)) | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f9058,f192])).
fof(f9058,plain,(
  p(s(bool,'F_0')) | s(sK0,sK1) = s(sK0,sK4(sK3,sK2,sK0)) | ($spl8 | $spl16)),
  inference(forward_demodulation,[],[f9006,f1011])).
fof(f9006,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),sK3)))) | s(sK0,sK1) = s(sK0,sK4(sK3,sK2,sK0)) | ($spl8 | $spl16)),
  inference(resolution,[],[f9002,f134])).
fof(f9002,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))))) | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f8983,f199])).
fof(f8983,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))))) | ($spl8 | $spl16)),
  inference(superposition,[],[f962,f238])).
fof(f962,plain,(
  ( ! [X0] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X0)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),X0))))) ) | $spl16),
  inference(resolution,[],[f899,f157])).
fof(f16457,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)),sK0)),s(fun(sK0,bool),sK3)))) | s(sK0,sK1) = s(sK0,sK4(sK3,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)),sK0)) | ($spl8 | $spl14 | $spl16)),
  inference(resolution,[],[f11410,f134])).
fof(f11410,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))))) | ($spl8 | $spl14 | $spl16)),
  inference(subsumption_resolution,[],[f11310,f191])).
fof(f11310,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))))) | ($spl8 | $spl14 | $spl16)),
  inference(superposition,[],[f158,f10978])).
fof(f16892,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK4(sK3,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)),sK0)))),s(fun(sK0,bool),sK3))) | ($spl8 | $spl14 | $spl16 | $spl30)),
  inference(trivial_inequality_removal,[],[f16886])).
fof(f16886,plain,(
  s(sK0,sK1) != s(sK0,sK1) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK4(sK3,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)),sK0)))),s(fun(sK0,bool),sK3))) | ($spl8 | $spl14 | $spl16 | $spl30)),
  inference(superposition,[],[f2911,f16510])).
fof(f16895,plain,(
  $false | ($spl8 | $spl10 | $spl16 | $spl30)),
  inference(subsumption_resolution,[],[f16894,f241])).
fof(f16894,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))) | ($spl8 | $spl16 | $spl30)),
  inference(forward_demodulation,[],[f16893,f9059])).
fof(f16893,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK4(sK3,sK2,sK0)))),s(fun(sK0,bool),sK3))) | ($spl8 | $spl16 | $spl30)),
  inference(trivial_inequality_removal,[],[f16885])).
fof(f16885,plain,(
  s(sK0,sK1) != s(sK0,sK1) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK4(sK3,sK2,sK0)))),s(fun(sK0,bool),sK3))) | ($spl8 | $spl16 | $spl30)),
  inference(superposition,[],[f2911,f9059])).
fof(f16775,plain,(
  ( ! [X31,X32] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X32)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X32)))))) | p(s(bool,X31))) ) | ($spl8 | $spl10)),
  inference(trivial_inequality_removal,[],[f16728])).
fof(f16728,plain,(
  ( ! [X31,X32] : (s(bool,X31) != s(bool,X31) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X32)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X32)))))) | p(s(bool,X31))) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f16394,f144])).
fof(f16394,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') != s(bool,X10) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X11)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X10),s(fun(bool,bool),X11))))))) ) | ($spl8 | $spl10)),
  inference(forward_demodulation,[],[f16393,f238])).
fof(f16393,plain,(
  ( ! [X10,X11] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X11)))) | s(bool,'F_0') != s(bool,X10) | ~p(s(bool,'const.bool.IN_2'(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X10),s(fun(bool,bool),X11))))))) ) | ($spl8 | $spl10)),
  inference(forward_demodulation,[],[f4904,f238])).
fof(f4904,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') != s(bool,X10) | p(s(bool,'const.bool.IN_2'(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))),s(fun(bool,bool),X11)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X10),s(fun(bool,bool),X11))))))) ) | $spl10),
  inference(superposition,[],[f3709,f134])).
fof(f16776,plain,(
  ( ! [X30,X29] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X30)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X30)))))) | p(s(bool,X29))) ) | ($spl8 | $spl10)),
  inference(trivial_inequality_removal,[],[f16727])).
fof(f16727,plain,(
  ( ! [X30,X29] : (s(bool,X29) != s(bool,X29) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X30)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X30)))))) | p(s(bool,X29))) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f16394,f146])).
fof(f16777,plain,(
  ( ! [X28,X27] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X28)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X28)))))) | s(bool,'T_0') = s(bool,X27)) ) | ($spl8 | $spl10)),
  inference(trivial_inequality_removal,[],[f16726])).
fof(f16726,plain,(
  ( ! [X28,X27] : (s(bool,X27) != s(bool,X27) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X28)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X28)))))) | s(bool,'T_0') = s(bool,X27)) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f16394,f151])).
fof(f16787,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X4) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X5)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X3),s(fun(bool,bool),X5)))))) | p(s(bool,X3))) ) | ($spl8 | $spl10)),
  inference(subsumption_resolution,[],[f16717,f143])).
fof(f16717,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X4) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X5)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X3),s(fun(bool,bool),X5)))))) | p(s(bool,X3)) | p(s(bool,X4))) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f16394,f128])).
fof(f16785,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X2)))))) | p(s(bool,X0))) ) | ($spl8 | $spl10)),
  inference(subsumption_resolution,[],[f16716,f143])).
fof(f16716,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X2)))))) | p(s(bool,X1)) | p(s(bool,X0))) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f16394,f128])).
fof(f16451,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK4(sK3,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)))),s(fun(sK0,bool),sK3))) | $spl101),
  inference(cnf_transformation,[],[f16451_D])).
fof(f16451_D,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK4(sK3,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)))),s(fun(sK0,bool),sK3))) <=> ~$spl101),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl101])])).
fof(f16396,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK4(sK3,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)))),s(fun(sK0,bool),sK3))) | ($spl8 | $spl14 | $spl16)),
  inference(resolution,[],[f11111,f4845])).
fof(f4845,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X4),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),X5)))))) | p(s(bool,'const.bool.IN_2'(s(sK0,X4),s(fun(sK0,bool),X5)))) | s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,X4))),s(fun(sK0,bool),sK3)))) ) | $spl8),
  inference(superposition,[],[f3196,f134])).
fof(f3196,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))) | $spl8),
  inference(backward_demodulation,[],[f238,f127])).
fof(f11111,plain,(
  ( ! [X5] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,'const.pred_set.INSERT_2'(s(sK0,X5),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X5),s(fun(sK0,bool),sK2))))))) ) | ($spl8 | $spl14 | $spl16)),
  inference(subsumption_resolution,[],[f11011,f191])).
fof(f11011,plain,(
  ( ! [X5] : (p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,'const.pred_set.INSERT_2'(s(sK0,X5),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X5),s(fun(sK0,bool),sK2))))))) ) | ($spl8 | $spl14 | $spl16)),
  inference(superposition,[],[f158,f10977])).
fof(f10977,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)))) ) | ($spl8 | $spl14 | $spl16)),
  inference(subsumption_resolution,[],[f10950,f199])).
fof(f10950,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),sK2))),s(fun(sK0,bool),sK3)))) ) | ($spl8 | $spl14 | $spl16)),
  inference(superposition,[],[f10102,f1955])).
fof(f1955,plain,(
  ( ! [X15] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X15),s(fun(sK0,bool),sK2)))))) ) | $spl14),
  inference(resolution,[],[f481,f148])).
fof(f481,plain,(
  ( ! [X3] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X3),s(fun(sK0,bool),sK2))))))) ) | $spl14),
  inference(resolution,[],[f377,f136])).
fof(f16447,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK3,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK3))))) | $spl97),
  inference(cnf_transformation,[],[f16447_D])).
fof(f16447_D,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK3,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK3))))) <=> ~$spl97),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl97])])).
fof(f16395,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2)))) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK3,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK3))))) | ($spl8 | $spl10 | $spl14 | $spl16)),
  inference(resolution,[],[f11111,f4880])).
fof(f4880,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.bool.IN_2'(s(sK0,X4),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),X5)))))) | p(s(bool,'const.bool.IN_2'(s(sK0,X4),s(fun(sK0,bool),X5)))) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X4),s(fun(sK0,bool),sK3)))))) ) | $spl10),
  inference(superposition,[],[f3709,f134])).
fof(f15631,plain,(
  ( ! [X28,X29] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X29)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X29)))))) | ~p(s(bool,X28))) ) | ($spl8 | $spl10)),
  inference(trivial_inequality_removal,[],[f15586])).
fof(f15586,plain,(
  ( ! [X28,X29] : (s(bool,X28) != s(bool,X28) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X29)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X29)))))) | ~p(s(bool,X28))) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f15434,f148])).
fof(f15434,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,X10) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X11)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X10),s(fun(bool,bool),X11))))))) ) | ($spl8 | $spl10)),
  inference(forward_demodulation,[],[f15433,f241])).
fof(f15433,plain,(
  ( ! [X10,X11] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X11)))) | s(bool,'T_0') != s(bool,X10) | ~p(s(bool,'const.bool.IN_2'(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X10),s(fun(bool,bool),X11))))))) ) | ($spl8 | $spl10)),
  inference(forward_demodulation,[],[f4868,f241])).
fof(f4868,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,X10) | p(s(bool,'const.bool.IN_2'(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))),s(fun(bool,bool),X11)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X10),s(fun(bool,bool),X11))))))) ) | $spl8),
  inference(superposition,[],[f3196,f134])).
fof(f15632,plain,(
  ( ! [X26,X27] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X27)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X27)))))) | ~p(s(bool,X26))) ) | ($spl8 | $spl10)),
  inference(trivial_inequality_removal,[],[f15585])).
fof(f15585,plain,(
  ( ! [X26,X27] : (s(bool,X26) != s(bool,X26) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X27)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X27)))))) | ~p(s(bool,X26))) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f15434,f150])).
fof(f15633,plain,(
  ( ! [X24,X25] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X25)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X25)))))) | s(bool,'F_0') = s(bool,X24)) ) | ($spl8 | $spl10)),
  inference(trivial_inequality_removal,[],[f15584])).
fof(f15584,plain,(
  ( ! [X24,X25] : (s(bool,X24) != s(bool,X24) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X25)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X25)))))) | s(bool,'F_0') = s(bool,X24)) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f15434,f151])).
fof(f15641,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,X10) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X11)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X9),s(fun(bool,bool),X11)))))) | ~p(s(bool,X9))) ) | ($spl8 | $spl10)),
  inference(subsumption_resolution,[],[f15577,f147])).
fof(f15577,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,X10) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X11)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X9),s(fun(bool,bool),X11)))))) | ~p(s(bool,X10)) | ~p(s(bool,X9))) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f15434,f131])).
fof(f15639,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X7) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X8)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X6),s(fun(bool,bool),X8)))))) | ~p(s(bool,X6))) ) | ($spl8 | $spl10)),
  inference(subsumption_resolution,[],[f15576,f147])).
fof(f15576,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X7) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X8)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X6),s(fun(bool,bool),X8)))))) | ~p(s(bool,X6)) | ~p(s(bool,X7))) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f15434,f131])).
fof(f14570,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1)),sK2,sK0)),s(fun(sK0,bool),sK3)))) | $spl87),
  inference(cnf_transformation,[],[f14570_D])).
fof(f14570_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1)),sK2,sK0)),s(fun(sK0,bool),sK3)))) <=> ~$spl87),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl87])])).
fof(f14576,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1)),sK2,sK0)))),s(fun(sK0,bool),sK3))) | $spl93),
  inference(cnf_transformation,[],[f14576_D])).
fof(f14576_D,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1)),sK2,sK0)))),s(fun(sK0,bool),sK3))) <=> ~$spl93),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl93])])).
fof(f14514,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1)),sK2,sK0)))),s(fun(sK0,bool),sK3))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1)),sK2,sK0)),s(fun(sK0,bool),sK3)))) | ($spl8 | $spl16)),
  inference(resolution,[],[f2127,f4843])).
fof(f4843,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X1),s(sK0,sK1)))))) | s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,X0))),s(fun(sK0,bool),sK3))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),X1))))) ) | $spl8),
  inference(superposition,[],[f3196,f139])).
fof(f2127,plain,(
  ( ! [X2] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,X2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,X2))))))) ) | $spl16),
  inference(resolution,[],[f901,f159])).
fof(f14574,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1)),sK2,sK0)),s(fun(sK0,bool),sK3))))) | $spl91),
  inference(cnf_transformation,[],[f14574_D])).
fof(f14574_D,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1)),sK2,sK0)),s(fun(sK0,bool),sK3))))) <=> ~$spl91),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl91])])).
fof(f14513,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1)),sK2,sK0)),s(fun(sK0,bool),sK3))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1)),sK2,sK0)),s(fun(sK0,bool),sK3)))) | ($spl10 | $spl16)),
  inference(resolution,[],[f2127,f4878])).
fof(f4878,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X1),s(sK0,sK1)))))) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),sK3))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),X1))))) ) | $spl10),
  inference(superposition,[],[f3709,f139])).
fof(f14572,plain,(
  s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1)),sK2,sK0)))),s(fun(sK0,bool),sK3))) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1)),sK2,sK0)),s(fun(sK0,bool),sK3))))) | $spl89),
  inference(cnf_transformation,[],[f14572_D])).
fof(f14572_D,plain,(
  s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1)),sK2,sK0)))),s(fun(sK0,bool),sK3))) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1)),sK2,sK0)),s(fun(sK0,bool),sK3))))) <=> ~$spl89),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl89])])).
fof(f14512,plain,(
  s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1)),sK2,sK0)))),s(fun(sK0,bool),sK3))) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1)),sK2,sK0)),s(fun(sK0,bool),sK3))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),sK3),s(sK0,sK1)),sK2,sK0)),s(fun(sK0,bool),sK3)))) | $spl16),
  inference(resolution,[],[f2127,f224])).
fof(f224,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X1),s(sK0,sK1)))))) | s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,X0))),s(fun(sK0,bool),sK3))) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),sK3))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),X1))))) )),
  inference(superposition,[],[f127,f139])).
fof(f14111,plain,(
  $false | ($spl1 | $spl85)),
  inference(subsumption_resolution,[],[f14104,f13957])).
fof(f13957,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)),sK0)),s(fun(sK0,bool),sK2)))) | $spl1),
  inference(resolution,[],[f4942,f137])).
fof(f14104,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)),sK0)),s(fun(sK0,bool),sK2)))) | $spl85),
  inference(trivial_inequality_removal,[],[f14096])).
fof(f14096,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)),sK0)),s(fun(sK0,bool),sK2)))) | $spl85),
  inference(superposition,[],[f14012,f150])).
fof(f14110,plain,(
  $false | ($spl1 | $spl85)),
  inference(subsumption_resolution,[],[f14105,f13957])).
fof(f14105,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)),sK0)),s(fun(sK0,bool),sK2)))) | $spl85),
  inference(trivial_inequality_removal,[],[f14095])).
fof(f14095,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)),sK0)),s(fun(sK0,bool),sK2)))) | $spl85),
  inference(superposition,[],[f14012,f148])).
fof(f14010,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK4(sK3,'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)),sK0))))))) | $spl82),
  inference(cnf_transformation,[],[f14010_D])).
fof(f14010_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK4(sK3,'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)),sK0))))))) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f14012,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)),sK0)),s(fun(sK0,bool),sK2))) | $spl85),
  inference(cnf_transformation,[],[f14012_D])).
fof(f14012_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)),sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl85),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl85])])).
fof(f14013,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)),sK0)),s(fun(sK0,bool),sK2))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK4(sK3,'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)),sK0))))))) | ($spl1 | $spl14)),
  inference(forward_demodulation,[],[f13989,f486])).
fof(f486,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))) | $spl14),
  inference(resolution,[],[f377,f148])).
fof(f13989,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK4(sK3,'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)),sK0))))))) | s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)),sK0)),s(fun(sK0,bool),sK2))) | $spl1),
  inference(superposition,[],[f4942,f132])).
fof(f132,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3))))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X2))))) | s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) )),
  inference(cnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) | s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3))))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X2)))))) & (s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3))))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X2))))) | s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))))),
  inference(nnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0,X1,X2,X3] : (s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) <=> s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3))))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(flattening,[],[f42])).
fof(f42,plain,(
  ! [X0] : ! [X1,X2,X3] : (s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) <=> s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3))))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6] : ! [X8,X7,X13] : (s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X8))) = s(bool,'const.bool.IN_2'(s(X6,X13),s(fun(X6,bool),X8))) <=> s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),'const.pred_set.DELETE_2'(s(fun(X6,bool),X8),s(X6,X13))))) = s(bool,'const.bool.IN_2'(s(X6,X13),s(fun(X6,bool),'const.pred_set.DELETE_2'(s(fun(X6,bool),X8),s(X6,X7))))))),
  file('thm.pred_set.SUBSET_INSERT_DELETE','thm.pred_set.IN_DELETE_EQ')).
fof(f14008,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)),sK0)),s(fun(sK0,bool),sK2))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK4(sK3,'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)),sK0))))))) | ($spl1 | $spl14)),
  inference(forward_demodulation,[],[f13988,f486])).
fof(f13988,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK4(sK3,'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)),sK0))))))) | s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)),sK0)),s(fun(sK0,bool),sK2))) | $spl1),
  inference(superposition,[],[f4942,f132])).
fof(f13087,plain,(
  $false | ($spl14 | $spl78)),
  inference(subsumption_resolution,[],[f13062,f198])).
fof(f13062,plain,(
  ~p(s(bool,'T_0')) | ($spl14 | $spl78)),
  inference(superposition,[],[f13012,f486])).
fof(f13085,plain,(
  $false | ($spl8 | $spl14 | $spl78)),
  inference(subsumption_resolution,[],[f13059,f198])).
fof(f13059,plain,(
  ~p(s(bool,'T_0')) | ($spl8 | $spl14 | $spl78)),
  inference(superposition,[],[f13012,f6054])).
fof(f13084,plain,(
  $false | ($spl14 | $spl78)),
  inference(subsumption_resolution,[],[f13058,f198])).
fof(f13058,plain,(
  ~p(s(bool,'T_0')) | ($spl14 | $spl78)),
  inference(superposition,[],[f13012,f1955])).
fof(f13023,plain,(
  $false | ($spl14 | $spl78)),
  inference(resolution,[],[f13012,f1950])).
fof(f1950,plain,(
  ( ! [X8,X7] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X7),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X8),s(fun(sK0,bool),sK2))))))))) ) | $spl14),
  inference(resolution,[],[f481,f136])).
fof(f13024,plain,(
  $false | ($spl8 | $spl14 | $spl78)),
  inference(resolution,[],[f13012,f6049])).
fof(f6049,plain,(
  ( ! [X4] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X4),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))))))) ) | ($spl8 | $spl14)),
  inference(resolution,[],[f6040,f136])).
fof(f13013,plain,(
  ( ! [X12] : (s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,X12))),s(fun(sK0,bool),sK3))) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X12),s(fun(sK0,bool),sK3))))) | s(sK0,sK1) != s(sK0,X12)) ) | $spl80),
  inference(cnf_transformation,[],[f13013_D])).
fof(f13013_D,plain,(
  ( ! [X12] : (s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,X12))),s(fun(sK0,bool),sK3))) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X12),s(fun(sK0,bool),sK3))))) | s(sK0,sK1) != s(sK0,X12)) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f12957,plain,(
  ( ! [X12,X13] : (s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,X12))),s(fun(sK0,bool),sK3))) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X12),s(fun(sK0,bool),sK3))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),X13)))) | s(sK0,sK1) != s(sK0,X12)) )),
  inference(resolution,[],[f225,f138])).
fof(f225,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X3),s(sK0,X2)))))) | s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,X2))),s(fun(sK0,bool),sK3))) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X2),s(fun(sK0,bool),sK3))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),X3))))) )),
  inference(superposition,[],[f127,f139])).
fof(f12472,plain,(
  ( ! [X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X9),s(fun(fun(sK0,bool),bool),X11)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X9)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X9),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))),s(fun(fun(sK0,bool),bool),X11))))))) ) | $spl76),
  inference(cnf_transformation,[],[f12472_D])).
fof(f12472_D,plain,(
  ( ! [X11,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X9),s(fun(fun(sK0,bool),bool),X11)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X9)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X9),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))),s(fun(fun(sK0,bool),bool),X11))))))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f12473,plain,(
  ( ! [X14,X12,X13] : (s(bool,'F_0') != s(bool,X13) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X12),s(fun(fun(sK0,bool),bool),X14)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X12),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))),s(fun(fun(sK0,bool),bool),X14)))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X12))))) ) | $spl10),
  inference(subsumption_resolution,[],[f12450,f143])).
fof(f12450,plain,(
  ( ! [X14,X12,X13] : (s(bool,'F_0') != s(bool,X13) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X12),s(fun(fun(sK0,bool),bool),X14)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X12),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))),s(fun(fun(sK0,bool),bool),X14)))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X12)))) | p(s(bool,X13))) ) | $spl10),
  inference(superposition,[],[f4888,f128])).
fof(f4888,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X4))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X4),s(fun(fun(sK0,bool),bool),X5)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X4),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))),s(fun(fun(sK0,bool),bool),X5))))))) ) | $spl10),
  inference(superposition,[],[f3709,f134])).
fof(f12471,plain,(
  ( ! [X10,X11,X9] : (s(bool,'F_0') != s(bool,X10) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X9),s(fun(fun(sK0,bool),bool),X11)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X9),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))),s(fun(fun(sK0,bool),bool),X11)))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X9))))) ) | $spl10),
  inference(subsumption_resolution,[],[f12449,f143])).
fof(f12449,plain,(
  ( ! [X10,X11,X9] : (s(bool,'F_0') != s(bool,X10) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X9),s(fun(fun(sK0,bool),bool),X11)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X9),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))),s(fun(fun(sK0,bool),bool),X11)))))) | p(s(bool,X10)) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),X9))))) ) | $spl10),
  inference(superposition,[],[f4888,f128])).
fof(f12286,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),X1))))) | s(fun(sK0,bool),sK3) != s(fun(sK0,bool),X1)) ) | $spl74),
  inference(cnf_transformation,[],[f12286_D])).
fof(f12286_D,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),X1))))) | s(fun(sK0,bool),sK3) != s(fun(sK0,bool),X1)) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f12232,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK3),s(fun(fun(sK0,bool),bool),X0)))) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),X1))))) | s(fun(sK0,bool),sK3) != s(fun(sK0,bool),X1)) ) | $spl10),
  inference(resolution,[],[f4885,f135])).
fof(f4885,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK3),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),X6),s(fun(fun(sK0,bool),bool),X7)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK3),s(fun(fun(sK0,bool),bool),X7)))) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),X6)))))) ) | $spl10),
  inference(superposition,[],[f3709,f134])).
fof(f10909,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | s(fun(sK0,bool),sK2) != s(fun(sK0,bool),X1)) ) | $spl72),
  inference(cnf_transformation,[],[f10909_D])).
fof(f10909_D,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | s(fun(sK0,bool),sK2) != s(fun(sK0,bool),X1)) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f10855,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),X0)))) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | s(fun(sK0,bool),sK2) != s(fun(sK0,bool),X1)) ) | $spl10),
  inference(resolution,[],[f4877,f135])).
fof(f4877,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),X6),s(fun(fun(sK0,bool),bool),X7)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),X7)))) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X6),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))))) ) | $spl10),
  inference(superposition,[],[f3709,f134])).
fof(f10476,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),X1))) | s(fun(sK0,bool),sK3) != s(fun(sK0,bool),X1)) ) | $spl68),
  inference(cnf_transformation,[],[f10476_D])).
fof(f10476_D,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),X1))) | s(fun(sK0,bool),sK3) != s(fun(sK0,bool),X1)) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f10422,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK3),s(fun(fun(sK0,bool),bool),X0)))) | s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),X1))) | s(fun(sK0,bool),sK3) != s(fun(sK0,bool),X1)) ) | $spl8),
  inference(resolution,[],[f4854,f135])).
fof(f4854,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK3),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),X6),s(fun(fun(sK0,bool),bool),X7)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK3),s(fun(fun(sK0,bool),bool),X7)))) | s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),X6)))) ) | $spl8),
  inference(superposition,[],[f3196,f134])).
fof(f9987,plain,(
  ( ! [X10,X8] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X8),s(fun(fun(sK0,bool),bool),X10)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X8),s(fun(sK0,bool),sK3)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X8),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),X10))))))) ) | $spl66),
  inference(cnf_transformation,[],[f9987_D])).
fof(f9987_D,plain,(
  ( ! [X10,X8] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X8),s(fun(fun(sK0,bool),bool),X10)))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X8),s(fun(sK0,bool),sK3)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X8),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),X10))))))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f9988,plain,(
  ( ! [X12,X13,X11] : (s(bool,'T_0') != s(bool,X12) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X11),s(fun(fun(sK0,bool),bool),X13)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X11),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),X13)))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X11),s(fun(sK0,bool),sK3))))) ) | $spl8),
  inference(subsumption_resolution,[],[f9970,f147])).
fof(f9970,plain,(
  ( ! [X12,X13,X11] : (s(bool,'T_0') != s(bool,X12) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X11),s(fun(fun(sK0,bool),bool),X13)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X11),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),X13)))))) | ~p(s(bool,X12)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X11),s(fun(sK0,bool),sK3))))) ) | $spl8),
  inference(superposition,[],[f4849,f131])).
fof(f4849,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X4),s(fun(sK0,bool),sK3))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X4),s(fun(fun(sK0,bool),bool),X5)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X4),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),X5))))))) ) | $spl8),
  inference(superposition,[],[f3196,f134])).
fof(f9986,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') != s(bool,X9) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X8),s(fun(fun(sK0,bool),bool),X10)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X8),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),X10)))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X8),s(fun(sK0,bool),sK3))))) ) | $spl8),
  inference(subsumption_resolution,[],[f9969,f147])).
fof(f9969,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') != s(bool,X9) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X8),s(fun(fun(sK0,bool),bool),X10)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),X8),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),X10)))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),X8),s(fun(sK0,bool),sK3)))) | ~p(s(bool,X9))) ) | $spl8),
  inference(superposition,[],[f4849,f131])).
fof(f9602,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X1),s(sK0,sK1))),s(fun(sK0,bool),sK3))) | s(fun(sK0,bool),sK2) != s(fun(sK0,bool),X1)) ) | $spl62),
  inference(cnf_transformation,[],[f9602_D])).
fof(f9602_D,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X1),s(sK0,sK1))),s(fun(sK0,bool),sK3))) | s(fun(sK0,bool),sK2) != s(fun(sK0,bool),X1)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f9548,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),X0)))) | s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X1),s(sK0,sK1))),s(fun(sK0,bool),sK3))) | s(fun(sK0,bool),sK2) != s(fun(sK0,bool),X1)) ) | $spl8),
  inference(resolution,[],[f4842,f135])).
fof(f4842,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),X6),s(fun(fun(sK0,bool),bool),X7)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),sK2),s(fun(fun(sK0,bool),bool),X7)))) | s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X6),s(sK0,sK1))),s(fun(sK0,bool),sK3)))) ) | $spl8),
  inference(superposition,[],[f3196,f134])).
fof(f8704,plain,(
  ( ! [X31,X32] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X32)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X32)))) | p(s(bool,X31))) ) | ($spl8 | $spl10)),
  inference(trivial_inequality_removal,[],[f8669])).
fof(f8669,plain,(
  ( ! [X31,X32] : (s(bool,X31) != s(bool,X31) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X32)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X32)))) | p(s(bool,X31))) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f8368,f144])).
fof(f8368,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,X8) | ~p(s(bool,'const.bool.IN_2'(s(bool,X8),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X9)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X8),s(fun(bool,bool),X9))))) ) | ($spl8 | $spl10)),
  inference(forward_demodulation,[],[f4903,f238])).
fof(f4903,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,X8) | p(s(bool,'const.bool.IN_2'(s(bool,X8),s(fun(bool,bool),X9)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X8),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))),s(fun(bool,bool),X9))))))) ) | $spl10),
  inference(superposition,[],[f3709,f134])).
fof(f8705,plain,(
  ( ! [X30,X29] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X30)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X30)))) | p(s(bool,X29))) ) | ($spl8 | $spl10)),
  inference(trivial_inequality_removal,[],[f8668])).
fof(f8668,plain,(
  ( ! [X30,X29] : (s(bool,X29) != s(bool,X29) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X30)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X30)))) | p(s(bool,X29))) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f8368,f146])).
fof(f8706,plain,(
  ( ! [X28,X27] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X28)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X28)))) | s(bool,'T_0') = s(bool,X27)) ) | ($spl8 | $spl10)),
  inference(trivial_inequality_removal,[],[f8667])).
fof(f8667,plain,(
  ( ! [X28,X27] : (s(bool,X27) != s(bool,X27) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X28)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X28)))) | s(bool,'T_0') = s(bool,X27)) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f8368,f151])).
fof(f8716,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X4) | ~p(s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X5)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),X5)))) | p(s(bool,X3))) ) | ($spl8 | $spl10)),
  inference(subsumption_resolution,[],[f8658,f143])).
fof(f8658,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X4) | ~p(s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X5)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),X5)))) | p(s(bool,X3)) | p(s(bool,X4))) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f8368,f128])).
fof(f8714,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,X1) | ~p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X2)))) | p(s(bool,X0))) ) | ($spl8 | $spl10)),
  inference(subsumption_resolution,[],[f8657,f143])).
fof(f8657,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,X1) | ~p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X2)))) | p(s(bool,X1)) | p(s(bool,X0))) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f8368,f128])).
fof(f8229,plain,(
  ( ! [X31,X32] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X32),s(bool,'T_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X32)))) | p(s(bool,X31))) ) | ($spl8 | $spl10)),
  inference(trivial_inequality_removal,[],[f8196])).
fof(f8196,plain,(
  ( ! [X31,X32] : (s(bool,X31) != s(bool,X31) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X32),s(bool,'T_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X32)))) | p(s(bool,X31))) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f8107,f144])).
fof(f8107,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X4) | p(s(bool,'const.bool.IN_2'(s(bool,X4),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X5),s(bool,'T_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X4),s(fun(bool,bool),X5))))) ) | ($spl8 | $spl10)),
  inference(forward_demodulation,[],[f4901,f238])).
fof(f4901,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X4) | p(s(bool,'const.bool.IN_2'(s(bool,X4),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X5),s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))))))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X4),s(fun(bool,bool),X5))))) ) | $spl10),
  inference(superposition,[],[f3709,f139])).
fof(f8230,plain,(
  ( ! [X30,X29] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X30),s(bool,'T_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X30)))) | p(s(bool,X29))) ) | ($spl8 | $spl10)),
  inference(trivial_inequality_removal,[],[f8195])).
fof(f8195,plain,(
  ( ! [X30,X29] : (s(bool,X29) != s(bool,X29) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X30),s(bool,'T_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X30)))) | p(s(bool,X29))) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f8107,f146])).
fof(f8231,plain,(
  ( ! [X28,X27] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X28),s(bool,'T_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X28)))) | s(bool,'T_0') = s(bool,X27)) ) | ($spl8 | $spl10)),
  inference(trivial_inequality_removal,[],[f8194])).
fof(f8194,plain,(
  ( ! [X28,X27] : (s(bool,X27) != s(bool,X27) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X28),s(bool,'T_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X28)))) | s(bool,'T_0') = s(bool,X27)) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f8107,f151])).
fof(f8241,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X4) | p(s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X5),s(bool,'T_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),X5)))) | p(s(bool,X3))) ) | ($spl8 | $spl10)),
  inference(subsumption_resolution,[],[f8185,f143])).
fof(f8185,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X4) | p(s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X5),s(bool,'T_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),X5)))) | p(s(bool,X3)) | p(s(bool,X4))) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f8107,f128])).
fof(f8239,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X2),s(bool,'T_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X2)))) | p(s(bool,X0))) ) | ($spl8 | $spl10)),
  inference(subsumption_resolution,[],[f8184,f143])).
fof(f8184,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X2),s(bool,'T_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X2)))) | p(s(bool,X1)) | p(s(bool,X0))) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f8107,f128])).
fof(f7839,plain,(
  ( ! [X1] : (s(sK0,sK1) != s(sK0,X1) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X1),s(fun(sK0,bool),sK3)))))) ) | $spl48),
  inference(cnf_transformation,[],[f7839_D])).
fof(f7839_D,plain,(
  ( ! [X1] : (s(sK0,sK1) != s(sK0,X1) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X1),s(fun(sK0,bool),sK3)))))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f7784,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),X0)))) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X1),s(fun(sK0,bool),sK3))))) | s(sK0,sK1) != s(sK0,X1)) ) | $spl10),
  inference(resolution,[],[f4881,f135])).
fof(f4881,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X6),s(fun(sK0,bool),X7)))))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),X7)))) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X6),s(fun(sK0,bool),sK3)))))) ) | $spl10),
  inference(superposition,[],[f3709,f134])).
fof(f6768,plain,(
  ( ! [X28,X29] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X29)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X29)))) | ~p(s(bool,X28))) ) | ($spl8 | $spl10)),
  inference(trivial_inequality_removal,[],[f6738])).
fof(f6738,plain,(
  ( ! [X28,X29] : (s(bool,X28) != s(bool,X28) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X29)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X29)))) | ~p(s(bool,X28))) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f6463,f148])).
fof(f6463,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,X8) | ~p(s(bool,'const.bool.IN_2'(s(bool,X8),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X9)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X8),s(fun(bool,bool),X9))))) ) | ($spl8 | $spl10)),
  inference(forward_demodulation,[],[f4867,f241])).
fof(f4867,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,X8) | p(s(bool,'const.bool.IN_2'(s(bool,X8),s(fun(bool,bool),X9)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X8),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))),s(fun(bool,bool),X9))))))) ) | $spl8),
  inference(superposition,[],[f3196,f134])).
fof(f6769,plain,(
  ( ! [X26,X27] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X27)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X27)))) | ~p(s(bool,X26))) ) | ($spl8 | $spl10)),
  inference(trivial_inequality_removal,[],[f6737])).
fof(f6737,plain,(
  ( ! [X26,X27] : (s(bool,X26) != s(bool,X26) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X27)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X27)))) | ~p(s(bool,X26))) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f6463,f150])).
fof(f6770,plain,(
  ( ! [X24,X25] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X25)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X25)))) | s(bool,'F_0') = s(bool,X24)) ) | ($spl8 | $spl10)),
  inference(trivial_inequality_removal,[],[f6736])).
fof(f6736,plain,(
  ( ! [X24,X25] : (s(bool,X24) != s(bool,X24) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X25)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X25)))) | s(bool,'F_0') = s(bool,X24)) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f6463,f151])).
fof(f6778,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'const.bool.IN_2'(s(bool,X9),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X11)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X9),s(fun(bool,bool),X11)))) | ~p(s(bool,X10)) | ~p(s(bool,X9))) ) | ($spl8 | $spl10)),
  inference(subsumption_resolution,[],[f6729,f148])).
fof(f6729,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,'const.bool.IN_2'(s(bool,X9),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X11)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X9),s(fun(bool,bool),X11)))) | ~p(s(bool,X10)) | ~p(s(bool,X9))) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f6463,f131])).
fof(f6776,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X8)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),X8)))) | ~p(s(bool,X6)) | ~p(s(bool,X7))) ) | ($spl8 | $spl10)),
  inference(subsumption_resolution,[],[f6728,f148])).
fof(f6728,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X7) | ~p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X8)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),X8)))) | ~p(s(bool,X6)) | ~p(s(bool,X7))) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f6463,f131])).
fof(f6246,plain,(
  $false | ($spl8 | $spl36)),
  inference(trivial_inequality_removal,[],[f6228])).
fof(f6228,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl8 | $spl36)),
  inference(superposition,[],[f4796,f238])).
fof(f6245,plain,(
  $false | ($spl14 | $spl36)),
  inference(trivial_inequality_removal,[],[f6234])).
fof(f6234,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl14 | $spl36)),
  inference(superposition,[],[f4796,f486])).
fof(f6244,plain,(
  $false | ($spl16 | $spl36)),
  inference(trivial_inequality_removal,[],[f6235])).
fof(f6235,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl16 | $spl36)),
  inference(superposition,[],[f4796,f970])).
fof(f970,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK4(sK3,sK2,sK0)),s(fun(sK0,bool),sK2))) | $spl16),
  inference(resolution,[],[f899,f148])).
fof(f6243,plain,(
  $false | ($spl14 | $spl36)),
  inference(trivial_inequality_removal,[],[f6237])).
fof(f6237,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl14 | $spl36)),
  inference(superposition,[],[f4796,f1955])).
fof(f6242,plain,(
  $false | $spl36),
  inference(equality_resolution,[],[f4796])).
fof(f4799,plain,(
  ( ! [X25] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X25)))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X25),s(bool,'F_0'))))))) ) | $spl40),
  inference(cnf_transformation,[],[f4799_D])).
fof(f4799_D,plain,(
  ( ! [X25] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X25)))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X25),s(bool,'F_0'))))))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f6176,plain,(
  ( ! [X28,X29] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X29),s(bool,'F_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X29)))) | ~p(s(bool,X28))) ) | ($spl8 | $spl10)),
  inference(trivial_inequality_removal,[],[f6149])).
fof(f6149,plain,(
  ( ! [X28,X29] : (s(bool,X28) != s(bool,X28) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X29),s(bool,'F_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X29)))) | ~p(s(bool,X28))) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f5993,f148])).
fof(f5993,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,'const.bool.IN_2'(s(bool,X4),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X5),s(bool,'F_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X4),s(fun(bool,bool),X5))))) ) | ($spl8 | $spl10)),
  inference(forward_demodulation,[],[f4865,f241])).
fof(f4865,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,'const.bool.IN_2'(s(bool,X4),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X5),s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3)))))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X4),s(fun(bool,bool),X5))))) ) | $spl8),
  inference(superposition,[],[f3196,f139])).
fof(f6177,plain,(
  ( ! [X26,X27] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X27),s(bool,'F_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X27)))) | ~p(s(bool,X26))) ) | ($spl8 | $spl10)),
  inference(trivial_inequality_removal,[],[f6148])).
fof(f6148,plain,(
  ( ! [X26,X27] : (s(bool,X26) != s(bool,X26) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X27),s(bool,'F_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X27)))) | ~p(s(bool,X26))) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f5993,f150])).
fof(f6178,plain,(
  ( ! [X24,X25] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X25),s(bool,'F_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X25)))) | s(bool,'F_0') = s(bool,X24)) ) | ($spl8 | $spl10)),
  inference(trivial_inequality_removal,[],[f6147])).
fof(f6147,plain,(
  ( ! [X24,X25] : (s(bool,X24) != s(bool,X24) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X25),s(bool,'F_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X25)))) | s(bool,'F_0') = s(bool,X24)) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f5993,f151])).
fof(f6185,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,X10) | p(s(bool,'const.bool.IN_2'(s(bool,X9),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X11),s(bool,'F_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X9),s(fun(bool,bool),X11)))) | ~p(s(bool,X9))) ) | ($spl8 | $spl10)),
  inference(subsumption_resolution,[],[f6140,f147])).
fof(f6140,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,X10) | p(s(bool,'const.bool.IN_2'(s(bool,X9),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X11),s(bool,'F_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X9),s(fun(bool,bool),X11)))) | ~p(s(bool,X10)) | ~p(s(bool,X9))) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f5993,f131])).
fof(f6184,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X7) | p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X8),s(bool,'F_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),X8)))) | ~p(s(bool,X6))) ) | ($spl8 | $spl10)),
  inference(subsumption_resolution,[],[f6139,f147])).
fof(f6139,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X7) | p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X8),s(bool,'F_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),X8)))) | ~p(s(bool,X6)) | ~p(s(bool,X7))) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f5993,f131])).
fof(f5988,plain,(
  ( ! [X1] : (s(sK0,sK1) != s(sK0,X1) | s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,X1))),s(fun(sK0,bool),sK3)))) ) | $spl42),
  inference(cnf_transformation,[],[f5988_D])).
fof(f5988_D,plain,(
  ( ! [X1] : (s(sK0,sK1) != s(sK0,X1) | s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,X1))),s(fun(sK0,bool),sK3)))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f5934,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),X0)))) | s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,X1))),s(fun(sK0,bool),sK3))) | s(sK0,sK1) != s(sK0,X1)) ) | $spl8),
  inference(resolution,[],[f4846,f135])).
fof(f4846,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X6),s(fun(sK0,bool),X7)))))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),X7)))) | s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,X6))),s(fun(sK0,bool),sK3)))) ) | $spl8),
  inference(superposition,[],[f3196,f134])).
fof(f241,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))) | $spl10),
  inference(cnf_transformation,[],[f241_D])).
fof(f241_D,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f239,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | $spl9),
  inference(cnf_transformation,[],[f239_D])).
fof(f239_D,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f4924,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3)))),
  inference(superposition,[],[f127,f151])).
fof(f4923,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))))),
  inference(superposition,[],[f127,f150])).
fof(f4922,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))))),
  inference(superposition,[],[f127,f148])).
fof(f4919,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | ~p(s(bool,X3)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))))) )),
  inference(superposition,[],[f127,f131])).
fof(f4918,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3)))) | ~p(s(bool,X2))) )),
  inference(superposition,[],[f127,f131])).
fof(f4917,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3)))) | p(s(bool,X1))) )),
  inference(superposition,[],[f127,f128])).
fof(f4916,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | p(s(bool,X0)) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))))) )),
  inference(superposition,[],[f127,f128])).
fof(f382,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))))) | $spl18),
  inference(cnf_transformation,[],[f382_D])).
fof(f382_D,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f4911,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))))) ) | $spl10),
  inference(subsumption_resolution,[],[f4892,f143])).
fof(f4892,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))))) | p(s(bool,X1))) ) | $spl10),
  inference(superposition,[],[f3709,f128])).
fof(f4910,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))))) ) | $spl10),
  inference(subsumption_resolution,[],[f4891,f143])).
fof(f4891,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,X0)) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))))) ) | $spl10),
  inference(superposition,[],[f3709,f128])).
fof(f4873,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))))) ) | $spl8),
  inference(subsumption_resolution,[],[f4858,f148])).
fof(f4858,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X3)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))))) ) | $spl8),
  inference(superposition,[],[f3196,f131])).
fof(f4872,plain,(
  ( ! [X2] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3)))) | ~p(s(bool,X2))) ) | $spl8),
  inference(subsumption_resolution,[],[f4857,f148])).
fof(f4857,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3)))) | ~p(s(bool,X2))) ) | $spl8),
  inference(superposition,[],[f3196,f131])).
fof(f3949,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))))),
  inference(superposition,[],[f127,f148])).
fof(f3950,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))))),
  inference(superposition,[],[f127,f150])).
fof(f3951,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3)))),
  inference(superposition,[],[f127,f151])).
fof(f4786,plain,(
  ( ! [X28,X29] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X29),s(bool,'F_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X29)))) | ~p(s(bool,X28))) ) | ($spl6 | $spl12)),
  inference(trivial_inequality_removal,[],[f4760])).
fof(f4760,plain,(
  ( ! [X28,X29] : (s(bool,X28) != s(bool,X28) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X29),s(bool,'F_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X29)))) | ~p(s(bool,X28))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f3859,f148])).
fof(f3859,plain,(
  ( ! [X8,X7] : (s(bool,'T_0') != s(bool,X7) | p(s(bool,'const.bool.IN_2'(s(bool,X7),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X8),s(bool,'F_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X7),s(fun(bool,bool),X8))))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f1756,f139])).
fof(f4787,plain,(
  ( ! [X26,X27] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X27),s(bool,'F_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X27)))) | ~p(s(bool,X26))) ) | ($spl6 | $spl12)),
  inference(trivial_inequality_removal,[],[f4759])).
fof(f4759,plain,(
  ( ! [X26,X27] : (s(bool,X26) != s(bool,X26) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X27),s(bool,'F_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X27)))) | ~p(s(bool,X26))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f3859,f150])).
fof(f4788,plain,(
  ( ! [X24,X25] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X25),s(bool,'F_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X25)))) | s(bool,'F_0') = s(bool,X24)) ) | ($spl6 | $spl12)),
  inference(trivial_inequality_removal,[],[f4758])).
fof(f4758,plain,(
  ( ! [X24,X25] : (s(bool,X24) != s(bool,X24) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X25),s(bool,'F_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X25)))) | s(bool,'F_0') = s(bool,X24)) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f3859,f151])).
fof(f4797,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,X10) | p(s(bool,'const.bool.IN_2'(s(bool,X9),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X11),s(bool,'F_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X9),s(fun(bool,bool),X11)))) | ~p(s(bool,X9))) ) | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f4751,f147])).
fof(f4751,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,X10) | p(s(bool,'const.bool.IN_2'(s(bool,X9),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X11),s(bool,'F_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X9),s(fun(bool,bool),X11)))) | ~p(s(bool,X10)) | ~p(s(bool,X9))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f3859,f131])).
fof(f4794,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X7) | p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X8),s(bool,'F_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),X8)))) | ~p(s(bool,X6))) ) | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f4750,f147])).
fof(f4750,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X7) | p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X8),s(bool,'F_0')))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),X8)))) | ~p(s(bool,X6)) | ~p(s(bool,X7))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f3859,f131])).
fof(f3952,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3)))),
  inference(superposition,[],[f127,f151])).
fof(f3948,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))))),
  inference(superposition,[],[f127,f146])).
fof(f3947,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))))),
  inference(superposition,[],[f127,f144])).
fof(f3946,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | ~p(s(bool,X3)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))))) )),
  inference(superposition,[],[f127,f131])).
fof(f3945,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3)))) | ~p(s(bool,X2))) )),
  inference(superposition,[],[f127,f131])).
fof(f3944,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3)))) | p(s(bool,X1))) )),
  inference(superposition,[],[f127,f128])).
fof(f3943,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | p(s(bool,X0)) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))))) )),
  inference(superposition,[],[f127,f128])).
fof(f3938,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))))) ) | $spl12),
  inference(subsumption_resolution,[],[f3923,f148])).
fof(f3923,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X3)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))))) ) | $spl12),
  inference(superposition,[],[f2917,f131])).
fof(f3937,plain,(
  ( ! [X2] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))))) | ~p(s(bool,X2))) ) | $spl12),
  inference(subsumption_resolution,[],[f3922,f148])).
fof(f3922,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))))) | ~p(s(bool,X2))) ) | $spl12),
  inference(superposition,[],[f2917,f131])).
fof(f3902,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))))) ) | $spl6),
  inference(subsumption_resolution,[],[f3885,f143])).
fof(f3885,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3)))) | p(s(bool,X1))) ) | $spl6),
  inference(superposition,[],[f1047,f128])).
fof(f3901,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))))) ) | $spl6),
  inference(subsumption_resolution,[],[f3884,f143])).
fof(f3884,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,X0)) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))))) ) | $spl6),
  inference(superposition,[],[f1047,f128])).
fof(f3867,plain,(
  $false | ($spl6 | $spl12 | $spl22)),
  inference(subsumption_resolution,[],[f385,f1756])).
fof(f3845,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f3802,f475])).
fof(f3802,plain,(
  ( ! [X57,X56,X55] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X55,bool),X56),s(fun(X55,bool),X57))))) ) | $spl28),
  inference(resolution,[],[f475,f158])).
fof(f3844,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f3790,f475])).
fof(f3790,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,'const.bool.IN_2'(s(X15,sK4(X16,X17,X15)),s(fun(X15,bool),X17))))) ) | $spl28),
  inference(resolution,[],[f475,f158])).
fof(f3843,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f3783,f475])).
fof(f3783,plain,(
  p(s(bool,'F_0')) | $spl28),
  inference(resolution,[],[f475,f195])).
fof(f195,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f126])).
fof(f126,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f81])).
fof(f81,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f31])).
fof(f31,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.pred_set.SUBSET_INSERT_DELETE','thm.bool.NOT_CLAUSES|split|2')).
fof(f3842,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f3782,f475])).
fof(f3782,plain,(
  p(s(bool,'F_0')) | $spl28),
  inference(resolution,[],[f475,f193])).
fof(f193,plain,(
  p(s(bool,'F_0')) | p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f125])).
fof(f125,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f80])).
fof(f80,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f33])).
fof(f33,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.pred_set.SUBSET_INSERT_DELETE','thm.bool.NOT_CLAUSES|split|1')).
fof(f3784,plain,(
  $false | $spl28),
  inference(resolution,[],[f475,f197])).
fof(f197,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f19])).
fof(f19,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.SUBSET_INSERT_DELETE','thm.bool.TRUTH')).
fof(f3785,plain,(
  $false | $spl28),
  inference(resolution,[],[f475,f198])).
fof(f3786,plain,(
  $false | $spl28),
  inference(resolution,[],[f475,f199])).
fof(f3793,plain,(
  $false | $spl28),
  inference(resolution,[],[f475,f156])).
fof(f156,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X2),s(X0,X1))),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1,X2] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X2),s(X0,X1))),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ! [X0] : ! [X1,X2] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X2),s(X0,X1))),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : ! [X7,X8] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),'const.pred_set.DELETE_2'(s(fun(X6,bool),X8),s(X6,X7))),s(fun(X6,bool),X8))))),
  file('thm.pred_set.SUBSET_INSERT_DELETE','thm.pred_set.DELETE_SUBSET')).
fof(f3803,plain,(
  $false | ($spl16 | $spl28)),
  inference(resolution,[],[f475,f965])).
fof(f3804,plain,(
  $false | ($spl14 | $spl28)),
  inference(resolution,[],[f475,f2076])).
fof(f3805,plain,(
  $false | ($spl16 | $spl28)),
  inference(resolution,[],[f475,f2126])).
fof(f3085,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))))),
  inference(superposition,[],[f127,f150])).
fof(f3084,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))))),
  inference(superposition,[],[f127,f148])).
fof(f3083,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))))),
  inference(superposition,[],[f127,f146])).
fof(f3082,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))))),
  inference(superposition,[],[f127,f144])).
fof(f3081,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | ~p(s(bool,X3)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))))) )),
  inference(superposition,[],[f127,f131])).
fof(f3080,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3)))) | ~p(s(bool,X2))) )),
  inference(superposition,[],[f127,f131])).
fof(f3079,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3)))) | p(s(bool,X1))) )),
  inference(superposition,[],[f127,f128])).
fof(f3078,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | p(s(bool,X0)) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))))) )),
  inference(superposition,[],[f127,f128])).
fof(f3767,plain,(
  $false | ($spl0 | $spl10)),
  inference(subsumption_resolution,[],[f3766,f192])).
fof(f3766,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl10)),
  inference(forward_demodulation,[],[f233,f241])).
fof(f3765,plain,(
  $false | ($spl6 | $spl10)),
  inference(subsumption_resolution,[],[f1047,f241])).
fof(f3764,plain,(
  $false | ($spl6 | $spl10)),
  inference(subsumption_resolution,[],[f3763,f192])).
fof(f3763,plain,(
  p(s(bool,'F_0')) | ($spl6 | $spl10)),
  inference(forward_demodulation,[],[f1523,f241])).
fof(f1523,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3)))) | $spl6),
  inference(trivial_inequality_removal,[],[f1513])).
fof(f1513,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3)))) | $spl6),
  inference(superposition,[],[f1047,f146])).
fof(f3762,plain,(
  $false | ($spl6 | $spl10)),
  inference(subsumption_resolution,[],[f3761,f192])).
fof(f3761,plain,(
  p(s(bool,'F_0')) | ($spl6 | $spl10)),
  inference(forward_demodulation,[],[f1524,f241])).
fof(f1524,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3)))) | $spl6),
  inference(trivial_inequality_removal,[],[f1512])).
fof(f1512,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3)))) | $spl6),
  inference(superposition,[],[f1047,f144])).
fof(f3693,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f3648,f3613])).
fof(f3613,plain,(
  ( ! [X2] : (~p(s(bool,X2))) ) | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f3612,f150])).
fof(f3612,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,X2))) ) | ($spl4 | $spl8)),
  inference(forward_demodulation,[],[f235,f238])).
fof(f3648,plain,(
  ( ! [X57,X56,X55] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X55,bool),X56),s(fun(X55,bool),X57))))) ) | ($spl4 | $spl8)),
  inference(resolution,[],[f3613,f158])).
fof(f3692,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f3636,f3613])).
fof(f3636,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,'const.bool.IN_2'(s(X15,sK4(X16,X17,X15)),s(fun(X15,bool),X17))))) ) | ($spl4 | $spl8)),
  inference(resolution,[],[f3613,f158])).
fof(f3691,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f3629,f3613])).
fof(f3629,plain,(
  p(s(bool,'F_0')) | ($spl4 | $spl8)),
  inference(resolution,[],[f3613,f195])).
fof(f3690,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f3628,f3613])).
fof(f3628,plain,(
  p(s(bool,'F_0')) | ($spl4 | $spl8)),
  inference(resolution,[],[f3613,f193])).
fof(f3630,plain,(
  $false | ($spl4 | $spl8)),
  inference(resolution,[],[f3613,f197])).
fof(f3631,plain,(
  $false | ($spl4 | $spl8)),
  inference(resolution,[],[f3613,f198])).
fof(f3632,plain,(
  $false | ($spl4 | $spl8)),
  inference(resolution,[],[f3613,f199])).
fof(f3639,plain,(
  $false | ($spl4 | $spl8)),
  inference(resolution,[],[f3613,f156])).
fof(f3649,plain,(
  $false | ($spl4 | $spl8 | $spl16)),
  inference(resolution,[],[f3613,f965])).
fof(f3650,plain,(
  $false | ($spl4 | $spl8 | $spl14)),
  inference(resolution,[],[f3613,f2076])).
fof(f3651,plain,(
  $false | ($spl4 | $spl8 | $spl16)),
  inference(resolution,[],[f3613,f2126])).
fof(f3087,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3)))),
  inference(superposition,[],[f127,f151])).
fof(f3086,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3)))),
  inference(superposition,[],[f127,f151])).
fof(f2968,plain,(
  $false | ($spl16 | $spl20)),
  inference(trivial_inequality_removal,[],[f2953])).
fof(f2953,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl16 | $spl20)),
  inference(superposition,[],[f383,f378])).
fof(f2967,plain,(
  $false | ($spl16 | $spl20)),
  inference(trivial_inequality_removal,[],[f2960])).
fof(f2960,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl16 | $spl20)),
  inference(superposition,[],[f383,f1011])).
fof(f2966,plain,(
  $false | $spl20),
  inference(equality_resolution,[],[f383])).
fof(f2921,plain,(
  $false | ($spl1 | $spl12)),
  inference(subsumption_resolution,[],[f2920,f198])).
fof(f2920,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl12)),
  inference(forward_demodulation,[],[f232,f243])).
fof(f2919,plain,(
  $false | ($spl10 | $spl12 | $spl23)),
  inference(subsumption_resolution,[],[f2918,f386])).
fof(f2918,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f241,f243])).
fof(f2857,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),X0)))) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,X1))),s(fun(sK0,bool),sK3))) | s(sK0,sK1) != s(sK0,X1)) ) | $spl6),
  inference(resolution,[],[f1499,f135])).
fof(f1499,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X6),s(fun(sK0,bool),X7)))))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),X7)))) | s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,X6))),s(fun(sK0,bool),sK3)))) ) | $spl6),
  inference(superposition,[],[f1047,f134])).
fof(f2072,plain,(
  $false | ($spl2 | $spl6)),
  inference(subsumption_resolution,[],[f2025,f2000])).
fof(f2000,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl2 | $spl6)),
  inference(subsumption_resolution,[],[f1999,f146])).
fof(f1999,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,X0))) ) | ($spl2 | $spl6)),
  inference(forward_demodulation,[],[f234,f236])).
fof(f2025,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X31,bool),X32),s(fun(X31,bool),'const.pred_set.DELETE_2'(s(fun(X31,bool),X33),s(X31,X34))))))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f2000,f152])).
fof(f2071,plain,(
  $false | ($spl2 | $spl6)),
  inference(subsumption_resolution,[],[f2022,f2000])).
fof(f2022,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,'const.bool.IN_2'(s(X21,X22),s(fun(X21,bool),X23))))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f2000,f152])).
fof(f2070,plain,(
  $false | ($spl2 | $spl6)),
  inference(subsumption_resolution,[],[f2019,f198])).
fof(f2019,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl6)),
  inference(resolution,[],[f2000,f196])).
fof(f196,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f126])).
fof(f2069,plain,(
  $false | ($spl2 | $spl6)),
  inference(subsumption_resolution,[],[f2018,f198])).
fof(f2018,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl6)),
  inference(resolution,[],[f2000,f194])).
fof(f194,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f125])).
fof(f2017,plain,(
  $false | ($spl2 | $spl6)),
  inference(resolution,[],[f2000,f192])).
fof(f477,plain,(
  $false | ($spl7 | $spl9)),
  inference(subsumption_resolution,[],[f469,f237])).
fof(f469,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | $spl9),
  inference(trivial_inequality_removal,[],[f463])).
fof(f463,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | $spl9),
  inference(superposition,[],[f239,f151])).
fof(f476,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))))) ) | $spl9),
  inference(subsumption_resolution,[],[f458,f148])).
fof(f458,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X3)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))))) ) | $spl9),
  inference(superposition,[],[f239,f131])).
fof(f474,plain,(
  ( ! [X2] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))))) | ~p(s(bool,X2))) ) | $spl9),
  inference(subsumption_resolution,[],[f457,f148])).
fof(f457,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))))) | ~p(s(bool,X2))) ) | $spl9),
  inference(superposition,[],[f239,f131])).
fof(f454,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)))) | $spl9),
  inference(superposition,[],[f239,f140])).
fof(f436,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))) | $spl24),
  inference(cnf_transformation,[],[f436_D])).
fof(f436_D,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f437,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) ) | $spl17),
  inference(subsumption_resolution,[],[f418,f143])).
fof(f418,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3)))) | p(s(bool,X1))) ) | $spl17),
  inference(superposition,[],[f379,f128])).
fof(f434,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) ) | $spl17),
  inference(subsumption_resolution,[],[f417,f143])).
fof(f417,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,X0)) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))))) ) | $spl17),
  inference(superposition,[],[f379,f128])).
fof(f366,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))))) | $spl7),
  inference(superposition,[],[f237,f150])).
fof(f365,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))))) | $spl7),
  inference(superposition,[],[f237,f148])).
fof(f384,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))))) ) | $spl7),
  inference(subsumption_resolution,[],[f360,f143])).
fof(f360,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3)))))) | p(s(bool,X1))) ) | $spl7),
  inference(superposition,[],[f237,f128])).
fof(f380,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))))) ) | $spl7),
  inference(subsumption_resolution,[],[f359,f143])).
fof(f359,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,X0)) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))))) ) | $spl7),
  inference(superposition,[],[f237,f128])).
fof(f379,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))) | $spl17),
  inference(cnf_transformation,[],[f379_D])).
fof(f379_D,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f358,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK3))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)))) | $spl7),
  inference(superposition,[],[f237,f140])).
fof(f215,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3)))),
  inference(superposition,[],[f127,f151])).
fof(f214,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3)))),
  inference(superposition,[],[f127,f151])).
fof(f213,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))))),
  inference(superposition,[],[f127,f150])).
fof(f212,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))))),
  inference(superposition,[],[f127,f148])).
fof(f211,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))))),
  inference(superposition,[],[f127,f146])).
fof(f210,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))))),
  inference(superposition,[],[f127,f144])).
fof(f209,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | ~p(s(bool,X3)) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))))) )),
  inference(superposition,[],[f127,f131])).
fof(f208,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3)))) | ~p(s(bool,X2))) )),
  inference(superposition,[],[f127,f131])).
fof(f207,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3)))) | p(s(bool,X1))) )),
  inference(superposition,[],[f127,f128])).
fof(f206,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK3))))) | p(s(bool,X0)) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))),s(fun(sK0,bool),sK3))))) )),
  inference(superposition,[],[f127,f128])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 441
% Passive clauses: 10849
% Generated clauses: 31708
% Final active clauses: 177
% Final passive clauses: 4638
% Input formulas: 37
% Initial clauses: 72
% 
% Selected by SInE selection: 33
% SInE iterations: 1
% 
% Duplicate literals: 454
% Trivial inequalities: 668
% Fw subsumption resolutions: 1855
% Fw demodulations: 5494
% Bw demodulations: 111
% 
% Simple tautologies: 885
% Equational tautologies: 8
% Forward subsumptions: 11760
% 
% Binary resolution: 3855
% Forward superposition: 12513
% Backward superposition: 6090
% Equality resolution: 19
% 
% Split clauses: 213
% Split components: 426
% SAT solver clauses: 276
% SAT solver unit clauses: 10
% SAT solver binary clauses: 86
% 
% Sat splits: 213
% Sat splitting refutations: 62
% 
% Memory used [KB]: 10874
% Time elapsed: 0.815 s
% ------------------------------
% ------------------------------
% Success in time 12.56 s
