Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f2764,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f2754,f80,f128,f126,f2634,f81,f129,f127,f574,f572,f173,f340,f124,f132,f84,f134])).
fof(f134,plain,(
  ~p(s(bool,'F_0')) | $spl9),
  inference(cnf_transformation,[],[f134_D])).
fof(f134_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f84,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f26])).
fof(f26,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pred_set.BIJ_EMPTY|split|1',aHL_FALSITY)).
fof(f132,plain,(
  p(s(bool,'T_0')) | $spl6),
  inference(cnf_transformation,[],[f132_D])).
fof(f132_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f124,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f109])).
fof(f109,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f38])).
fof(f38,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X7] : ! [X11] : (s(X7,X11) = s(X7,X11) <=> p(s(bool,'T_0')))),
  file('thm.pred_set.BIJ_EMPTY|split|1','thm.bool.REFL_CLAUSE')).
fof(f340,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(X22,X23),X24),s(fun(X22,bool),'const.pred_set.EMPTY_0'),s(fun(X23,bool),X25)))) ) | $spl12),
  inference(cnf_transformation,[],[f340_D])).
fof(f340_D,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(X22,X23),X24),s(fun(X22,bool),'const.pred_set.EMPTY_0'),s(fun(X23,bool),X25)))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f173,plain,(
  ( ! [X6,X8,X7,X5] : (s(bool,'T_0') = s(bool,'const.pred_set.INJ_3'(s(fun(X5,X6),X7),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(X6,bool),X8)))) )),
  inference(resolution,[],[f113,f87])).
fof(f87,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.pred_set.BIJ_EMPTY|split|1','thm.bool.EQ_CLAUSES|split|0')).
fof(f113,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X1,bool),X3))))) )),
  inference(cnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0,X1,X2,X3] : p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X1,bool),X3))))),
  inference(flattening,[],[f41])).
fof(f41,plain,(
  ! [X0,X1] : ! [X2,X3] : p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X1,bool),X3))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X7,X6] : ! [X8,X9] : p(s(bool,'const.pred_set.INJ_3'(s(fun(X7,X6),X8),s(fun(X7,bool),'const.pred_set.EMPTY_0'),s(fun(X6,bool),X9))))),
  file('thm.pred_set.BIJ_EMPTY|split|1','thm.pred_set.INJ_EMPTY|split|0')).
fof(f572,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun(sK1,X6),X7),s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(X6,bool),'const.pred_set.EMPTY_0')))) ) | $spl20),
  inference(cnf_transformation,[],[f572_D])).
fof(f572_D,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun(sK1,X6),X7),s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(X6,bool),'const.pred_set.EMPTY_0')))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f574,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun(sK1,X12),X13),s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(X12,bool),'const.pred_set.EMPTY_0')))) ) | ($spl2 | $spl9)),
  inference(subsumption_resolution,[],[f570,f134])).
fof(f570,plain,(
  ( ! [X12,X13] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun(sK1,X12),X13),s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(X12,bool),'const.pred_set.EMPTY_0')))) ) | $spl2),
  inference(superposition,[],[f561,f85])).
fof(f85,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.BIJ_EMPTY|split|1',aHL_BOOL_CASES)).
fof(f561,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK1,X0),X1),s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(forward_demodulation,[],[f560,f128])).
fof(f560,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK1,X0),X1),s(fun(sK1,bool),sK3),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(trivial_inequality_removal,[],[f558])).
fof(f558,plain,(
  ( ! [X0,X1] : (s(fun(sK1,bool),'const.pred_set.EMPTY_0') != s(fun(sK1,bool),'const.pred_set.EMPTY_0') | p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK1,X0),X1),s(fun(sK1,bool),sK3),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl2),
  inference(superposition,[],[f118,f128])).
fof(f118,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X1,bool),'const.pred_set.EMPTY_0') != s(fun(X1,bool),X3) | p(s(bool,'const.pred_set.SURJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) | s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3)) & (s(fun(X1,bool),'const.pred_set.EMPTY_0') != s(fun(X1,bool),X3) | p(s(bool,'const.pred_set.SURJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))))),
  inference(nnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) <=> s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3))),
  inference(flattening,[],[f47])).
fof(f47,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) <=> s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6,X7] : ! [X8,X9] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X7,X6),X8),s(fun(X7,bool),X9),s(fun(X6,bool),'const.pred_set.EMPTY_0')))) <=> s(fun(X7,bool),'const.pred_set.EMPTY_0') = s(fun(X7,bool),X9))),
  file('thm.pred_set.BIJ_EMPTY|split|1','thm.pred_set.SURJ_EMPTY|split|1')).
fof(f127,plain,(
  ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK1,sK0),sK2),s(fun(sK1,bool),sK3),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | $spl1),
  inference(cnf_transformation,[],[f127_D])).
fof(f127_D,plain,(
  ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK1,sK0),sK2),s(fun(sK1,bool),sK3),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f129,plain,(
  s(fun(sK1,bool),'const.pred_set.EMPTY_0') != s(fun(sK1,bool),sK3) | $spl3),
  inference(cnf_transformation,[],[f129_D])).
fof(f129_D,plain,(
  s(fun(sK1,bool),'const.pred_set.EMPTY_0') != s(fun(sK1,bool),sK3) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f81,plain,(
  s(fun(sK1,bool),'const.pred_set.EMPTY_0') != s(fun(sK1,bool),sK3) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK1,sK0),sK2),s(fun(sK1,bool),sK3),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))),
  inference(cnf_transformation,[],[f56])).
fof(f56,plain,(
  (p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK1,sK0),sK2),s(fun(sK1,bool),sK3),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | s(fun(sK1,bool),'const.pred_set.EMPTY_0') = s(fun(sK1,bool),sK3)) & (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK1,sK0),sK2),s(fun(sK1,bool),sK3),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | s(fun(sK1,bool),'const.pred_set.EMPTY_0') != s(fun(sK1,bool),sK3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f55])).
fof(f55,plain,(
  ? [X0,X1,X2,X3] : ((p(s(bool,'const.pred_set.BIJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) | s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3)) & (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) | s(fun(X1,bool),'const.pred_set.EMPTY_0') != s(fun(X1,bool),X3)))),
  inference(nnf_transformation,[],[f53])).
fof(f53,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) <~> s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3))),
  inference(ennf_transformation,[],[f25])).
fof(f25,plain,(
  ~! [X0,X1,X2,X3] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) <=> s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3))),
  inference(flattening,[],[f24])).
fof(f24,plain,(
  ~! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) <=> s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X7,X6),X8),s(fun(X7,bool),X9),s(fun(X6,bool),'const.pred_set.EMPTY_0')))) <=> s(fun(X7,bool),'const.pred_set.EMPTY_0') = s(fun(X7,bool),X9))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X7,X6),X8),s(fun(X7,bool),X9),s(fun(X6,bool),'const.pred_set.EMPTY_0')))) <=> s(fun(X7,bool),'const.pred_set.EMPTY_0') = s(fun(X7,bool),X9))),
  file('thm.pred_set.BIJ_EMPTY|split|1','thm.pred_set.BIJ_EMPTY|split|1')).
fof(f2634,plain,(
  $false | ($spl1 | $spl2 | $spl6 | $spl9 | $spl12 | $spl20)),
  inference(subsumption_resolution,[],[f2623,f134])).
fof(f2623,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2 | $spl6 | $spl12 | $spl20)),
  inference(superposition,[],[f1614,f169])).
fof(f169,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.BIJ_3'(s(fun(sK1,sK0),sK2),s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ($spl1 | $spl2)),
  inference(resolution,[],[f163,f91])).
fof(f91,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f29])).
fof(f29,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.pred_set.BIJ_EMPTY|split|1','thm.bool.EQ_CLAUSES|split|3')).
fof(f163,plain,(
  ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK1,sK0),sK2),s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f127,f128])).
fof(f1614,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK1,X0),X1),s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl2 | $spl6 | $spl12 | $spl20)),
  inference(subsumption_resolution,[],[f1599,f132])).
fof(f1599,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK1,X0),X1),s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl2 | $spl6 | $spl12 | $spl20)),
  inference(superposition,[],[f1288,f572])).
fof(f1288,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK1,X0),X1),s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(X0,bool),X2)))) | p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK1,X0),X1),s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(X0,bool),X2))))) ) | ($spl2 | $spl6 | $spl12)),
  inference(forward_demodulation,[],[f1287,f128])).
fof(f1287,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK1,X0),X1),s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(X0,bool),X2)))) | p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK1,X0),X1),s(fun(sK1,bool),sK3),s(fun(X0,bool),X2))))) ) | ($spl2 | $spl6 | $spl12)),
  inference(forward_demodulation,[],[f1286,f128])).
fof(f1286,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK1,X0),X1),s(fun(sK1,bool),sK3),s(fun(X0,bool),X2)))) | p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK1,X0),X1),s(fun(sK1,bool),sK3),s(fun(X0,bool),X2))))) ) | ($spl2 | $spl6 | $spl12)),
  inference(subsumption_resolution,[],[f1285,f132])).
fof(f1285,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK1,X0),X1),s(fun(sK1,bool),sK3),s(fun(X0,bool),X2)))) | p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK1,X0),X1),s(fun(sK1,bool),sK3),s(fun(X0,bool),X2))))) ) | ($spl2 | $spl12)),
  inference(forward_demodulation,[],[f1273,f340])).
fof(f1273,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK1,X0),X1),s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(sK1,X0),X1),s(fun(sK1,bool),sK3),s(fun(X0,bool),X2)))) | p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK1,X0),X1),s(fun(sK1,bool),sK3),s(fun(X0,bool),X2))))) ) | $spl2),
  inference(superposition,[],[f123,f128])).
fof(f123,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | (p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) & p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))))) & (~p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(flattening,[],[f78])).
fof(f78,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | (p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) & p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))))) & ((~p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))) | p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(nnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) <=> (p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) & p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(flattening,[],[f51])).
fof(f51,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) <=> (p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) & p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X7,X6] : ! [X8,X9,X10] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X7,X6),X8),s(fun(X7,bool),X9),s(fun(X6,bool),X10)))) <=> (p(s(bool,'const.pred_set.INJ_3'(s(fun(X7,X6),X8),s(fun(X7,bool),X9),s(fun(X6,bool),X10)))) & p(s(bool,'const.pred_set.SURJ_3'(s(fun(X7,X6),X8),s(fun(X7,bool),X9),s(fun(X6,bool),X10))))))),
  file('thm.pred_set.BIJ_EMPTY|split|1','thm.pred_set.BIJ_DEF')).
fof(f126,plain,(
  p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK1,sK0),sK2),s(fun(sK1,bool),sK3),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | $spl0),
  inference(cnf_transformation,[],[f126_D])).
fof(f126_D,plain,(
  p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK1,sK0),sK2),s(fun(sK1,bool),sK3),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f128,plain,(
  s(fun(sK1,bool),'const.pred_set.EMPTY_0') = s(fun(sK1,bool),sK3) | $spl2),
  inference(cnf_transformation,[],[f128_D])).
fof(f128_D,plain,(
  s(fun(sK1,bool),'const.pred_set.EMPTY_0') = s(fun(sK1,bool),sK3) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f80,plain,(
  s(fun(sK1,bool),'const.pred_set.EMPTY_0') = s(fun(sK1,bool),sK3) | p(s(bool,'const.pred_set.BIJ_3'(s(fun(sK1,sK0),sK2),s(fun(sK1,bool),sK3),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))),
  inference(cnf_transformation,[],[f56])).
fof(f2754,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f2742,f129])).
fof(f2742,plain,(
  s(fun(sK1,bool),'const.pred_set.EMPTY_0') = s(fun(sK1,bool),sK3) | $spl0),
  inference(resolution,[],[f126,f991])).
fof(f991,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),'const.pred_set.EMPTY_0')))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X3)) )),
  inference(resolution,[],[f121,f115])).
fof(f115,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) | s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3)) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) | s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3)) & (s(fun(X1,bool),'const.pred_set.EMPTY_0') != s(fun(X1,bool),X3) | p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))))),
  inference(nnf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) <=> s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3))),
  inference(flattening,[],[f45])).
fof(f45,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) <=> s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X3))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6,X7] : ! [X8,X9] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X7,X6),X8),s(fun(X7,bool),X9),s(fun(X6,bool),'const.pred_set.EMPTY_0')))) <=> s(fun(X7,bool),'const.pred_set.EMPTY_0') = s(fun(X7,bool),X9))),
  file('thm.pred_set.BIJ_EMPTY|split|1','thm.pred_set.INJ_EMPTY|split|1')).
fof(f121,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))) )),
  inference(cnf_transformation,[],[f79])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 81
% Passive clauses: 382
% Generated clauses: 2633
% Final active clauses: 66
% Final passive clauses: 210
% Input formulas: 22
% Initial clauses: 44
% 
% Duplicate literals: 181
% Trivial inequalities: 59
% Fw subsumption resolutions: 380
% Fw demodulations: 196
% Bw demodulations: 7
% 
% Simple tautologies: 127
% Equational tautologies: 23
% Forward subsumptions: 1113
% 
% Binary resolution: 140
% Forward superposition: 943
% Backward superposition: 485
% Self superposition: 29
% Equality factoring: 30
% Equality resolution: 15
% 
% Split clauses: 154
% Split components: 317
% SAT solver clauses: 248
% SAT solver unit clauses: 16
% SAT solver binary clauses: 47
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 1
% 
% Sat splits: 154
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 239
% TWLsolver calls for satisfiability: 64
% 
% Memory used [KB]: 1918
% Time elapsed: 0.070 s
% ------------------------------
% ------------------------------
% Success in time 0.074 s
